# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_5, plain, (sorti1(esk1_0))).
#
#cnf(i_0_4, plain, (sorti1(esk2_0))).
#
#cnf(i_0_11, negated_conjecture, (sorti1(j(X1))|~sorti2(X1))).
#
#cnf(i_0_3, plain, (op1(esk1_0,esk2_0)!=op1(esk2_0,esk1_0))).
#
#cnf(i_0_12, negated_conjecture, (sorti2(h(X1))|~sorti1(X1))).
#
#cnf(i_0_8, negated_conjecture, (h(j(X1))=X1|~sorti2(X1))).
#
#cnf(i_0_7, negated_conjecture, (j(h(X1))=X1|~sorti1(X1))).
#
#cnf(i_0_1, plain, (sorti1(op1(X1,X2))|~sorti1(X2)|~sorti1(X1))).
#
#cnf(i_0_2, plain, (sorti2(op2(X1,X2))|~sorti2(X2)|~sorti2(X1))).
#
#cnf(i_0_6, plain, (op2(X1,X2)=op2(X2,X1)|~sorti2(X2)|~sorti2(X1))).
###
#cnf(i_0_10, negated_conjecture, (op2(h(X1),h(X2))=h(op1(X1,X2))|~sorti1(X2)|~sorti1(X1))).
#
#cnf(i_0_9, negated_conjecture, (op1(j(X1),j(X2))=j(op2(X1,X2))|~sorti2(X2)|~sorti2(X1))).
#
#cnf(i_0_18, negated_conjecture, (sorti2(h(op1(X1,X2)))|~sorti2(h(X2))|~sorti2(h(X1))|~sorti1(X2)|~sorti1(X1))).
#
#cnf(i_0_22, negated_conjecture, (sorti1(j(op2(X1,X2)))|~sorti2(X2)|~sorti2(X1)|~sorti1(j(X2))|~sorti1(j(X1)))).
###
#cnf(i_0_20, negated_conjecture, (h(op1(X1,j(X2)))=op2(h(X1),X2)|~sorti2(X2)|~sorti1(j(X2))|~sorti1(X1))).
#
#cnf(i_0_21, negated_conjecture, (h(op1(j(X1),X2))=op2(X1,h(X2))|~sorti2(X1)|~sorti1(j(X1))|~sorti1(X2))).
#
#cnf(i_0_23, negated_conjecture, (j(op2(X1,h(X2)))=op1(j(X1),X2)|~sorti2(h(X2))|~sorti2(X1)|~sorti1(X2))).
#
#cnf(i_0_24, negated_conjecture, (j(op2(h(X1),X2))=op1(X1,j(X2))|~sorti2(h(X1))|~sorti2(X2)|~sorti1(X1))).
#
#cnf(i_0_55, negated_conjecture, (j(op2(h(X2),X1))=op1(j(X1),X2)|~sorti2(h(X2))|~sorti2(X1)|~sorti1(X2))).
##
#cnf(i_0_68, negated_conjecture, (j(op2(X2,h(X1)))=op1(X1,j(X2))|~sorti2(h(X1))|~sorti2(X2)|~sorti1(X1))).
##
#cnf(i_0_79, negated_conjecture, (op1(j(X2),X1)=op1(X1,j(X2))|~sorti2(h(X1))|~sorti2(X2)|~sorti1(X1))).
########
#cnf(i_0_29, negated_conjecture, (sorti2(op2(h(X1),X2))|~sorti2(X2)|~sorti1(op1(X1,j(X2)))|~sorti1(j(X2))|~sorti1(X1))).
#
#cnf(i_0_130, negated_conjecture, (sorti2(op2(h(X1),X2))|~sorti2(X2)|~sorti1(j(X2))|~sorti1(X1))).
####
#cnf(i_0_136, negated_conjecture, (sorti2(h(op1(X1,X2)))|~sorti2(h(X2))|~sorti1(j(h(X2)))|~sorti1(X1)|~sorti1(X2))).
#
#cnf(i_0_36, negated_conjecture, (sorti2(op2(X1,h(X2)))|~sorti2(X1)|~sorti1(op1(j(X1),X2))|~sorti1(j(X1))|~sorti1(X2))).
#
#cnf(i_0_150, negated_conjecture, (sorti2(op2(X1,h(X2)))|~sorti2(X1)|~sorti1(j(X1))|~sorti1(X2))).
####
#cnf(i_0_154, negated_conjecture, (sorti2(h(op1(X1,X2)))|~sorti2(h(X1))|~sorti1(j(h(X1)))|~sorti1(X2)|~sorti1(X1))).
#
#cnf(i_0_46, negated_conjecture, (sorti1(op1(j(X1),X2))|~sorti2(op2(X1,h(X2)))|~sorti2(h(X2))|~sorti2(X1)|~sorti1(X2))).
#
#cnf(i_0_59, negated_conjecture, (sorti1(op1(X1,j(X2)))|~sorti2(op2(h(X1),X2))|~sorti2(h(X1))|~sorti2(X2)|~sorti1(X1))).
#
#cnf(i_0_73, negated_conjecture, (sorti1(op1(j(X2),X1))|~sorti2(op2(h(X1),X2))|~sorti2(h(X1))|~sorti2(X2)|~sorti1(X1))).
#
#cnf(i_0_88, negated_conjecture, (sorti1(op1(X2,j(X1)))|~sorti2(op2(X1,h(X2)))|~sorti2(h(X2))|~sorti2(X1)|~sorti1(X2))).
#####
#cnf(i_0_112, negated_conjecture, (op1(j(h(X1)),X2)=op1(X2,X1)|~sorti2(h(X2))|~sorti2(h(X1))|~sorti1(X2)|~sorti1(X1))).
#
#cnf(i_0_210, negated_conjecture, (op1(X1,X2)=op1(X2,X1)|~sorti2(h(X2))|~sorti2(h(X1))|~sorti1(X2)|~sorti1(X1))).
####
#cnf(i_0_211, negated_conjecture, (~sorti2(h(esk1_0))|~sorti2(h(esk2_0)))).
##
#cnf(i_0_117, negated_conjecture, (op1(X2,j(h(X1)))=op1(X1,X2)|~sorti2(h(X2))|~sorti2(h(X1))|~sorti1(X2)|~sorti1(X1))).
##
#cnf(i_0_237, negated_conjecture, (~sorti2(h(esk2_0)))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(co1, conjecture, ((![X1]:(sorti1(X1)=>sorti2(h(X1)))&![X2]:(sorti2(X2)=>sorti1(j(X2))))=>~((((![X3]:(sorti1(X3)=>![X4]:(sorti1(X4)=>h(op1(X3,X4))=op2(h(X3),h(X4))))&![X5]:(sorti2(X5)=>![X6]:(sorti2(X6)=>j(op2(X5,X6))=op1(j(X5),j(X6)))))&![X7]:(sorti2(X7)=>h(j(X7))=X7))&![X8]:(sorti1(X8)=>j(h(X8))=X8)))), file('./p/ALG030+1.tptp', co1)).
fof(ax4, axiom, ~(~(![X1]:(sorti2(X1)=>![X2]:(sorti2(X2)=>op2(X1,X2)=op2(X2,X1))))), file('./p/ALG030+1.tptp', ax4)).
fof(ax3, axiom, ~(![X1]:(sorti1(X1)=>![X2]:(sorti1(X2)=>op1(X1,X2)=op1(X2,X1)))), file('./p/ALG030+1.tptp', ax3)).
fof(c_0_3, negated_conjecture, ~(((![X1]:(sorti1(X1)=>sorti2(h(X1)))&![X2]:(sorti2(X2)=>sorti1(j(X2))))=>~((((![X3]:(sorti1(X3)=>![X4]:(sorti1(X4)=>h(op1(X3,X4))=op2(h(X3),h(X4))))&![X5]:(sorti2(X5)=>![X6]:(sorti2(X6)=>j(op2(X5,X6))=op1(j(X5),j(X6)))))&![X7]:(sorti2(X7)=>h(j(X7))=X7))&![X8]:(sorti1(X8)=>j(h(X8))=X8))))), inference(assume_negation,[status(cth)],[co1])).
fof(c_0_4, negated_conjecture, ![X17, X18, X19, X20, X21, X22, X23, X24]:(((~sorti1(X17)|sorti2(h(X17)))&(~sorti2(X18)|sorti1(j(X18))))&((((~sorti1(X19)|(~sorti1(X20)|h(op1(X19,X20))=op2(h(X19),h(X20))))&(~sorti2(X21)|(~sorti2(X22)|j(op2(X21,X22))=op1(j(X21),j(X22)))))&(~sorti2(X23)|h(j(X23))=X23))&(~sorti1(X24)|j(h(X24))=X24))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
cnf(c_0_5, negated_conjecture, (j(op2(X1,X2))=op1(j(X1),j(X2))|~sorti2(X1)|~sorti2(X2)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_6, negated_conjecture, (j(h(X1))=X1|~sorti1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_7, plain, ![X15, X16]:(~sorti2(X15)|(~sorti2(X16)|op2(X15,X16)=op2(X16,X15))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax4])])])).
cnf(c_0_8, negated_conjecture, (j(op2(X1,h(X2)))=op1(j(X1),X2)|~sorti2(h(X2))|~sorti2(X1)|~sorti1(X2)), inference(pm,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_9, plain, (op2(X1,X2)=op2(X2,X1)|~sorti2(X1)|~sorti2(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, negated_conjecture, (j(op2(h(X1),X2))=op1(X1,j(X2))|~sorti2(h(X1))|~sorti2(X2)|~sorti1(X1)), inference(pm,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_11, negated_conjecture, (j(op2(h(X1),X2))=op1(j(X2),X1)|~sorti2(h(X1))|~sorti2(X2)|~sorti1(X1)), inference(pm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12, negated_conjecture, (op1(j(X1),X2)=op1(X2,j(X1))|~sorti2(h(X2))|~sorti2(X1)|~sorti1(X2)), inference(pm,[status(thm)],[c_0_10, c_0_11])).
fof(c_0_13, plain, (sorti1(esk1_0)&(sorti1(esk2_0)&op1(esk1_0,esk2_0)!=op1(esk2_0,esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax3])])])).
cnf(c_0_14, negated_conjecture, (op1(j(h(X1)),X2)=op1(X2,X1)|~sorti2(h(X2))|~sorti2(h(X1))|~sorti1(X2)|~sorti1(X1)), inference(pm,[status(thm)],[c_0_12, c_0_6])).
cnf(c_0_15, plain, (op1(esk1_0,esk2_0)!=op1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_16, negated_conjecture, (op1(X1,X2)=op1(X2,X1)|~sorti2(h(X2))|~sorti2(h(X1))|~sorti1(X2)|~sorti1(X1)), inference(pm,[status(thm)],[c_0_14, c_0_6])).
cnf(c_0_17, plain, (sorti1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (sorti1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, negated_conjecture, (~sorti2(h(esk1_0))|~sorti2(h(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_15, c_0_16]), c_0_17]), c_0_18])])).
cnf(c_0_20, negated_conjecture, (sorti2(h(X1))|~sorti1(X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_21, negated_conjecture, (~sorti2(h(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_19, c_0_20]), c_0_17])])).
cnf(c_0_22, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_21, c_0_20]), c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 23
# Proof object clause steps            : 16
# Proof object formula steps           : 7
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 12
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 12
# Processed clauses                    : 64
# ...of these trivial                  : 0
# ...subsumed                          : 28
# ...remaining for further processing  : 36
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 0
# Generated clauses                    : 246
# ...of the previous two non-trivial   : 243
# Contextual simplify-reflections      : 0
# Paramodulations                      : 246
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 34
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 30
# Current number of unprocessed clauses: 191
# ...number of literals in the above   : 1254
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 140
# Rec. Clause-clause subsumption calls : 40
# Non-unit clause-clause subsumptions  : 30
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5431
