# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_106, plain, (op(e0,e0)=e0)).
#
#cnf(i_0_105, plain, (op(e1,e1)=e1)).
#
#cnf(i_0_104, plain, (op(e2,e2)=e2)).
#
#cnf(i_0_102, plain, (e1!=e0)).
#
#cnf(i_0_103, plain, (op(e3,e3)=e3)).
###
#cnf(i_0_101, plain, (e2!=e0)).
##
#cnf(i_0_46, plain, (op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1)).
#
#cnf(i_0_45, plain, (op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1)).
#
#cnf(i_0_99, plain, (e2!=e1)).
#
#cnf(i_0_44, plain, (op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2)).
#
#cnf(i_0_43, plain, (op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2)).
#
#cnf(i_0_42, plain, (op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|e3=e0)).
#
#cnf(i_0_100, plain, (e3!=e0)).
#
#cnf(i_0_42, plain, (op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3)).
#
#cnf(i_0_41, plain, (op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3)).
#
#cnf(i_0_15, plain, (op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3)).
#
#cnf(i_0_98, plain, (e3!=e1)).
###
#cnf(i_0_141, plain, (op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|e3!=e2)).
#
#cnf(i_0_97, plain, (e3!=e2)).
###
#cnf(i_0_39, plain, (op(e0,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0)).
#
#cnf(i_0_72, plain, (op(e0,e1)!=e0)).
##
#cnf(i_0_39, plain, (op(e2,e1)=e0|op(e3,e1)=e0)).
#
#cnf(i_0_15, plain, (op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3)).
#
#cnf(i_0_71, plain, (op(e0,e2)!=e0)).
####
#cnf(i_0_69, plain, (op(e0,e2)!=op(e0,e1))).
##
#cnf(i_0_35, plain, (op(e0,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2)).
#
#cnf(i_0_150, plain, (op(e2,e1)=e0|op(e0,e1)=e2|op(e2,e1)=e2)).
#
#cnf(i_0_70, plain, (op(e0,e3)!=e0)).
####
#cnf(i_0_68, plain, (op(e0,e3)!=op(e0,e1))).
##
#cnf(i_0_33, plain, (op(e0,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3)).
#
#cnf(i_0_159, plain, (op(e2,e1)=e0|op(e0,e1)=e3|op(e2,e1)=e3)).
#
#cnf(i_0_67, plain, (op(e0,e3)!=op(e0,e2))).
##
#cnf(i_0_161, plain, (op(e2,e1)=e0|op(e0,e1)=e2|op(e0,e1)=e3)).
##
#cnf(i_0_96, plain, (op(e1,e0)!=e0)).
##
#cnf(i_0_14, plain, (op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3)).
##
#cnf(i_0_90, plain, (op(e0,e1)!=e1)).
#
#cnf(i_0_46, plain, (op(e0,e2)=e1|op(e0,e3)=e1)).
####
#cnf(i_0_15, plain, (op(e0,e1)=e2|op(e0,e1)=e3)).
###
#cnf(i_0_66, plain, (op(e1,e0)!=e1)).
##
#cnf(i_0_45, plain, (op(e2,e0)=e1|op(e3,e0)=e1)).
##
#cnf(i_0_84, plain, (op(e1,e2)!=op(e0,e2))).
##
#cnf(i_0_170, plain, (op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e1)!=e3)).
##
#cnf(i_0_65, plain, (op(e1,e2)!=op(e1,e0))).
##
#cnf(i_0_185, plain, (op(e0,e2)=e1|op(e0,e1)!=e2|op(e0,e1)!=e3)).
##
#cnf(i_0_63, plain, (op(e1,e2)!=e1)).
##
#cnf(i_0_175, plain, (op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e2)=e1)).
##
#cnf(i_0_78, plain, (op(e1,e3)!=op(e0,e3))).
##
#cnf(i_0_176, plain, (op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e2)=e1)).
##
#cnf(i_0_64, plain, (op(e1,e3)!=op(e1,e0))).
##
#cnf(i_0_181, plain, (op(e1,e0)=e2|op(e2,e0)=e2|op(e2,e0)=e1)).
##
#cnf(i_0_62, plain, (op(e1,e3)!=e1)).
#
#cnf(i_0_182, plain, (op(e1,e0)=e3|op(e2,e0)=e3|op(e2,e0)=e1)).
###
#cnf(i_0_61, plain, (op(e1,e3)!=op(e1,e2))).
##
#cnf(i_0_204, plain, (op(e1,e0)=e2|op(e2,e0)=e1|op(e1,e0)=e3)).
##
#cnf(i_0_95, plain, (op(e2,e0)!=e0)).
####
#cnf(i_0_93, plain, (op(e2,e0)!=op(e1,e0))).
####
#cnf(i_0_89, plain, (op(e2,e1)!=op(e0,e1))).
#
#cnf(i_0_31, plain, (op(e1,e2)=e0|op(e3,e2)=e0)).
#
#cnf(i_0_29, plain, (op(e0,e2)=e1|op(e3,e2)=e1)).
#
#cnf(i_0_215, plain, (op(e1,e2)=e0|op(e0,e2)=e1)).
#
#cnf(i_0_87, plain, (op(e2,e1)!=e1)).
####
#cnf(i_0_60, plain, (op(e2,e1)!=op(e2,e0))).
###
#cnf(i_0_221, plain, (op(e2,e1)=e0|op(e0,e1)=e3|op(e2,e0)!=e3)).
#
#cnf(i_0_83, plain, (op(e0,e2)!=e2)).
#
#cnf(i_0_228, plain, (op(e0,e2)=e1|op(e0,e1)!=e3)).
#
#cnf(i_0_44, plain, (op(e0,e1)=e2|op(e0,e3)=e2)).
##
#cnf(i_0_81, plain, (op(e1,e2)!=e2)).
###
#cnf(i_0_14, plain, (op(e0,e2)=e1|op(e0,e2)=e3)).
#
#cnf(i_0_59, plain, (op(e2,e0)!=e2)).
####
#cnf(i_0_57, plain, (op(e2,e1)!=e2)).
###
#cnf(i_0_175, plain, (op(e0,e1)=e2|op(e0,e2)=e1)).
#
#cnf(i_0_77, plain, (op(e2,e3)!=op(e0,e3))).
####
#cnf(i_0_75, plain, (op(e2,e3)!=op(e1,e3))).
##
#cnf(i_0_43, plain, (op(e1,e0)=e2|op(e3,e0)=e2)).
#
#cnf(i_0_181, plain, (op(e1,e0)=e2|op(e2,e0)=e1)).
#
#cnf(i_0_58, plain, (op(e2,e3)!=op(e2,e0))).
####
#cnf(i_0_56, plain, (op(e2,e3)!=op(e2,e1))).
#
#cnf(i_0_35, plain, (op(e0,e1)=e2|op(e3,e1)=e2)).
#
#cnf(i_0_150, plain, (op(e2,e1)=e0|op(e0,e1)=e2)).
##
#cnf(i_0_55, plain, (op(e2,e3)!=e2)).
####
#cnf(i_0_94, plain, (op(e3,e0)!=e0)).
####
#cnf(i_0_92, plain, (op(e3,e0)!=op(e1,e0))).
####
#cnf(i_0_91, plain, (op(e3,e0)!=op(e2,e0))).
####
#cnf(i_0_88, plain, (op(e3,e1)!=op(e0,e1))).
####
#cnf(i_0_86, plain, (op(e3,e1)!=e1)).
####
#cnf(i_0_85, plain, (op(e3,e1)!=op(e2,e1))).
####
#cnf(i_0_54, plain, (op(e3,e1)!=op(e3,e0))).
##
#cnf(i_0_287, plain, (op(e0,e1)=e2|op(e3,e0)!=e2)).
#
#cnf(i_0_290, plain, (op(e0,e1)=e2|op(e1,e0)=e2)).
#
#cnf(i_0_82, plain, (op(e3,e2)!=op(e0,e2))).
####
#cnf(i_0_80, plain, (op(e3,e2)!=op(e1,e2))).
####
#cnf(i_0_79, plain, (op(e3,e2)!=e2)).
###
#cnf(i_0_285, plain, (op(e0,e1)=e3|op(e2,e1)=e3|op(e3,e0)!=e3)).
#
#cnf(i_0_53, plain, (op(e3,e2)!=op(e3,e0))).
##
#cnf(i_0_303, plain, (op(e0,e2)=e1|op(e3,e0)!=e1)).
#
#cnf(i_0_305, plain, (op(e0,e2)=e1|op(e2,e0)=e1)).
#
#cnf(i_0_51, plain, (op(e3,e2)!=op(e3,e1))).
####
#cnf(i_0_76, plain, (op(e0,e3)!=e3)).
###
#cnf(i_0_311, plain, (op(e1,e2)=e0|op(e3,e1)!=e0)).
#
#cnf(i_0_74, plain, (op(e1,e3)!=e3)).
##
#cnf(i_0_42, plain, (op(e0,e1)=e3|op(e0,e2)=e3)).
##
#cnf(i_0_73, plain, (op(e2,e3)!=e3)).
##
#cnf(i_0_316, plain, (op(e1,e2)=e0|op(e2,e1)=e0)).
##
#cnf(i_0_52, plain, (op(e3,e0)!=e3)).
####
#cnf(i_0_50, plain, (op(e3,e1)!=e3)).
####
#cnf(i_0_49, plain, (op(e3,e2)!=e3)).
######
#cnf(i_0_41, plain, (op(e1,e0)=e3|op(e2,e0)=e3)).
####
#cnf(i_0_33, plain, (op(e0,e1)=e3|op(e2,e1)=e3)).
####
#cnf(i_0_341, plain, (op(e0,e1)=e3|op(e2,e0)!=e3)).
#
#cnf(i_0_337, plain, (op(e1,e0)=e2|op(e1,e0)=e3)).
#
#cnf(i_0_111, plain, (op(op(e1,e2),e2)!=e3|op(e1,e2)!=e0)).
####
#cnf(i_0_109, plain, (op(op(e1,e3),e3)!=e2|op(e1,e3)!=e0)).
#
#cnf(i_0_338, plain, (op(e0,e2)=e1|op(e1,e0)=e3)).
###
#cnf(i_0_112, plain, (op(op(e2,e1),e1)!=e3|op(e2,e1)!=e0)).
##
#cnf(i_0_345, plain, (op(e1,e2)=e0|op(e0,e1)=e3)).
##
#cnf(i_0_107, plain, (op(op(e2,e3),e3)!=e1|op(e2,e3)!=e0)).
####
#cnf(i_0_110, plain, (op(op(e3,e1),e1)!=e2|op(e3,e1)!=e0)).
#
#cnf(i_0_348, plain, (op(e0,e1)=e3|op(e1,e0)=e3)).
###
#cnf(i_0_108, plain, (op(op(e3,e2),e2)!=e1|op(e3,e2)!=e0)).
####
#cnf(i_0_117, plain, (op(op(e0,e2),e2)!=e3|op(e0,e2)!=e1)).
####
#cnf(i_0_115, plain, (op(op(e0,e3),e3)!=e2|op(e0,e3)!=e1)).
##
#cnf(i_0_353, plain, (op(e0,e2)=e1|op(e0,e2)!=e3|op(e1,e2)!=e0)).
#
#cnf(i_0_378, plain, (op(e0,e2)=e1|op(e0,e2)!=e3)).
#
#cnf(i_0_118, plain, (op(op(e2,e0),e0)!=e3|op(e2,e0)!=e1)).
#
#cnf(i_0_380, plain, (op(e0,e2)=e1)).
#
#cnf(i_0_387, plain, (op(e0,e1)=e3)).
#
#cnf(i_0_35, plain, (op(e3,e1)=e2)).
##
#cnf(i_0_290, plain, (op(e1,e0)=e2)).
#
#cnf(i_0_150, plain, (op(e2,e1)=e0)).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
fof(ax3, axiom, (((((((((((((((((((((((((((((((((((((((((((((((op(e0,e0)!=op(e1,e0)&op(e0,e0)!=op(e2,e0))&op(e0,e0)!=op(e3,e0))&op(e1,e0)!=op(e2,e0))&op(e1,e0)!=op(e3,e0))&op(e2,e0)!=op(e3,e0))&op(e0,e1)!=op(e1,e1))&op(e0,e1)!=op(e2,e1))&op(e0,e1)!=op(e3,e1))&op(e1,e1)!=op(e2,e1))&op(e1,e1)!=op(e3,e1))&op(e2,e1)!=op(e3,e1))&op(e0,e2)!=op(e1,e2))&op(e0,e2)!=op(e2,e2))&op(e0,e2)!=op(e3,e2))&op(e1,e2)!=op(e2,e2))&op(e1,e2)!=op(e3,e2))&op(e2,e2)!=op(e3,e2))&op(e0,e3)!=op(e1,e3))&op(e0,e3)!=op(e2,e3))&op(e0,e3)!=op(e3,e3))&op(e1,e3)!=op(e2,e3))&op(e1,e3)!=op(e3,e3))&op(e2,e3)!=op(e3,e3))&op(e0,e0)!=op(e0,e1))&op(e0,e0)!=op(e0,e2))&op(e0,e0)!=op(e0,e3))&op(e0,e1)!=op(e0,e2))&op(e0,e1)!=op(e0,e3))&op(e0,e2)!=op(e0,e3))&op(e1,e0)!=op(e1,e1))&op(e1,e0)!=op(e1,e2))&op(e1,e0)!=op(e1,e3))&op(e1,e1)!=op(e1,e2))&op(e1,e1)!=op(e1,e3))&op(e1,e2)!=op(e1,e3))&op(e2,e0)!=op(e2,e1))&op(e2,e0)!=op(e2,e2))&op(e2,e0)!=op(e2,e3))&op(e2,e1)!=op(e2,e2))&op(e2,e1)!=op(e2,e3))&op(e2,e2)!=op(e2,e3))&op(e3,e0)!=op(e3,e1))&op(e3,e0)!=op(e3,e2))&op(e3,e0)!=op(e3,e3))&op(e3,e1)!=op(e3,e2))&op(e3,e1)!=op(e3,e3))&op(e3,e2)!=op(e3,e3)), file('./p/ALG138+1.tptp', ax3)).
fof(ax5, axiom, (((op(e0,e0)=e0&op(e1,e1)=e1)&op(e2,e2)=e2)&op(e3,e3)=e3), file('./p/ALG138+1.tptp', ax5)).
fof(ax2, axiom, ((((((((((((((((((((((((((((((((op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0)&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3))&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3))&(op(e3,e0)=e0|op(e3,e1)=e0|op(e3,e2)=e0|op(e3,e3)=e0))&(op(e0,e3)=e0|op(e1,e3)=e0|op(e2,e3)=e0|op(e3,e3)=e0))&(op(e3,e0)=e1|op(e3,e1)=e1|op(e3,e2)=e1|op(e3,e3)=e1))&(op(e0,e3)=e1|op(e1,e3)=e1|op(e2,e3)=e1|op(e3,e3)=e1))&(op(e3,e0)=e2|op(e3,e1)=e2|op(e3,e2)=e2|op(e3,e3)=e2))&(op(e0,e3)=e2|op(e1,e3)=e2|op(e2,e3)=e2|op(e3,e3)=e2))&(op(e3,e0)=e3|op(e3,e1)=e3|op(e3,e2)=e3|op(e3,e3)=e3))&(op(e0,e3)=e3|op(e1,e3)=e3|op(e2,e3)=e3|op(e3,e3)=e3)), file('./p/ALG138+1.tptp', ax2)).
fof(ax4, axiom, (((((e0!=e1&e0!=e2)&e0!=e3)&e1!=e2)&e1!=e3)&e2!=e3), file('./p/ALG138+1.tptp', ax4)).
fof(ax10, axiom, ~((e0=op(e1,e2)&e3=op(op(e1,e2),e2))), file('./p/ALG138+1.tptp', ax10)).
fof(ax1, axiom, ((((((((((((((((op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3)&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3)), file('./p/ALG138+1.tptp', ax1)).
fof(ax11, axiom, ~((e0=op(e2,e1)&e3=op(op(e2,e1),e1))), file('./p/ALG138+1.tptp', ax11)).
cnf(c_0_7, plain, (op(e0,e0)!=op(e0,e2)), inference(split_conjunct,[status(thm)],[ax3])).
cnf(c_0_8, plain, (op(e0,e0)=e0), inference(split_conjunct,[status(thm)],[ax5])).
cnf(c_0_9, plain, (op(e1,e1)!=op(e1,e2)), inference(split_conjunct,[status(thm)],[ax3])).
cnf(c_0_10, plain, (op(e1,e1)=e1), inference(split_conjunct,[status(thm)],[ax5])).
cnf(c_0_11, plain, (op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_12, plain, (op(e2,e2)=e2), inference(split_conjunct,[status(thm)],[ax5])).
cnf(c_0_13, plain, (op(e0,e2)!=e0), inference(rw,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_14, plain, (e0!=e2), inference(split_conjunct,[status(thm)],[ax4])).
cnf(c_0_15, plain, (op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_16, plain, (op(e1,e2)!=e1), inference(rw,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_17, plain, (e1!=e2), inference(split_conjunct,[status(thm)],[ax4])).
fof(c_0_18, plain, (e0!=op(e1,e2)|e3!=op(op(e1,e2),e2)), inference(fof_nnf,[status(thm)],[ax10])).
cnf(c_0_19, plain, (op(e1,e2)=e0|op(e3,e2)=e0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_12]), c_0_13]), c_0_14])).
cnf(c_0_20, plain, (op(e0,e2)=e1|op(e3,e2)=e1), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_12]), c_0_16]), c_0_17])).
cnf(c_0_21, plain, (e0!=e1), inference(split_conjunct,[status(thm)],[ax4])).
cnf(c_0_22, plain, (op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_23, plain, (op(e0,e0)!=op(e0,e1)), inference(split_conjunct,[status(thm)],[ax3])).
cnf(c_0_24, plain, (op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_25, plain, (e0!=op(e1,e2)|e3!=op(op(e1,e2),e2)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (op(e1,e2)=e0|op(e0,e2)=e1), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_27, plain, (op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3), inference(split_conjunct,[status(thm)],[ax1])).
cnf(c_0_28, plain, (op(e0,e2)!=op(e2,e2)), inference(split_conjunct,[status(thm)],[ax3])).
cnf(c_0_29, plain, (op(e0,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_10]), c_0_21])).
cnf(c_0_30, plain, (op(e0,e1)!=e0), inference(rw,[status(thm)],[c_0_23, c_0_8])).
cnf(c_0_31, plain, (op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_32, plain, (op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|e3=e0), inference(rw,[status(thm)],[c_0_24, c_0_8])).
cnf(c_0_33, plain, (e0!=e3), inference(split_conjunct,[status(thm)],[ax4])).
cnf(c_0_34, plain, (op(e0,e3)!=op(e3,e3)), inference(split_conjunct,[status(thm)],[ax3])).
cnf(c_0_35, plain, (op(e3,e3)=e3), inference(split_conjunct,[status(thm)],[ax5])).
cnf(c_0_36, plain, (op(e0,e2)=e1|op(e0,e2)!=e3|op(e1,e2)!=e0), inference(pm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_37, plain, (op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3), inference(sr,[status(thm)],[c_0_27, c_0_13])).
cnf(c_0_38, plain, (op(e0,e2)!=e2), inference(rw,[status(thm)],[c_0_28, c_0_12])).
cnf(c_0_39, plain, (op(e2,e1)=e0|op(e3,e1)=e0), inference(sr,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_40, plain, (op(e0,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_10]), c_0_17])).
cnf(c_0_41, plain, (op(e2,e1)!=op(e2,e2)), inference(split_conjunct,[status(thm)],[ax3])).
cnf(c_0_42, plain, (op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3), inference(sr,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_43, plain, (op(e0,e3)!=e3), inference(rw,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_44, plain, (op(e0,e2)=e1|op(e0,e2)!=e3), inference(pm,[status(thm)],[c_0_36, c_0_26])).
cnf(c_0_45, plain, (op(e0,e2)=e1|op(e0,e2)=e3), inference(sr,[status(thm)],[c_0_37, c_0_38])).
cnf(c_0_46, plain, (op(e2,e1)=e0|op(e0,e1)=e2|op(e2,e1)=e2), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_39, c_0_40]), c_0_14])).
cnf(c_0_47, plain, (op(e2,e1)!=e2), inference(rw,[status(thm)],[c_0_41, c_0_12])).
cnf(c_0_48, plain, (op(e0,e1)=e3|op(e0,e2)=e3), inference(sr,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_49, plain, (op(e0,e2)=e1), inference(pm,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_50, plain, (e1!=e3), inference(split_conjunct,[status(thm)],[ax4])).
fof(c_0_51, plain, (e0!=op(e2,e1)|e3!=op(op(e2,e1),e1)), inference(fof_nnf,[status(thm)],[ax11])).
cnf(c_0_52, plain, (op(e2,e1)=e0|op(e0,e1)=e2), inference(sr,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_53, plain, (op(e0,e1)=e3), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_48, c_0_49]), c_0_50])).
cnf(c_0_54, plain, (e2!=e3), inference(split_conjunct,[status(thm)],[ax4])).
cnf(c_0_55, plain, (e0!=op(e2,e1)|e3!=op(op(e2,e1),e1)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_56, plain, (op(e2,e1)=e0), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_53]), c_0_54])).
cnf(c_0_57, plain, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_55, c_0_56]), c_0_53]), c_0_56])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 49
# Proof object formula steps           : 9
# Proof object conjectures             : 0
# Proof object clause conjectures      : 0
# Proof object formula conjectures     : 0
# Proof object initial clauses used    : 24
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 32
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 29
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 130
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 130
# Processed clauses                    : 274
# ...of these trivial                  : 4
# ...subsumed                          : 141
# ...remaining for further processing  : 129
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 13
# Backward-rewritten                   : 55
# Generated clauses                    : 274
# ...of the previous two non-trivial   : 269
# Contextual simplify-reflections      : 0
# Paramodulations                      : 239
# Factorizations                       : 19
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 45
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 28
#    Non-unit-clauses                  : 8
# Current number of unprocessed clauses: 122
# ...number of literals in the above   : 351
# Current number of archived formulas  : 0
# Current number of archived clauses   : 84
# Clause-clause subsumption calls (NU) : 268
# Rec. Clause-clause subsumption calls : 162
# Non-unit clause-clause subsumptions  : 30
# Unit Clause-clause subsumption calls : 127
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 9195
