# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_35, plain, (op(e0,e0)=e0)).
#
#cnf(i_0_34, plain, (op(e0,e1)=e3)).
#
#cnf(i_0_33, plain, (op(e0,e2)=e4)).
#
#cnf(i_0_10, plain, (e1!=e0)).
#
#cnf(i_0_32, plain, (op(e0,e3)=e1)).
#
#cnf(i_0_31, plain, (op(e0,e4)=e2)).
#
#cnf(i_0_30, plain, (op(e1,e0)=e4)).
#
#cnf(i_0_9, plain, (e2!=e0)).
#
#cnf(i_0_29, plain, (op(e1,e1)=e1)).
#
#cnf(i_0_28, plain, (op(e1,e2)=e3)).
#
#cnf(i_0_27, plain, (op(e1,e3)=e2)).
#
#cnf(i_0_6, plain, (e2!=e1)).
#
#cnf(i_0_26, plain, (op(e1,e4)=e0)).
#
#cnf(i_0_25, plain, (op(e2,e0)=e1)).
#
#cnf(i_0_24, plain, (op(e2,e1)=e0)).
#
#cnf(i_0_8, plain, (e3!=e0)).
#
#cnf(i_0_23, plain, (op(e2,e2)=e2)).
#
#cnf(i_0_22, plain, (op(e2,e3)=e4)).
#
#cnf(i_0_21, plain, (op(e2,e4)=e3)).
#
#cnf(i_0_5, plain, (e3!=e1)).
#
#cnf(i_0_20, plain, (op(e3,e0)=e2)).
#
#cnf(i_0_19, plain, (op(e3,e1)=e4)).
#
#cnf(i_0_18, plain, (op(e3,e2)=e0)).
#
#cnf(i_0_3, plain, (e2!=e3)).
#
#cnf(i_0_17, plain, (op(e3,e3)=e3)).
#
#cnf(i_0_16, plain, (op(e3,e4)=e1)).
#
#cnf(i_0_15, plain, (op(e4,e0)=e3)).
#
#cnf(i_0_7, plain, (e4!=e0)).
#
#cnf(i_0_14, plain, (op(e4,e1)=e2)).
#
#cnf(i_0_13, plain, (op(e4,e2)=e1)).
#
#cnf(i_0_12, plain, (op(e4,e3)=e0)).
#
#cnf(i_0_4, plain, (e4!=e1)).
#
#cnf(i_0_11, plain, (op(e4,e4)=e4)).
###
#cnf(i_0_2, plain, (e4!=e2)).
####
#cnf(i_0_1, plain, (e4!=e3)).
####################
#cnf(i_0_36, negated_conjecture, (~epred37_0)).
##
#cnf(i_0_69, plain, (epred2_0|~epred1_0)).
##################
#cnf(i_0_101, plain, (epred3_0|~epred2_0)).
###################
#cnf(i_0_132, plain, (epred4_0|~epred3_0)).
###########################
#cnf(i_0_155, plain, (epred5_0|~epred4_0)).
####################
#cnf(i_0_185, plain, (epred6_0|~epred5_0)).
##################################
#cnf(i_0_201, plain, (epred7_0|~epred6_0)).
############################
#cnf(i_0_223, plain, (epred8_0|~epred7_0)).
############################
#cnf(i_0_245, plain, (epred9_0|~epred8_0)).
#####################
#cnf(i_0_274, plain, (epred10_0|~epred9_0)).
##########################
#cnf(i_0_298, plain, (epred11_0|~epred10_0)).
#############
#cnf(i_0_335, plain, (epred12_0|~epred11_0)).
########################
#cnf(i_0_345, plain, (epred13_0|~epred12_0)).
########################
#cnf(i_0_387, plain, (epred14_0|~epred13_0)).
############################################
#cnf(i_0_401, plain, (epred15_0|~epred14_0)).
##############
#cnf(i_0_424, plain, (epred16_0|~epred15_0)).
##############################
#cnf(i_0_443, plain, (epred17_0|~epred16_0)).
#####################
#cnf(i_0_481, plain, (epred18_0|~epred17_0)).
##########################
#cnf(i_0_504, plain, (epred19_0|~epred18_0)).
#######################
#cnf(i_0_532, plain, (epred20_0|~epred19_0)).
################################
#cnf(i_0_543, plain, (epred21_0|~epred20_0)).
#######
#cnf(i_0_590, plain, (epred22_0|~epred21_0)).
#######################################
#cnf(i_0_599, plain, (epred23_0|~epred22_0)).
######################
#cnf(i_0_625, plain, (epred24_0|~epred23_0)).
##################################
#cnf(i_0_643, plain, (epred25_0|~epred24_0)).
#######################
#cnf(i_0_676, plain, (epred26_0|~epred25_0)).
###############
#cnf(i_0_707, plain, (epred27_0|~epred26_0)).
###################
#cnf(i_0_738, plain, (epred28_0|~epred27_0)).
###############################################
#cnf(i_0_741, plain, (epred29_0|~epred28_0)).
#
#cnf(i_0_790, plain, (epred30_0|~epred29_0)).
#################################
#cnf(i_0_807, plain, (epred31_0|~epred30_0)).
##################################
#cnf(i_0_824, plain, (epred32_0|~epred31_0)).
#####################
#cnf(i_0_852, plain, (epred33_0|~epred32_0)).
###################
#cnf(i_0_883, plain, (epred34_0|~epred33_0)).
##################################
#cnf(i_0_899, plain, (epred35_0|~epred34_0)).
#########
#cnf(i_0_940, plain, (epred36_0|~epred35_0)).
#######################################
#cnf(i_0_951, plain, (~epred36_0)).
##
#cnf(i_0_966, plain, (~epred35_0)).
####
#cnf(i_0_967, plain, (~epred34_0)).
####
#cnf(i_0_968, plain, (~epred33_0)).
####
#cnf(i_0_969, plain, (~epred32_0)).
####
#cnf(i_0_970, plain, (~epred31_0)).
#
#cnf(i_0_62, plain, (epred1_0)).
#
#cnf(i_0_69, plain, (epred2_0)).
#
#cnf(i_0_101, plain, (epred3_0)).
#
#cnf(i_0_971, plain, (~epred30_0)).
#
#cnf(i_0_132, plain, (epred4_0)).
#
#cnf(i_0_155, plain, (epred5_0)).
#
#cnf(i_0_185, plain, (epred6_0)).
#
#cnf(i_0_972, plain, (~epred29_0)).
#
#cnf(i_0_201, plain, (epred7_0)).
#
#cnf(i_0_223, plain, (epred8_0)).
#
#cnf(i_0_245, plain, (epred9_0)).
#
#cnf(i_0_973, plain, (~epred28_0)).
#
#cnf(i_0_274, plain, (epred10_0)).
#
#cnf(i_0_298, plain, (epred11_0)).
#
#cnf(i_0_335, plain, (epred12_0)).
#
#cnf(i_0_974, plain, (~epred27_0)).
#
#cnf(i_0_345, plain, (epred13_0)).
#
#cnf(i_0_387, plain, (epred14_0)).
#
#cnf(i_0_401, plain, (epred15_0)).
#
#cnf(i_0_975, plain, (~epred26_0)).
#
#cnf(i_0_424, plain, (epred16_0)).
#
#cnf(i_0_443, plain, (epred17_0)).
#
#cnf(i_0_481, plain, (epred18_0)).
#
#cnf(i_0_976, plain, (~epred25_0)).
#
#cnf(i_0_504, plain, (epred19_0)).
#
#cnf(i_0_532, plain, (epred20_0)).
#
#cnf(i_0_543, plain, (epred21_0)).
#
#cnf(i_0_977, plain, (~epred24_0)).
#
#cnf(i_0_590, plain, (epred22_0)).
#
#cnf(i_0_599, plain, (epred23_0)).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(ax2, axiom, ((((((((((((((((((((((((op(e0,e0)=e0&op(e0,e1)=e3)&op(e0,e2)=e4)&op(e0,e3)=e1)&op(e0,e4)=e2)&op(e1,e0)=e4)&op(e1,e1)=e1)&op(e1,e2)=e3)&op(e1,e3)=e2)&op(e1,e4)=e0)&op(e2,e0)=e1)&op(e2,e1)=e0)&op(e2,e2)=e2)&op(e2,e3)=e4)&op(e2,e4)=e3)&op(e3,e0)=e2)&op(e3,e1)=e4)&op(e3,e2)=e0)&op(e3,e3)=e3)&op(e3,e4)=e1)&op(e4,e0)=e3)&op(e4,e1)=e2)&op(e4,e2)=e1)&op(e4,e3)=e0)&op(e4,e4)=e4), file('./p/ALG188+1.tptp', ax2)).
fof(co1, conjecture, ((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3|op(e4,e2)=e3))&(op(e2,e0)=e4|op(e2,e1)=e4|op(e2,e2)=e4|op(e2,e3)=e4|op(e2,e4)=e4))&(op(e0,e2)=e4|op(e1,e2)=e4|op(e2,e2)=e4|op(e3,e2)=e4|op(e4,e2)=e4))&(op(e3,e0)=e0|op(e3,e1)=e0|op(e3,e2)=e0|op(e3,e3)=e0|op(e3,e4)=e0))&(op(e0,e3)=e0|op(e1,e3)=e0|op(e2,e3)=e0|op(e3,e3)=e0|op(e4,e3)=e0))&(op(e3,e0)=e1|op(e3,e1)=e1|op(e3,e2)=e1|op(e3,e3)=e1|op(e3,e4)=e1))&(op(e0,e3)=e1|op(e1,e3)=e1|op(e2,e3)=e1|op(e3,e3)=e1|op(e4,e3)=e1))&(op(e3,e0)=e2|op(e3,e1)=e2|op(e3,e2)=e2|op(e3,e3)=e2|op(e3,e4)=e2))&(op(e0,e3)=e2|op(e1,e3)=e2|op(e2,e3)=e2|op(e3,e3)=e2|op(e4,e3)=e2))&(op(e3,e0)=e3|op(e3,e1)=e3|op(e3,e2)=e3|op(e3,e3)=e3|op(e3,e4)=e3))&(op(e0,e3)=e3|op(e1,e3)=e3|op(e2,e3)=e3|op(e3,e3)=e3|op(e4,e3)=e3))&(op(e3,e0)=e4|op(e3,e1)=e4|op(e3,e2)=e4|op(e3,e3)=e4|op(e3,e4)=e4))&(op(e0,e3)=e4|op(e1,e3)=e4|op(e2,e3)=e4|op(e3,e3)=e4|op(e4,e3)=e4))&(op(e4,e0)=e0|op(e4,e1)=e0|op(e4,e2)=e0|op(e4,e3)=e0|op(e4,e4)=e0))&(op(e0,e4)=e0|op(e1,e4)=e0|op(e2,e4)=e0|op(e3,e4)=e0|op(e4,e4)=e0))&(op(e4,e0)=e1|op(e4,e1)=e1|op(e4,e2)=e1|op(e4,e3)=e1|op(e4,e4)=e1))&(op(e0,e4)=e1|op(e1,e4)=e1|op(e2,e4)=e1|op(e3,e4)=e1|op(e4,e4)=e1))&(op(e4,e0)=e2|op(e4,e1)=e2|op(e4,e2)=e2|op(e4,e3)=e2|op(e4,e4)=e2))&(op(e0,e4)=e2|op(e1,e4)=e2|op(e2,e4)=e2|op(e3,e4)=e2|op(e4,e4)=e2))&(op(e4,e0)=e3|op(e4,e1)=e3|op(e4,e2)=e3|op(e4,e3)=e3|op(e4,e4)=e3))&(op(e0,e4)=e3|op(e1,e4)=e3|op(e2,e4)=e3|op(e3,e4)=e3|op(e4,e4)=e3))&(op(e4,e0)=e4|op(e4,e1)=e4|op(e4,e2)=e4|op(e4,e3)=e4|op(e4,e4)=e4))&(op(e0,e4)=e4|op(e1,e4)=e4|op(e2,e4)=e4|op(e3,e4)=e4|op(e4,e4)=e4))&op(op(e0,e0),op(e0,e0))=e0)&op(op(e1,e0),op(e0,e1))=e0)&op(op(e2,e0),op(e0,e2))=e0)&op(op(e3,e0),op(e0,e3))=e0)&op(op(e4,e0),op(e0,e4))=e0)&op(op(e0,e1),op(e1,e0))=e1)&op(op(e1,e1),op(e1,e1))=e1)&op(op(e2,e1),op(e1,e2))=e1)&op(op(e3,e1),op(e1,e3))=e1)&op(op(e4,e1),op(e1,e4))=e1)&op(op(e0,e2),op(e2,e0))=e2)&op(op(e1,e2),op(e2,e1))=e2)&op(op(e2,e2),op(e2,e2))=e2)&op(op(e3,e2),op(e2,e3))=e2)&op(op(e4,e2),op(e2,e4))=e2)&op(op(e0,e3),op(e3,e0))=e3)&op(op(e1,e3),op(e3,e1))=e3)&op(op(e2,e3),op(e3,e2))=e3)&op(op(e3,e3),op(e3,e3))=e3)&op(op(e4,e3),op(e3,e4))=e3)&op(op(e0,e4),op(e4,e0))=e4)&op(op(e1,e4),op(e4,e1))=e4)&op(op(e2,e4),op(e4,e2))=e4)&op(op(e3,e4),op(e4,e3))=e4)&op(op(e4,e4),op(e4,e4))=e4), file('./p/ALG188+1.tptp', co1)).
fof(c_0_2, plain, (epred2_0<=>((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))), introduced(definition)).
fof(c_0_3, plain, (epred1_0<=>((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))), introduced(definition)).
fof(c_0_4, plain, (epred3_0<=>((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))), introduced(definition)).
fof(c_0_5, plain, (((epred1_0&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))=>epred2_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_2]), c_0_3])).
fof(c_0_6, plain, (((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))=>epred1_0), inference(split_equiv,[status(thm)],[c_0_3])).
fof(c_0_7, plain, (epred4_0<=>((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))), introduced(definition)).
fof(c_0_8, plain, (((epred2_0&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))=>epred3_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_4]), c_0_2])).
fof(c_0_9, plain, (((((((((op(e0,e3)!=e0|(op(e0,e2)!=e0|~epred1_0)|epred2_0)&(op(e0,e3)!=e1|(op(e0,e2)!=e0|~epred1_0)|epred2_0))&(op(e0,e3)!=e2|(op(e0,e2)!=e0|~epred1_0)|epred2_0))&(op(e0,e3)!=e3|(op(e0,e2)!=e0|~epred1_0)|epred2_0))&(op(e0,e3)!=e4|(op(e0,e2)!=e0|~epred1_0)|epred2_0))&(((((op(e0,e3)!=e0|(op(e0,e2)!=e1|~epred1_0)|epred2_0)&(op(e0,e3)!=e1|(op(e0,e2)!=e1|~epred1_0)|epred2_0))&(op(e0,e3)!=e2|(op(e0,e2)!=e1|~epred1_0)|epred2_0))&(op(e0,e3)!=e3|(op(e0,e2)!=e1|~epred1_0)|epred2_0))&(op(e0,e3)!=e4|(op(e0,e2)!=e1|~epred1_0)|epred2_0)))&(((((op(e0,e3)!=e0|(op(e0,e2)!=e2|~epred1_0)|epred2_0)&(op(e0,e3)!=e1|(op(e0,e2)!=e2|~epred1_0)|epred2_0))&(op(e0,e3)!=e2|(op(e0,e2)!=e2|~epred1_0)|epred2_0))&(op(e0,e3)!=e3|(op(e0,e2)!=e2|~epred1_0)|epred2_0))&(op(e0,e3)!=e4|(op(e0,e2)!=e2|~epred1_0)|epred2_0)))&(((((op(e0,e3)!=e0|(op(e0,e2)!=e3|~epred1_0)|epred2_0)&(op(e0,e3)!=e1|(op(e0,e2)!=e3|~epred1_0)|epred2_0))&(op(e0,e3)!=e2|(op(e0,e2)!=e3|~epred1_0)|epred2_0))&(op(e0,e3)!=e3|(op(e0,e2)!=e3|~epred1_0)|epred2_0))&(op(e0,e3)!=e4|(op(e0,e2)!=e3|~epred1_0)|epred2_0)))&(((((op(e0,e3)!=e0|(op(e0,e2)!=e4|~epred1_0)|epred2_0)&(op(e0,e3)!=e1|(op(e0,e2)!=e4|~epred1_0)|epred2_0))&(op(e0,e3)!=e2|(op(e0,e2)!=e4|~epred1_0)|epred2_0))&(op(e0,e3)!=e3|(op(e0,e2)!=e4|~epred1_0)|epred2_0))&(op(e0,e3)!=e4|(op(e0,e2)!=e4|~epred1_0)|epred2_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])).
fof(c_0_10, plain, (((((((((op(e0,e1)!=e0|(op(e0,e0)!=e0|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0)&(op(e0,e1)!=e1|(op(e0,e0)!=e0|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e2|(op(e0,e0)!=e0|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e3|(op(e0,e0)!=e0|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e4|(op(e0,e0)!=e0|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(((((op(e0,e1)!=e0|(op(e0,e0)!=e1|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0)&(op(e0,e1)!=e1|(op(e0,e0)!=e1|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e2|(op(e0,e0)!=e1|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e3|(op(e0,e0)!=e1|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e4|(op(e0,e0)!=e1|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0)))&(((((op(e0,e1)!=e0|(op(e0,e0)!=e2|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0)&(op(e0,e1)!=e1|(op(e0,e0)!=e2|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e2|(op(e0,e0)!=e2|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e3|(op(e0,e0)!=e2|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e4|(op(e0,e0)!=e2|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0)))&(((((op(e0,e1)!=e0|(op(e0,e0)!=e3|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0)&(op(e0,e1)!=e1|(op(e0,e0)!=e3|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e2|(op(e0,e0)!=e3|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e3|(op(e0,e0)!=e3|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e4|(op(e0,e0)!=e3|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0)))&(((((op(e0,e1)!=e0|(op(e0,e0)!=e4|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0)&(op(e0,e1)!=e1|(op(e0,e0)!=e4|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e2|(op(e0,e0)!=e4|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e3|(op(e0,e0)!=e4|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))&(op(e0,e1)!=e4|(op(e0,e0)!=e4|(op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4))|epred1_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])).
fof(c_0_11, plain, (epred5_0<=>((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))), introduced(definition)).
fof(c_0_12, plain, (((epred3_0&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))=>epred4_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_7]), c_0_4])).
fof(c_0_13, plain, (((((((((op(e1,e0)!=e0|(op(e0,e4)!=e0|~epred2_0)|epred3_0)&(op(e1,e0)!=e1|(op(e0,e4)!=e0|~epred2_0)|epred3_0))&(op(e1,e0)!=e2|(op(e0,e4)!=e0|~epred2_0)|epred3_0))&(op(e1,e0)!=e3|(op(e0,e4)!=e0|~epred2_0)|epred3_0))&(op(e1,e0)!=e4|(op(e0,e4)!=e0|~epred2_0)|epred3_0))&(((((op(e1,e0)!=e0|(op(e0,e4)!=e1|~epred2_0)|epred3_0)&(op(e1,e0)!=e1|(op(e0,e4)!=e1|~epred2_0)|epred3_0))&(op(e1,e0)!=e2|(op(e0,e4)!=e1|~epred2_0)|epred3_0))&(op(e1,e0)!=e3|(op(e0,e4)!=e1|~epred2_0)|epred3_0))&(op(e1,e0)!=e4|(op(e0,e4)!=e1|~epred2_0)|epred3_0)))&(((((op(e1,e0)!=e0|(op(e0,e4)!=e2|~epred2_0)|epred3_0)&(op(e1,e0)!=e1|(op(e0,e4)!=e2|~epred2_0)|epred3_0))&(op(e1,e0)!=e2|(op(e0,e4)!=e2|~epred2_0)|epred3_0))&(op(e1,e0)!=e3|(op(e0,e4)!=e2|~epred2_0)|epred3_0))&(op(e1,e0)!=e4|(op(e0,e4)!=e2|~epred2_0)|epred3_0)))&(((((op(e1,e0)!=e0|(op(e0,e4)!=e3|~epred2_0)|epred3_0)&(op(e1,e0)!=e1|(op(e0,e4)!=e3|~epred2_0)|epred3_0))&(op(e1,e0)!=e2|(op(e0,e4)!=e3|~epred2_0)|epred3_0))&(op(e1,e0)!=e3|(op(e0,e4)!=e3|~epred2_0)|epred3_0))&(op(e1,e0)!=e4|(op(e0,e4)!=e3|~epred2_0)|epred3_0)))&(((((op(e1,e0)!=e0|(op(e0,e4)!=e4|~epred2_0)|epred3_0)&(op(e1,e0)!=e1|(op(e0,e4)!=e4|~epred2_0)|epred3_0))&(op(e1,e0)!=e2|(op(e0,e4)!=e4|~epred2_0)|epred3_0))&(op(e1,e0)!=e3|(op(e0,e4)!=e4|~epred2_0)|epred3_0))&(op(e1,e0)!=e4|(op(e0,e4)!=e4|~epred2_0)|epred3_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_8])])).
cnf(c_0_14, plain, (epred2_0|op(e0,e3)!=e1|op(e0,e2)!=e4|~epred1_0), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (op(e0,e2)=e4), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_16, plain, (op(e0,e3)=e1), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_17, plain, (epred1_0|op(e0,e1)!=e3|op(e0,e0)!=e0|op(e0,op(e0,e0))!=e0|op(e0,op(e0,e1))!=e1|op(e0,op(e0,e2))!=e2|op(e0,op(e0,e3))!=e3|op(e0,op(e0,e4))!=e4|op(e1,op(e1,e0))!=e0|op(e1,op(e1,e1))!=e1|op(e1,op(e1,e2))!=e2|op(e1,op(e1,e3))!=e3|op(e1,op(e1,e4))!=e4|op(e2,op(e2,e0))!=e0|op(e2,op(e2,e1))!=e1|op(e2,op(e2,e2))!=e2|op(e2,op(e2,e3))!=e3|op(e2,op(e2,e4))!=e4|op(e3,op(e3,e0))!=e0|op(e3,op(e3,e1))!=e1|op(e3,op(e3,e2))!=e2|op(e3,op(e3,e3))!=e3|op(e3,op(e3,e4))!=e4|op(e4,op(e4,e0))!=e0|op(e4,op(e4,e1))!=e1|op(e4,op(e4,e2))!=e2|op(e4,op(e4,e3))!=e3|op(e4,op(e4,e4))!=e4), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, plain, (op(e0,e0)=e0), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_19, plain, (op(e0,e1)=e3), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_20, plain, (op(e0,e4)=e2), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_21, plain, (op(e1,e0)=e4), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_22, plain, (op(e1,e4)=e0), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_23, plain, (op(e1,e1)=e1), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_24, plain, (op(e1,e2)=e3), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_25, plain, (op(e1,e3)=e2), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_26, plain, (op(e2,e0)=e1), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_27, plain, (op(e2,e1)=e0), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_28, plain, (op(e2,e2)=e2), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_29, plain, (op(e2,e3)=e4), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_30, plain, (op(e2,e4)=e3), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_31, plain, (op(e3,e0)=e2), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_32, plain, (op(e3,e2)=e0), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_33, plain, (op(e3,e1)=e4), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_34, plain, (op(e3,e4)=e1), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_35, plain, (op(e3,e3)=e3), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_36, plain, (op(e4,e0)=e3), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_37, plain, (op(e4,e3)=e0), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_38, plain, (op(e4,e1)=e2), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_39, plain, (op(e4,e2)=e1), inference(split_conjunct,[status(thm)],[ax2])).
cnf(c_0_40, plain, (op(e4,e4)=e4), inference(split_conjunct,[status(thm)],[ax2])).
fof(c_0_41, plain, (epred6_0<=>((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))), introduced(definition)).
fof(c_0_42, plain, (((epred4_0&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))=>epred5_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_11]), c_0_7])).
fof(c_0_43, plain, (((((((((op(e1,e2)!=e0|(op(e1,e1)!=e0|~epred3_0)|epred4_0)&(op(e1,e2)!=e1|(op(e1,e1)!=e0|~epred3_0)|epred4_0))&(op(e1,e2)!=e2|(op(e1,e1)!=e0|~epred3_0)|epred4_0))&(op(e1,e2)!=e3|(op(e1,e1)!=e0|~epred3_0)|epred4_0))&(op(e1,e2)!=e4|(op(e1,e1)!=e0|~epred3_0)|epred4_0))&(((((op(e1,e2)!=e0|(op(e1,e1)!=e1|~epred3_0)|epred4_0)&(op(e1,e2)!=e1|(op(e1,e1)!=e1|~epred3_0)|epred4_0))&(op(e1,e2)!=e2|(op(e1,e1)!=e1|~epred3_0)|epred4_0))&(op(e1,e2)!=e3|(op(e1,e1)!=e1|~epred3_0)|epred4_0))&(op(e1,e2)!=e4|(op(e1,e1)!=e1|~epred3_0)|epred4_0)))&(((((op(e1,e2)!=e0|(op(e1,e1)!=e2|~epred3_0)|epred4_0)&(op(e1,e2)!=e1|(op(e1,e1)!=e2|~epred3_0)|epred4_0))&(op(e1,e2)!=e2|(op(e1,e1)!=e2|~epred3_0)|epred4_0))&(op(e1,e2)!=e3|(op(e1,e1)!=e2|~epred3_0)|epred4_0))&(op(e1,e2)!=e4|(op(e1,e1)!=e2|~epred3_0)|epred4_0)))&(((((op(e1,e2)!=e0|(op(e1,e1)!=e3|~epred3_0)|epred4_0)&(op(e1,e2)!=e1|(op(e1,e1)!=e3|~epred3_0)|epred4_0))&(op(e1,e2)!=e2|(op(e1,e1)!=e3|~epred3_0)|epred4_0))&(op(e1,e2)!=e3|(op(e1,e1)!=e3|~epred3_0)|epred4_0))&(op(e1,e2)!=e4|(op(e1,e1)!=e3|~epred3_0)|epred4_0)))&(((((op(e1,e2)!=e0|(op(e1,e1)!=e4|~epred3_0)|epred4_0)&(op(e1,e2)!=e1|(op(e1,e1)!=e4|~epred3_0)|epred4_0))&(op(e1,e2)!=e2|(op(e1,e1)!=e4|~epred3_0)|epred4_0))&(op(e1,e2)!=e3|(op(e1,e1)!=e4|~epred3_0)|epred4_0))&(op(e1,e2)!=e4|(op(e1,e1)!=e4|~epred3_0)|epred4_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])).
cnf(c_0_44, plain, (epred3_0|op(e1,e0)!=e4|op(e0,e4)!=e2|~epred2_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_45, plain, (epred2_0|~epred1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_15]), c_0_16])])).
cnf(c_0_46, plain, (epred1_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_18]), c_0_18]), c_0_19]), c_0_16]), c_0_15]), c_0_20]), c_0_16]), c_0_19]), c_0_20]), c_0_15]), c_0_21]), c_0_22]), c_0_23]), c_0_23]), c_0_24]), c_0_25]), c_0_25]), c_0_24]), c_0_22]), c_0_21]), c_0_26]), c_0_27]), c_0_27]), c_0_26]), c_0_28]), c_0_28]), c_0_29]), c_0_30]), c_0_30]), c_0_29]), c_0_31]), c_0_32]), c_0_33]), c_0_34]), c_0_32]), c_0_31]), c_0_35]), c_0_35]), c_0_34]), c_0_33]), c_0_36]), c_0_37]), c_0_38]), c_0_39]), c_0_39]), c_0_38]), c_0_37]), c_0_36]), c_0_40]), c_0_40])])).
fof(c_0_47, plain, (epred7_0<=>((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))), introduced(definition)).
fof(c_0_48, plain, (((epred5_0&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))=>epred6_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_41]), c_0_11])).
fof(c_0_49, plain, (((((((((op(e1,e4)!=e0|(op(e1,e3)!=e0|~epred4_0)|epred5_0)&(op(e1,e4)!=e1|(op(e1,e3)!=e0|~epred4_0)|epred5_0))&(op(e1,e4)!=e2|(op(e1,e3)!=e0|~epred4_0)|epred5_0))&(op(e1,e4)!=e3|(op(e1,e3)!=e0|~epred4_0)|epred5_0))&(op(e1,e4)!=e4|(op(e1,e3)!=e0|~epred4_0)|epred5_0))&(((((op(e1,e4)!=e0|(op(e1,e3)!=e1|~epred4_0)|epred5_0)&(op(e1,e4)!=e1|(op(e1,e3)!=e1|~epred4_0)|epred5_0))&(op(e1,e4)!=e2|(op(e1,e3)!=e1|~epred4_0)|epred5_0))&(op(e1,e4)!=e3|(op(e1,e3)!=e1|~epred4_0)|epred5_0))&(op(e1,e4)!=e4|(op(e1,e3)!=e1|~epred4_0)|epred5_0)))&(((((op(e1,e4)!=e0|(op(e1,e3)!=e2|~epred4_0)|epred5_0)&(op(e1,e4)!=e1|(op(e1,e3)!=e2|~epred4_0)|epred5_0))&(op(e1,e4)!=e2|(op(e1,e3)!=e2|~epred4_0)|epred5_0))&(op(e1,e4)!=e3|(op(e1,e3)!=e2|~epred4_0)|epred5_0))&(op(e1,e4)!=e4|(op(e1,e3)!=e2|~epred4_0)|epred5_0)))&(((((op(e1,e4)!=e0|(op(e1,e3)!=e3|~epred4_0)|epred5_0)&(op(e1,e4)!=e1|(op(e1,e3)!=e3|~epred4_0)|epred5_0))&(op(e1,e4)!=e2|(op(e1,e3)!=e3|~epred4_0)|epred5_0))&(op(e1,e4)!=e3|(op(e1,e3)!=e3|~epred4_0)|epred5_0))&(op(e1,e4)!=e4|(op(e1,e3)!=e3|~epred4_0)|epred5_0)))&(((((op(e1,e4)!=e0|(op(e1,e3)!=e4|~epred4_0)|epred5_0)&(op(e1,e4)!=e1|(op(e1,e3)!=e4|~epred4_0)|epred5_0))&(op(e1,e4)!=e2|(op(e1,e3)!=e4|~epred4_0)|epred5_0))&(op(e1,e4)!=e3|(op(e1,e3)!=e4|~epred4_0)|epred5_0))&(op(e1,e4)!=e4|(op(e1,e3)!=e4|~epred4_0)|epred5_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_42])])).
cnf(c_0_50, plain, (epred4_0|op(e1,e2)!=e3|op(e1,e1)!=e1|~epred3_0), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_51, plain, (epred3_0|~epred2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_20]), c_0_21])])).
cnf(c_0_52, plain, (epred2_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46])])).
fof(c_0_53, plain, (epred8_0<=>((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))), introduced(definition)).
fof(c_0_54, plain, (((epred6_0&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))=>epred7_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_47]), c_0_41])).
fof(c_0_55, plain, (((((((((op(e2,e1)!=e0|(op(e2,e0)!=e0|~epred5_0)|epred6_0)&(op(e2,e1)!=e1|(op(e2,e0)!=e0|~epred5_0)|epred6_0))&(op(e2,e1)!=e2|(op(e2,e0)!=e0|~epred5_0)|epred6_0))&(op(e2,e1)!=e3|(op(e2,e0)!=e0|~epred5_0)|epred6_0))&(op(e2,e1)!=e4|(op(e2,e0)!=e0|~epred5_0)|epred6_0))&(((((op(e2,e1)!=e0|(op(e2,e0)!=e1|~epred5_0)|epred6_0)&(op(e2,e1)!=e1|(op(e2,e0)!=e1|~epred5_0)|epred6_0))&(op(e2,e1)!=e2|(op(e2,e0)!=e1|~epred5_0)|epred6_0))&(op(e2,e1)!=e3|(op(e2,e0)!=e1|~epred5_0)|epred6_0))&(op(e2,e1)!=e4|(op(e2,e0)!=e1|~epred5_0)|epred6_0)))&(((((op(e2,e1)!=e0|(op(e2,e0)!=e2|~epred5_0)|epred6_0)&(op(e2,e1)!=e1|(op(e2,e0)!=e2|~epred5_0)|epred6_0))&(op(e2,e1)!=e2|(op(e2,e0)!=e2|~epred5_0)|epred6_0))&(op(e2,e1)!=e3|(op(e2,e0)!=e2|~epred5_0)|epred6_0))&(op(e2,e1)!=e4|(op(e2,e0)!=e2|~epred5_0)|epred6_0)))&(((((op(e2,e1)!=e0|(op(e2,e0)!=e3|~epred5_0)|epred6_0)&(op(e2,e1)!=e1|(op(e2,e0)!=e3|~epred5_0)|epred6_0))&(op(e2,e1)!=e2|(op(e2,e0)!=e3|~epred5_0)|epred6_0))&(op(e2,e1)!=e3|(op(e2,e0)!=e3|~epred5_0)|epred6_0))&(op(e2,e1)!=e4|(op(e2,e0)!=e3|~epred5_0)|epred6_0)))&(((((op(e2,e1)!=e0|(op(e2,e0)!=e4|~epred5_0)|epred6_0)&(op(e2,e1)!=e1|(op(e2,e0)!=e4|~epred5_0)|epred6_0))&(op(e2,e1)!=e2|(op(e2,e0)!=e4|~epred5_0)|epred6_0))&(op(e2,e1)!=e3|(op(e2,e0)!=e4|~epred5_0)|epred6_0))&(op(e2,e1)!=e4|(op(e2,e0)!=e4|~epred5_0)|epred6_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_48])])).
cnf(c_0_56, plain, (epred5_0|op(e1,e4)!=e0|op(e1,e3)!=e2|~epred4_0), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_57, plain, (epred4_0|~epred3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_23]), c_0_24])])).
cnf(c_0_58, plain, (epred3_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_52])])).
fof(c_0_59, plain, (epred9_0<=>((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))), introduced(definition)).
fof(c_0_60, plain, (((epred7_0&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))=>epred8_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_53]), c_0_47])).
fof(c_0_61, plain, (((((((((op(e2,e3)!=e0|(op(e2,e2)!=e0|~epred6_0)|epred7_0)&(op(e2,e3)!=e1|(op(e2,e2)!=e0|~epred6_0)|epred7_0))&(op(e2,e3)!=e2|(op(e2,e2)!=e0|~epred6_0)|epred7_0))&(op(e2,e3)!=e3|(op(e2,e2)!=e0|~epred6_0)|epred7_0))&(op(e2,e3)!=e4|(op(e2,e2)!=e0|~epred6_0)|epred7_0))&(((((op(e2,e3)!=e0|(op(e2,e2)!=e1|~epred6_0)|epred7_0)&(op(e2,e3)!=e1|(op(e2,e2)!=e1|~epred6_0)|epred7_0))&(op(e2,e3)!=e2|(op(e2,e2)!=e1|~epred6_0)|epred7_0))&(op(e2,e3)!=e3|(op(e2,e2)!=e1|~epred6_0)|epred7_0))&(op(e2,e3)!=e4|(op(e2,e2)!=e1|~epred6_0)|epred7_0)))&(((((op(e2,e3)!=e0|(op(e2,e2)!=e2|~epred6_0)|epred7_0)&(op(e2,e3)!=e1|(op(e2,e2)!=e2|~epred6_0)|epred7_0))&(op(e2,e3)!=e2|(op(e2,e2)!=e2|~epred6_0)|epred7_0))&(op(e2,e3)!=e3|(op(e2,e2)!=e2|~epred6_0)|epred7_0))&(op(e2,e3)!=e4|(op(e2,e2)!=e2|~epred6_0)|epred7_0)))&(((((op(e2,e3)!=e0|(op(e2,e2)!=e3|~epred6_0)|epred7_0)&(op(e2,e3)!=e1|(op(e2,e2)!=e3|~epred6_0)|epred7_0))&(op(e2,e3)!=e2|(op(e2,e2)!=e3|~epred6_0)|epred7_0))&(op(e2,e3)!=e3|(op(e2,e2)!=e3|~epred6_0)|epred7_0))&(op(e2,e3)!=e4|(op(e2,e2)!=e3|~epred6_0)|epred7_0)))&(((((op(e2,e3)!=e0|(op(e2,e2)!=e4|~epred6_0)|epred7_0)&(op(e2,e3)!=e1|(op(e2,e2)!=e4|~epred6_0)|epred7_0))&(op(e2,e3)!=e2|(op(e2,e2)!=e4|~epred6_0)|epred7_0))&(op(e2,e3)!=e3|(op(e2,e2)!=e4|~epred6_0)|epred7_0))&(op(e2,e3)!=e4|(op(e2,e2)!=e4|~epred6_0)|epred7_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_54])])).
cnf(c_0_62, plain, (epred6_0|op(e2,e1)!=e0|op(e2,e0)!=e1|~epred5_0), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_63, plain, (epred5_0|~epred4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_25]), c_0_22])])).
cnf(c_0_64, plain, (epred4_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_58])])).
fof(c_0_65, plain, (epred10_0<=>((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))), introduced(definition)).
fof(c_0_66, plain, (((epred8_0&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))=>epred9_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_59]), c_0_53])).
fof(c_0_67, plain, (((((((((op(e3,e0)!=e0|(op(e2,e4)!=e0|~epred7_0)|epred8_0)&(op(e3,e0)!=e1|(op(e2,e4)!=e0|~epred7_0)|epred8_0))&(op(e3,e0)!=e2|(op(e2,e4)!=e0|~epred7_0)|epred8_0))&(op(e3,e0)!=e3|(op(e2,e4)!=e0|~epred7_0)|epred8_0))&(op(e3,e0)!=e4|(op(e2,e4)!=e0|~epred7_0)|epred8_0))&(((((op(e3,e0)!=e0|(op(e2,e4)!=e1|~epred7_0)|epred8_0)&(op(e3,e0)!=e1|(op(e2,e4)!=e1|~epred7_0)|epred8_0))&(op(e3,e0)!=e2|(op(e2,e4)!=e1|~epred7_0)|epred8_0))&(op(e3,e0)!=e3|(op(e2,e4)!=e1|~epred7_0)|epred8_0))&(op(e3,e0)!=e4|(op(e2,e4)!=e1|~epred7_0)|epred8_0)))&(((((op(e3,e0)!=e0|(op(e2,e4)!=e2|~epred7_0)|epred8_0)&(op(e3,e0)!=e1|(op(e2,e4)!=e2|~epred7_0)|epred8_0))&(op(e3,e0)!=e2|(op(e2,e4)!=e2|~epred7_0)|epred8_0))&(op(e3,e0)!=e3|(op(e2,e4)!=e2|~epred7_0)|epred8_0))&(op(e3,e0)!=e4|(op(e2,e4)!=e2|~epred7_0)|epred8_0)))&(((((op(e3,e0)!=e0|(op(e2,e4)!=e3|~epred7_0)|epred8_0)&(op(e3,e0)!=e1|(op(e2,e4)!=e3|~epred7_0)|epred8_0))&(op(e3,e0)!=e2|(op(e2,e4)!=e3|~epred7_0)|epred8_0))&(op(e3,e0)!=e3|(op(e2,e4)!=e3|~epred7_0)|epred8_0))&(op(e3,e0)!=e4|(op(e2,e4)!=e3|~epred7_0)|epred8_0)))&(((((op(e3,e0)!=e0|(op(e2,e4)!=e4|~epred7_0)|epred8_0)&(op(e3,e0)!=e1|(op(e2,e4)!=e4|~epred7_0)|epred8_0))&(op(e3,e0)!=e2|(op(e2,e4)!=e4|~epred7_0)|epred8_0))&(op(e3,e0)!=e3|(op(e2,e4)!=e4|~epred7_0)|epred8_0))&(op(e3,e0)!=e4|(op(e2,e4)!=e4|~epred7_0)|epred8_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_60])])).
cnf(c_0_68, plain, (epred7_0|op(e2,e3)!=e4|op(e2,e2)!=e2|~epred6_0), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_69, plain, (epred6_0|~epred5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_62, c_0_26]), c_0_27])])).
cnf(c_0_70, plain, (epred5_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_64])])).
fof(c_0_71, plain, (epred11_0<=>((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))), introduced(definition)).
fof(c_0_72, plain, (((epred9_0&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))=>epred10_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_65]), c_0_59])).
fof(c_0_73, plain, (((((((((op(e3,e2)!=e0|(op(e3,e1)!=e0|~epred8_0)|epred9_0)&(op(e3,e2)!=e1|(op(e3,e1)!=e0|~epred8_0)|epred9_0))&(op(e3,e2)!=e2|(op(e3,e1)!=e0|~epred8_0)|epred9_0))&(op(e3,e2)!=e3|(op(e3,e1)!=e0|~epred8_0)|epred9_0))&(op(e3,e2)!=e4|(op(e3,e1)!=e0|~epred8_0)|epred9_0))&(((((op(e3,e2)!=e0|(op(e3,e1)!=e1|~epred8_0)|epred9_0)&(op(e3,e2)!=e1|(op(e3,e1)!=e1|~epred8_0)|epred9_0))&(op(e3,e2)!=e2|(op(e3,e1)!=e1|~epred8_0)|epred9_0))&(op(e3,e2)!=e3|(op(e3,e1)!=e1|~epred8_0)|epred9_0))&(op(e3,e2)!=e4|(op(e3,e1)!=e1|~epred8_0)|epred9_0)))&(((((op(e3,e2)!=e0|(op(e3,e1)!=e2|~epred8_0)|epred9_0)&(op(e3,e2)!=e1|(op(e3,e1)!=e2|~epred8_0)|epred9_0))&(op(e3,e2)!=e2|(op(e3,e1)!=e2|~epred8_0)|epred9_0))&(op(e3,e2)!=e3|(op(e3,e1)!=e2|~epred8_0)|epred9_0))&(op(e3,e2)!=e4|(op(e3,e1)!=e2|~epred8_0)|epred9_0)))&(((((op(e3,e2)!=e0|(op(e3,e1)!=e3|~epred8_0)|epred9_0)&(op(e3,e2)!=e1|(op(e3,e1)!=e3|~epred8_0)|epred9_0))&(op(e3,e2)!=e2|(op(e3,e1)!=e3|~epred8_0)|epred9_0))&(op(e3,e2)!=e3|(op(e3,e1)!=e3|~epred8_0)|epred9_0))&(op(e3,e2)!=e4|(op(e3,e1)!=e3|~epred8_0)|epred9_0)))&(((((op(e3,e2)!=e0|(op(e3,e1)!=e4|~epred8_0)|epred9_0)&(op(e3,e2)!=e1|(op(e3,e1)!=e4|~epred8_0)|epred9_0))&(op(e3,e2)!=e2|(op(e3,e1)!=e4|~epred8_0)|epred9_0))&(op(e3,e2)!=e3|(op(e3,e1)!=e4|~epred8_0)|epred9_0))&(op(e3,e2)!=e4|(op(e3,e1)!=e4|~epred8_0)|epred9_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_66])])).
cnf(c_0_74, plain, (epred8_0|op(e3,e0)!=e2|op(e2,e4)!=e3|~epred7_0), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, plain, (epred7_0|~epred6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_28]), c_0_29])])).
cnf(c_0_76, plain, (epred6_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_69, c_0_70])])).
fof(c_0_77, plain, (epred37_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3|op(e4,e2)=e3))&(op(e2,e0)=e4|op(e2,e1)=e4|op(e2,e2)=e4|op(e2,e3)=e4|op(e2,e4)=e4))&(op(e0,e2)=e4|op(e1,e2)=e4|op(e2,e2)=e4|op(e3,e2)=e4|op(e4,e2)=e4))&(op(e3,e0)=e0|op(e3,e1)=e0|op(e3,e2)=e0|op(e3,e3)=e0|op(e3,e4)=e0))&(op(e0,e3)=e0|op(e1,e3)=e0|op(e2,e3)=e0|op(e3,e3)=e0|op(e4,e3)=e0))&(op(e3,e0)=e1|op(e3,e1)=e1|op(e3,e2)=e1|op(e3,e3)=e1|op(e3,e4)=e1))&(op(e0,e3)=e1|op(e1,e3)=e1|op(e2,e3)=e1|op(e3,e3)=e1|op(e4,e3)=e1))&(op(e3,e0)=e2|op(e3,e1)=e2|op(e3,e2)=e2|op(e3,e3)=e2|op(e3,e4)=e2))&(op(e0,e3)=e2|op(e1,e3)=e2|op(e2,e3)=e2|op(e3,e3)=e2|op(e4,e3)=e2))&(op(e3,e0)=e3|op(e3,e1)=e3|op(e3,e2)=e3|op(e3,e3)=e3|op(e3,e4)=e3))&(op(e0,e3)=e3|op(e1,e3)=e3|op(e2,e3)=e3|op(e3,e3)=e3|op(e4,e3)=e3))&(op(e3,e0)=e4|op(e3,e1)=e4|op(e3,e2)=e4|op(e3,e3)=e4|op(e3,e4)=e4))&(op(e0,e3)=e4|op(e1,e3)=e4|op(e2,e3)=e4|op(e3,e3)=e4|op(e4,e3)=e4))&(op(e4,e0)=e0|op(e4,e1)=e0|op(e4,e2)=e0|op(e4,e3)=e0|op(e4,e4)=e0))&(op(e0,e4)=e0|op(e1,e4)=e0|op(e2,e4)=e0|op(e3,e4)=e0|op(e4,e4)=e0))&(op(e4,e0)=e1|op(e4,e1)=e1|op(e4,e2)=e1|op(e4,e3)=e1|op(e4,e4)=e1))&(op(e0,e4)=e1|op(e1,e4)=e1|op(e2,e4)=e1|op(e3,e4)=e1|op(e4,e4)=e1))&(op(e4,e0)=e2|op(e4,e1)=e2|op(e4,e2)=e2|op(e4,e3)=e2|op(e4,e4)=e2))&(op(e0,e4)=e2|op(e1,e4)=e2|op(e2,e4)=e2|op(e3,e4)=e2|op(e4,e4)=e2))&(op(e4,e0)=e3|op(e4,e1)=e3|op(e4,e2)=e3|op(e4,e3)=e3|op(e4,e4)=e3))&(op(e0,e4)=e3|op(e1,e4)=e3|op(e2,e4)=e3|op(e3,e4)=e3|op(e4,e4)=e3))&(op(e4,e0)=e4|op(e4,e1)=e4|op(e4,e2)=e4|op(e4,e3)=e4|op(e4,e4)=e4))), introduced(definition)).
fof(c_0_78, plain, (epred12_0<=>((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))), introduced(definition)).
fof(c_0_79, plain, (((epred10_0&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))=>epred11_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_71]), c_0_65])).
fof(c_0_80, plain, (((((((((op(e3,e4)!=e0|(op(e3,e3)!=e0|~epred9_0)|epred10_0)&(op(e3,e4)!=e1|(op(e3,e3)!=e0|~epred9_0)|epred10_0))&(op(e3,e4)!=e2|(op(e3,e3)!=e0|~epred9_0)|epred10_0))&(op(e3,e4)!=e3|(op(e3,e3)!=e0|~epred9_0)|epred10_0))&(op(e3,e4)!=e4|(op(e3,e3)!=e0|~epred9_0)|epred10_0))&(((((op(e3,e4)!=e0|(op(e3,e3)!=e1|~epred9_0)|epred10_0)&(op(e3,e4)!=e1|(op(e3,e3)!=e1|~epred9_0)|epred10_0))&(op(e3,e4)!=e2|(op(e3,e3)!=e1|~epred9_0)|epred10_0))&(op(e3,e4)!=e3|(op(e3,e3)!=e1|~epred9_0)|epred10_0))&(op(e3,e4)!=e4|(op(e3,e3)!=e1|~epred9_0)|epred10_0)))&(((((op(e3,e4)!=e0|(op(e3,e3)!=e2|~epred9_0)|epred10_0)&(op(e3,e4)!=e1|(op(e3,e3)!=e2|~epred9_0)|epred10_0))&(op(e3,e4)!=e2|(op(e3,e3)!=e2|~epred9_0)|epred10_0))&(op(e3,e4)!=e3|(op(e3,e3)!=e2|~epred9_0)|epred10_0))&(op(e3,e4)!=e4|(op(e3,e3)!=e2|~epred9_0)|epred10_0)))&(((((op(e3,e4)!=e0|(op(e3,e3)!=e3|~epred9_0)|epred10_0)&(op(e3,e4)!=e1|(op(e3,e3)!=e3|~epred9_0)|epred10_0))&(op(e3,e4)!=e2|(op(e3,e3)!=e3|~epred9_0)|epred10_0))&(op(e3,e4)!=e3|(op(e3,e3)!=e3|~epred9_0)|epred10_0))&(op(e3,e4)!=e4|(op(e3,e3)!=e3|~epred9_0)|epred10_0)))&(((((op(e3,e4)!=e0|(op(e3,e3)!=e4|~epred9_0)|epred10_0)&(op(e3,e4)!=e1|(op(e3,e3)!=e4|~epred9_0)|epred10_0))&(op(e3,e4)!=e2|(op(e3,e3)!=e4|~epred9_0)|epred10_0))&(op(e3,e4)!=e3|(op(e3,e3)!=e4|~epred9_0)|epred10_0))&(op(e3,e4)!=e4|(op(e3,e3)!=e4|~epred9_0)|epred10_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_72])])).
cnf(c_0_81, plain, (epred9_0|op(e3,e2)!=e0|op(e3,e1)!=e4|~epred8_0), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_82, plain, (epred8_0|~epred7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_30]), c_0_31])])).
cnf(c_0_83, plain, (epred7_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_76])])).
fof(c_0_84, plain, (epred36_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3|op(e4,e2)=e3))&(op(e2,e0)=e4|op(e2,e1)=e4|op(e2,e2)=e4|op(e2,e3)=e4|op(e2,e4)=e4))&(op(e0,e2)=e4|op(e1,e2)=e4|op(e2,e2)=e4|op(e3,e2)=e4|op(e4,e2)=e4))&(op(e3,e0)=e0|op(e3,e1)=e0|op(e3,e2)=e0|op(e3,e3)=e0|op(e3,e4)=e0))&(op(e0,e3)=e0|op(e1,e3)=e0|op(e2,e3)=e0|op(e3,e3)=e0|op(e4,e3)=e0))&(op(e3,e0)=e1|op(e3,e1)=e1|op(e3,e2)=e1|op(e3,e3)=e1|op(e3,e4)=e1))&(op(e0,e3)=e1|op(e1,e3)=e1|op(e2,e3)=e1|op(e3,e3)=e1|op(e4,e3)=e1))&(op(e3,e0)=e2|op(e3,e1)=e2|op(e3,e2)=e2|op(e3,e3)=e2|op(e3,e4)=e2))&(op(e0,e3)=e2|op(e1,e3)=e2|op(e2,e3)=e2|op(e3,e3)=e2|op(e4,e3)=e2))&(op(e3,e0)=e3|op(e3,e1)=e3|op(e3,e2)=e3|op(e3,e3)=e3|op(e3,e4)=e3))&(op(e0,e3)=e3|op(e1,e3)=e3|op(e2,e3)=e3|op(e3,e3)=e3|op(e4,e3)=e3))&(op(e3,e0)=e4|op(e3,e1)=e4|op(e3,e2)=e4|op(e3,e3)=e4|op(e3,e4)=e4))&(op(e0,e3)=e4|op(e1,e3)=e4|op(e2,e3)=e4|op(e3,e3)=e4|op(e4,e3)=e4))&(op(e4,e0)=e0|op(e4,e1)=e0|op(e4,e2)=e0|op(e4,e3)=e0|op(e4,e4)=e0))&(op(e0,e4)=e0|op(e1,e4)=e0|op(e2,e4)=e0|op(e3,e4)=e0|op(e4,e4)=e0))&(op(e4,e0)=e1|op(e4,e1)=e1|op(e4,e2)=e1|op(e4,e3)=e1|op(e4,e4)=e1))&(op(e0,e4)=e1|op(e1,e4)=e1|op(e2,e4)=e1|op(e3,e4)=e1|op(e4,e4)=e1))&(op(e4,e0)=e2|op(e4,e1)=e2|op(e4,e2)=e2|op(e4,e3)=e2|op(e4,e4)=e2))&(op(e0,e4)=e2|op(e1,e4)=e2|op(e2,e4)=e2|op(e3,e4)=e2|op(e4,e4)=e2))&(op(e4,e0)=e3|op(e4,e1)=e3|op(e4,e2)=e3|op(e4,e3)=e3|op(e4,e4)=e3))), introduced(definition)).
fof(c_0_85, negated_conjecture, ~(((((((((((((((((((((((((((epred37_0&(op(e0,e4)=e4|op(e1,e4)=e4|op(e2,e4)=e4|op(e3,e4)=e4|op(e4,e4)=e4))&op(op(e0,e0),op(e0,e0))=e0)&op(op(e1,e0),op(e0,e1))=e0)&op(op(e2,e0),op(e0,e2))=e0)&op(op(e3,e0),op(e0,e3))=e0)&op(op(e4,e0),op(e0,e4))=e0)&op(op(e0,e1),op(e1,e0))=e1)&op(op(e1,e1),op(e1,e1))=e1)&op(op(e2,e1),op(e1,e2))=e1)&op(op(e3,e1),op(e1,e3))=e1)&op(op(e4,e1),op(e1,e4))=e1)&op(op(e0,e2),op(e2,e0))=e2)&op(op(e1,e2),op(e2,e1))=e2)&op(op(e2,e2),op(e2,e2))=e2)&op(op(e3,e2),op(e2,e3))=e2)&op(op(e4,e2),op(e2,e4))=e2)&op(op(e0,e3),op(e3,e0))=e3)&op(op(e1,e3),op(e3,e1))=e3)&op(op(e2,e3),op(e3,e2))=e3)&op(op(e3,e3),op(e3,e3))=e3)&op(op(e4,e3),op(e3,e4))=e3)&op(op(e0,e4),op(e4,e0))=e4)&op(op(e1,e4),op(e4,e1))=e4)&op(op(e2,e4),op(e4,e2))=e4)&op(op(e3,e4),op(e4,e3))=e4)&op(op(e4,e4),op(e4,e4))=e4)), inference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[co1]), c_0_77])).
fof(c_0_86, plain, (epred35_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3|op(e4,e2)=e3))&(op(e2,e0)=e4|op(e2,e1)=e4|op(e2,e2)=e4|op(e2,e3)=e4|op(e2,e4)=e4))&(op(e0,e2)=e4|op(e1,e2)=e4|op(e2,e2)=e4|op(e3,e2)=e4|op(e4,e2)=e4))&(op(e3,e0)=e0|op(e3,e1)=e0|op(e3,e2)=e0|op(e3,e3)=e0|op(e3,e4)=e0))&(op(e0,e3)=e0|op(e1,e3)=e0|op(e2,e3)=e0|op(e3,e3)=e0|op(e4,e3)=e0))&(op(e3,e0)=e1|op(e3,e1)=e1|op(e3,e2)=e1|op(e3,e3)=e1|op(e3,e4)=e1))&(op(e0,e3)=e1|op(e1,e3)=e1|op(e2,e3)=e1|op(e3,e3)=e1|op(e4,e3)=e1))&(op(e3,e0)=e2|op(e3,e1)=e2|op(e3,e2)=e2|op(e3,e3)=e2|op(e3,e4)=e2))&(op(e0,e3)=e2|op(e1,e3)=e2|op(e2,e3)=e2|op(e3,e3)=e2|op(e4,e3)=e2))&(op(e3,e0)=e3|op(e3,e1)=e3|op(e3,e2)=e3|op(e3,e3)=e3|op(e3,e4)=e3))&(op(e0,e3)=e3|op(e1,e3)=e3|op(e2,e3)=e3|op(e3,e3)=e3|op(e4,e3)=e3))&(op(e3,e0)=e4|op(e3,e1)=e4|op(e3,e2)=e4|op(e3,e3)=e4|op(e3,e4)=e4))&(op(e0,e3)=e4|op(e1,e3)=e4|op(e2,e3)=e4|op(e3,e3)=e4|op(e4,e3)=e4))&(op(e4,e0)=e0|op(e4,e1)=e0|op(e4,e2)=e0|op(e4,e3)=e0|op(e4,e4)=e0))&(op(e0,e4)=e0|op(e1,e4)=e0|op(e2,e4)=e0|op(e3,e4)=e0|op(e4,e4)=e0))&(op(e4,e0)=e1|op(e4,e1)=e1|op(e4,e2)=e1|op(e4,e3)=e1|op(e4,e4)=e1))&(op(e0,e4)=e1|op(e1,e4)=e1|op(e2,e4)=e1|op(e3,e4)=e1|op(e4,e4)=e1))&(op(e4,e0)=e2|op(e4,e1)=e2|op(e4,e2)=e2|op(e4,e3)=e2|op(e4,e4)=e2))), introduced(definition)).
fof(c_0_87, plain, (epred13_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))), introduced(definition)).
fof(c_0_88, plain, (((epred11_0&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))=>epred12_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_78]), c_0_71])).
fof(c_0_89, plain, (((((((((op(e4,e1)!=e0|(op(e4,e0)!=e0|~epred10_0)|epred11_0)&(op(e4,e1)!=e1|(op(e4,e0)!=e0|~epred10_0)|epred11_0))&(op(e4,e1)!=e2|(op(e4,e0)!=e0|~epred10_0)|epred11_0))&(op(e4,e1)!=e3|(op(e4,e0)!=e0|~epred10_0)|epred11_0))&(op(e4,e1)!=e4|(op(e4,e0)!=e0|~epred10_0)|epred11_0))&(((((op(e4,e1)!=e0|(op(e4,e0)!=e1|~epred10_0)|epred11_0)&(op(e4,e1)!=e1|(op(e4,e0)!=e1|~epred10_0)|epred11_0))&(op(e4,e1)!=e2|(op(e4,e0)!=e1|~epred10_0)|epred11_0))&(op(e4,e1)!=e3|(op(e4,e0)!=e1|~epred10_0)|epred11_0))&(op(e4,e1)!=e4|(op(e4,e0)!=e1|~epred10_0)|epred11_0)))&(((((op(e4,e1)!=e0|(op(e4,e0)!=e2|~epred10_0)|epred11_0)&(op(e4,e1)!=e1|(op(e4,e0)!=e2|~epred10_0)|epred11_0))&(op(e4,e1)!=e2|(op(e4,e0)!=e2|~epred10_0)|epred11_0))&(op(e4,e1)!=e3|(op(e4,e0)!=e2|~epred10_0)|epred11_0))&(op(e4,e1)!=e4|(op(e4,e0)!=e2|~epred10_0)|epred11_0)))&(((((op(e4,e1)!=e0|(op(e4,e0)!=e3|~epred10_0)|epred11_0)&(op(e4,e1)!=e1|(op(e4,e0)!=e3|~epred10_0)|epred11_0))&(op(e4,e1)!=e2|(op(e4,e0)!=e3|~epred10_0)|epred11_0))&(op(e4,e1)!=e3|(op(e4,e0)!=e3|~epred10_0)|epred11_0))&(op(e4,e1)!=e4|(op(e4,e0)!=e3|~epred10_0)|epred11_0)))&(((((op(e4,e1)!=e0|(op(e4,e0)!=e4|~epred10_0)|epred11_0)&(op(e4,e1)!=e1|(op(e4,e0)!=e4|~epred10_0)|epred11_0))&(op(e4,e1)!=e2|(op(e4,e0)!=e4|~epred10_0)|epred11_0))&(op(e4,e1)!=e3|(op(e4,e0)!=e4|~epred10_0)|epred11_0))&(op(e4,e1)!=e4|(op(e4,e0)!=e4|~epred10_0)|epred11_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_79])])).
cnf(c_0_90, plain, (epred10_0|op(e3,e4)!=e1|op(e3,e3)!=e3|~epred9_0), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_91, plain, (epred9_0|~epred8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_81, c_0_33]), c_0_32])])).
cnf(c_0_92, plain, (epred8_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_83])])).
fof(c_0_93, plain, (((epred36_0&(op(e0,e4)=e3|op(e1,e4)=e3|op(e2,e4)=e3|op(e3,e4)=e3|op(e4,e4)=e3))&(op(e4,e0)=e4|op(e4,e1)=e4|op(e4,e2)=e4|op(e4,e3)=e4|op(e4,e4)=e4))=>epred37_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_77]), c_0_84])).
fof(c_0_94, negated_conjecture, (((((op(e0,e4)!=e4|~epred37_0|op(op(e0,e0),op(e0,e0))!=e0|op(op(e1,e0),op(e0,e1))!=e0|op(op(e2,e0),op(e0,e2))!=e0|op(op(e3,e0),op(e0,e3))!=e0|op(op(e4,e0),op(e0,e4))!=e0|op(op(e0,e1),op(e1,e0))!=e1|op(op(e1,e1),op(e1,e1))!=e1|op(op(e2,e1),op(e1,e2))!=e1|op(op(e3,e1),op(e1,e3))!=e1|op(op(e4,e1),op(e1,e4))!=e1|op(op(e0,e2),op(e2,e0))!=e2|op(op(e1,e2),op(e2,e1))!=e2|op(op(e2,e2),op(e2,e2))!=e2|op(op(e3,e2),op(e2,e3))!=e2|op(op(e4,e2),op(e2,e4))!=e2|op(op(e0,e3),op(e3,e0))!=e3|op(op(e1,e3),op(e3,e1))!=e3|op(op(e2,e3),op(e3,e2))!=e3|op(op(e3,e3),op(e3,e3))!=e3|op(op(e4,e3),op(e3,e4))!=e3|op(op(e0,e4),op(e4,e0))!=e4|op(op(e1,e4),op(e4,e1))!=e4|op(op(e2,e4),op(e4,e2))!=e4|op(op(e3,e4),op(e4,e3))!=e4|op(op(e4,e4),op(e4,e4))!=e4)&(op(e1,e4)!=e4|~epred37_0|op(op(e0,e0),op(e0,e0))!=e0|op(op(e1,e0),op(e0,e1))!=e0|op(op(e2,e0),op(e0,e2))!=e0|op(op(e3,e0),op(e0,e3))!=e0|op(op(e4,e0),op(e0,e4))!=e0|op(op(e0,e1),op(e1,e0))!=e1|op(op(e1,e1),op(e1,e1))!=e1|op(op(e2,e1),op(e1,e2))!=e1|op(op(e3,e1),op(e1,e3))!=e1|op(op(e4,e1),op(e1,e4))!=e1|op(op(e0,e2),op(e2,e0))!=e2|op(op(e1,e2),op(e2,e1))!=e2|op(op(e2,e2),op(e2,e2))!=e2|op(op(e3,e2),op(e2,e3))!=e2|op(op(e4,e2),op(e2,e4))!=e2|op(op(e0,e3),op(e3,e0))!=e3|op(op(e1,e3),op(e3,e1))!=e3|op(op(e2,e3),op(e3,e2))!=e3|op(op(e3,e3),op(e3,e3))!=e3|op(op(e4,e3),op(e3,e4))!=e3|op(op(e0,e4),op(e4,e0))!=e4|op(op(e1,e4),op(e4,e1))!=e4|op(op(e2,e4),op(e4,e2))!=e4|op(op(e3,e4),op(e4,e3))!=e4|op(op(e4,e4),op(e4,e4))!=e4))&(op(e2,e4)!=e4|~epred37_0|op(op(e0,e0),op(e0,e0))!=e0|op(op(e1,e0),op(e0,e1))!=e0|op(op(e2,e0),op(e0,e2))!=e0|op(op(e3,e0),op(e0,e3))!=e0|op(op(e4,e0),op(e0,e4))!=e0|op(op(e0,e1),op(e1,e0))!=e1|op(op(e1,e1),op(e1,e1))!=e1|op(op(e2,e1),op(e1,e2))!=e1|op(op(e3,e1),op(e1,e3))!=e1|op(op(e4,e1),op(e1,e4))!=e1|op(op(e0,e2),op(e2,e0))!=e2|op(op(e1,e2),op(e2,e1))!=e2|op(op(e2,e2),op(e2,e2))!=e2|op(op(e3,e2),op(e2,e3))!=e2|op(op(e4,e2),op(e2,e4))!=e2|op(op(e0,e3),op(e3,e0))!=e3|op(op(e1,e3),op(e3,e1))!=e3|op(op(e2,e3),op(e3,e2))!=e3|op(op(e3,e3),op(e3,e3))!=e3|op(op(e4,e3),op(e3,e4))!=e3|op(op(e0,e4),op(e4,e0))!=e4|op(op(e1,e4),op(e4,e1))!=e4|op(op(e2,e4),op(e4,e2))!=e4|op(op(e3,e4),op(e4,e3))!=e4|op(op(e4,e4),op(e4,e4))!=e4))&(op(e3,e4)!=e4|~epred37_0|op(op(e0,e0),op(e0,e0))!=e0|op(op(e1,e0),op(e0,e1))!=e0|op(op(e2,e0),op(e0,e2))!=e0|op(op(e3,e0),op(e0,e3))!=e0|op(op(e4,e0),op(e0,e4))!=e0|op(op(e0,e1),op(e1,e0))!=e1|op(op(e1,e1),op(e1,e1))!=e1|op(op(e2,e1),op(e1,e2))!=e1|op(op(e3,e1),op(e1,e3))!=e1|op(op(e4,e1),op(e1,e4))!=e1|op(op(e0,e2),op(e2,e0))!=e2|op(op(e1,e2),op(e2,e1))!=e2|op(op(e2,e2),op(e2,e2))!=e2|op(op(e3,e2),op(e2,e3))!=e2|op(op(e4,e2),op(e2,e4))!=e2|op(op(e0,e3),op(e3,e0))!=e3|op(op(e1,e3),op(e3,e1))!=e3|op(op(e2,e3),op(e3,e2))!=e3|op(op(e3,e3),op(e3,e3))!=e3|op(op(e4,e3),op(e3,e4))!=e3|op(op(e0,e4),op(e4,e0))!=e4|op(op(e1,e4),op(e4,e1))!=e4|op(op(e2,e4),op(e4,e2))!=e4|op(op(e3,e4),op(e4,e3))!=e4|op(op(e4,e4),op(e4,e4))!=e4))&(op(e4,e4)!=e4|~epred37_0|op(op(e0,e0),op(e0,e0))!=e0|op(op(e1,e0),op(e0,e1))!=e0|op(op(e2,e0),op(e0,e2))!=e0|op(op(e3,e0),op(e0,e3))!=e0|op(op(e4,e0),op(e0,e4))!=e0|op(op(e0,e1),op(e1,e0))!=e1|op(op(e1,e1),op(e1,e1))!=e1|op(op(e2,e1),op(e1,e2))!=e1|op(op(e3,e1),op(e1,e3))!=e1|op(op(e4,e1),op(e1,e4))!=e1|op(op(e0,e2),op(e2,e0))!=e2|op(op(e1,e2),op(e2,e1))!=e2|op(op(e2,e2),op(e2,e2))!=e2|op(op(e3,e2),op(e2,e3))!=e2|op(op(e4,e2),op(e2,e4))!=e2|op(op(e0,e3),op(e3,e0))!=e3|op(op(e1,e3),op(e3,e1))!=e3|op(op(e2,e3),op(e3,e2))!=e3|op(op(e3,e3),op(e3,e3))!=e3|op(op(e4,e3),op(e3,e4))!=e3|op(op(e0,e4),op(e4,e0))!=e4|op(op(e1,e4),op(e4,e1))!=e4|op(op(e2,e4),op(e4,e2))!=e4|op(op(e3,e4),op(e4,e3))!=e4|op(op(e4,e4),op(e4,e4))!=e4)), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_85])])).
fof(c_0_95, plain, (((epred35_0&(op(e0,e4)=e2|op(e1,e4)=e2|op(e2,e4)=e2|op(e3,e4)=e2|op(e4,e4)=e2))&(op(e4,e0)=e3|op(e4,e1)=e3|op(e4,e2)=e3|op(e4,e3)=e3|op(e4,e4)=e3))=>epred36_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_84]), c_0_86])).
fof(c_0_96, plain, (epred34_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3|op(e4,e2)=e3))&(op(e2,e0)=e4|op(e2,e1)=e4|op(e2,e2)=e4|op(e2,e3)=e4|op(e2,e4)=e4))&(op(e0,e2)=e4|op(e1,e2)=e4|op(e2,e2)=e4|op(e3,e2)=e4|op(e4,e2)=e4))&(op(e3,e0)=e0|op(e3,e1)=e0|op(e3,e2)=e0|op(e3,e3)=e0|op(e3,e4)=e0))&(op(e0,e3)=e0|op(e1,e3)=e0|op(e2,e3)=e0|op(e3,e3)=e0|op(e4,e3)=e0))&(op(e3,e0)=e1|op(e3,e1)=e1|op(e3,e2)=e1|op(e3,e3)=e1|op(e3,e4)=e1))&(op(e0,e3)=e1|op(e1,e3)=e1|op(e2,e3)=e1|op(e3,e3)=e1|op(e4,e3)=e1))&(op(e3,e0)=e2|op(e3,e1)=e2|op(e3,e2)=e2|op(e3,e3)=e2|op(e3,e4)=e2))&(op(e0,e3)=e2|op(e1,e3)=e2|op(e2,e3)=e2|op(e3,e3)=e2|op(e4,e3)=e2))&(op(e3,e0)=e3|op(e3,e1)=e3|op(e3,e2)=e3|op(e3,e3)=e3|op(e3,e4)=e3))&(op(e0,e3)=e3|op(e1,e3)=e3|op(e2,e3)=e3|op(e3,e3)=e3|op(e4,e3)=e3))&(op(e3,e0)=e4|op(e3,e1)=e4|op(e3,e2)=e4|op(e3,e3)=e4|op(e3,e4)=e4))&(op(e0,e3)=e4|op(e1,e3)=e4|op(e2,e3)=e4|op(e3,e3)=e4|op(e4,e3)=e4))&(op(e4,e0)=e0|op(e4,e1)=e0|op(e4,e2)=e0|op(e4,e3)=e0|op(e4,e4)=e0))&(op(e0,e4)=e0|op(e1,e4)=e0|op(e2,e4)=e0|op(e3,e4)=e0|op(e4,e4)=e0))&(op(e4,e0)=e1|op(e4,e1)=e1|op(e4,e2)=e1|op(e4,e3)=e1|op(e4,e4)=e1))), introduced(definition)).
fof(c_0_97, plain, (epred14_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))), introduced(definition)).
fof(c_0_98, plain, (((epred12_0&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))=>epred13_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_87]), c_0_78])).
fof(c_0_99, plain, (((((((((op(e4,e3)!=e0|(op(e4,e2)!=e0|~epred11_0)|epred12_0)&(op(e4,e3)!=e1|(op(e4,e2)!=e0|~epred11_0)|epred12_0))&(op(e4,e3)!=e2|(op(e4,e2)!=e0|~epred11_0)|epred12_0))&(op(e4,e3)!=e3|(op(e4,e2)!=e0|~epred11_0)|epred12_0))&(op(e4,e3)!=e4|(op(e4,e2)!=e0|~epred11_0)|epred12_0))&(((((op(e4,e3)!=e0|(op(e4,e2)!=e1|~epred11_0)|epred12_0)&(op(e4,e3)!=e1|(op(e4,e2)!=e1|~epred11_0)|epred12_0))&(op(e4,e3)!=e2|(op(e4,e2)!=e1|~epred11_0)|epred12_0))&(op(e4,e3)!=e3|(op(e4,e2)!=e1|~epred11_0)|epred12_0))&(op(e4,e3)!=e4|(op(e4,e2)!=e1|~epred11_0)|epred12_0)))&(((((op(e4,e3)!=e0|(op(e4,e2)!=e2|~epred11_0)|epred12_0)&(op(e4,e3)!=e1|(op(e4,e2)!=e2|~epred11_0)|epred12_0))&(op(e4,e3)!=e2|(op(e4,e2)!=e2|~epred11_0)|epred12_0))&(op(e4,e3)!=e3|(op(e4,e2)!=e2|~epred11_0)|epred12_0))&(op(e4,e3)!=e4|(op(e4,e2)!=e2|~epred11_0)|epred12_0)))&(((((op(e4,e3)!=e0|(op(e4,e2)!=e3|~epred11_0)|epred12_0)&(op(e4,e3)!=e1|(op(e4,e2)!=e3|~epred11_0)|epred12_0))&(op(e4,e3)!=e2|(op(e4,e2)!=e3|~epred11_0)|epred12_0))&(op(e4,e3)!=e3|(op(e4,e2)!=e3|~epred11_0)|epred12_0))&(op(e4,e3)!=e4|(op(e4,e2)!=e3|~epred11_0)|epred12_0)))&(((((op(e4,e3)!=e0|(op(e4,e2)!=e4|~epred11_0)|epred12_0)&(op(e4,e3)!=e1|(op(e4,e2)!=e4|~epred11_0)|epred12_0))&(op(e4,e3)!=e2|(op(e4,e2)!=e4|~epred11_0)|epred12_0))&(op(e4,e3)!=e3|(op(e4,e2)!=e4|~epred11_0)|epred12_0))&(op(e4,e3)!=e4|(op(e4,e2)!=e4|~epred11_0)|epred12_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_88])])).
cnf(c_0_100, plain, (epred11_0|op(e4,e1)!=e2|op(e4,e0)!=e3|~epred10_0), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_101, plain, (epred10_0|~epred9_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_35]), c_0_34])])).
cnf(c_0_102, plain, (epred9_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_92])])).
fof(c_0_103, plain, (((((((((op(e4,e0)!=e4|(op(e0,e4)!=e3|~epred36_0)|epred37_0)&(op(e4,e1)!=e4|(op(e0,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e2)!=e4|(op(e0,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e3)!=e4|(op(e0,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e4)!=e4|(op(e0,e4)!=e3|~epred36_0)|epred37_0))&(((((op(e4,e0)!=e4|(op(e1,e4)!=e3|~epred36_0)|epred37_0)&(op(e4,e1)!=e4|(op(e1,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e2)!=e4|(op(e1,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e3)!=e4|(op(e1,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e4)!=e4|(op(e1,e4)!=e3|~epred36_0)|epred37_0)))&(((((op(e4,e0)!=e4|(op(e2,e4)!=e3|~epred36_0)|epred37_0)&(op(e4,e1)!=e4|(op(e2,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e2)!=e4|(op(e2,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e3)!=e4|(op(e2,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e4)!=e4|(op(e2,e4)!=e3|~epred36_0)|epred37_0)))&(((((op(e4,e0)!=e4|(op(e3,e4)!=e3|~epred36_0)|epred37_0)&(op(e4,e1)!=e4|(op(e3,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e2)!=e4|(op(e3,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e3)!=e4|(op(e3,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e4)!=e4|(op(e3,e4)!=e3|~epred36_0)|epred37_0)))&(((((op(e4,e0)!=e4|(op(e4,e4)!=e3|~epred36_0)|epred37_0)&(op(e4,e1)!=e4|(op(e4,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e2)!=e4|(op(e4,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e3)!=e4|(op(e4,e4)!=e3|~epred36_0)|epred37_0))&(op(e4,e4)!=e4|(op(e4,e4)!=e3|~epred36_0)|epred37_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_93])])).
cnf(c_0_104, negated_conjecture, (op(e4,e4)!=e4|~epred37_0|op(op(e0,e0),op(e0,e0))!=e0|op(op(e1,e0),op(e0,e1))!=e0|op(op(e2,e0),op(e0,e2))!=e0|op(op(e3,e0),op(e0,e3))!=e0|op(op(e4,e0),op(e0,e4))!=e0|op(op(e0,e1),op(e1,e0))!=e1|op(op(e1,e1),op(e1,e1))!=e1|op(op(e2,e1),op(e1,e2))!=e1|op(op(e3,e1),op(e1,e3))!=e1|op(op(e4,e1),op(e1,e4))!=e1|op(op(e0,e2),op(e2,e0))!=e2|op(op(e1,e2),op(e2,e1))!=e2|op(op(e2,e2),op(e2,e2))!=e2|op(op(e3,e2),op(e2,e3))!=e2|op(op(e4,e2),op(e2,e4))!=e2|op(op(e0,e3),op(e3,e0))!=e3|op(op(e1,e3),op(e3,e1))!=e3|op(op(e2,e3),op(e3,e2))!=e3|op(op(e3,e3),op(e3,e3))!=e3|op(op(e4,e3),op(e3,e4))!=e3|op(op(e0,e4),op(e4,e0))!=e4|op(op(e1,e4),op(e4,e1))!=e4|op(op(e2,e4),op(e4,e2))!=e4|op(op(e3,e4),op(e4,e3))!=e4|op(op(e4,e4),op(e4,e4))!=e4), inference(split_conjunct,[status(thm)],[c_0_94])).
fof(c_0_105, plain, (((((((((op(e4,e0)!=e3|(op(e0,e4)!=e2|~epred35_0)|epred36_0)&(op(e4,e1)!=e3|(op(e0,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e2)!=e3|(op(e0,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e3)!=e3|(op(e0,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e4)!=e3|(op(e0,e4)!=e2|~epred35_0)|epred36_0))&(((((op(e4,e0)!=e3|(op(e1,e4)!=e2|~epred35_0)|epred36_0)&(op(e4,e1)!=e3|(op(e1,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e2)!=e3|(op(e1,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e3)!=e3|(op(e1,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e4)!=e3|(op(e1,e4)!=e2|~epred35_0)|epred36_0)))&(((((op(e4,e0)!=e3|(op(e2,e4)!=e2|~epred35_0)|epred36_0)&(op(e4,e1)!=e3|(op(e2,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e2)!=e3|(op(e2,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e3)!=e3|(op(e2,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e4)!=e3|(op(e2,e4)!=e2|~epred35_0)|epred36_0)))&(((((op(e4,e0)!=e3|(op(e3,e4)!=e2|~epred35_0)|epred36_0)&(op(e4,e1)!=e3|(op(e3,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e2)!=e3|(op(e3,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e3)!=e3|(op(e3,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e4)!=e3|(op(e3,e4)!=e2|~epred35_0)|epred36_0)))&(((((op(e4,e0)!=e3|(op(e4,e4)!=e2|~epred35_0)|epred36_0)&(op(e4,e1)!=e3|(op(e4,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e2)!=e3|(op(e4,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e3)!=e3|(op(e4,e4)!=e2|~epred35_0)|epred36_0))&(op(e4,e4)!=e3|(op(e4,e4)!=e2|~epred35_0)|epred36_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_95])])).
fof(c_0_106, plain, (((epred34_0&(op(e0,e4)=e1|op(e1,e4)=e1|op(e2,e4)=e1|op(e3,e4)=e1|op(e4,e4)=e1))&(op(e4,e0)=e2|op(e4,e1)=e2|op(e4,e2)=e2|op(e4,e3)=e2|op(e4,e4)=e2))=>epred35_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_86]), c_0_96])).
fof(c_0_107, plain, (epred33_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3|op(e4,e2)=e3))&(op(e2,e0)=e4|op(e2,e1)=e4|op(e2,e2)=e4|op(e2,e3)=e4|op(e2,e4)=e4))&(op(e0,e2)=e4|op(e1,e2)=e4|op(e2,e2)=e4|op(e3,e2)=e4|op(e4,e2)=e4))&(op(e3,e0)=e0|op(e3,e1)=e0|op(e3,e2)=e0|op(e3,e3)=e0|op(e3,e4)=e0))&(op(e0,e3)=e0|op(e1,e3)=e0|op(e2,e3)=e0|op(e3,e3)=e0|op(e4,e3)=e0))&(op(e3,e0)=e1|op(e3,e1)=e1|op(e3,e2)=e1|op(e3,e3)=e1|op(e3,e4)=e1))&(op(e0,e3)=e1|op(e1,e3)=e1|op(e2,e3)=e1|op(e3,e3)=e1|op(e4,e3)=e1))&(op(e3,e0)=e2|op(e3,e1)=e2|op(e3,e2)=e2|op(e3,e3)=e2|op(e3,e4)=e2))&(op(e0,e3)=e2|op(e1,e3)=e2|op(e2,e3)=e2|op(e3,e3)=e2|op(e4,e3)=e2))&(op(e3,e0)=e3|op(e3,e1)=e3|op(e3,e2)=e3|op(e3,e3)=e3|op(e3,e4)=e3))&(op(e0,e3)=e3|op(e1,e3)=e3|op(e2,e3)=e3|op(e3,e3)=e3|op(e4,e3)=e3))&(op(e3,e0)=e4|op(e3,e1)=e4|op(e3,e2)=e4|op(e3,e3)=e4|op(e3,e4)=e4))&(op(e0,e3)=e4|op(e1,e3)=e4|op(e2,e3)=e4|op(e3,e3)=e4|op(e4,e3)=e4))&(op(e4,e0)=e0|op(e4,e1)=e0|op(e4,e2)=e0|op(e4,e3)=e0|op(e4,e4)=e0))), introduced(definition)).
fof(c_0_108, plain, (epred15_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))), introduced(definition)).
fof(c_0_109, plain, (((epred13_0&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))=>epred14_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_97]), c_0_87])).
fof(c_0_110, plain, (((((((((op(e0,e0)!=e0|(op(e4,e4)!=e0|~epred12_0)|epred13_0)&(op(e0,e1)!=e0|(op(e4,e4)!=e0|~epred12_0)|epred13_0))&(op(e0,e2)!=e0|(op(e4,e4)!=e0|~epred12_0)|epred13_0))&(op(e0,e3)!=e0|(op(e4,e4)!=e0|~epred12_0)|epred13_0))&(op(e0,e4)!=e0|(op(e4,e4)!=e0|~epred12_0)|epred13_0))&(((((op(e0,e0)!=e0|(op(e4,e4)!=e1|~epred12_0)|epred13_0)&(op(e0,e1)!=e0|(op(e4,e4)!=e1|~epred12_0)|epred13_0))&(op(e0,e2)!=e0|(op(e4,e4)!=e1|~epred12_0)|epred13_0))&(op(e0,e3)!=e0|(op(e4,e4)!=e1|~epred12_0)|epred13_0))&(op(e0,e4)!=e0|(op(e4,e4)!=e1|~epred12_0)|epred13_0)))&(((((op(e0,e0)!=e0|(op(e4,e4)!=e2|~epred12_0)|epred13_0)&(op(e0,e1)!=e0|(op(e4,e4)!=e2|~epred12_0)|epred13_0))&(op(e0,e2)!=e0|(op(e4,e4)!=e2|~epred12_0)|epred13_0))&(op(e0,e3)!=e0|(op(e4,e4)!=e2|~epred12_0)|epred13_0))&(op(e0,e4)!=e0|(op(e4,e4)!=e2|~epred12_0)|epred13_0)))&(((((op(e0,e0)!=e0|(op(e4,e4)!=e3|~epred12_0)|epred13_0)&(op(e0,e1)!=e0|(op(e4,e4)!=e3|~epred12_0)|epred13_0))&(op(e0,e2)!=e0|(op(e4,e4)!=e3|~epred12_0)|epred13_0))&(op(e0,e3)!=e0|(op(e4,e4)!=e3|~epred12_0)|epred13_0))&(op(e0,e4)!=e0|(op(e4,e4)!=e3|~epred12_0)|epred13_0)))&(((((op(e0,e0)!=e0|(op(e4,e4)!=e4|~epred12_0)|epred13_0)&(op(e0,e1)!=e0|(op(e4,e4)!=e4|~epred12_0)|epred13_0))&(op(e0,e2)!=e0|(op(e4,e4)!=e4|~epred12_0)|epred13_0))&(op(e0,e3)!=e0|(op(e4,e4)!=e4|~epred12_0)|epred13_0))&(op(e0,e4)!=e0|(op(e4,e4)!=e4|~epred12_0)|epred13_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_98])])).
cnf(c_0_111, plain, (epred12_0|op(e4,e3)!=e0|op(e4,e2)!=e1|~epred11_0), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_112, plain, (epred11_0|~epred10_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_100, c_0_36]), c_0_38])])).
cnf(c_0_113, plain, (epred10_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_101, c_0_102])])).
cnf(c_0_114, plain, (epred37_0|op(e4,e4)!=e4|op(e2,e4)!=e3|~epred36_0), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_115, negated_conjecture, (~epred37_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_40]), c_0_18]), c_0_18]), c_0_18]), c_0_19]), c_0_21]), c_0_34]), c_0_15]), c_0_26]), c_0_38]), c_0_16]), c_0_31]), c_0_24]), c_0_20]), c_0_36]), c_0_29]), c_0_21]), c_0_19]), c_0_37]), c_0_23]), c_0_23]), c_0_23]), c_0_24]), c_0_27]), c_0_31]), c_0_25]), c_0_33]), c_0_30]), c_0_22]), c_0_38]), c_0_15]), c_0_26]), c_0_15]), c_0_22]), c_0_27]), c_0_24]), c_0_16]), c_0_28]), c_0_28]), c_0_28]), c_0_29]), c_0_32]), c_0_36]), c_0_30]), c_0_39]), c_0_33]), c_0_31]), c_0_16]), c_0_27]), c_0_33]), c_0_25]), c_0_39]), c_0_32]), c_0_29]), c_0_20]), c_0_35]), c_0_35]), c_0_35]), c_0_34]), c_0_37]), c_0_21]), c_0_36]), c_0_20]), c_0_32]), c_0_38]), c_0_22]), c_0_26]), c_0_39]), c_0_30]), c_0_25]), c_0_37]), c_0_34]), c_0_19]), c_0_40]), c_0_40]), c_0_40])])).
cnf(c_0_116, plain, (epred36_0|op(e4,e0)!=e3|op(e0,e4)!=e2|~epred35_0), inference(split_conjunct,[status(thm)],[c_0_105])).
fof(c_0_117, plain, (((((((((op(e4,e0)!=e2|(op(e0,e4)!=e1|~epred34_0)|epred35_0)&(op(e4,e1)!=e2|(op(e0,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e2)!=e2|(op(e0,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e3)!=e2|(op(e0,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e4)!=e2|(op(e0,e4)!=e1|~epred34_0)|epred35_0))&(((((op(e4,e0)!=e2|(op(e1,e4)!=e1|~epred34_0)|epred35_0)&(op(e4,e1)!=e2|(op(e1,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e2)!=e2|(op(e1,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e3)!=e2|(op(e1,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e4)!=e2|(op(e1,e4)!=e1|~epred34_0)|epred35_0)))&(((((op(e4,e0)!=e2|(op(e2,e4)!=e1|~epred34_0)|epred35_0)&(op(e4,e1)!=e2|(op(e2,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e2)!=e2|(op(e2,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e3)!=e2|(op(e2,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e4)!=e2|(op(e2,e4)!=e1|~epred34_0)|epred35_0)))&(((((op(e4,e0)!=e2|(op(e3,e4)!=e1|~epred34_0)|epred35_0)&(op(e4,e1)!=e2|(op(e3,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e2)!=e2|(op(e3,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e3)!=e2|(op(e3,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e4)!=e2|(op(e3,e4)!=e1|~epred34_0)|epred35_0)))&(((((op(e4,e0)!=e2|(op(e4,e4)!=e1|~epred34_0)|epred35_0)&(op(e4,e1)!=e2|(op(e4,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e2)!=e2|(op(e4,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e3)!=e2|(op(e4,e4)!=e1|~epred34_0)|epred35_0))&(op(e4,e4)!=e2|(op(e4,e4)!=e1|~epred34_0)|epred35_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_106])])).
fof(c_0_118, plain, (((epred33_0&(op(e0,e4)=e0|op(e1,e4)=e0|op(e2,e4)=e0|op(e3,e4)=e0|op(e4,e4)=e0))&(op(e4,e0)=e1|op(e4,e1)=e1|op(e4,e2)=e1|op(e4,e3)=e1|op(e4,e4)=e1))=>epred34_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_96]), c_0_107])).
fof(c_0_119, plain, (epred32_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3|op(e4,e2)=e3))&(op(e2,e0)=e4|op(e2,e1)=e4|op(e2,e2)=e4|op(e2,e3)=e4|op(e2,e4)=e4))&(op(e0,e2)=e4|op(e1,e2)=e4|op(e2,e2)=e4|op(e3,e2)=e4|op(e4,e2)=e4))&(op(e3,e0)=e0|op(e3,e1)=e0|op(e3,e2)=e0|op(e3,e3)=e0|op(e3,e4)=e0))&(op(e0,e3)=e0|op(e1,e3)=e0|op(e2,e3)=e0|op(e3,e3)=e0|op(e4,e3)=e0))&(op(e3,e0)=e1|op(e3,e1)=e1|op(e3,e2)=e1|op(e3,e3)=e1|op(e3,e4)=e1))&(op(e0,e3)=e1|op(e1,e3)=e1|op(e2,e3)=e1|op(e3,e3)=e1|op(e4,e3)=e1))&(op(e3,e0)=e2|op(e3,e1)=e2|op(e3,e2)=e2|op(e3,e3)=e2|op(e3,e4)=e2))&(op(e0,e3)=e2|op(e1,e3)=e2|op(e2,e3)=e2|op(e3,e3)=e2|op(e4,e3)=e2))&(op(e3,e0)=e3|op(e3,e1)=e3|op(e3,e2)=e3|op(e3,e3)=e3|op(e3,e4)=e3))&(op(e0,e3)=e3|op(e1,e3)=e3|op(e2,e3)=e3|op(e3,e3)=e3|op(e4,e3)=e3))&(op(e3,e0)=e4|op(e3,e1)=e4|op(e3,e2)=e4|op(e3,e3)=e4|op(e3,e4)=e4))), introduced(definition)).
fof(c_0_120, plain, (epred16_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))), introduced(definition)).
fof(c_0_121, plain, (((epred14_0&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))=>epred15_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_108]), c_0_97])).
fof(c_0_122, plain, (((((((((op(e0,e0)!=e1|(op(e0,e0)!=e0|~epred13_0)|epred14_0)&(op(e0,e1)!=e1|(op(e0,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e2)!=e1|(op(e0,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e3)!=e1|(op(e0,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e4)!=e1|(op(e0,e0)!=e0|~epred13_0)|epred14_0))&(((((op(e0,e0)!=e1|(op(e1,e0)!=e0|~epred13_0)|epred14_0)&(op(e0,e1)!=e1|(op(e1,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e2)!=e1|(op(e1,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e3)!=e1|(op(e1,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e4)!=e1|(op(e1,e0)!=e0|~epred13_0)|epred14_0)))&(((((op(e0,e0)!=e1|(op(e2,e0)!=e0|~epred13_0)|epred14_0)&(op(e0,e1)!=e1|(op(e2,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e2)!=e1|(op(e2,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e3)!=e1|(op(e2,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e4)!=e1|(op(e2,e0)!=e0|~epred13_0)|epred14_0)))&(((((op(e0,e0)!=e1|(op(e3,e0)!=e0|~epred13_0)|epred14_0)&(op(e0,e1)!=e1|(op(e3,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e2)!=e1|(op(e3,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e3)!=e1|(op(e3,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e4)!=e1|(op(e3,e0)!=e0|~epred13_0)|epred14_0)))&(((((op(e0,e0)!=e1|(op(e4,e0)!=e0|~epred13_0)|epred14_0)&(op(e0,e1)!=e1|(op(e4,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e2)!=e1|(op(e4,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e3)!=e1|(op(e4,e0)!=e0|~epred13_0)|epred14_0))&(op(e0,e4)!=e1|(op(e4,e0)!=e0|~epred13_0)|epred14_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_109])])).
cnf(c_0_123, plain, (epred13_0|op(e0,e0)!=e0|op(e4,e4)!=e4|~epred12_0), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_124, plain, (epred12_0|~epred11_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_111, c_0_39]), c_0_37])])).
cnf(c_0_125, plain, (epred11_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_112, c_0_113])])).
cnf(c_0_126, plain, (~epred36_0), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_114, c_0_30]), c_0_40])]), c_0_115])).
cnf(c_0_127, plain, (epred36_0|~epred35_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_116, c_0_20]), c_0_36])])).
cnf(c_0_128, plain, (epred35_0|op(e4,e1)!=e2|op(e3,e4)!=e1|~epred34_0), inference(split_conjunct,[status(thm)],[c_0_117])).
fof(c_0_129, plain, (((((((((op(e4,e0)!=e1|(op(e0,e4)!=e0|~epred33_0)|epred34_0)&(op(e4,e1)!=e1|(op(e0,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e2)!=e1|(op(e0,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e3)!=e1|(op(e0,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e4)!=e1|(op(e0,e4)!=e0|~epred33_0)|epred34_0))&(((((op(e4,e0)!=e1|(op(e1,e4)!=e0|~epred33_0)|epred34_0)&(op(e4,e1)!=e1|(op(e1,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e2)!=e1|(op(e1,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e3)!=e1|(op(e1,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e4)!=e1|(op(e1,e4)!=e0|~epred33_0)|epred34_0)))&(((((op(e4,e0)!=e1|(op(e2,e4)!=e0|~epred33_0)|epred34_0)&(op(e4,e1)!=e1|(op(e2,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e2)!=e1|(op(e2,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e3)!=e1|(op(e2,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e4)!=e1|(op(e2,e4)!=e0|~epred33_0)|epred34_0)))&(((((op(e4,e0)!=e1|(op(e3,e4)!=e0|~epred33_0)|epred34_0)&(op(e4,e1)!=e1|(op(e3,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e2)!=e1|(op(e3,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e3)!=e1|(op(e3,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e4)!=e1|(op(e3,e4)!=e0|~epred33_0)|epred34_0)))&(((((op(e4,e0)!=e1|(op(e4,e4)!=e0|~epred33_0)|epred34_0)&(op(e4,e1)!=e1|(op(e4,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e2)!=e1|(op(e4,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e3)!=e1|(op(e4,e4)!=e0|~epred33_0)|epred34_0))&(op(e4,e4)!=e1|(op(e4,e4)!=e0|~epred33_0)|epred34_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_118])])).
fof(c_0_130, plain, (((epred32_0&(op(e0,e3)=e4|op(e1,e3)=e4|op(e2,e3)=e4|op(e3,e3)=e4|op(e4,e3)=e4))&(op(e4,e0)=e0|op(e4,e1)=e0|op(e4,e2)=e0|op(e4,e3)=e0|op(e4,e4)=e0))=>epred33_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_107]), c_0_119])).
fof(c_0_131, plain, (epred31_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3|op(e4,e2)=e3))&(op(e2,e0)=e4|op(e2,e1)=e4|op(e2,e2)=e4|op(e2,e3)=e4|op(e2,e4)=e4))&(op(e0,e2)=e4|op(e1,e2)=e4|op(e2,e2)=e4|op(e3,e2)=e4|op(e4,e2)=e4))&(op(e3,e0)=e0|op(e3,e1)=e0|op(e3,e2)=e0|op(e3,e3)=e0|op(e3,e4)=e0))&(op(e0,e3)=e0|op(e1,e3)=e0|op(e2,e3)=e0|op(e3,e3)=e0|op(e4,e3)=e0))&(op(e3,e0)=e1|op(e3,e1)=e1|op(e3,e2)=e1|op(e3,e3)=e1|op(e3,e4)=e1))&(op(e0,e3)=e1|op(e1,e3)=e1|op(e2,e3)=e1|op(e3,e3)=e1|op(e4,e3)=e1))&(op(e3,e0)=e2|op(e3,e1)=e2|op(e3,e2)=e2|op(e3,e3)=e2|op(e3,e4)=e2))&(op(e0,e3)=e2|op(e1,e3)=e2|op(e2,e3)=e2|op(e3,e3)=e2|op(e4,e3)=e2))&(op(e3,e0)=e3|op(e3,e1)=e3|op(e3,e2)=e3|op(e3,e3)=e3|op(e3,e4)=e3))), introduced(definition)).
fof(c_0_132, plain, (epred17_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))), introduced(definition)).
fof(c_0_133, plain, (((epred15_0&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))=>epred16_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_120]), c_0_108])).
fof(c_0_134, plain, (((((((((op(e0,e0)!=e2|(op(e0,e0)!=e1|~epred14_0)|epred15_0)&(op(e0,e1)!=e2|(op(e0,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e2)!=e2|(op(e0,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e3)!=e2|(op(e0,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e4)!=e2|(op(e0,e0)!=e1|~epred14_0)|epred15_0))&(((((op(e0,e0)!=e2|(op(e1,e0)!=e1|~epred14_0)|epred15_0)&(op(e0,e1)!=e2|(op(e1,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e2)!=e2|(op(e1,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e3)!=e2|(op(e1,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e4)!=e2|(op(e1,e0)!=e1|~epred14_0)|epred15_0)))&(((((op(e0,e0)!=e2|(op(e2,e0)!=e1|~epred14_0)|epred15_0)&(op(e0,e1)!=e2|(op(e2,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e2)!=e2|(op(e2,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e3)!=e2|(op(e2,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e4)!=e2|(op(e2,e0)!=e1|~epred14_0)|epred15_0)))&(((((op(e0,e0)!=e2|(op(e3,e0)!=e1|~epred14_0)|epred15_0)&(op(e0,e1)!=e2|(op(e3,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e2)!=e2|(op(e3,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e3)!=e2|(op(e3,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e4)!=e2|(op(e3,e0)!=e1|~epred14_0)|epred15_0)))&(((((op(e0,e0)!=e2|(op(e4,e0)!=e1|~epred14_0)|epred15_0)&(op(e0,e1)!=e2|(op(e4,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e2)!=e2|(op(e4,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e3)!=e2|(op(e4,e0)!=e1|~epred14_0)|epred15_0))&(op(e0,e4)!=e2|(op(e4,e0)!=e1|~epred14_0)|epred15_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_121])])).
cnf(c_0_135, plain, (epred14_0|op(e0,e3)!=e1|op(e0,e0)!=e0|~epred13_0), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_136, plain, (epred13_0|~epred12_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_123, c_0_18]), c_0_40])])).
cnf(c_0_137, plain, (epred12_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_124, c_0_125])])).
cnf(c_0_138, plain, (~epred35_0), inference(pm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_139, plain, (epred35_0|~epred34_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_34]), c_0_38])])).
cnf(c_0_140, plain, (epred34_0|op(e4,e2)!=e1|op(e1,e4)!=e0|~epred33_0), inference(split_conjunct,[status(thm)],[c_0_129])).
fof(c_0_141, plain, (((((((((op(e4,e0)!=e0|(op(e0,e3)!=e4|~epred32_0)|epred33_0)&(op(e4,e1)!=e0|(op(e0,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e2)!=e0|(op(e0,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e3)!=e0|(op(e0,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e4)!=e0|(op(e0,e3)!=e4|~epred32_0)|epred33_0))&(((((op(e4,e0)!=e0|(op(e1,e3)!=e4|~epred32_0)|epred33_0)&(op(e4,e1)!=e0|(op(e1,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e2)!=e0|(op(e1,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e3)!=e0|(op(e1,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e4)!=e0|(op(e1,e3)!=e4|~epred32_0)|epred33_0)))&(((((op(e4,e0)!=e0|(op(e2,e3)!=e4|~epred32_0)|epred33_0)&(op(e4,e1)!=e0|(op(e2,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e2)!=e0|(op(e2,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e3)!=e0|(op(e2,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e4)!=e0|(op(e2,e3)!=e4|~epred32_0)|epred33_0)))&(((((op(e4,e0)!=e0|(op(e3,e3)!=e4|~epred32_0)|epred33_0)&(op(e4,e1)!=e0|(op(e3,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e2)!=e0|(op(e3,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e3)!=e0|(op(e3,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e4)!=e0|(op(e3,e3)!=e4|~epred32_0)|epred33_0)))&(((((op(e4,e0)!=e0|(op(e4,e3)!=e4|~epred32_0)|epred33_0)&(op(e4,e1)!=e0|(op(e4,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e2)!=e0|(op(e4,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e3)!=e0|(op(e4,e3)!=e4|~epred32_0)|epred33_0))&(op(e4,e4)!=e0|(op(e4,e3)!=e4|~epred32_0)|epred33_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_130])])).
fof(c_0_142, plain, (((epred31_0&(op(e0,e3)=e3|op(e1,e3)=e3|op(e2,e3)=e3|op(e3,e3)=e3|op(e4,e3)=e3))&(op(e3,e0)=e4|op(e3,e1)=e4|op(e3,e2)=e4|op(e3,e3)=e4|op(e3,e4)=e4))=>epred32_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_119]), c_0_131])).
fof(c_0_143, plain, (epred30_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3|op(e4,e2)=e3))&(op(e2,e0)=e4|op(e2,e1)=e4|op(e2,e2)=e4|op(e2,e3)=e4|op(e2,e4)=e4))&(op(e0,e2)=e4|op(e1,e2)=e4|op(e2,e2)=e4|op(e3,e2)=e4|op(e4,e2)=e4))&(op(e3,e0)=e0|op(e3,e1)=e0|op(e3,e2)=e0|op(e3,e3)=e0|op(e3,e4)=e0))&(op(e0,e3)=e0|op(e1,e3)=e0|op(e2,e3)=e0|op(e3,e3)=e0|op(e4,e3)=e0))&(op(e3,e0)=e1|op(e3,e1)=e1|op(e3,e2)=e1|op(e3,e3)=e1|op(e3,e4)=e1))&(op(e0,e3)=e1|op(e1,e3)=e1|op(e2,e3)=e1|op(e3,e3)=e1|op(e4,e3)=e1))&(op(e3,e0)=e2|op(e3,e1)=e2|op(e3,e2)=e2|op(e3,e3)=e2|op(e3,e4)=e2))), introduced(definition)).
fof(c_0_144, plain, (epred18_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))), introduced(definition)).
fof(c_0_145, plain, (((epred16_0&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))=>epred17_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_132]), c_0_120])).
fof(c_0_146, plain, (((((((((op(e0,e0)!=e3|(op(e0,e0)!=e2|~epred15_0)|epred16_0)&(op(e0,e1)!=e3|(op(e0,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e2)!=e3|(op(e0,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e3)!=e3|(op(e0,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e4)!=e3|(op(e0,e0)!=e2|~epred15_0)|epred16_0))&(((((op(e0,e0)!=e3|(op(e1,e0)!=e2|~epred15_0)|epred16_0)&(op(e0,e1)!=e3|(op(e1,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e2)!=e3|(op(e1,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e3)!=e3|(op(e1,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e4)!=e3|(op(e1,e0)!=e2|~epred15_0)|epred16_0)))&(((((op(e0,e0)!=e3|(op(e2,e0)!=e2|~epred15_0)|epred16_0)&(op(e0,e1)!=e3|(op(e2,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e2)!=e3|(op(e2,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e3)!=e3|(op(e2,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e4)!=e3|(op(e2,e0)!=e2|~epred15_0)|epred16_0)))&(((((op(e0,e0)!=e3|(op(e3,e0)!=e2|~epred15_0)|epred16_0)&(op(e0,e1)!=e3|(op(e3,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e2)!=e3|(op(e3,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e3)!=e3|(op(e3,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e4)!=e3|(op(e3,e0)!=e2|~epred15_0)|epred16_0)))&(((((op(e0,e0)!=e3|(op(e4,e0)!=e2|~epred15_0)|epred16_0)&(op(e0,e1)!=e3|(op(e4,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e2)!=e3|(op(e4,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e3)!=e3|(op(e4,e0)!=e2|~epred15_0)|epred16_0))&(op(e0,e4)!=e3|(op(e4,e0)!=e2|~epred15_0)|epred16_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_133])])).
cnf(c_0_147, plain, (epred15_0|op(e0,e4)!=e2|op(e2,e0)!=e1|~epred14_0), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_148, plain, (epred14_0|~epred13_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_135, c_0_18]), c_0_16])])).
cnf(c_0_149, plain, (epred13_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_136, c_0_137])])).
cnf(c_0_150, plain, (~epred34_0), inference(pm,[status(thm)],[c_0_138, c_0_139])).
cnf(c_0_151, plain, (epred34_0|~epred33_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_140, c_0_22]), c_0_39])])).
cnf(c_0_152, plain, (epred33_0|op(e4,e3)!=e0|op(e2,e3)!=e4|~epred32_0), inference(split_conjunct,[status(thm)],[c_0_141])).
fof(c_0_153, plain, (((((((((op(e3,e0)!=e4|(op(e0,e3)!=e3|~epred31_0)|epred32_0)&(op(e3,e1)!=e4|(op(e0,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e2)!=e4|(op(e0,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e3)!=e4|(op(e0,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e4)!=e4|(op(e0,e3)!=e3|~epred31_0)|epred32_0))&(((((op(e3,e0)!=e4|(op(e1,e3)!=e3|~epred31_0)|epred32_0)&(op(e3,e1)!=e4|(op(e1,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e2)!=e4|(op(e1,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e3)!=e4|(op(e1,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e4)!=e4|(op(e1,e3)!=e3|~epred31_0)|epred32_0)))&(((((op(e3,e0)!=e4|(op(e2,e3)!=e3|~epred31_0)|epred32_0)&(op(e3,e1)!=e4|(op(e2,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e2)!=e4|(op(e2,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e3)!=e4|(op(e2,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e4)!=e4|(op(e2,e3)!=e3|~epred31_0)|epred32_0)))&(((((op(e3,e0)!=e4|(op(e3,e3)!=e3|~epred31_0)|epred32_0)&(op(e3,e1)!=e4|(op(e3,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e2)!=e4|(op(e3,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e3)!=e4|(op(e3,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e4)!=e4|(op(e3,e3)!=e3|~epred31_0)|epred32_0)))&(((((op(e3,e0)!=e4|(op(e4,e3)!=e3|~epred31_0)|epred32_0)&(op(e3,e1)!=e4|(op(e4,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e2)!=e4|(op(e4,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e3)!=e4|(op(e4,e3)!=e3|~epred31_0)|epred32_0))&(op(e3,e4)!=e4|(op(e4,e3)!=e3|~epred31_0)|epred32_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_142])])).
fof(c_0_154, plain, (((epred30_0&(op(e0,e3)=e2|op(e1,e3)=e2|op(e2,e3)=e2|op(e3,e3)=e2|op(e4,e3)=e2))&(op(e3,e0)=e3|op(e3,e1)=e3|op(e3,e2)=e3|op(e3,e3)=e3|op(e3,e4)=e3))=>epred31_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_131]), c_0_143])).
fof(c_0_155, plain, (epred29_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3|op(e4,e2)=e3))&(op(e2,e0)=e4|op(e2,e1)=e4|op(e2,e2)=e4|op(e2,e3)=e4|op(e2,e4)=e4))&(op(e0,e2)=e4|op(e1,e2)=e4|op(e2,e2)=e4|op(e3,e2)=e4|op(e4,e2)=e4))&(op(e3,e0)=e0|op(e3,e1)=e0|op(e3,e2)=e0|op(e3,e3)=e0|op(e3,e4)=e0))&(op(e0,e3)=e0|op(e1,e3)=e0|op(e2,e3)=e0|op(e3,e3)=e0|op(e4,e3)=e0))&(op(e3,e0)=e1|op(e3,e1)=e1|op(e3,e2)=e1|op(e3,e3)=e1|op(e3,e4)=e1))), introduced(definition)).
fof(c_0_156, plain, (epred19_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))), introduced(definition)).
fof(c_0_157, plain, (((epred17_0&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))=>epred18_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_144]), c_0_132])).
fof(c_0_158, plain, (((((((((op(e0,e0)!=e4|(op(e0,e0)!=e3|~epred16_0)|epred17_0)&(op(e0,e1)!=e4|(op(e0,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e2)!=e4|(op(e0,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e3)!=e4|(op(e0,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e4)!=e4|(op(e0,e0)!=e3|~epred16_0)|epred17_0))&(((((op(e0,e0)!=e4|(op(e1,e0)!=e3|~epred16_0)|epred17_0)&(op(e0,e1)!=e4|(op(e1,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e2)!=e4|(op(e1,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e3)!=e4|(op(e1,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e4)!=e4|(op(e1,e0)!=e3|~epred16_0)|epred17_0)))&(((((op(e0,e0)!=e4|(op(e2,e0)!=e3|~epred16_0)|epred17_0)&(op(e0,e1)!=e4|(op(e2,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e2)!=e4|(op(e2,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e3)!=e4|(op(e2,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e4)!=e4|(op(e2,e0)!=e3|~epred16_0)|epred17_0)))&(((((op(e0,e0)!=e4|(op(e3,e0)!=e3|~epred16_0)|epred17_0)&(op(e0,e1)!=e4|(op(e3,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e2)!=e4|(op(e3,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e3)!=e4|(op(e3,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e4)!=e4|(op(e3,e0)!=e3|~epred16_0)|epred17_0)))&(((((op(e0,e0)!=e4|(op(e4,e0)!=e3|~epred16_0)|epred17_0)&(op(e0,e1)!=e4|(op(e4,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e2)!=e4|(op(e4,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e3)!=e4|(op(e4,e0)!=e3|~epred16_0)|epred17_0))&(op(e0,e4)!=e4|(op(e4,e0)!=e3|~epred16_0)|epred17_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_145])])).
cnf(c_0_159, plain, (epred16_0|op(e0,e1)!=e3|op(e3,e0)!=e2|~epred15_0), inference(split_conjunct,[status(thm)],[c_0_146])).
cnf(c_0_160, plain, (epred15_0|~epred14_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_147, c_0_20]), c_0_26])])).
cnf(c_0_161, plain, (epred14_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_148, c_0_149])])).
cnf(c_0_162, plain, (~epred33_0), inference(pm,[status(thm)],[c_0_150, c_0_151])).
cnf(c_0_163, plain, (epred33_0|~epred32_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_152, c_0_29]), c_0_37])])).
cnf(c_0_164, plain, (epred32_0|op(e3,e1)!=e4|op(e3,e3)!=e3|~epred31_0), inference(split_conjunct,[status(thm)],[c_0_153])).
fof(c_0_165, plain, (((((((((op(e3,e0)!=e3|(op(e0,e3)!=e2|~epred30_0)|epred31_0)&(op(e3,e1)!=e3|(op(e0,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e2)!=e3|(op(e0,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e3)!=e3|(op(e0,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e4)!=e3|(op(e0,e3)!=e2|~epred30_0)|epred31_0))&(((((op(e3,e0)!=e3|(op(e1,e3)!=e2|~epred30_0)|epred31_0)&(op(e3,e1)!=e3|(op(e1,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e2)!=e3|(op(e1,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e3)!=e3|(op(e1,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e4)!=e3|(op(e1,e3)!=e2|~epred30_0)|epred31_0)))&(((((op(e3,e0)!=e3|(op(e2,e3)!=e2|~epred30_0)|epred31_0)&(op(e3,e1)!=e3|(op(e2,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e2)!=e3|(op(e2,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e3)!=e3|(op(e2,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e4)!=e3|(op(e2,e3)!=e2|~epred30_0)|epred31_0)))&(((((op(e3,e0)!=e3|(op(e3,e3)!=e2|~epred30_0)|epred31_0)&(op(e3,e1)!=e3|(op(e3,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e2)!=e3|(op(e3,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e3)!=e3|(op(e3,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e4)!=e3|(op(e3,e3)!=e2|~epred30_0)|epred31_0)))&(((((op(e3,e0)!=e3|(op(e4,e3)!=e2|~epred30_0)|epred31_0)&(op(e3,e1)!=e3|(op(e4,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e2)!=e3|(op(e4,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e3)!=e3|(op(e4,e3)!=e2|~epred30_0)|epred31_0))&(op(e3,e4)!=e3|(op(e4,e3)!=e2|~epred30_0)|epred31_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_154])])).
fof(c_0_166, plain, (((epred29_0&(op(e0,e3)=e1|op(e1,e3)=e1|op(e2,e3)=e1|op(e3,e3)=e1|op(e4,e3)=e1))&(op(e3,e0)=e2|op(e3,e1)=e2|op(e3,e2)=e2|op(e3,e3)=e2|op(e3,e4)=e2))=>epred30_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_143]), c_0_155])).
fof(c_0_167, plain, (epred28_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3|op(e4,e2)=e3))&(op(e2,e0)=e4|op(e2,e1)=e4|op(e2,e2)=e4|op(e2,e3)=e4|op(e2,e4)=e4))&(op(e0,e2)=e4|op(e1,e2)=e4|op(e2,e2)=e4|op(e3,e2)=e4|op(e4,e2)=e4))&(op(e3,e0)=e0|op(e3,e1)=e0|op(e3,e2)=e0|op(e3,e3)=e0|op(e3,e4)=e0))), introduced(definition)).
fof(c_0_168, plain, (epred20_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))), introduced(definition)).
fof(c_0_169, plain, (((epred18_0&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))=>epred19_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_156]), c_0_144])).
fof(c_0_170, plain, (((((((((op(e1,e0)!=e0|(op(e0,e0)!=e4|~epred17_0)|epred18_0)&(op(e1,e1)!=e0|(op(e0,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e2)!=e0|(op(e0,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e3)!=e0|(op(e0,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e4)!=e0|(op(e0,e0)!=e4|~epred17_0)|epred18_0))&(((((op(e1,e0)!=e0|(op(e1,e0)!=e4|~epred17_0)|epred18_0)&(op(e1,e1)!=e0|(op(e1,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e2)!=e0|(op(e1,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e3)!=e0|(op(e1,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e4)!=e0|(op(e1,e0)!=e4|~epred17_0)|epred18_0)))&(((((op(e1,e0)!=e0|(op(e2,e0)!=e4|~epred17_0)|epred18_0)&(op(e1,e1)!=e0|(op(e2,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e2)!=e0|(op(e2,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e3)!=e0|(op(e2,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e4)!=e0|(op(e2,e0)!=e4|~epred17_0)|epred18_0)))&(((((op(e1,e0)!=e0|(op(e3,e0)!=e4|~epred17_0)|epred18_0)&(op(e1,e1)!=e0|(op(e3,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e2)!=e0|(op(e3,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e3)!=e0|(op(e3,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e4)!=e0|(op(e3,e0)!=e4|~epred17_0)|epred18_0)))&(((((op(e1,e0)!=e0|(op(e4,e0)!=e4|~epred17_0)|epred18_0)&(op(e1,e1)!=e0|(op(e4,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e2)!=e0|(op(e4,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e3)!=e0|(op(e4,e0)!=e4|~epred17_0)|epred18_0))&(op(e1,e4)!=e0|(op(e4,e0)!=e4|~epred17_0)|epred18_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_157])])).
cnf(c_0_171, plain, (epred17_0|op(e0,e2)!=e4|op(e4,e0)!=e3|~epred16_0), inference(split_conjunct,[status(thm)],[c_0_158])).
cnf(c_0_172, plain, (epred16_0|~epred15_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_159, c_0_19]), c_0_31])])).
cnf(c_0_173, plain, (epred15_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_160, c_0_161])])).
cnf(c_0_174, plain, (~epred32_0), inference(pm,[status(thm)],[c_0_162, c_0_163])).
cnf(c_0_175, plain, (epred32_0|~epred31_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_164, c_0_33]), c_0_35])])).
cnf(c_0_176, plain, (epred31_0|op(e3,e3)!=e3|op(e1,e3)!=e2|~epred30_0), inference(split_conjunct,[status(thm)],[c_0_165])).
fof(c_0_177, plain, (((((((((op(e3,e0)!=e2|(op(e0,e3)!=e1|~epred29_0)|epred30_0)&(op(e3,e1)!=e2|(op(e0,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e2)!=e2|(op(e0,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e3)!=e2|(op(e0,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e4)!=e2|(op(e0,e3)!=e1|~epred29_0)|epred30_0))&(((((op(e3,e0)!=e2|(op(e1,e3)!=e1|~epred29_0)|epred30_0)&(op(e3,e1)!=e2|(op(e1,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e2)!=e2|(op(e1,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e3)!=e2|(op(e1,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e4)!=e2|(op(e1,e3)!=e1|~epred29_0)|epred30_0)))&(((((op(e3,e0)!=e2|(op(e2,e3)!=e1|~epred29_0)|epred30_0)&(op(e3,e1)!=e2|(op(e2,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e2)!=e2|(op(e2,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e3)!=e2|(op(e2,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e4)!=e2|(op(e2,e3)!=e1|~epred29_0)|epred30_0)))&(((((op(e3,e0)!=e2|(op(e3,e3)!=e1|~epred29_0)|epred30_0)&(op(e3,e1)!=e2|(op(e3,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e2)!=e2|(op(e3,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e3)!=e2|(op(e3,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e4)!=e2|(op(e3,e3)!=e1|~epred29_0)|epred30_0)))&(((((op(e3,e0)!=e2|(op(e4,e3)!=e1|~epred29_0)|epred30_0)&(op(e3,e1)!=e2|(op(e4,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e2)!=e2|(op(e4,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e3)!=e2|(op(e4,e3)!=e1|~epred29_0)|epred30_0))&(op(e3,e4)!=e2|(op(e4,e3)!=e1|~epred29_0)|epred30_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_166])])).
fof(c_0_178, plain, (((epred28_0&(op(e0,e3)=e0|op(e1,e3)=e0|op(e2,e3)=e0|op(e3,e3)=e0|op(e4,e3)=e0))&(op(e3,e0)=e1|op(e3,e1)=e1|op(e3,e2)=e1|op(e3,e3)=e1|op(e3,e4)=e1))=>epred29_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_155]), c_0_167])).
fof(c_0_179, plain, (epred27_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3|op(e4,e2)=e3))&(op(e2,e0)=e4|op(e2,e1)=e4|op(e2,e2)=e4|op(e2,e3)=e4|op(e2,e4)=e4))), introduced(definition)).
fof(c_0_180, plain, (epred21_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))), introduced(definition)).
fof(c_0_181, plain, (((epred19_0&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))=>epred20_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_168]), c_0_156])).
fof(c_0_182, plain, (((((((((op(e1,e0)!=e1|(op(e0,e1)!=e0|~epred18_0)|epred19_0)&(op(e1,e1)!=e1|(op(e0,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e2)!=e1|(op(e0,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e3)!=e1|(op(e0,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e4)!=e1|(op(e0,e1)!=e0|~epred18_0)|epred19_0))&(((((op(e1,e0)!=e1|(op(e1,e1)!=e0|~epred18_0)|epred19_0)&(op(e1,e1)!=e1|(op(e1,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e2)!=e1|(op(e1,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e3)!=e1|(op(e1,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e4)!=e1|(op(e1,e1)!=e0|~epred18_0)|epred19_0)))&(((((op(e1,e0)!=e1|(op(e2,e1)!=e0|~epred18_0)|epred19_0)&(op(e1,e1)!=e1|(op(e2,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e2)!=e1|(op(e2,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e3)!=e1|(op(e2,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e4)!=e1|(op(e2,e1)!=e0|~epred18_0)|epred19_0)))&(((((op(e1,e0)!=e1|(op(e3,e1)!=e0|~epred18_0)|epred19_0)&(op(e1,e1)!=e1|(op(e3,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e2)!=e1|(op(e3,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e3)!=e1|(op(e3,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e4)!=e1|(op(e3,e1)!=e0|~epred18_0)|epred19_0)))&(((((op(e1,e0)!=e1|(op(e4,e1)!=e0|~epred18_0)|epred19_0)&(op(e1,e1)!=e1|(op(e4,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e2)!=e1|(op(e4,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e3)!=e1|(op(e4,e1)!=e0|~epred18_0)|epred19_0))&(op(e1,e4)!=e1|(op(e4,e1)!=e0|~epred18_0)|epred19_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_169])])).
cnf(c_0_183, plain, (epred18_0|op(e1,e4)!=e0|op(e1,e0)!=e4|~epred17_0), inference(split_conjunct,[status(thm)],[c_0_170])).
cnf(c_0_184, plain, (epred17_0|~epred16_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_171, c_0_15]), c_0_36])])).
cnf(c_0_185, plain, (epred16_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_172, c_0_173])])).
cnf(c_0_186, plain, (~epred31_0), inference(pm,[status(thm)],[c_0_174, c_0_175])).
cnf(c_0_187, plain, (epred31_0|~epred30_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_176, c_0_25]), c_0_35])])).
cnf(c_0_188, plain, (epred30_0|op(e3,e0)!=e2|op(e0,e3)!=e1|~epred29_0), inference(split_conjunct,[status(thm)],[c_0_177])).
fof(c_0_189, plain, (((((((((op(e3,e0)!=e1|(op(e0,e3)!=e0|~epred28_0)|epred29_0)&(op(e3,e1)!=e1|(op(e0,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e2)!=e1|(op(e0,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e3)!=e1|(op(e0,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e4)!=e1|(op(e0,e3)!=e0|~epred28_0)|epred29_0))&(((((op(e3,e0)!=e1|(op(e1,e3)!=e0|~epred28_0)|epred29_0)&(op(e3,e1)!=e1|(op(e1,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e2)!=e1|(op(e1,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e3)!=e1|(op(e1,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e4)!=e1|(op(e1,e3)!=e0|~epred28_0)|epred29_0)))&(((((op(e3,e0)!=e1|(op(e2,e3)!=e0|~epred28_0)|epred29_0)&(op(e3,e1)!=e1|(op(e2,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e2)!=e1|(op(e2,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e3)!=e1|(op(e2,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e4)!=e1|(op(e2,e3)!=e0|~epred28_0)|epred29_0)))&(((((op(e3,e0)!=e1|(op(e3,e3)!=e0|~epred28_0)|epred29_0)&(op(e3,e1)!=e1|(op(e3,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e2)!=e1|(op(e3,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e3)!=e1|(op(e3,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e4)!=e1|(op(e3,e3)!=e0|~epred28_0)|epred29_0)))&(((((op(e3,e0)!=e1|(op(e4,e3)!=e0|~epred28_0)|epred29_0)&(op(e3,e1)!=e1|(op(e4,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e2)!=e1|(op(e4,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e3)!=e1|(op(e4,e3)!=e0|~epred28_0)|epred29_0))&(op(e3,e4)!=e1|(op(e4,e3)!=e0|~epred28_0)|epred29_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_178])])).
fof(c_0_190, plain, (((epred27_0&(op(e0,e2)=e4|op(e1,e2)=e4|op(e2,e2)=e4|op(e3,e2)=e4|op(e4,e2)=e4))&(op(e3,e0)=e0|op(e3,e1)=e0|op(e3,e2)=e0|op(e3,e3)=e0|op(e3,e4)=e0))=>epred28_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_167]), c_0_179])).
fof(c_0_191, plain, (epred26_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))), introduced(definition)).
fof(c_0_192, plain, (epred22_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))), introduced(definition)).
fof(c_0_193, plain, (((epred20_0&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))=>epred21_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_180]), c_0_168])).
fof(c_0_194, plain, (((((((((op(e1,e0)!=e2|(op(e0,e1)!=e1|~epred19_0)|epred20_0)&(op(e1,e1)!=e2|(op(e0,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e2)!=e2|(op(e0,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e3)!=e2|(op(e0,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e4)!=e2|(op(e0,e1)!=e1|~epred19_0)|epred20_0))&(((((op(e1,e0)!=e2|(op(e1,e1)!=e1|~epred19_0)|epred20_0)&(op(e1,e1)!=e2|(op(e1,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e2)!=e2|(op(e1,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e3)!=e2|(op(e1,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e4)!=e2|(op(e1,e1)!=e1|~epred19_0)|epred20_0)))&(((((op(e1,e0)!=e2|(op(e2,e1)!=e1|~epred19_0)|epred20_0)&(op(e1,e1)!=e2|(op(e2,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e2)!=e2|(op(e2,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e3)!=e2|(op(e2,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e4)!=e2|(op(e2,e1)!=e1|~epred19_0)|epred20_0)))&(((((op(e1,e0)!=e2|(op(e3,e1)!=e1|~epred19_0)|epred20_0)&(op(e1,e1)!=e2|(op(e3,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e2)!=e2|(op(e3,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e3)!=e2|(op(e3,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e4)!=e2|(op(e3,e1)!=e1|~epred19_0)|epred20_0)))&(((((op(e1,e0)!=e2|(op(e4,e1)!=e1|~epred19_0)|epred20_0)&(op(e1,e1)!=e2|(op(e4,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e2)!=e2|(op(e4,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e3)!=e2|(op(e4,e1)!=e1|~epred19_0)|epred20_0))&(op(e1,e4)!=e2|(op(e4,e1)!=e1|~epred19_0)|epred20_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_181])])).
cnf(c_0_195, plain, (epred19_0|op(e1,e1)!=e1|op(e2,e1)!=e0|~epred18_0), inference(split_conjunct,[status(thm)],[c_0_182])).
cnf(c_0_196, plain, (epred18_0|~epred17_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_183, c_0_21]), c_0_22])])).
cnf(c_0_197, plain, (epred17_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_184, c_0_185])])).
cnf(c_0_198, plain, (~epred30_0), inference(pm,[status(thm)],[c_0_186, c_0_187])).
cnf(c_0_199, plain, (epred30_0|~epred29_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_188, c_0_16]), c_0_31])])).
cnf(c_0_200, plain, (epred29_0|op(e3,e4)!=e1|op(e4,e3)!=e0|~epred28_0), inference(split_conjunct,[status(thm)],[c_0_189])).
fof(c_0_201, plain, (((((((((op(e3,e0)!=e0|(op(e0,e2)!=e4|~epred27_0)|epred28_0)&(op(e3,e1)!=e0|(op(e0,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e2)!=e0|(op(e0,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e3)!=e0|(op(e0,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e4)!=e0|(op(e0,e2)!=e4|~epred27_0)|epred28_0))&(((((op(e3,e0)!=e0|(op(e1,e2)!=e4|~epred27_0)|epred28_0)&(op(e3,e1)!=e0|(op(e1,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e2)!=e0|(op(e1,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e3)!=e0|(op(e1,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e4)!=e0|(op(e1,e2)!=e4|~epred27_0)|epred28_0)))&(((((op(e3,e0)!=e0|(op(e2,e2)!=e4|~epred27_0)|epred28_0)&(op(e3,e1)!=e0|(op(e2,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e2)!=e0|(op(e2,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e3)!=e0|(op(e2,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e4)!=e0|(op(e2,e2)!=e4|~epred27_0)|epred28_0)))&(((((op(e3,e0)!=e0|(op(e3,e2)!=e4|~epred27_0)|epred28_0)&(op(e3,e1)!=e0|(op(e3,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e2)!=e0|(op(e3,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e3)!=e0|(op(e3,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e4)!=e0|(op(e3,e2)!=e4|~epred27_0)|epred28_0)))&(((((op(e3,e0)!=e0|(op(e4,e2)!=e4|~epred27_0)|epred28_0)&(op(e3,e1)!=e0|(op(e4,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e2)!=e0|(op(e4,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e3)!=e0|(op(e4,e2)!=e4|~epred27_0)|epred28_0))&(op(e3,e4)!=e0|(op(e4,e2)!=e4|~epred27_0)|epred28_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_190])])).
fof(c_0_202, plain, (((epred26_0&(op(e0,e2)=e3|op(e1,e2)=e3|op(e2,e2)=e3|op(e3,e2)=e3|op(e4,e2)=e3))&(op(e2,e0)=e4|op(e2,e1)=e4|op(e2,e2)=e4|op(e2,e3)=e4|op(e2,e4)=e4))=>epred27_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_179]), c_0_191])).
fof(c_0_203, plain, (epred25_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))), introduced(definition)).
fof(c_0_204, plain, (epred23_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))), introduced(definition)).
fof(c_0_205, plain, (((epred21_0&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))=>epred22_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_192]), c_0_180])).
fof(c_0_206, plain, (((((((((op(e1,e0)!=e3|(op(e0,e1)!=e2|~epred20_0)|epred21_0)&(op(e1,e1)!=e3|(op(e0,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e2)!=e3|(op(e0,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e3)!=e3|(op(e0,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e4)!=e3|(op(e0,e1)!=e2|~epred20_0)|epred21_0))&(((((op(e1,e0)!=e3|(op(e1,e1)!=e2|~epred20_0)|epred21_0)&(op(e1,e1)!=e3|(op(e1,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e2)!=e3|(op(e1,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e3)!=e3|(op(e1,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e4)!=e3|(op(e1,e1)!=e2|~epred20_0)|epred21_0)))&(((((op(e1,e0)!=e3|(op(e2,e1)!=e2|~epred20_0)|epred21_0)&(op(e1,e1)!=e3|(op(e2,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e2)!=e3|(op(e2,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e3)!=e3|(op(e2,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e4)!=e3|(op(e2,e1)!=e2|~epred20_0)|epred21_0)))&(((((op(e1,e0)!=e3|(op(e3,e1)!=e2|~epred20_0)|epred21_0)&(op(e1,e1)!=e3|(op(e3,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e2)!=e3|(op(e3,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e3)!=e3|(op(e3,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e4)!=e3|(op(e3,e1)!=e2|~epred20_0)|epred21_0)))&(((((op(e1,e0)!=e3|(op(e4,e1)!=e2|~epred20_0)|epred21_0)&(op(e1,e1)!=e3|(op(e4,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e2)!=e3|(op(e4,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e3)!=e3|(op(e4,e1)!=e2|~epred20_0)|epred21_0))&(op(e1,e4)!=e3|(op(e4,e1)!=e2|~epred20_0)|epred21_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_193])])).
cnf(c_0_207, plain, (epred20_0|op(e1,e3)!=e2|op(e1,e1)!=e1|~epred19_0), inference(split_conjunct,[status(thm)],[c_0_194])).
cnf(c_0_208, plain, (epred19_0|~epred18_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_195, c_0_23]), c_0_27])])).
cnf(c_0_209, plain, (epred18_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_196, c_0_197])])).
cnf(c_0_210, plain, (~epred29_0), inference(pm,[status(thm)],[c_0_198, c_0_199])).
cnf(c_0_211, plain, (epred29_0|~epred28_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_200, c_0_34]), c_0_37])])).
cnf(c_0_212, plain, (epred28_0|op(e3,e2)!=e0|op(e0,e2)!=e4|~epred27_0), inference(split_conjunct,[status(thm)],[c_0_201])).
fof(c_0_213, plain, (((((((((op(e2,e0)!=e4|(op(e0,e2)!=e3|~epred26_0)|epred27_0)&(op(e2,e1)!=e4|(op(e0,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e2)!=e4|(op(e0,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e3)!=e4|(op(e0,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e4)!=e4|(op(e0,e2)!=e3|~epred26_0)|epred27_0))&(((((op(e2,e0)!=e4|(op(e1,e2)!=e3|~epred26_0)|epred27_0)&(op(e2,e1)!=e4|(op(e1,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e2)!=e4|(op(e1,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e3)!=e4|(op(e1,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e4)!=e4|(op(e1,e2)!=e3|~epred26_0)|epred27_0)))&(((((op(e2,e0)!=e4|(op(e2,e2)!=e3|~epred26_0)|epred27_0)&(op(e2,e1)!=e4|(op(e2,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e2)!=e4|(op(e2,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e3)!=e4|(op(e2,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e4)!=e4|(op(e2,e2)!=e3|~epred26_0)|epred27_0)))&(((((op(e2,e0)!=e4|(op(e3,e2)!=e3|~epred26_0)|epred27_0)&(op(e2,e1)!=e4|(op(e3,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e2)!=e4|(op(e3,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e3)!=e4|(op(e3,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e4)!=e4|(op(e3,e2)!=e3|~epred26_0)|epred27_0)))&(((((op(e2,e0)!=e4|(op(e4,e2)!=e3|~epred26_0)|epred27_0)&(op(e2,e1)!=e4|(op(e4,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e2)!=e4|(op(e4,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e3)!=e4|(op(e4,e2)!=e3|~epred26_0)|epred27_0))&(op(e2,e4)!=e4|(op(e4,e2)!=e3|~epred26_0)|epred27_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_202])])).
fof(c_0_214, plain, (((epred25_0&(op(e0,e2)=e2|op(e1,e2)=e2|op(e2,e2)=e2|op(e3,e2)=e2|op(e4,e2)=e2))&(op(e2,e0)=e3|op(e2,e1)=e3|op(e2,e2)=e3|op(e2,e3)=e3|op(e2,e4)=e3))=>epred26_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_191]), c_0_203])).
fof(c_0_215, plain, (epred24_0<=>((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((op(e0,op(e0,e0))=e0&op(e0,op(e0,e1))=e1)&op(e0,op(e0,e2))=e2)&op(e0,op(e0,e3))=e3)&op(e0,op(e0,e4))=e4)&op(e1,op(e1,e0))=e0)&op(e1,op(e1,e1))=e1)&op(e1,op(e1,e2))=e2)&op(e1,op(e1,e3))=e3)&op(e1,op(e1,e4))=e4)&op(e2,op(e2,e0))=e0)&op(e2,op(e2,e1))=e1)&op(e2,op(e2,e2))=e2)&op(e2,op(e2,e3))=e3)&op(e2,op(e2,e4))=e4)&op(e3,op(e3,e0))=e0)&op(e3,op(e3,e1))=e1)&op(e3,op(e3,e2))=e2)&op(e3,op(e3,e3))=e3)&op(e3,op(e3,e4))=e4)&op(e4,op(e4,e0))=e0)&op(e4,op(e4,e1))=e1)&op(e4,op(e4,e2))=e2)&op(e4,op(e4,e3))=e3)&op(e4,op(e4,e4))=e4)&(op(e0,e0)=e0|op(e0,e0)=e1|op(e0,e0)=e2|op(e0,e0)=e3|op(e0,e0)=e4))&(op(e0,e1)=e0|op(e0,e1)=e1|op(e0,e1)=e2|op(e0,e1)=e3|op(e0,e1)=e4))&(op(e0,e2)=e0|op(e0,e2)=e1|op(e0,e2)=e2|op(e0,e2)=e3|op(e0,e2)=e4))&(op(e0,e3)=e0|op(e0,e3)=e1|op(e0,e3)=e2|op(e0,e3)=e3|op(e0,e3)=e4))&(op(e0,e4)=e0|op(e0,e4)=e1|op(e0,e4)=e2|op(e0,e4)=e3|op(e0,e4)=e4))&(op(e1,e0)=e0|op(e1,e0)=e1|op(e1,e0)=e2|op(e1,e0)=e3|op(e1,e0)=e4))&(op(e1,e1)=e0|op(e1,e1)=e1|op(e1,e1)=e2|op(e1,e1)=e3|op(e1,e1)=e4))&(op(e1,e2)=e0|op(e1,e2)=e1|op(e1,e2)=e2|op(e1,e2)=e3|op(e1,e2)=e4))&(op(e1,e3)=e0|op(e1,e3)=e1|op(e1,e3)=e2|op(e1,e3)=e3|op(e1,e3)=e4))&(op(e1,e4)=e0|op(e1,e4)=e1|op(e1,e4)=e2|op(e1,e4)=e3|op(e1,e4)=e4))&(op(e2,e0)=e0|op(e2,e0)=e1|op(e2,e0)=e2|op(e2,e0)=e3|op(e2,e0)=e4))&(op(e2,e1)=e0|op(e2,e1)=e1|op(e2,e1)=e2|op(e2,e1)=e3|op(e2,e1)=e4))&(op(e2,e2)=e0|op(e2,e2)=e1|op(e2,e2)=e2|op(e2,e2)=e3|op(e2,e2)=e4))&(op(e2,e3)=e0|op(e2,e3)=e1|op(e2,e3)=e2|op(e2,e3)=e3|op(e2,e3)=e4))&(op(e2,e4)=e0|op(e2,e4)=e1|op(e2,e4)=e2|op(e2,e4)=e3|op(e2,e4)=e4))&(op(e3,e0)=e0|op(e3,e0)=e1|op(e3,e0)=e2|op(e3,e0)=e3|op(e3,e0)=e4))&(op(e3,e1)=e0|op(e3,e1)=e1|op(e3,e1)=e2|op(e3,e1)=e3|op(e3,e1)=e4))&(op(e3,e2)=e0|op(e3,e2)=e1|op(e3,e2)=e2|op(e3,e2)=e3|op(e3,e2)=e4))&(op(e3,e3)=e0|op(e3,e3)=e1|op(e3,e3)=e2|op(e3,e3)=e3|op(e3,e3)=e4))&(op(e3,e4)=e0|op(e3,e4)=e1|op(e3,e4)=e2|op(e3,e4)=e3|op(e3,e4)=e4))&(op(e4,e0)=e0|op(e4,e0)=e1|op(e4,e0)=e2|op(e4,e0)=e3|op(e4,e0)=e4))&(op(e4,e1)=e0|op(e4,e1)=e1|op(e4,e1)=e2|op(e4,e1)=e3|op(e4,e1)=e4))&(op(e4,e2)=e0|op(e4,e2)=e1|op(e4,e2)=e2|op(e4,e2)=e3|op(e4,e2)=e4))&(op(e4,e3)=e0|op(e4,e3)=e1|op(e4,e3)=e2|op(e4,e3)=e3|op(e4,e3)=e4))&(op(e4,e4)=e0|op(e4,e4)=e1|op(e4,e4)=e2|op(e4,e4)=e3|op(e4,e4)=e4))&(op(e0,e0)=e0|op(e0,e1)=e0|op(e0,e2)=e0|op(e0,e3)=e0|op(e0,e4)=e0))&(op(e0,e0)=e0|op(e1,e0)=e0|op(e2,e0)=e0|op(e3,e0)=e0|op(e4,e0)=e0))&(op(e0,e0)=e1|op(e0,e1)=e1|op(e0,e2)=e1|op(e0,e3)=e1|op(e0,e4)=e1))&(op(e0,e0)=e1|op(e1,e0)=e1|op(e2,e0)=e1|op(e3,e0)=e1|op(e4,e0)=e1))&(op(e0,e0)=e2|op(e0,e1)=e2|op(e0,e2)=e2|op(e0,e3)=e2|op(e0,e4)=e2))&(op(e0,e0)=e2|op(e1,e0)=e2|op(e2,e0)=e2|op(e3,e0)=e2|op(e4,e0)=e2))&(op(e0,e0)=e3|op(e0,e1)=e3|op(e0,e2)=e3|op(e0,e3)=e3|op(e0,e4)=e3))&(op(e0,e0)=e3|op(e1,e0)=e3|op(e2,e0)=e3|op(e3,e0)=e3|op(e4,e0)=e3))&(op(e0,e0)=e4|op(e0,e1)=e4|op(e0,e2)=e4|op(e0,e3)=e4|op(e0,e4)=e4))&(op(e0,e0)=e4|op(e1,e0)=e4|op(e2,e0)=e4|op(e3,e0)=e4|op(e4,e0)=e4))&(op(e1,e0)=e0|op(e1,e1)=e0|op(e1,e2)=e0|op(e1,e3)=e0|op(e1,e4)=e0))&(op(e0,e1)=e0|op(e1,e1)=e0|op(e2,e1)=e0|op(e3,e1)=e0|op(e4,e1)=e0))&(op(e1,e0)=e1|op(e1,e1)=e1|op(e1,e2)=e1|op(e1,e3)=e1|op(e1,e4)=e1))&(op(e0,e1)=e1|op(e1,e1)=e1|op(e2,e1)=e1|op(e3,e1)=e1|op(e4,e1)=e1))&(op(e1,e0)=e2|op(e1,e1)=e2|op(e1,e2)=e2|op(e1,e3)=e2|op(e1,e4)=e2))&(op(e0,e1)=e2|op(e1,e1)=e2|op(e2,e1)=e2|op(e3,e1)=e2|op(e4,e1)=e2))&(op(e1,e0)=e3|op(e1,e1)=e3|op(e1,e2)=e3|op(e1,e3)=e3|op(e1,e4)=e3))&(op(e0,e1)=e3|op(e1,e1)=e3|op(e2,e1)=e3|op(e3,e1)=e3|op(e4,e1)=e3))&(op(e1,e0)=e4|op(e1,e1)=e4|op(e1,e2)=e4|op(e1,e3)=e4|op(e1,e4)=e4))&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))), introduced(definition)).
fof(c_0_216, plain, (((epred22_0&(op(e0,e1)=e4|op(e1,e1)=e4|op(e2,e1)=e4|op(e3,e1)=e4|op(e4,e1)=e4))&(op(e2,e0)=e0|op(e2,e1)=e0|op(e2,e2)=e0|op(e2,e3)=e0|op(e2,e4)=e0))=>epred23_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_204]), c_0_192])).
fof(c_0_217, plain, (((((((((op(e1,e0)!=e4|(op(e0,e1)!=e3|~epred21_0)|epred22_0)&(op(e1,e1)!=e4|(op(e0,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e2)!=e4|(op(e0,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e3)!=e4|(op(e0,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e4)!=e4|(op(e0,e1)!=e3|~epred21_0)|epred22_0))&(((((op(e1,e0)!=e4|(op(e1,e1)!=e3|~epred21_0)|epred22_0)&(op(e1,e1)!=e4|(op(e1,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e2)!=e4|(op(e1,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e3)!=e4|(op(e1,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e4)!=e4|(op(e1,e1)!=e3|~epred21_0)|epred22_0)))&(((((op(e1,e0)!=e4|(op(e2,e1)!=e3|~epred21_0)|epred22_0)&(op(e1,e1)!=e4|(op(e2,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e2)!=e4|(op(e2,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e3)!=e4|(op(e2,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e4)!=e4|(op(e2,e1)!=e3|~epred21_0)|epred22_0)))&(((((op(e1,e0)!=e4|(op(e3,e1)!=e3|~epred21_0)|epred22_0)&(op(e1,e1)!=e4|(op(e3,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e2)!=e4|(op(e3,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e3)!=e4|(op(e3,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e4)!=e4|(op(e3,e1)!=e3|~epred21_0)|epred22_0)))&(((((op(e1,e0)!=e4|(op(e4,e1)!=e3|~epred21_0)|epred22_0)&(op(e1,e1)!=e4|(op(e4,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e2)!=e4|(op(e4,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e3)!=e4|(op(e4,e1)!=e3|~epred21_0)|epred22_0))&(op(e1,e4)!=e4|(op(e4,e1)!=e3|~epred21_0)|epred22_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_205])])).
cnf(c_0_218, plain, (epred21_0|op(e1,e2)!=e3|op(e4,e1)!=e2|~epred20_0), inference(split_conjunct,[status(thm)],[c_0_206])).
cnf(c_0_219, plain, (epred20_0|~epred19_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_207, c_0_23]), c_0_25])])).
cnf(c_0_220, plain, (epred19_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_208, c_0_209])])).
cnf(c_0_221, plain, (~epred28_0), inference(pm,[status(thm)],[c_0_210, c_0_211])).
cnf(c_0_222, plain, (epred28_0|~epred27_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_212, c_0_15]), c_0_32])])).
cnf(c_0_223, plain, (epred27_0|op(e2,e3)!=e4|op(e1,e2)!=e3|~epred26_0), inference(split_conjunct,[status(thm)],[c_0_213])).
fof(c_0_224, plain, (((((((((op(e2,e0)!=e3|(op(e0,e2)!=e2|~epred25_0)|epred26_0)&(op(e2,e1)!=e3|(op(e0,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e2)!=e3|(op(e0,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e3)!=e3|(op(e0,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e4)!=e3|(op(e0,e2)!=e2|~epred25_0)|epred26_0))&(((((op(e2,e0)!=e3|(op(e1,e2)!=e2|~epred25_0)|epred26_0)&(op(e2,e1)!=e3|(op(e1,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e2)!=e3|(op(e1,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e3)!=e3|(op(e1,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e4)!=e3|(op(e1,e2)!=e2|~epred25_0)|epred26_0)))&(((((op(e2,e0)!=e3|(op(e2,e2)!=e2|~epred25_0)|epred26_0)&(op(e2,e1)!=e3|(op(e2,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e2)!=e3|(op(e2,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e3)!=e3|(op(e2,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e4)!=e3|(op(e2,e2)!=e2|~epred25_0)|epred26_0)))&(((((op(e2,e0)!=e3|(op(e3,e2)!=e2|~epred25_0)|epred26_0)&(op(e2,e1)!=e3|(op(e3,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e2)!=e3|(op(e3,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e3)!=e3|(op(e3,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e4)!=e3|(op(e3,e2)!=e2|~epred25_0)|epred26_0)))&(((((op(e2,e0)!=e3|(op(e4,e2)!=e2|~epred25_0)|epred26_0)&(op(e2,e1)!=e3|(op(e4,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e2)!=e3|(op(e4,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e3)!=e3|(op(e4,e2)!=e2|~epred25_0)|epred26_0))&(op(e2,e4)!=e3|(op(e4,e2)!=e2|~epred25_0)|epred26_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_214])])).
fof(c_0_225, plain, (((epred24_0&(op(e0,e2)=e1|op(e1,e2)=e1|op(e2,e2)=e1|op(e3,e2)=e1|op(e4,e2)=e1))&(op(e2,e0)=e2|op(e2,e1)=e2|op(e2,e2)=e2|op(e2,e3)=e2|op(e2,e4)=e2))=>epred25_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_203]), c_0_215])).
fof(c_0_226, plain, (((epred23_0&(op(e0,e2)=e0|op(e1,e2)=e0|op(e2,e2)=e0|op(e3,e2)=e0|op(e4,e2)=e0))&(op(e2,e0)=e1|op(e2,e1)=e1|op(e2,e2)=e1|op(e2,e3)=e1|op(e2,e4)=e1))=>epred24_0), inference(apply_def,[status(thm)],[inference(split_equiv,[status(thm)],[c_0_215]), c_0_204])).
fof(c_0_227, plain, (((((((((op(e2,e0)!=e0|(op(e0,e1)!=e4|~epred22_0)|epred23_0)&(op(e2,e1)!=e0|(op(e0,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e2)!=e0|(op(e0,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e3)!=e0|(op(e0,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e4)!=e0|(op(e0,e1)!=e4|~epred22_0)|epred23_0))&(((((op(e2,e0)!=e0|(op(e1,e1)!=e4|~epred22_0)|epred23_0)&(op(e2,e1)!=e0|(op(e1,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e2)!=e0|(op(e1,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e3)!=e0|(op(e1,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e4)!=e0|(op(e1,e1)!=e4|~epred22_0)|epred23_0)))&(((((op(e2,e0)!=e0|(op(e2,e1)!=e4|~epred22_0)|epred23_0)&(op(e2,e1)!=e0|(op(e2,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e2)!=e0|(op(e2,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e3)!=e0|(op(e2,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e4)!=e0|(op(e2,e1)!=e4|~epred22_0)|epred23_0)))&(((((op(e2,e0)!=e0|(op(e3,e1)!=e4|~epred22_0)|epred23_0)&(op(e2,e1)!=e0|(op(e3,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e2)!=e0|(op(e3,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e3)!=e0|(op(e3,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e4)!=e0|(op(e3,e1)!=e4|~epred22_0)|epred23_0)))&(((((op(e2,e0)!=e0|(op(e4,e1)!=e4|~epred22_0)|epred23_0)&(op(e2,e1)!=e0|(op(e4,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e2)!=e0|(op(e4,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e3)!=e0|(op(e4,e1)!=e4|~epred22_0)|epred23_0))&(op(e2,e4)!=e0|(op(e4,e1)!=e4|~epred22_0)|epred23_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_216])])).
cnf(c_0_228, plain, (epred22_0|op(e1,e0)!=e4|op(e0,e1)!=e3|~epred21_0), inference(split_conjunct,[status(thm)],[c_0_217])).
cnf(c_0_229, plain, (epred21_0|~epred20_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_218, c_0_24]), c_0_38])])).
cnf(c_0_230, plain, (epred20_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_219, c_0_220])])).
cnf(c_0_231, plain, (~epred27_0), inference(pm,[status(thm)],[c_0_221, c_0_222])).
cnf(c_0_232, plain, (epred27_0|~epred26_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_223, c_0_24]), c_0_29])])).
cnf(c_0_233, plain, (epred26_0|op(e2,e4)!=e3|op(e2,e2)!=e2|~epred25_0), inference(split_conjunct,[status(thm)],[c_0_224])).
fof(c_0_234, plain, (((((((((op(e2,e0)!=e2|(op(e0,e2)!=e1|~epred24_0)|epred25_0)&(op(e2,e1)!=e2|(op(e0,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e2)!=e2|(op(e0,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e3)!=e2|(op(e0,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e4)!=e2|(op(e0,e2)!=e1|~epred24_0)|epred25_0))&(((((op(e2,e0)!=e2|(op(e1,e2)!=e1|~epred24_0)|epred25_0)&(op(e2,e1)!=e2|(op(e1,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e2)!=e2|(op(e1,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e3)!=e2|(op(e1,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e4)!=e2|(op(e1,e2)!=e1|~epred24_0)|epred25_0)))&(((((op(e2,e0)!=e2|(op(e2,e2)!=e1|~epred24_0)|epred25_0)&(op(e2,e1)!=e2|(op(e2,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e2)!=e2|(op(e2,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e3)!=e2|(op(e2,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e4)!=e2|(op(e2,e2)!=e1|~epred24_0)|epred25_0)))&(((((op(e2,e0)!=e2|(op(e3,e2)!=e1|~epred24_0)|epred25_0)&(op(e2,e1)!=e2|(op(e3,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e2)!=e2|(op(e3,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e3)!=e2|(op(e3,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e4)!=e2|(op(e3,e2)!=e1|~epred24_0)|epred25_0)))&(((((op(e2,e0)!=e2|(op(e4,e2)!=e1|~epred24_0)|epred25_0)&(op(e2,e1)!=e2|(op(e4,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e2)!=e2|(op(e4,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e3)!=e2|(op(e4,e2)!=e1|~epred24_0)|epred25_0))&(op(e2,e4)!=e2|(op(e4,e2)!=e1|~epred24_0)|epred25_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_225])])).
fof(c_0_235, plain, (((((((((op(e2,e0)!=e1|(op(e0,e2)!=e0|~epred23_0)|epred24_0)&(op(e2,e1)!=e1|(op(e0,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e2)!=e1|(op(e0,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e3)!=e1|(op(e0,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e4)!=e1|(op(e0,e2)!=e0|~epred23_0)|epred24_0))&(((((op(e2,e0)!=e1|(op(e1,e2)!=e0|~epred23_0)|epred24_0)&(op(e2,e1)!=e1|(op(e1,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e2)!=e1|(op(e1,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e3)!=e1|(op(e1,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e4)!=e1|(op(e1,e2)!=e0|~epred23_0)|epred24_0)))&(((((op(e2,e0)!=e1|(op(e2,e2)!=e0|~epred23_0)|epred24_0)&(op(e2,e1)!=e1|(op(e2,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e2)!=e1|(op(e2,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e3)!=e1|(op(e2,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e4)!=e1|(op(e2,e2)!=e0|~epred23_0)|epred24_0)))&(((((op(e2,e0)!=e1|(op(e3,e2)!=e0|~epred23_0)|epred24_0)&(op(e2,e1)!=e1|(op(e3,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e2)!=e1|(op(e3,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e3)!=e1|(op(e3,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e4)!=e1|(op(e3,e2)!=e0|~epred23_0)|epred24_0)))&(((((op(e2,e0)!=e1|(op(e4,e2)!=e0|~epred23_0)|epred24_0)&(op(e2,e1)!=e1|(op(e4,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e2)!=e1|(op(e4,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e3)!=e1|(op(e4,e2)!=e0|~epred23_0)|epred24_0))&(op(e2,e4)!=e1|(op(e4,e2)!=e0|~epred23_0)|epred24_0))), inference(distribute,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_226])])).
cnf(c_0_236, plain, (epred23_0|op(e2,e1)!=e0|op(e3,e1)!=e4|~epred22_0), inference(split_conjunct,[status(thm)],[c_0_227])).
cnf(c_0_237, plain, (epred22_0|~epred21_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_228, c_0_19]), c_0_21])])).
cnf(c_0_238, plain, (epred21_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_229, c_0_230])])).
cnf(c_0_239, plain, (~epred26_0), inference(pm,[status(thm)],[c_0_231, c_0_232])).
cnf(c_0_240, plain, (epred26_0|~epred25_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_233, c_0_28]), c_0_30])])).
cnf(c_0_241, plain, (epred25_0|op(e2,e2)!=e2|op(e4,e2)!=e1|~epred24_0), inference(split_conjunct,[status(thm)],[c_0_234])).
cnf(c_0_242, plain, (epred24_0|op(e2,e0)!=e1|op(e3,e2)!=e0|~epred23_0), inference(split_conjunct,[status(thm)],[c_0_235])).
cnf(c_0_243, plain, (epred23_0|~epred22_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_236, c_0_27]), c_0_33])])).
cnf(c_0_244, plain, (epred22_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_237, c_0_238])])).
cnf(c_0_245, plain, (~epred25_0), inference(pm,[status(thm)],[c_0_239, c_0_240])).
cnf(c_0_246, plain, (epred25_0|~epred24_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_241, c_0_28]), c_0_39])])).
cnf(c_0_247, plain, (epred24_0|~epred23_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_242, c_0_26]), c_0_32])])).
cnf(c_0_248, plain, (epred23_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_243, c_0_244])])).
cnf(c_0_249, plain, (~epred24_0), inference(pm,[status(thm)],[c_0_245, c_0_246])).
cnf(c_0_250, plain, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_247, c_0_248])]), c_0_249]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 251
# Proof object clause steps            : 136
# Proof object formula steps           : 115
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 63
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 12
# Proof object simplifying inferences  : 286
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 965
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 965
# Processed clauses                    : 978
# ...of these trivial                  : 0
# ...subsumed                          : 871
# ...remaining for further processing  : 107
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 23
# Generated clauses                    : 13
# ...of the previous two non-trivial   : 35
# Contextual simplify-reflections      : 0
# Paramodulations                      : 13
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 84
#    Positive orientable unit clauses  : 48
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 24
#    Non-unit-clauses                  : 12
# Current number of unprocessed clauses: 22
# ...number of literals in the above   : 589
# Current number of archived formulas  : 0
# Current number of archived clauses   : 23
# Clause-clause subsumption calls (NU) : 566
# Rec. Clause-clause subsumption calls : 566
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 548
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 23
# BW rewrite match successes           : 23
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 51777
