# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_8, negated_conjecture, (class_Ring__and__Field_Oordered__idom(t_b))).
#
#cnf(i_0_12, plain, (class_OrderedGroup_Oab__semigroup__mult(X1)|~class_Ring__and__Field_Oordered__idom(X1))).
#
#cnf(i_0_9, plain, (c_times(X2,X3,X1)=c_times(X3,X2,X1)|~class_OrderedGroup_Oab__semigroup__mult(X1))).
#
#cnf(i_0_7, negated_conjecture, (c_times(c_times(v_c,c_HOL_Oabs(v_f(v_x),t_b),t_b),c_times(v_ca,c_HOL_Oabs(v_g(v_x),t_b),t_b),t_b)!=c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b))).
#
#cnf(i_0_11, plain, (c_times(c_HOL_Oabs(X2,X1),c_HOL_Oabs(X3,X1),X1)=c_HOL_Oabs(c_times(X2,X3,X1),X1)|~class_Ring__and__Field_Oordered__idom(X1))).
#
#cnf(i_0_10, plain, (c_times(X2,c_times(X3,X4,X1),X1)=c_times(X3,c_times(X2,X4,X1),X1)|~class_OrderedGroup_Oab__semigroup__mult(X1))).
#
#cnf(i_0_18, plain, (c_times(c_times(X2,X3,X4),X1,X4)=c_times(X2,c_times(X1,X3,X4),X4)|~class_OrderedGroup_Oab__semigroup__mult(X4))).
#
#cnf(i_0_13, negated_conjecture, (c_times(c_times(v_ca,c_HOL_Oabs(v_g(v_x),t_b),t_b),c_times(v_c,c_HOL_Oabs(v_f(v_x),t_b),t_b),t_b)!=c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b)|~class_OrderedGroup_Oab__semigroup__mult(t_b))).
#####
#cnf(i_0_44, plain, (c_times(c_times(X2,X1,X3),X4,X3)=c_times(X1,c_times(X4,X2,X3),X3)|~class_OrderedGroup_Oab__semigroup__mult(X3))).
###
#cnf(i_0_17, negated_conjecture, (c_times(v_ca,c_times(c_times(v_c,c_HOL_Oabs(v_f(v_x),t_b),t_b),c_HOL_Oabs(v_g(v_x),t_b),t_b),t_b)!=c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b)|~class_OrderedGroup_Oab__semigroup__mult(t_b))).
##
#cnf(i_0_28, plain, (c_times(X2,c_times(X3,X1,X4),X4)=c_times(X1,c_times(X2,X3,X4),X4)|~class_OrderedGroup_Oab__semigroup__mult(X4))).
#
#cnf(i_0_74, plain, (c_times(c_times(X2,X3,X4),X1,X4)=c_times(X2,c_times(X3,X1,X4),X4)|~class_OrderedGroup_Oab__semigroup__mult(X4))).
#####
#cnf(i_0_40, negated_conjecture, (c_times(v_c,c_times(c_times(v_ca,c_HOL_Oabs(v_g(v_x),t_b),t_b),c_HOL_Oabs(v_f(v_x),t_b),t_b),t_b)!=c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b)|~class_OrderedGroup_Oab__semigroup__mult(t_b))).
#
#cnf(i_0_139, plain, (c_times(c_times(X2,X1,X3),X4,X3)=c_times(X1,c_times(X2,X4,X3),X3)|~class_OrderedGroup_Oab__semigroup__mult(X3))).
##############
#cnf(i_0_59, plain, (c_times(X2,c_times(X4,X1,X3),X3)=c_times(X1,c_times(X4,X2,X3),X3)|~class_OrderedGroup_Oab__semigroup__mult(X3))).
#
#cnf(i_0_57, negated_conjecture, (c_times(c_HOL_Oabs(v_g(v_x),t_b),c_times(c_times(v_c,c_HOL_Oabs(v_f(v_x),t_b),t_b),v_ca,t_b),t_b)!=c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b)|~class_OrderedGroup_Oab__semigroup__mult(t_b))).
####
#cnf(i_0_58, negated_conjecture, (c_times(c_HOL_Oabs(v_f(v_x),t_b),c_times(c_times(v_ca,c_HOL_Oabs(v_g(v_x),t_b),t_b),v_c,t_b),t_b)!=c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b)|~class_OrderedGroup_Oab__semigroup__mult(t_b))).
##
#cnf(i_0_82, plain, (c_times(X2,c_times(X3,X1,X4),X4)=c_times(X2,c_times(X1,X3,X4),X4)|~class_OrderedGroup_Oab__semigroup__mult(X4))).
##
#cnf(i_0_66, negated_conjecture, (c_times(v_ca,c_times(c_HOL_Oabs(v_g(v_x),t_b),c_times(v_c,c_HOL_Oabs(v_f(v_x),t_b),t_b),t_b),t_b)!=c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b)|~class_OrderedGroup_Oab__semigroup__mult(t_b))).
########
#cnf(i_0_68, negated_conjecture, (c_times(v_ca,c_times(v_c,c_times(c_HOL_Oabs(v_g(v_x),t_b),c_HOL_Oabs(v_f(v_x),t_b),t_b),t_b),t_b)!=c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b)|~class_OrderedGroup_Oab__semigroup__mult(t_b))).
####
#cnf(i_0_342, negated_conjecture, (c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b)!=c_times(v_ca,c_times(v_c,c_HOL_Oabs(c_times(v_g(v_x),v_f(v_x),t_b),t_b),t_b),t_b)|~class_OrderedGroup_Oab__semigroup__mult(t_b))).
####
#cnf(i_0_346, negated_conjecture, (c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_g(v_x),v_f(v_x),t_b),t_b),t_b)!=c_times(v_ca,c_times(v_c,c_HOL_Oabs(c_times(v_g(v_x),v_f(v_x),t_b),t_b),t_b),t_b)|~class_OrderedGroup_Oab__semigroup__mult(t_b))).
####
#cnf(i_0_355, negated_conjecture, (~class_OrderedGroup_Oab__semigroup__mult(t_b))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(cls_conjecture_4, negated_conjecture, (c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b)!=c_times(c_times(v_c,c_HOL_Oabs(v_f(v_x),t_b),t_b),c_times(v_ca,c_HOL_Oabs(v_g(v_x),t_b),t_b),t_b)), file('./p/ANA035-2.tptp', cls_conjecture_4)).
cnf(cls_OrderedGroup_Omult__ac__3_0, axiom, (c_times(X2,c_times(X3,X4,X1),X1)=c_times(X3,c_times(X2,X4,X1),X1)|~class_OrderedGroup_Oab__semigroup__mult(X1)), file('./p/ANA035-2.tptp', cls_OrderedGroup_Omult__ac__3_0)).
cnf(cls_OrderedGroup_Omult__ac__2_0, axiom, (c_times(X2,X3,X1)=c_times(X3,X2,X1)|~class_OrderedGroup_Oab__semigroup__mult(X1)), file('./p/ANA035-2.tptp', cls_OrderedGroup_Omult__ac__2_0)).
cnf(cls_Ring__and__Field_Oabs__mult_0, axiom, (c_HOL_Oabs(c_times(X2,X3,X1),X1)=c_times(c_HOL_Oabs(X2,X1),c_HOL_Oabs(X3,X1),X1)|~class_Ring__and__Field_Oordered__idom(X1)), file('./p/ANA035-2.tptp', cls_Ring__and__Field_Oabs__mult_0)).
cnf(tfree_tcs, negated_conjecture, (class_Ring__and__Field_Oordered__idom(t_b)), file('./p/ANA035-2.tptp', tfree_tcs)).
cnf(clsrel_Ring__and__Field_Oordered__idom_17, axiom, (class_OrderedGroup_Oab__semigroup__mult(X1)|~class_Ring__and__Field_Oordered__idom(X1)), file('./p/ANA035-2.tptp', clsrel_Ring__and__Field_Oordered__idom_17)).
cnf(c_0_6, negated_conjecture, (c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b)!=c_times(c_times(v_c,c_HOL_Oabs(v_f(v_x),t_b),t_b),c_times(v_ca,c_HOL_Oabs(v_g(v_x),t_b),t_b),t_b)), cls_conjecture_4).
cnf(c_0_7, axiom, (c_times(X2,c_times(X3,X4,X1),X1)=c_times(X3,c_times(X2,X4,X1),X1)|~class_OrderedGroup_Oab__semigroup__mult(X1)), cls_OrderedGroup_Omult__ac__3_0).
cnf(c_0_8, axiom, (c_times(X2,X3,X1)=c_times(X3,X2,X1)|~class_OrderedGroup_Oab__semigroup__mult(X1)), cls_OrderedGroup_Omult__ac__2_0).
cnf(c_0_9, negated_conjecture, (c_times(v_ca,c_times(c_times(v_c,c_HOL_Oabs(v_f(v_x),t_b),t_b),c_HOL_Oabs(v_g(v_x),t_b),t_b),t_b)!=c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b)|~class_OrderedGroup_Oab__semigroup__mult(t_b)), inference(pm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_10, plain, (c_times(c_times(X1,X2,X3),X4,X3)=c_times(X1,c_times(X4,X2,X3),X3)|~class_OrderedGroup_Oab__semigroup__mult(X3)), inference(pm,[status(thm)],[c_0_8, c_0_7])).
cnf(c_0_11, negated_conjecture, (c_times(v_ca,c_times(v_c,c_times(c_HOL_Oabs(v_g(v_x),t_b),c_HOL_Oabs(v_f(v_x),t_b),t_b),t_b),t_b)!=c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b)|~class_OrderedGroup_Oab__semigroup__mult(t_b)), inference(pm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_12, axiom, (c_HOL_Oabs(c_times(X2,X3,X1),X1)=c_times(c_HOL_Oabs(X2,X1),c_HOL_Oabs(X3,X1),X1)|~class_Ring__and__Field_Oordered__idom(X1)), cls_Ring__and__Field_Oabs__mult_0).
cnf(c_0_13, negated_conjecture, (class_Ring__and__Field_Oordered__idom(t_b)), tfree_tcs).
cnf(c_0_14, plain, (c_times(X1,c_times(X2,X3,X4),X4)=c_times(X3,c_times(X1,X2,X4),X4)|~class_OrderedGroup_Oab__semigroup__mult(X4)), inference(pm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_15, negated_conjecture, (c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_f(v_x),v_g(v_x),t_b),t_b),t_b)!=c_times(v_ca,c_times(v_c,c_HOL_Oabs(c_times(v_g(v_x),v_f(v_x),t_b),t_b),t_b),t_b)|~class_OrderedGroup_Oab__semigroup__mult(t_b)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_11, c_0_12]), c_0_13])])).
cnf(c_0_16, plain, (c_times(c_times(X1,X2,X3),X4,X3)=c_times(X1,c_times(X2,X4,X3),X3)|~class_OrderedGroup_Oab__semigroup__mult(X3)), inference(pm,[status(thm)],[c_0_8, c_0_14])).
cnf(c_0_17, negated_conjecture, (c_times(c_times(v_c,v_ca,t_b),c_HOL_Oabs(c_times(v_g(v_x),v_f(v_x),t_b),t_b),t_b)!=c_times(v_ca,c_times(v_c,c_HOL_Oabs(c_times(v_g(v_x),v_f(v_x),t_b),t_b),t_b),t_b)|~class_OrderedGroup_Oab__semigroup__mult(t_b)), inference(pm,[status(thm)],[c_0_15, c_0_8])).
cnf(c_0_18, plain, (c_times(c_times(X1,X2,X3),X4,X3)=c_times(X2,c_times(X1,X4,X3),X3)|~class_OrderedGroup_Oab__semigroup__mult(X3)), inference(pm,[status(thm)],[c_0_16, c_0_8])).
cnf(c_0_19, negated_conjecture, (~class_OrderedGroup_Oab__semigroup__mult(t_b)), inference(pm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_20, axiom, (class_OrderedGroup_Oab__semigroup__mult(X1)|~class_Ring__and__Field_Oordered__idom(X1)), clsrel_Ring__and__Field_Oordered__idom_17).
cnf(c_0_21, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_19, c_0_20]), c_0_13])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 22
# Proof object clause steps            : 16
# Proof object formula steps           : 6
# Proof object conjectures             : 10
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 2
# Proof object initial clauses used    : 6
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 6
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 6
# Processed clauses                    : 68
# ...of these trivial                  : 0
# ...subsumed                          : 45
# ...remaining for further processing  : 23
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 344
# ...of the previous two non-trivial   : 343
# Contextual simplify-reflections      : 0
# Paramodulations                      : 344
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 23
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 20
# Current number of unprocessed clauses: 281
# ...number of literals in the above   : 576
# Current number of archived formulas  : 0
# Current number of archived clauses   : 0
# Clause-clause subsumption calls (NU) : 310
# Rec. Clause-clause subsumption calls : 310
# Non-unit clause-clause subsumptions  : 45
# Unit Clause-clause subsumption calls : 9
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3679
