# Initializing proof state
# Scanning for AC axioms
# add is commutative
# multiply is commutative
# AC handling enabled
#
#cnf(i_0_14, plain, (add(X1,additive_identity)=X1)).
#
#cnf(i_0_15, plain, (multiply(X1,multiplicative_identity)=X1)).
#
#cnf(i_0_16, plain, (add(X1,inverse(X1))=multiplicative_identity)).
#
#cnf(i_0_18, negated_conjecture, (add(a,a)!=a)).
#
#cnf(i_0_17, plain, (multiply(X1,inverse(X1))=additive_identity)).
#
#cnf(i_0_10, plain, (add(X1,X2)=add(X2,X1))).
#
#cnf(i_0_20, plain, (add(additive_identity,X1)=X1)).
##
#cnf(i_0_23, plain, (inverse(additive_identity)=multiplicative_identity)).
##
#cnf(i_0_11, plain, (multiply(X1,X2)=multiply(X2,X1))).
##
#cnf(i_0_28, plain, (multiply(multiplicative_identity,X1)=X1)).
#
#cnf(i_0_30, plain, (inverse(multiplicative_identity)=additive_identity)).
##
#cnf(i_0_12, plain, (multiply(add(X1,X2),add(X1,X3))=add(X1,multiply(X2,X3)))).
#
#cnf(i_0_38, plain, (add(X1,multiply(X2,inverse(X1)))=add(X1,X2))).
#
#cnf(i_0_55, plain, (add(X1,X1)=X1)).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(distributivity1, axiom, (add(X1,multiply(X2,X3))=multiply(add(X1,X2),add(X1,X3))), file('./p/BOO004-4.tptp', distributivity1)).
cnf(additive_inverse1, axiom, (add(X1,inverse(X1))=multiplicative_identity), file('./p/BOO004-4.tptp', additive_inverse1)).
cnf(multiplicative_id1, axiom, (multiply(X1,multiplicative_identity)=X1), file('./p/BOO004-4.tptp', multiplicative_id1)).
cnf(multiplicative_inverse1, axiom, (multiply(X1,inverse(X1))=additive_identity), file('./p/BOO004-4.tptp', multiplicative_inverse1)).
cnf(additive_id1, axiom, (add(X1,additive_identity)=X1), file('./p/BOO004-4.tptp', additive_id1)).
cnf(prove_a_plus_a_is_a, negated_conjecture, (add(a,a)!=a), file('./p/BOO004-4.tptp', prove_a_plus_a_is_a)).
cnf(c_0_6, axiom, (add(X1,multiply(X2,X3))=multiply(add(X1,X2),add(X1,X3))), distributivity1).
cnf(c_0_7, axiom, (add(X1,inverse(X1))=multiplicative_identity), additive_inverse1).
cnf(c_0_8, axiom, (multiply(X1,multiplicative_identity)=X1), multiplicative_id1).
cnf(c_0_9, plain, (add(X1,multiply(X2,inverse(X1)))=add(X1,X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_6, c_0_7]), c_0_8])).
cnf(c_0_10, axiom, (multiply(X1,inverse(X1))=additive_identity), multiplicative_inverse1).
cnf(c_0_11, axiom, (add(X1,additive_identity)=X1), additive_id1).
cnf(c_0_12, negated_conjecture, (add(a,a)!=a), prove_a_plus_a_is_a).
cnf(c_0_13, plain, (add(X1,X1)=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_9, c_0_10]), c_0_11])).
cnf(c_0_14, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_13])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 15
# Proof object clause steps            : 9
# Proof object formula steps           : 6
# Proof object conjectures             : 3
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 6
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 9
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 9
# Processed clauses                    : 18
# ...of these trivial                  : 2
# ...subsumed                          : 2
# ...remaining for further processing  : 14
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 41
# ...of the previous two non-trivial   : 25
# Contextual simplify-reflections      : 0
# Paramodulations                      : 41
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 13
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 16
# ...number of literals in the above   : 16
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 19
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 339
