# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_7, plain, (apply(apply(w,X1),X2)=apply(apply(X1,X2),X2))).
#
#cnf(i_0_5, plain, (apply(apply(apply(b,X1),X2),X3)=apply(X1,apply(X2,X3)))).
#
#cnf(i_0_6, plain, (apply(apply(apply(q,X1),X2),X3)=apply(X2,apply(X1,X3)))).
#
#cnf(i_0_8, negated_conjecture, (apply(apply(apply(w,apply(apply(w,apply(q,apply(q,q))),apply(q,x))),y),z)!=apply(apply(x,y),apply(apply(x,y),z)))).
#
#cnf(i_0_10, plain, (apply(apply(w,w),X1)=apply(apply(X1,X1),X1))).
####
#cnf(i_0_14, plain, (apply(X1,apply(X2,X2))=apply(apply(w,apply(b,X1)),X2))).
#
#cnf(i_0_15, plain, (apply(apply(apply(w,b),X1),X2)=apply(X1,apply(X1,X2)))).
#
#cnf(i_0_17, plain, (apply(X2,apply(X1,X2))=apply(apply(w,apply(q,X1)),X2))).
#
#cnf(i_0_8, negated_conjecture, (apply(apply(apply(w,apply(apply(q,x),apply(apply(q,q),apply(q,x)))),y),z)!=apply(apply(x,y),apply(apply(x,y),z)))).
#
#cnf(i_0_18, plain, (apply(apply(apply(w,q),X1),X2)=apply(X1,apply(X1,X2)))).
#
#cnf(i_0_59, plain, (apply(X1,apply(X1,X1))=apply(apply(w,apply(w,b)),X1))).
#
#cnf(i_0_83, plain, (apply(apply(w,apply(q,X1)),X1)=apply(apply(w,apply(b,X1)),X1))).
##
#cnf(i_0_107, plain, (apply(apply(apply(w,q),X1),X2)=apply(apply(apply(w,b),X1),X2))).
###
#cnf(i_0_174, plain, (apply(apply(w,apply(w,q)),X1)=apply(apply(w,apply(w,b)),X1))).
#
#cnf(i_0_92, plain, (apply(apply(w,apply(q,apply(X2,X1))),X1)=apply(X1,apply(apply(w,X2),X1)))).
##
#cnf(i_0_9, plain, (apply(apply(w,apply(w,X1)),X2)=apply(apply(w,apply(X1,X2)),X2))).
######
#cnf(i_0_20, plain, (apply(X1,apply(apply(b,X1),apply(b,X1)))=apply(apply(w,w),apply(b,X1)))).
#
#cnf(i_0_21, plain, (apply(apply(q,X1),apply(X1,apply(q,X1)))=apply(apply(w,w),apply(q,X1)))).
###
#cnf(i_0_19, plain, (apply(apply(w,apply(w,w)),X1)=apply(apply(w,apply(X1,X1)),X1))).
#####
#cnf(i_0_13, plain, (apply(apply(X1,apply(X2,X3)),X3)=apply(apply(w,apply(apply(b,X1),X2)),X3))).
#
#cnf(i_0_334, plain, (apply(apply(w,apply(apply(b,w),X1)),X2)=apply(apply(w,apply(w,X1)),X2))).
####
#cnf(i_0_16, plain, (apply(apply(X2,apply(X1,X3)),X3)=apply(apply(w,apply(apply(q,X1),X2)),X3))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(prove_p_combinator, negated_conjecture, (apply(apply(apply(apply(apply(apply(b,apply(w,apply(q,apply(q,q)))),q),x),y),y),z)!=apply(apply(x,y),apply(apply(x,y),z))), file('./p/COL066-3.tptp', prove_p_combinator)).
cnf(b_definition, axiom, (apply(apply(apply(b,X1),X2),X3)=apply(X1,apply(X2,X3))), file('./p/COL066-3.tptp', b_definition)).
cnf(w_definition, axiom, (apply(apply(w,X1),X2)=apply(apply(X1,X2),X2)), file('./p/COL066-3.tptp', w_definition)).
cnf(q_definition, axiom, (apply(apply(apply(q,X1),X2),X3)=apply(X2,apply(X1,X3))), file('./p/COL066-3.tptp', q_definition)).
cnf(c_0_4, negated_conjecture, (apply(apply(apply(apply(apply(apply(b,apply(w,apply(q,apply(q,q)))),q),x),y),y),z)!=apply(apply(x,y),apply(apply(x,y),z))), prove_p_combinator).
cnf(c_0_5, axiom, (apply(apply(apply(b,X1),X2),X3)=apply(X1,apply(X2,X3))), b_definition).
cnf(c_0_6, axiom, (apply(apply(w,X1),X2)=apply(apply(X1,X2),X2)), w_definition).
cnf(c_0_7, axiom, (apply(apply(apply(q,X1),X2),X3)=apply(X2,apply(X1,X3))), q_definition).
cnf(c_0_8, negated_conjecture, (apply(apply(apply(w,apply(apply(w,apply(q,apply(q,q))),apply(q,x))),y),z)!=apply(apply(x,y),apply(apply(x,y),z))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_5]), c_0_6])).
cnf(c_0_9, plain, (apply(X1,apply(X2,X1))=apply(apply(w,apply(q,X2)),X1)), inference(pm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_10, plain, (apply(apply(apply(w,b),X1),X2)=apply(X1,apply(X1,X2))), inference(pm,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_11, plain, (apply(apply(apply(w,q),X1),X2)=apply(X1,apply(X1,X2))), inference(pm,[status(thm)],[c_0_7, c_0_6])).
cnf(c_0_12, negated_conjecture, (apply(apply(apply(w,apply(apply(q,x),apply(apply(q,q),apply(q,x)))),y),z)!=apply(apply(x,y),apply(apply(x,y),z))), inference(rw,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_13, plain, (apply(apply(X1,apply(X2,X3)),X3)=apply(apply(w,apply(apply(q,X2),X1)),X3)), inference(pm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_14, plain, (apply(apply(apply(w,q),X1),X2)=apply(apply(apply(w,b),X1),X2)), inference(pm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_13]), c_0_7]), c_0_7]), c_0_6]), c_0_14]), c_0_10])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 16
# Proof object clause steps            : 12
# Proof object formula steps           : 4
# Proof object conjectures             : 5
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 4
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 4
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 4
# Processed clauses                    : 43
# ...of these trivial                  : 2
# ...subsumed                          : 19
# ...remaining for further processing  : 22
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 519
# ...of the previous two non-trivial   : 409
# Contextual simplify-reflections      : 0
# Paramodulations                      : 519
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 19
#    Positive orientable unit clauses  : 9
#    Positive unorientable unit clauses: 10
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 368
# ...number of literals in the above   : 368
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 7
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 202
# BW rewrite match successes           : 22
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7346
