# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_63, hypothesis, (aElement0(xa))).
#
#cnf(i_0_62, hypothesis, (aElement0(xb))).
#
#cnf(i_0_61, hypothesis, (aElement0(xc))).
#
#cnf(i_0_351, negated_conjecture, (X1!=xb|~aReductOfIn0(X1,xa,xR)|~aElement0(X1))).
#
#cnf(i_0_44, hypothesis, (aRewritingSystem0(xR))).
#
#cnf(i_0_49, hypothesis, (isLocallyConfluent0(xR))).
#
#cnf(i_0_45, hypothesis, (isTerminating0(xR))).
#
#cnf(i_0_40, plain, (~aNormalFormOfIn0(X2,X4,X3)|~aReductOfIn0(X1,X2,X3)|~aRewritingSystem0(X3)|~aElement0(X4))).
#
#cnf(i_0_343, hypothesis, (sdtmndtplgtdt0(xa,xR,xb))).
#
#cnf(i_0_339, hypothesis, (sdtmndtplgtdt0(xa,xR,xc))).
#
#cnf(i_0_36, plain, (isTerminating0(X1)|aElement0(esk11_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_348, negated_conjecture, (~sdtmndtplgtdt0(X1,xR,xb)|~aReductOfIn0(X1,xa,xR)|~aElement0(X1))).
#
#cnf(i_0_21, plain, (isConfluent0(X1)|aElement0(esk3_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_20, plain, (isConfluent0(X1)|aElement0(esk4_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_19, plain, (isConfluent0(X1)|aElement0(esk5_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_347, negated_conjecture, (~sdtmndtasgtdt0(X1,xR,xb)|~aReductOfIn0(X1,xa,xR)|~aElement0(X1))).
#
#cnf(i_0_30, plain, (isLocallyConfluent0(X1)|aElement0(esk7_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_29, plain, (isLocallyConfluent0(X1)|aElement0(esk8_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_28, plain, (isLocallyConfluent0(X1)|aElement0(esk9_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_350, negated_conjecture, (~aReductOfIn0(X1,xa,xR)|~aReductOfIn0(xb,X1,xR)|~aElement0(X1))).
#
#cnf(i_0_37, plain, (isTerminating0(X1)|aElement0(esk10_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_346, hypothesis, (aReductOfIn0(xb,xa,xR)|aElement0(esk20_0))).
#
#cnf(i_0_352, negated_conjecture, (aElement0(esk20_0))).
#
#cnf(i_0_349, negated_conjecture, (~sdtmndtplgtdt0(X1,xR,xb)|~aReductOfIn0(X2,xa,xR)|~aReductOfIn0(X1,X2,xR)|~aElement0(X2)|~aElement0(X1))).
#
#cnf(i_0_342, hypothesis, (aReductOfIn0(xc,xa,xR)|aElement0(esk21_0))).
#
#cnf(i_0_355, negated_conjecture, (aElement0(esk21_0)|xc!=xb)).
#
#cnf(i_0_356, negated_conjecture, (aElement0(esk21_0)|~sdtmndtplgtdt0(xc,xR,xb))).
#
#cnf(i_0_357, negated_conjecture, (aElement0(esk21_0)|~aReductOfIn0(xb,xc,xR))).
#
#cnf(i_0_3, plain, (aElement0(X3)|~aReductOfIn0(X3,X1,X2)|~aRewritingSystem0(X2)|~aElement0(X1))).
#
#cnf(i_0_42, plain, (aElement0(X1)|~aNormalFormOfIn0(X1,X2,X3)|~aRewritingSystem0(X3)|~aElement0(X2))).
#
#cnf(i_0_34, plain, (isTerminating0(X1)|~iLess0(esk11_1(X1),esk10_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_18, plain, (isConfluent0(X1)|sdtmndtasgtdt0(esk3_1(X1),X1,esk4_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_48, hypothesis, (iLess0(X1,X2)|~aReductOfIn0(X1,X2,xR)|~aElement0(X2)|~aElement0(X1))).
#
#cnf(i_0_360, hypothesis, (iLess0(xc,xa)|aElement0(esk21_0))).
#
#cnf(i_0_46, hypothesis, (iLess0(X2,X1)|~sdtmndtplgtdt0(X1,xR,X2)|~aElement0(X2)|~aElement0(X1))).
#
#cnf(i_0_361, hypothesis, (iLess0(xb,xa))).
#
#cnf(i_0_362, hypothesis, (iLess0(xc,xa))).
#
#cnf(i_0_41, plain, (sdtmndtasgtdt0(X1,X2,X3)|~aNormalFormOfIn0(X3,X1,X2)|~aRewritingSystem0(X2)|~aElement0(X1))).
#
#cnf(i_0_17, plain, (isConfluent0(X1)|sdtmndtasgtdt0(esk3_1(X1),X1,esk5_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_35, plain, (isTerminating0(X1)|sdtmndtplgtdt0(esk10_1(X1),X1,esk11_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_27, plain, (isLocallyConfluent0(X1)|aReductOfIn0(esk8_1(X1),esk7_1(X1),X1)|~aRewritingSystem0(X1))).
#
#cnf(i_0_365, plain, (isLocallyConfluent0(X1)|~aNormalFormOfIn0(esk7_1(X1),X2,X1)|~aRewritingSystem0(X1)|~aElement0(X2))).
#
#cnf(i_0_26, plain, (isLocallyConfluent0(X1)|aReductOfIn0(esk9_1(X1),esk7_1(X1),X1)|~aRewritingSystem0(X1))).
####
#cnf(i_0_13, plain, (sdtmndtasgtdt0(X1,X3,X2)|X1!=X2|~aRewritingSystem0(X3)|~aElement0(X2)|~aElement0(X1))).
##
#cnf(i_0_345, hypothesis, (aReductOfIn0(xb,xa,xR)|aReductOfIn0(esk20_0,xa,xR))).
#
#cnf(i_0_376, negated_conjecture, (aReductOfIn0(xb,xa,xR)|esk20_0!=xb)).
#
#cnf(i_0_374, hypothesis, (iLess0(esk20_0,xa)|aReductOfIn0(xb,xa,xR))).
#
#cnf(i_0_382, negated_conjecture, (esk20_0!=xb)).
#
#cnf(i_0_388, negated_conjecture, (iLess0(esk20_0,xa))).
#
#cnf(i_0_378, negated_conjecture, (aReductOfIn0(xb,xa,xR)|~aReductOfIn0(xb,esk20_0,xR))).
#
#cnf(i_0_377, negated_conjecture, (aReductOfIn0(xb,xa,xR)|~sdtmndtplgtdt0(esk20_0,xR,xb))).
#
#cnf(i_0_394, negated_conjecture, (~sdtmndtplgtdt0(esk20_0,xR,xb))).
#
#cnf(i_0_344, hypothesis, (aReductOfIn0(xb,xa,xR))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(m__, conjecture, ?[X1]:((aElement0(X1)&aReductOfIn0(X1,xa,xR))&(X1=xb|aReductOfIn0(xb,X1,xR)|?[X2]:((aElement0(X2)&aReductOfIn0(X2,X1,xR))&sdtmndtplgtdt0(X2,xR,xb))|sdtmndtplgtdt0(X1,xR,xb)|sdtmndtasgtdt0(X1,xR,xb))), file('./p/COM016+4.tptp', m__)).
fof(m__731_02, hypothesis, ((((aReductOfIn0(xb,xa,xR)|?[X1]:((aElement0(X1)&aReductOfIn0(X1,xa,xR))&sdtmndtplgtdt0(X1,xR,xb)))&sdtmndtplgtdt0(xa,xR,xb))&(aReductOfIn0(xc,xa,xR)|?[X1]:((aElement0(X1)&aReductOfIn0(X1,xa,xR))&sdtmndtplgtdt0(X1,xR,xc))))&sdtmndtplgtdt0(xa,xR,xc)), file('./p/COM016+4.tptp', m__731_02)).
fof(m__731, hypothesis, ((aElement0(xa)&aElement0(xb))&aElement0(xc)), file('./p/COM016+4.tptp', m__731)).
fof(c_0_3, negated_conjecture, ~(?[X1]:((aElement0(X1)&aReductOfIn0(X1,xa,xR))&(X1=xb|aReductOfIn0(xb,X1,xR)|?[X2]:((aElement0(X2)&aReductOfIn0(X2,X1,xR))&sdtmndtplgtdt0(X2,xR,xb))|sdtmndtplgtdt0(X1,xR,xb)|sdtmndtasgtdt0(X1,xR,xb)))), inference(assume_negation,[status(cth)],[m__])).
fof(c_0_4, negated_conjecture, ![X76, X77]:(((((X76!=xb|(~aElement0(X76)|~aReductOfIn0(X76,xa,xR)))&(~aReductOfIn0(xb,X76,xR)|(~aElement0(X76)|~aReductOfIn0(X76,xa,xR))))&(~aElement0(X77)|~aReductOfIn0(X77,X76,xR)|~sdtmndtplgtdt0(X77,xR,xb)|(~aElement0(X76)|~aReductOfIn0(X76,xa,xR))))&(~sdtmndtplgtdt0(X76,xR,xb)|(~aElement0(X76)|~aReductOfIn0(X76,xa,xR))))&(~sdtmndtasgtdt0(X76,xR,xb)|(~aElement0(X76)|~aReductOfIn0(X76,xa,xR)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])])).
fof(c_0_5, hypothesis, ((((((aElement0(esk20_0)|aReductOfIn0(xb,xa,xR))&(aReductOfIn0(esk20_0,xa,xR)|aReductOfIn0(xb,xa,xR)))&(sdtmndtplgtdt0(esk20_0,xR,xb)|aReductOfIn0(xb,xa,xR)))&sdtmndtplgtdt0(xa,xR,xb))&(((aElement0(esk21_0)|aReductOfIn0(xc,xa,xR))&(aReductOfIn0(esk21_0,xa,xR)|aReductOfIn0(xc,xa,xR)))&(sdtmndtplgtdt0(esk21_0,xR,xc)|aReductOfIn0(xc,xa,xR))))&sdtmndtplgtdt0(xa,xR,xc)), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[m__731_02])])])).
cnf(c_0_6, negated_conjecture, (X1!=xb|~aElement0(X1)|~aReductOfIn0(X1,xa,xR)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, hypothesis, (aElement0(esk20_0)|aReductOfIn0(xb,xa,xR)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, hypothesis, (aElement0(xb)), inference(split_conjunct,[status(thm)],[m__731])).
cnf(c_0_9, negated_conjecture, (~sdtmndtplgtdt0(X1,xR,xb)|~aElement0(X1)|~aReductOfIn0(X1,xa,xR)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_10, hypothesis, (aReductOfIn0(esk20_0,xa,xR)|aReductOfIn0(xb,xa,xR)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_11, negated_conjecture, (aElement0(esk20_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_6, c_0_7]), c_0_8])])).
cnf(c_0_12, negated_conjecture, (aReductOfIn0(xb,xa,xR)|~sdtmndtplgtdt0(esk20_0,xR,xb)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_9, c_0_10]), c_0_11])])).
cnf(c_0_13, hypothesis, (sdtmndtplgtdt0(esk20_0,xR,xb)|aReductOfIn0(xb,xa,xR)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14, negated_conjecture, (~sdtmndtplgtdt0(esk20_0,xR,xb)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_6, c_0_12]), c_0_8])])).
cnf(c_0_15, hypothesis, (aReductOfIn0(xb,xa,xR)), inference(sr,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_16, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_6, c_0_15]), c_0_8])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 17
# Proof object clause steps            : 11
# Proof object formula steps           : 6
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 6
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 9
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 20
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 351
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 347
# Processed clauses                    : 57
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 53
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 4
# Generated clauses                    : 52
# ...of the previous two non-trivial   : 33
# Contextual simplify-reflections      : 0
# Paramodulations                      : 52
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 46
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 31
# Current number of unprocessed clauses: 319
# ...number of literals in the above   : 2745
# Current number of archived formulas  : 0
# Current number of archived clauses   : 7
# Clause-clause subsumption calls (NU) : 117
# Rec. Clause-clause subsumption calls : 77
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 51
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 5
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 20370
