# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_62, negated_conjecture, (aElement0(esk13_0))).
#
#cnf(i_0_47, hypothesis, (aRewritingSystem0(xR))).
#
#cnf(i_0_43, hypothesis, (isTerminating0(xR))).
#
#cnf(i_0_48, negated_conjecture, (~aNormalFormOfIn0(X1,esk13_0,xR))).
#
#cnf(i_0_21, plain, (isConfluent0(X1)|aElement0(esk3_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_20, plain, (isConfluent0(X1)|aElement0(esk4_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_19, plain, (isConfluent0(X1)|aElement0(esk5_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_55, negated_conjecture, (~iLess0(X2,esk13_0)|~aReductOfIn0(X1,esk14_1(X2),xR)|~aElement0(X2))).
#
#cnf(i_0_30, plain, (isLocallyConfluent0(X1)|aElement0(esk7_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_29, plain, (isLocallyConfluent0(X1)|aElement0(esk8_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_28, plain, (isLocallyConfluent0(X1)|aElement0(esk9_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_40, plain, (~aNormalFormOfIn0(X2,X4,X3)|~aReductOfIn0(X1,X2,X3)|~aRewritingSystem0(X3)|~aElement0(X4))).
#
#cnf(i_0_37, plain, (isTerminating0(X1)|aElement0(esk10_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_36, plain, (isTerminating0(X1)|aElement0(esk11_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_61, negated_conjecture, (aElement0(esk14_1(X1))|~iLess0(X1,esk13_0)|~aElement0(X1))).
#
#cnf(i_0_3, plain, (aElement0(X3)|~aReductOfIn0(X3,X1,X2)|~aRewritingSystem0(X2)|~aElement0(X1))).
#
#cnf(i_0_42, plain, (aElement0(X1)|~aNormalFormOfIn0(X1,X2,X3)|~aRewritingSystem0(X3)|~aElement0(X2))).
#
#cnf(i_0_34, plain, (isTerminating0(X1)|~iLess0(esk11_1(X1),esk10_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_46, hypothesis, (iLess0(X1,X2)|~aReductOfIn0(X1,X2,xR)|~aElement0(X2)|~aElement0(X1))).
#
#cnf(i_0_18, plain, (isConfluent0(X1)|sdtmndtasgtdt0(esk3_1(X1),X1,esk4_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_44, hypothesis, (iLess0(X2,X1)|~sdtmndtplgtdt0(X1,xR,X2)|~aElement0(X2)|~aElement0(X1))).
#
#cnf(i_0_41, plain, (sdtmndtasgtdt0(X1,X2,X3)|~aNormalFormOfIn0(X3,X1,X2)|~aRewritingSystem0(X2)|~aElement0(X1))).
#
#cnf(i_0_53, negated_conjecture, (aReductOfIn0(esk16_1(X1),X1,xR)|esk13_0!=X1|~aElement0(X1))).
#
#cnf(i_0_66, negated_conjecture, (esk13_0!=X1|~aNormalFormOfIn0(X1,X2,xR)|~aElement0(X2)|~aElement0(X1))).
#
#cnf(i_0_64, negated_conjecture, (aElement0(esk16_1(X1))|esk13_0!=X1|~aElement0(X1))).
#
#cnf(i_0_17, plain, (isConfluent0(X1)|sdtmndtasgtdt0(esk3_1(X1),X1,esk5_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_35, plain, (isTerminating0(X1)|sdtmndtplgtdt0(esk10_1(X1),X1,esk11_1(X1))|~aRewritingSystem0(X1))).
#
#cnf(i_0_63, negated_conjecture, (esk14_1(X1)!=esk13_0|~iLess0(X1,esk13_0)|~aElement0(esk14_1(X1))|~aElement0(X1))).
#
#cnf(i_0_27, plain, (isLocallyConfluent0(X1)|aReductOfIn0(esk8_1(X1),esk7_1(X1),X1)|~aRewritingSystem0(X1))).
#
#cnf(i_0_71, plain, (isLocallyConfluent0(X1)|~aNormalFormOfIn0(esk7_1(X1),X2,X1)|~aRewritingSystem0(X1)|~aElement0(X2))).
#
#cnf(i_0_26, plain, (isLocallyConfluent0(X1)|aReductOfIn0(esk9_1(X1),esk7_1(X1),X1)|~aRewritingSystem0(X1))).
#
#cnf(i_0_68, negated_conjecture, (esk14_1(X1)!=esk13_0|~iLess0(X1,esk13_0)|~aElement0(X1))).
####
#cnf(i_0_56, negated_conjecture, (sdtmndtasgtdt0(X1,xR,esk14_1(X1))|~iLess0(X1,esk13_0)|~aElement0(X1))).
#
#cnf(i_0_54, negated_conjecture, (aNormalFormOfIn0(esk14_1(X1),X1,xR)|~iLess0(X1,esk13_0)|~aElement0(X1))).
###
#cnf(i_0_75, negated_conjecture, (~iLess0(esk13_0,esk13_0))).
#
#cnf(i_0_13, plain, (sdtmndtasgtdt0(X1,X3,X2)|X1!=X2|~aRewritingSystem0(X3)|~aElement0(X2)|~aElement0(X1))).
#
#cnf(i_0_65, hypothesis, (iLess0(esk16_1(X1),X1)|esk13_0!=X1|~aElement0(esk16_1(X1))|~aElement0(X1))).
#
#cnf(i_0_52, negated_conjecture, (aReductOfIn0(esk16_1(X1),X1,xR)|~aReductOfIn0(X1,esk13_0,xR)|~aElement0(X1))).
#
#cnf(i_0_80, negated_conjecture, (esk14_1(esk16_1(esk13_0))!=esk13_0|~aElement0(esk16_1(esk13_0)))).
#
#cnf(i_0_82, negated_conjecture, (aElement0(esk16_1(X1))|~aReductOfIn0(X1,esk13_0,xR)|~aElement0(X1))).
#
#cnf(i_0_85, negated_conjecture, (aElement0(esk16_1(esk16_1(esk13_0)))|~aElement0(esk16_1(esk13_0)))).
#
#cnf(i_0_50, negated_conjecture, (aReductOfIn0(esk16_1(X1),X1,xR)|~sdtmndtplgtdt0(esk13_0,xR,X1)|~aElement0(X1))).
##
#cnf(i_0_89, negated_conjecture, (aElement0(esk16_1(X1))|~sdtmndtplgtdt0(esk13_0,xR,X1)|~aElement0(X1))).
#
#cnf(i_0_49, negated_conjecture, (aReductOfIn0(esk16_1(X1),X1,xR)|~sdtmndtasgtdt0(esk13_0,xR,X1)|~aElement0(X1))).
#
#cnf(i_0_94, negated_conjecture, (aElement0(esk16_1(X1))|~sdtmndtasgtdt0(esk13_0,xR,X1)|~aElement0(X1))).
###
#cnf(i_0_7, plain, (sdtmndtplgtdt0(X2,X3,X1)|~aReductOfIn0(X1,X2,X3)|~aRewritingSystem0(X3)|~aElement0(X2)|~aElement0(X1))).
#
#cnf(i_0_12, plain, (sdtmndtasgtdt0(X1,X2,X3)|~sdtmndtplgtdt0(X1,X2,X3)|~aRewritingSystem0(X2)|~aElement0(X3)|~aElement0(X1))).
#
#cnf(i_0_84, negated_conjecture, (~aNormalFormOfIn0(X1,X2,xR)|~aReductOfIn0(X1,esk13_0,xR)|~aElement0(X2)|~aElement0(X1))).
##
#cnf(i_0_38, plain, (iLess0(X3,X2)|~isTerminating0(X1)|~sdtmndtplgtdt0(X2,X1,X3)|~aRewritingSystem0(X1)|~aElement0(X3)|~aElement0(X2))).
#
#cnf(i_0_99, negated_conjecture, (sdtmndtplgtdt0(X1,xR,esk16_1(X1))|esk13_0!=X1|~aElement0(esk16_1(X1))|~aElement0(X1))).
#
#cnf(i_0_91, negated_conjecture, (~aNormalFormOfIn0(X1,X2,xR)|~sdtmndtplgtdt0(esk13_0,xR,X1)|~aElement0(X2)|~aElement0(X1))).
####
#cnf(i_0_96, negated_conjecture, (~aNormalFormOfIn0(X1,X2,xR)|~sdtmndtasgtdt0(esk13_0,xR,X1)|~aElement0(X2)|~aElement0(X1))).
#
#cnf(i_0_109, negated_conjecture, (sdtmndtasgtdt0(X1,xR,esk16_1(X1))|esk13_0!=X1|~aElement0(esk16_1(X1))|~aElement0(X1))).
##
#cnf(i_0_83, hypothesis, (iLess0(esk16_1(X1),X1)|~aReductOfIn0(X1,esk13_0,xR)|~aElement0(esk16_1(X1))|~aElement0(X1))).
#
#cnf(i_0_81, negated_conjecture, (~iLess0(X1,esk13_0)|~aReductOfIn0(esk14_1(X1),esk13_0,xR)|~aElement0(esk14_1(X1))|~aElement0(X1))).
#
#cnf(i_0_90, hypothesis, (iLess0(esk16_1(X1),X1)|~sdtmndtplgtdt0(esk13_0,xR,X1)|~aElement0(esk16_1(X1))|~aElement0(X1))).
#
#cnf(i_0_95, hypothesis, (iLess0(esk16_1(X1),X1)|~sdtmndtasgtdt0(esk13_0,xR,X1)|~aElement0(esk16_1(X1))|~aElement0(X1))).
##
#cnf(i_0_87, negated_conjecture, (~sdtmndtplgtdt0(esk13_0,xR,esk14_1(X1))|~iLess0(X1,esk13_0)|~aElement0(esk14_1(X1))|~aElement0(X1))).
#
#cnf(i_0_57, negated_conjecture, (esk14_1(X1)=X1|sdtmndtplgtdt0(X1,xR,esk14_1(X1))|~iLess0(X1,esk13_0)|~aElement0(X1))).
###
#cnf(i_0_92, negated_conjecture, (~sdtmndtasgtdt0(esk13_0,xR,esk14_1(X1))|~iLess0(X1,esk13_0)|~aElement0(esk14_1(X1))|~aElement0(X1))).
###
#cnf(i_0_14, plain, (X1=X3|sdtmndtplgtdt0(X1,X2,X3)|~sdtmndtasgtdt0(X1,X2,X3)|~aRewritingSystem0(X2)|~aElement0(X3)|~aElement0(X1))).
##
#cnf(i_0_126, hypothesis, (X2=X1|iLess0(X1,X2)|~sdtmndtasgtdt0(X2,xR,X1)|~aElement0(X1)|~aElement0(X2))).
##
#cnf(i_0_16, plain, (isConfluent0(X2)|~sdtmndtasgtdt0(esk4_1(X2),X2,X1)|~sdtmndtasgtdt0(esk5_1(X2),X2,X1)|~aRewritingSystem0(X2)|~aElement0(X1))).
##
#cnf(i_0_25, plain, (isLocallyConfluent0(X2)|~sdtmndtasgtdt0(esk8_1(X2),X2,X1)|~sdtmndtasgtdt0(esk9_1(X2),X2,X1)|~aRewritingSystem0(X2)|~aElement0(X1))).
#
#cnf(i_0_100, negated_conjecture, (sdtmndtplgtdt0(X1,xR,esk16_1(X1))|~aReductOfIn0(X1,esk13_0,xR)|~aElement0(esk16_1(X1))|~aElement0(X1))).
#####
#cnf(i_0_101, negated_conjecture, (sdtmndtplgtdt0(X1,xR,esk16_1(X1))|~sdtmndtplgtdt0(esk13_0,xR,X1)|~aElement0(esk16_1(X1))|~aElement0(X1))).
######
#cnf(i_0_102, negated_conjecture, (sdtmndtplgtdt0(X1,xR,esk16_1(X1))|~sdtmndtasgtdt0(esk13_0,xR,X1)|~aElement0(esk16_1(X1))|~aElement0(X1))).
####
#cnf(i_0_141, negated_conjecture, (sdtmndtasgtdt0(X1,xR,esk16_1(X1))|~aReductOfIn0(X1,esk13_0,xR)|~aElement0(esk16_1(X1))|~aElement0(X1))).
##
#cnf(i_0_145, negated_conjecture, (sdtmndtasgtdt0(X1,xR,esk16_1(X1))|~sdtmndtplgtdt0(esk13_0,xR,X1)|~aElement0(esk16_1(X1))|~aElement0(X1))).
##
#cnf(i_0_149, negated_conjecture, (sdtmndtasgtdt0(X1,xR,esk16_1(X1))|~sdtmndtasgtdt0(esk13_0,xR,X1)|~aElement0(esk16_1(X1))|~aElement0(X1))).
###
#cnf(i_0_119, negated_conjecture, (esk14_1(X1)=X1|iLess0(esk14_1(X1),X1)|~iLess0(X1,esk13_0)|~aElement0(esk14_1(X1))|~aElement0(X1))).
#
#cnf(i_0_128, plain, (X2=X1|iLess0(X1,X2)|~isTerminating0(X3)|~sdtmndtasgtdt0(X2,X3,X1)|~aRewritingSystem0(X3)|~aElement0(X1)|~aElement0(X2))).
#####
#cnf(i_0_45, hypothesis, (iLess0(X3,X2)|~sdtmndtplgtdt0(X1,xR,X3)|~aReductOfIn0(X1,X2,xR)|~aElement0(X3)|~aElement0(X2)|~aElement0(X1))).
#
#cnf(i_0_103, plain, (isLocallyConfluent0(X1)|sdtmndtplgtdt0(esk7_1(X1),X1,esk9_1(X1))|~aRewritingSystem0(X1)|~aElement0(esk7_1(X1))|~aElement0(esk9_1(X1)))).
#
#cnf(i_0_104, plain, (isLocallyConfluent0(X1)|sdtmndtplgtdt0(esk7_1(X1),X1,esk8_1(X1))|~aRewritingSystem0(X1)|~aElement0(esk7_1(X1))|~aElement0(esk8_1(X1)))).
#
#cnf(i_0_105, plain, (isTerminating0(X1)|sdtmndtasgtdt0(esk10_1(X1),X1,esk11_1(X1))|~aRewritingSystem0(X1)|~aElement0(esk11_1(X1))|~aElement0(esk10_1(X1)))).
#
#cnf(i_0_189, plain, (isLocallyConfluent0(X1)|sdtmndtasgtdt0(esk7_1(X1),X1,esk9_1(X1))|~aRewritingSystem0(X1)|~aElement0(esk9_1(X1))|~aElement0(esk7_1(X1)))).
#
#cnf(i_0_193, plain, (isLocallyConfluent0(X1)|sdtmndtasgtdt0(esk7_1(X1),X1,esk8_1(X1))|~aRewritingSystem0(X1)|~aElement0(esk8_1(X1))|~aElement0(esk7_1(X1)))).
##
#cnf(i_0_136, plain, (isConfluent0(X1)|esk5_1(X1)!=X2|~sdtmndtasgtdt0(esk4_1(X1),X1,X2)|~aRewritingSystem0(X1)|~aElement0(esk5_1(X1))|~aElement0(X2))).
#
#cnf(i_0_139, plain, (isLocallyConfluent0(X1)|esk9_1(X1)!=X2|~sdtmndtasgtdt0(esk8_1(X1),X1,X2)|~aRewritingSystem0(X1)|~aElement0(esk9_1(X1))|~aElement0(X2))).
#
#cnf(i_0_10, plain, (aReductOfIn0(X3,X1,X2)|aElement0(esk1_3(X1,X2,X3))|~sdtmndtplgtdt0(X1,X2,X3)|~aRewritingSystem0(X2)|~aElement0(X3)|~aElement0(X1))).
#
#cnf(i_0_60, negated_conjecture, (esk14_1(X1)=X1|aReductOfIn0(esk14_1(X1),X1,xR)|aElement0(esk15_1(X1))|~iLess0(X1,esk13_0)|~aElement0(X1))).
####
#cnf(i_0_70, hypothesis, (isLocallyConfluent0(xR)|iLess0(esk8_1(xR),esk7_1(xR))|~aElement0(esk7_1(xR))|~aElement0(esk8_1(xR)))).
####
#cnf(i_0_73, hypothesis, (isLocallyConfluent0(xR)|iLess0(esk9_1(xR),esk7_1(xR))|~aElement0(esk7_1(xR))|~aElement0(esk9_1(xR)))).
######
#cnf(i_0_174, negated_conjecture, (iLess0(X1,X2)|esk13_0!=X2|~sdtmndtplgtdt0(esk16_1(X2),xR,X1)|~aElement0(esk16_1(X2))|~aElement0(X1)|~aElement0(X2))).
#
#cnf(i_0_182, negated_conjecture, (iLess0(esk16_1(X1),X2)|esk13_0!=X1|~aReductOfIn0(X1,X2,xR)|~aElement0(esk16_1(X1))|~aElement0(X2)|~aElement0(X1))).
##
#cnf(i_0_219, negated_conjecture, (esk14_1(X1)=X1|aElement0(esk15_1(X1))|~aNormalFormOfIn0(X1,X2,xR)|~iLess0(X1,esk13_0)|~aElement0(X2)|~aElement0(X1))).
##
#cnf(i_0_190, plain, (isLocallyConfluent0(X1)|iLess0(esk9_1(X1),esk7_1(X1))|~isTerminating0(X1)|~aRewritingSystem0(X1)|~aElement0(esk9_1(X1))|~aElement0(esk7_1(X1)))).
#
#cnf(i_0_194, plain, (isLocallyConfluent0(X1)|iLess0(esk8_1(X1),esk7_1(X1))|~isTerminating0(X1)|~aRewritingSystem0(X1)|~aElement0(esk8_1(X1))|~aElement0(esk7_1(X1)))).
#
#cnf(i_0_6, plain, (sdtmndtplgtdt0(X2,X3,X4)|~sdtmndtplgtdt0(X1,X3,X4)|~aReductOfIn0(X1,X2,X3)|~aRewritingSystem0(X3)|~aElement0(X4)|~aElement0(X2)|~aElement0(X1))).
#
#cnf(i_0_237, negated_conjecture, (aElement0(esk16_1(X1))|~sdtmndtplgtdt0(X2,xR,X1)|~aReductOfIn0(X2,esk13_0,xR)|~aElement0(X1)|~aElement0(X2))).
#
#cnf(i_0_271, negated_conjecture, (aElement0(esk16_1(esk16_1(X1)))|~aReductOfIn0(X1,esk13_0,xR)|~aElement0(esk16_1(X1))|~aElement0(X1))).
#
#cnf(i_0_236, negated_conjecture, (~sdtmndtplgtdt0(X2,xR,esk14_1(X1))|~iLess0(X1,esk13_0)|~aReductOfIn0(X2,esk13_0,xR)|~aElement0(esk14_1(X1))|~aElement0(X1)|~aElement0(X2))).
#
#cnf(i_0_260, negated_conjecture, (aElement0(esk16_1(X1))|~sdtmndtplgtdt0(esk16_1(esk13_0),xR,X1)|~aElement0(esk16_1(esk13_0))|~aElement0(X1))).
#
#cnf(i_0_277, negated_conjecture, (esk14_1(X1)=X1|~iLess0(X1,esk13_0)|~aReductOfIn0(X1,esk13_0,xR)|~aElement0(esk14_1(X1))|~aElement0(X1))).
##
#cnf(i_0_275, negated_conjecture, (~sdtmndtplgtdt0(X3,xR,esk14_1(X2))|~iLess0(X2,esk13_0)|~aReductOfIn0(X1,esk13_0,xR)|~aReductOfIn0(X3,X1,xR)|~aElement0(esk14_1(X2))|~aElement0(X2)|~aElement0(X1)|~aElement0(X3))).
##
#cnf(i_0_11, plain, (sdtmndtplgtdt0(X1,X2,X4)|~sdtmndtplgtdt0(X3,X2,X4)|~sdtmndtplgtdt0(X1,X2,X3)|~aRewritingSystem0(X2)|~aElement0(X4)|~aElement0(X3)|~aElement0(X1))).
#
#cnf(i_0_308, negated_conjecture, (aElement0(esk16_1(X1))|~sdtmndtplgtdt0(esk13_0,xR,X2)|~sdtmndtplgtdt0(X2,xR,X1)|~aElement0(X1)|~aElement0(X2))).
#
#cnf(i_0_307, negated_conjecture, (~sdtmndtplgtdt0(X2,xR,esk14_1(X1))|~sdtmndtplgtdt0(esk13_0,xR,X2)|~iLess0(X1,esk13_0)|~aElement0(esk14_1(X1))|~aElement0(X1)|~aElement0(X2))).
#
#cnf(i_0_348, negated_conjecture, (aElement0(esk16_1(esk16_1(X1)))|~sdtmndtplgtdt0(esk13_0,xR,X1)|~aElement0(esk16_1(X1))|~aElement0(X1))).
#
#cnf(i_0_357, negated_conjecture, (esk14_1(X1)=X1|~sdtmndtplgtdt0(esk13_0,xR,X1)|~iLess0(X1,esk13_0)|~aElement0(esk14_1(X1))|~aElement0(X1))).
#
#cnf(i_0_310, hypothesis, (iLess0(X1,X2)|~sdtmndtplgtdt0(X3,xR,X1)|~sdtmndtplgtdt0(X2,xR,X3)|~aElement0(X1)|~aElement0(X2)|~aElement0(X3))).
#
#cnf(i_0_291, negated_conjecture, (~sdtmndtplgtdt0(esk16_1(X1),xR,esk14_1(X2))|~iLess0(X2,esk13_0)|~aReductOfIn0(X1,esk13_0,xR)|~aElement0(esk14_1(X2))|~aElement0(esk16_1(X1))|~aElement0(X2)|~aElement0(X1))).
###
#cnf(i_0_384, negated_conjecture, (iLess0(esk16_1(X1),X2)|esk13_0!=X1|~sdtmndtplgtdt0(X2,xR,X1)|~aElement0(esk16_1(X1))|~aElement0(X2)|~aElement0(X1))).
#
#cnf(i_0_297, negated_conjecture, (~sdtmndtplgtdt0(X1,xR,esk14_1(X2))|~iLess0(X2,esk13_0)|~aReductOfIn0(X1,esk16_1(esk13_0),xR)|~aElement0(esk16_1(esk13_0))|~aElement0(esk14_1(X2))|~aElement0(X2)|~aElement0(X1))).
#
#cnf(i_0_15, plain, (sdtmndtasgtdt0(X1,X2,X4)|~sdtmndtasgtdt0(X3,X2,X4)|~sdtmndtasgtdt0(X1,X2,X3)|~aRewritingSystem0(X2)|~aElement0(X4)|~aElement0(X3)|~aElement0(X1))).
#
#cnf(i_0_414, negated_conjecture, (aElement0(esk16_1(X1))|~sdtmndtasgtdt0(esk13_0,xR,X2)|~sdtmndtasgtdt0(X2,xR,X1)|~aElement0(X1)|~aElement0(X2))).
#
#cnf(i_0_458, negated_conjecture, (aElement0(esk16_1(X1))|esk13_0!=X2|~sdtmndtasgtdt0(X2,xR,X1)|~aElement0(X1)|~aElement0(X2))).
#
#cnf(i_0_413, negated_conjecture, (~sdtmndtasgtdt0(X2,xR,esk14_1(X1))|~sdtmndtasgtdt0(esk13_0,xR,X2)|~iLess0(X1,esk13_0)|~aElement0(esk14_1(X1))|~aElement0(X1)|~aElement0(X2))).
#
#cnf(i_0_465, negated_conjecture, (aElement0(esk16_1(X1))|esk13_0!=X2|X2!=X1|~aElement0(X1)|~aElement0(X2))).
##
#cnf(i_0_469, negated_conjecture, (aElement0(esk16_1(esk16_1(X1)))|esk13_0!=X1|~aElement0(esk16_1(X1))|~aElement0(X1))).
#
#cnf(i_0_474, negated_conjecture, (~sdtmndtasgtdt0(esk13_0,xR,X1)|~iLess0(X1,esk13_0)|~aElement0(esk14_1(X1))|~aElement0(X1))).
#
#cnf(i_0_451, negated_conjecture, (aElement0(esk16_1(X1))|X2!=X1|~sdtmndtasgtdt0(esk13_0,xR,X2)|~aElement0(X1)|~aElement0(X2))).
###
#cnf(i_0_492, negated_conjecture, (esk13_0!=X1|~iLess0(X1,esk13_0)|~aElement0(esk14_1(X1))|~aElement0(X1))).
#
#cnf(i_0_505, negated_conjecture, (aElement0(esk16_1(X1))|esk16_1(esk13_0)!=X1|~aElement0(esk16_1(esk13_0))|~aElement0(X1))).
#
#cnf(i_0_510, negated_conjecture, (aElement0(esk16_1(X1))|esk16_1(esk13_0)!=X1|~aElement0(X1))).
#
#cnf(i_0_453, negated_conjecture, (aElement0(esk16_1(esk16_1(X1)))|~sdtmndtasgtdt0(esk13_0,xR,X1)|~aElement0(esk16_1(X1))|~aElement0(X1))).
#
#cnf(i_0_508, negated_conjecture, (esk13_0!=X1|~iLess0(X1,esk13_0)|~aElement0(X1))).
####
#cnf(i_0_514, hypothesis, (esk16_1(esk13_0)!=esk13_0|~aElement0(esk16_1(esk13_0)))).
#
#cnf(i_0_482, negated_conjecture, (isTerminating0(X2)|aElement0(esk16_1(X1))|esk11_1(X2)!=esk13_0|esk11_1(X2)!=X1|~aRewritingSystem0(X2)|~aElement0(X1))).
#
#cnf(i_0_521, hypothesis, (isTerminating0(X1)|esk16_1(esk13_0)!=esk13_0|esk11_1(X1)!=esk13_0|~aRewritingSystem0(X1))).
##
#cnf(i_0_516, negated_conjecture, (esk16_1(esk13_0)!=esk13_0)).
#
#cnf(i_0_483, negated_conjecture, (isTerminating0(X2)|aElement0(esk16_1(X1))|esk10_1(X2)!=esk13_0|esk10_1(X2)!=X1|~aRewritingSystem0(X2)|~aElement0(X1))).
####
#cnf(i_0_484, negated_conjecture, (isLocallyConfluent0(X2)|aElement0(esk16_1(X1))|esk9_1(X2)!=esk13_0|esk9_1(X2)!=X1|~aRewritingSystem0(X2)|~aElement0(X1))).
####
#cnf(i_0_485, negated_conjecture, (isLocallyConfluent0(X2)|aElement0(esk16_1(X1))|esk8_1(X2)!=esk13_0|esk8_1(X2)!=X1|~aRewritingSystem0(X2)|~aElement0(X1))).
###
#cnf(i_0_496, negated_conjecture, (~iLess0(esk16_1(esk13_0),esk13_0)|~aElement0(esk14_1(esk16_1(esk13_0)))|~aElement0(esk16_1(esk13_0)))).
#
#cnf(i_0_486, negated_conjecture, (isLocallyConfluent0(X2)|aElement0(esk16_1(X1))|esk7_1(X2)!=esk13_0|esk7_1(X2)!=X1|~aRewritingSystem0(X2)|~aElement0(X1))).
###
#cnf(i_0_531, negated_conjecture, (~iLess0(esk16_1(esk13_0),esk13_0)|~aElement0(esk16_1(esk13_0)))).
#
#cnf(i_0_487, negated_conjecture, (isConfluent0(X2)|aElement0(esk16_1(X1))|esk5_1(X2)!=esk13_0|esk5_1(X2)!=X1|~aRewritingSystem0(X2)|~aElement0(X1))).
###
#cnf(i_0_535, hypothesis, (~aElement0(esk16_1(esk13_0)))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(m__, conjecture, ![X1]:(aElement0(X1)=>(![X2]:(aElement0(X2)=>(iLess0(X2,X1)=>?[X3]:((((aElement0(X3)&(X2=X3|((aReductOfIn0(X3,X2,xR)|?[X4]:((aElement0(X4)&aReductOfIn0(X4,X2,xR))&sdtmndtplgtdt0(X4,xR,X3)))&sdtmndtplgtdt0(X2,xR,X3))))&sdtmndtasgtdt0(X2,xR,X3))&~(?[X4]:aReductOfIn0(X4,X3,xR)))&aNormalFormOfIn0(X3,X2,xR))))=>?[X2]:(((aElement0(X2)&(X1=X2|aReductOfIn0(X2,X1,xR)|?[X3]:((aElement0(X3)&aReductOfIn0(X3,X1,xR))&sdtmndtplgtdt0(X3,xR,X2))|sdtmndtplgtdt0(X1,xR,X2)|sdtmndtasgtdt0(X1,xR,X2)))&~(?[X3]:aReductOfIn0(X3,X2,xR)))|aNormalFormOfIn0(X2,X1,xR)))), file('./p/COM013+4.tptp', m__)).
fof(mTCRTrans, axiom, ![X1, X2, X3, X4]:((((aElement0(X1)&aRewritingSystem0(X2))&aElement0(X3))&aElement0(X4))=>((sdtmndtasgtdt0(X1,X2,X3)&sdtmndtasgtdt0(X3,X2,X4))=>sdtmndtasgtdt0(X1,X2,X4))), file('./p/COM013+4.tptp', mTCRTrans)).
fof(m__587, hypothesis, ((aRewritingSystem0(xR)&![X1, X2]:((aElement0(X1)&aElement0(X2))=>((aReductOfIn0(X2,X1,xR)|?[X3]:((aElement0(X3)&aReductOfIn0(X3,X1,xR))&sdtmndtplgtdt0(X3,xR,X2))|sdtmndtplgtdt0(X1,xR,X2))=>iLess0(X2,X1))))&isTerminating0(xR)), file('./p/COM013+4.tptp', m__587)).
fof(mTCDef, axiom, ![X1, X2, X3]:(((aElement0(X1)&aRewritingSystem0(X2))&aElement0(X3))=>(sdtmndtplgtdt0(X1,X2,X3)<=>(aReductOfIn0(X3,X1,X2)|?[X4]:((aElement0(X4)&aReductOfIn0(X4,X1,X2))&sdtmndtplgtdt0(X4,X2,X3))))), file('./p/COM013+4.tptp', mTCDef)).
fof(mTCRDef, axiom, ![X1, X2, X3]:(((aElement0(X1)&aRewritingSystem0(X2))&aElement0(X3))=>(sdtmndtasgtdt0(X1,X2,X3)<=>(X1=X3|sdtmndtplgtdt0(X1,X2,X3)))), file('./p/COM013+4.tptp', mTCRDef)).
fof(mReduct, axiom, ![X1, X2]:((aElement0(X1)&aRewritingSystem0(X2))=>![X3]:(aReductOfIn0(X3,X1,X2)=>aElement0(X3))), file('./p/COM013+4.tptp', mReduct)).
fof(c_0_6, negated_conjecture, ~(![X1]:(aElement0(X1)=>(![X2]:(aElement0(X2)=>(iLess0(X2,X1)=>?[X3]:((((aElement0(X3)&(X2=X3|((aReductOfIn0(X3,X2,xR)|?[X4]:((aElement0(X4)&aReductOfIn0(X4,X2,xR))&sdtmndtplgtdt0(X4,xR,X3)))&sdtmndtplgtdt0(X2,xR,X3))))&sdtmndtasgtdt0(X2,xR,X3))&~(?[X4]:aReductOfIn0(X4,X3,xR)))&aNormalFormOfIn0(X3,X2,xR))))=>?[X2]:(((aElement0(X2)&(X1=X2|aReductOfIn0(X2,X1,xR)|?[X3]:((aElement0(X3)&aReductOfIn0(X3,X1,xR))&sdtmndtplgtdt0(X3,xR,X2))|sdtmndtplgtdt0(X1,xR,X2)|sdtmndtasgtdt0(X1,xR,X2)))&~(?[X3]:aReductOfIn0(X3,X2,xR)))|aNormalFormOfIn0(X2,X1,xR))))), inference(assume_negation,[status(cth)],[m__])).
fof(c_0_7, negated_conjecture, ![X58, X61, X62, X63]:(aElement0(esk13_0)&((((((aElement0(esk14_1(X58))|~iLess0(X58,esk13_0)|~aElement0(X58))&((((aElement0(esk15_1(X58))|aReductOfIn0(esk14_1(X58),X58,xR)|X58=esk14_1(X58)|~iLess0(X58,esk13_0)|~aElement0(X58))&(aReductOfIn0(esk15_1(X58),X58,xR)|aReductOfIn0(esk14_1(X58),X58,xR)|X58=esk14_1(X58)|~iLess0(X58,esk13_0)|~aElement0(X58)))&(sdtmndtplgtdt0(esk15_1(X58),xR,esk14_1(X58))|aReductOfIn0(esk14_1(X58),X58,xR)|X58=esk14_1(X58)|~iLess0(X58,esk13_0)|~aElement0(X58)))&(sdtmndtplgtdt0(X58,xR,esk14_1(X58))|X58=esk14_1(X58)|~iLess0(X58,esk13_0)|~aElement0(X58))))&(sdtmndtasgtdt0(X58,xR,esk14_1(X58))|~iLess0(X58,esk13_0)|~aElement0(X58)))&(~aReductOfIn0(X61,esk14_1(X58),xR)|~iLess0(X58,esk13_0)|~aElement0(X58)))&(aNormalFormOfIn0(esk14_1(X58),X58,xR)|~iLess0(X58,esk13_0)|~aElement0(X58)))&((((((esk13_0!=X62|~aElement0(X62)|aReductOfIn0(esk16_1(X62),X62,xR))&(~aReductOfIn0(X62,esk13_0,xR)|~aElement0(X62)|aReductOfIn0(esk16_1(X62),X62,xR)))&(~aElement0(X63)|~aReductOfIn0(X63,esk13_0,xR)|~sdtmndtplgtdt0(X63,xR,X62)|~aElement0(X62)|aReductOfIn0(esk16_1(X62),X62,xR)))&(~sdtmndtplgtdt0(esk13_0,xR,X62)|~aElement0(X62)|aReductOfIn0(esk16_1(X62),X62,xR)))&(~sdtmndtasgtdt0(esk13_0,xR,X62)|~aElement0(X62)|aReductOfIn0(esk16_1(X62),X62,xR)))&~aNormalFormOfIn0(X62,esk13_0,xR)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])])])).
cnf(c_0_8, negated_conjecture, (~aReductOfIn0(X1,esk14_1(X2),xR)|~iLess0(X2,esk13_0)|~aElement0(X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_9, negated_conjecture, (aReductOfIn0(esk16_1(X1),X1,xR)|~sdtmndtasgtdt0(esk13_0,xR,X1)|~aElement0(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X21, X22, X23, X24]:(~aElement0(X21)|~aRewritingSystem0(X22)|~aElement0(X23)|~aElement0(X24)|(~sdtmndtasgtdt0(X21,X22,X23)|~sdtmndtasgtdt0(X23,X22,X24)|sdtmndtasgtdt0(X21,X22,X24))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mTCRTrans])])).
fof(c_0_11, hypothesis, ![X54, X55, X56]:((aRewritingSystem0(xR)&(((~aReductOfIn0(X55,X54,xR)|iLess0(X55,X54)|(~aElement0(X54)|~aElement0(X55)))&(~aElement0(X56)|~aReductOfIn0(X56,X54,xR)|~sdtmndtplgtdt0(X56,xR,X55)|iLess0(X55,X54)|(~aElement0(X54)|~aElement0(X55))))&(~sdtmndtplgtdt0(X54,xR,X55)|iLess0(X55,X54)|(~aElement0(X54)|~aElement0(X55)))))&isTerminating0(xR)), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[m__587])])])])).
fof(c_0_12, plain, ![X9, X10, X11, X13]:((((aElement0(esk1_3(X9,X10,X11))|aReductOfIn0(X11,X9,X10)|~sdtmndtplgtdt0(X9,X10,X11)|(~aElement0(X9)|~aRewritingSystem0(X10)|~aElement0(X11)))&(aReductOfIn0(esk1_3(X9,X10,X11),X9,X10)|aReductOfIn0(X11,X9,X10)|~sdtmndtplgtdt0(X9,X10,X11)|(~aElement0(X9)|~aRewritingSystem0(X10)|~aElement0(X11))))&(sdtmndtplgtdt0(esk1_3(X9,X10,X11),X10,X11)|aReductOfIn0(X11,X9,X10)|~sdtmndtplgtdt0(X9,X10,X11)|(~aElement0(X9)|~aRewritingSystem0(X10)|~aElement0(X11))))&((~aReductOfIn0(X11,X9,X10)|sdtmndtplgtdt0(X9,X10,X11)|(~aElement0(X9)|~aRewritingSystem0(X10)|~aElement0(X11)))&(~aElement0(X13)|~aReductOfIn0(X13,X9,X10)|~sdtmndtplgtdt0(X13,X10,X11)|sdtmndtplgtdt0(X9,X10,X11)|(~aElement0(X9)|~aRewritingSystem0(X10)|~aElement0(X11))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mTCDef])])])])])).
cnf(c_0_13, negated_conjecture, (~sdtmndtasgtdt0(esk13_0,xR,esk14_1(X1))|~iLess0(X1,esk13_0)|~aElement0(esk14_1(X1))|~aElement0(X1)), inference(pm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_14, plain, (sdtmndtasgtdt0(X1,X2,X4)|~aElement0(X1)|~aRewritingSystem0(X2)|~aElement0(X3)|~aElement0(X4)|~sdtmndtasgtdt0(X1,X2,X3)|~sdtmndtasgtdt0(X3,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, hypothesis, (aRewritingSystem0(xR)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (aElement0(esk13_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_17, plain, ![X18, X19, X20]:((~sdtmndtasgtdt0(X18,X19,X20)|(X18=X20|sdtmndtplgtdt0(X18,X19,X20))|(~aElement0(X18)|~aRewritingSystem0(X19)|~aElement0(X20)))&((X18!=X20|sdtmndtasgtdt0(X18,X19,X20)|(~aElement0(X18)|~aRewritingSystem0(X19)|~aElement0(X20)))&(~sdtmndtplgtdt0(X18,X19,X20)|sdtmndtasgtdt0(X18,X19,X20)|(~aElement0(X18)|~aRewritingSystem0(X19)|~aElement0(X20))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mTCRDef])])])).
cnf(c_0_18, plain, (sdtmndtplgtdt0(X2,X3,X1)|~aReductOfIn0(X1,X2,X3)|~aElement0(X2)|~aRewritingSystem0(X3)|~aElement0(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_19, negated_conjecture, (aReductOfIn0(esk16_1(X1),X1,xR)|esk13_0!=X1|~aElement0(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_20, negated_conjecture, (~sdtmndtasgtdt0(X1,xR,esk14_1(X2))|~sdtmndtasgtdt0(esk13_0,xR,X1)|~iLess0(X2,esk13_0)|~aElement0(esk14_1(X2))|~aElement0(X2)|~aElement0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_13, c_0_14]), c_0_15]), c_0_16])])).
cnf(c_0_21, negated_conjecture, (sdtmndtasgtdt0(X1,xR,esk14_1(X1))|~iLess0(X1,esk13_0)|~aElement0(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_22, plain, (sdtmndtasgtdt0(X1,X2,X3)|~sdtmndtplgtdt0(X1,X2,X3)|~aElement0(X1)|~aRewritingSystem0(X2)|~aElement0(X3)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_23, negated_conjecture, (sdtmndtplgtdt0(X1,xR,esk16_1(X1))|esk13_0!=X1|~aElement0(esk16_1(X1))|~aElement0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_18, c_0_19]), c_0_15])])).
cnf(c_0_24, negated_conjecture, (~sdtmndtasgtdt0(esk13_0,xR,X1)|~iLess0(X1,esk13_0)|~aElement0(esk14_1(X1))|~aElement0(X1)), inference(pm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_25, negated_conjecture, (sdtmndtasgtdt0(X1,xR,esk16_1(X1))|esk13_0!=X1|~aElement0(esk16_1(X1))|~aElement0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_22, c_0_23]), c_0_15])])).
cnf(c_0_26, negated_conjecture, (~iLess0(esk16_1(esk13_0),esk13_0)|~aElement0(esk14_1(esk16_1(esk13_0)))|~aElement0(esk16_1(esk13_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_24, c_0_25]), c_0_16])])).
cnf(c_0_27, negated_conjecture, (aElement0(esk14_1(X1))|~iLess0(X1,esk13_0)|~aElement0(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_28, hypothesis, (iLess0(X1,X2)|~aReductOfIn0(X1,X2,xR)|~aElement0(X2)|~aElement0(X1)), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_29, plain, ![X6, X7, X8]:(~aElement0(X6)|~aRewritingSystem0(X7)|(~aReductOfIn0(X8,X6,X7)|aElement0(X8))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mReduct])])])).
cnf(c_0_30, negated_conjecture, (~iLess0(esk16_1(esk13_0),esk13_0)|~aElement0(esk16_1(esk13_0))), inference(pm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31, hypothesis, (iLess0(esk16_1(X1),X1)|esk13_0!=X1|~aElement0(esk16_1(X1))|~aElement0(X1)), inference(pm,[status(thm)],[c_0_28, c_0_19])).
cnf(c_0_32, plain, (aElement0(X3)|~aElement0(X1)|~aRewritingSystem0(X2)|~aReductOfIn0(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_29])).
cnf(c_0_33, hypothesis, (~aElement0(esk16_1(esk13_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_30, c_0_31]), c_0_16])])).
cnf(c_0_34, negated_conjecture, (aElement0(esk16_1(X1))|esk13_0!=X1|~aElement0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_32, c_0_19]), c_0_15])])).
cnf(c_0_35, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_33, c_0_34]), c_0_16])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 36
# Proof object clause steps            : 23
# Proof object formula steps           : 13
# Proof object conjectures             : 18
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 15
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 62
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 58
# Processed clauses                    : 208
# ...of these trivial                  : 2
# ...subsumed                          : 76
# ...remaining for further processing  : 130
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 4
# Backward-rewritten                   : 0
# Generated clauses                    : 485
# ...of the previous two non-trivial   : 460
# Contextual simplify-reflections      : 0
# Paramodulations                      : 477
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 126
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 119
# Current number of unprocessed clauses: 309
# ...number of literals in the above   : 2019
# Current number of archived formulas  : 0
# Current number of archived clauses   : 4
# Clause-clause subsumption calls (NU) : 704
# Rec. Clause-clause subsumption calls : 216
# Non-unit clause-clause subsumptions  : 67
# Unit Clause-clause subsumption calls : 91
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 12246
