# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_50, negated_conjecture, (distinct_points(esk1_0,esk2_0))).
#
#cnf(i_0_49, negated_conjecture, (convergent_lines(esk3_0,esk4_0))).
#
#cnf(i_0_15, plain, (convergent_lines(X1,X2)|~distinct_lines(X1,X2))).
#
#cnf(i_0_1, plain, (~distinct_points(X1,X1))).
#
#cnf(i_0_37, plain, (equal_points(X1,X2)|distinct_points(X1,X2))).
#
#cnf(i_0_39, plain, (equal_lines(X1,X2)|distinct_lines(X1,X2))).
#
#cnf(i_0_19, plain, (unorthogonal_lines(X1,X2)|convergent_lines(X1,X2))).
#
#cnf(i_0_2, plain, (~distinct_lines(X1,X1))).
##
#cnf(i_0_41, plain, (parallel_lines(X1,X2)|convergent_lines(X1,X2))).
#
#cnf(i_0_43, plain, (incident_point_and_line(X1,X2)|apart_point_and_line(X1,X2))).
#
#cnf(i_0_3, plain, (~convergent_lines(X1,X1))).
#
#cnf(i_0_45, plain, (orthogonal_lines(X1,X2)|unorthogonal_lines(X1,X2))).
#
#cnf(i_0_35, plain, (line(parallel_through_point(X1,X2))|~line(X1)|~point(X2))).
#
#cnf(i_0_36, plain, (line(orthogonal_through_point(X1,X2))|~line(X1)|~point(X2))).
#
#cnf(i_0_17, plain, (~apart_point_and_line(X1,parallel_through_point(X2,X1)))).
#
#cnf(i_0_48, negated_conjecture, (apart_point_and_line(esk1_0,esk3_0)|apart_point_and_line(esk2_0,esk3_0))).
#
#cnf(i_0_4, plain, (distinct_points(X2,X3)|distinct_points(X1,X3)|~distinct_points(X1,X2))).
##
#cnf(i_0_25, plain, (~apart_point_and_line(X1,orthogonal_through_point(X2,X1)))).
##
#cnf(i_0_53, plain, (distinct_points(X2,X1)|~distinct_points(X1,X2))).
#
#cnf(i_0_60, negated_conjecture, (distinct_points(esk2_0,esk1_0))).
#
#cnf(i_0_16, plain, (~convergent_lines(parallel_through_point(X1,X2),X1))).
#
#cnf(i_0_57, negated_conjecture, (distinct_points(esk2_0,X1)|distinct_points(esk1_0,X1))).
#
#cnf(i_0_65, negated_conjecture, (distinct_points(X1,esk1_0)|distinct_points(esk2_0,X1))).
#
#cnf(i_0_67, negated_conjecture, (distinct_points(X1,esk2_0)|distinct_points(X1,esk1_0))).
#
#cnf(i_0_24, plain, (~unorthogonal_lines(orthogonal_through_point(X1,X2),X1))).
#
#cnf(i_0_74, plain, (convergent_lines(orthogonal_through_point(X1,X2),X1))).
##
#cnf(i_0_73, negated_conjecture, (distinct_points(esk1_0,X1)|distinct_points(X1,esk2_0))).
#
#cnf(i_0_38, plain, (~equal_points(X1,X2)|~distinct_points(X1,X2))).
###
#cnf(i_0_12, plain, (apart_point_and_line(X3,X2)|distinct_points(X1,X3)|~apart_point_and_line(X1,X2))).
#
#cnf(i_0_40, plain, (~equal_lines(X1,X2)|~distinct_lines(X1,X2))).
#
#cnf(i_0_84, plain, (distinct_points(X3,X1)|~apart_point_and_line(X3,orthogonal_through_point(X2,X1)))).
#
#cnf(i_0_85, plain, (distinct_points(X3,X1)|~apart_point_and_line(X3,parallel_through_point(X2,X1)))).
#
#cnf(i_0_5, plain, (distinct_lines(X2,X3)|distinct_lines(X1,X3)|~distinct_lines(X1,X2))).
#
#cnf(i_0_42, plain, (~parallel_lines(X1,X2)|~convergent_lines(X1,X2))).
###
#cnf(i_0_93, plain, (distinct_lines(X2,X1)|~distinct_lines(X1,X2))).
#
#cnf(i_0_44, plain, (~incident_point_and_line(X1,X2)|~apart_point_and_line(X1,X2))).
#
#cnf(i_0_100, plain, (convergent_lines(X1,X2)|~distinct_lines(X2,X1))).
##
#cnf(i_0_14, plain, (convergent_lines(X1,X3)|distinct_lines(X2,X3)|~convergent_lines(X1,X2))).
#
#cnf(i_0_46, plain, (~orthogonal_lines(X1,X2)|~unorthogonal_lines(X1,X2))).
#
#cnf(i_0_115, plain, (distinct_lines(X2,X1)|~convergent_lines(X1,X2))).
#
#cnf(i_0_123, plain, (distinct_lines(X1,X2)|~convergent_lines(X1,X2))).
###
#cnf(i_0_120, plain, (convergent_lines(X1,X2)|~convergent_lines(X2,X1))).
#
#cnf(i_0_132, plain, (convergent_lines(X1,orthogonal_through_point(X1,X2)))).
###
#cnf(i_0_114, plain, (distinct_lines(X3,X1)|~convergent_lines(parallel_through_point(X1,X2),X3))).
#
#cnf(i_0_135, plain, (distinct_lines(orthogonal_through_point(parallel_through_point(X1,X2),X3),X1))).
#
#cnf(i_0_137, plain, (distinct_lines(X1,X2)|~convergent_lines(X1,parallel_through_point(X2,X3)))).
#
#cnf(i_0_129, plain, (~convergent_lines(X1,parallel_through_point(X1,X2)))).
#
#cnf(i_0_138, plain, (convergent_lines(orthogonal_through_point(parallel_through_point(X1,X2),X3),X1))).
#
#cnf(i_0_140, plain, (convergent_lines(X1,orthogonal_through_point(parallel_through_point(X1,X2),X3)))).
#
#cnf(i_0_141, plain, (distinct_lines(X1,orthogonal_through_point(parallel_through_point(X1,X2),X3)))).
#
#cnf(i_0_47, negated_conjecture, (~distinct_lines(esk3_0,line_connecting(esk1_0,esk2_0)))).
###
#cnf(i_0_145, plain, (distinct_lines(X3,parallel_through_point(X1,X2))|~convergent_lines(X1,X3))).
##
#cnf(i_0_160, plain, (convergent_lines(X1,parallel_through_point(X2,X3))|~convergent_lines(X2,X1))).
##
#cnf(i_0_162, plain, (convergent_lines(parallel_through_point(X1,X2),X3)|~convergent_lines(X1,X3))).
###
#cnf(i_0_163, plain, (distinct_lines(parallel_through_point(X1,X2),X3)|~convergent_lines(X1,X3))).
########
#cnf(i_0_117, plain, (convergent_lines(orthogonal_through_point(X1,X2),X3)|distinct_lines(X1,X3))).
#
#cnf(i_0_185, plain, (distinct_lines(X1,orthogonal_through_point(X1,X2)))).
##
#cnf(i_0_183, plain, (distinct_lines(X1,parallel_through_point(orthogonal_through_point(X1,X2),X3)))).
#
#cnf(i_0_191, plain, (convergent_lines(X1,parallel_through_point(orthogonal_through_point(X1,X2),X3)))).
#
#cnf(i_0_193, plain, (convergent_lines(parallel_through_point(orthogonal_through_point(X1,X2),X3),X1))).
##
#cnf(i_0_134, plain, (convergent_lines(X1,X2)|distinct_lines(orthogonal_through_point(X1,X3),X2))).
#
#cnf(i_0_187, plain, (convergent_lines(X1,orthogonal_through_point(X2,X3))|distinct_lines(X2,X1))).
#
#cnf(i_0_189, plain, (distinct_lines(orthogonal_through_point(X1,X2),X3)|distinct_lines(X1,X3))).
#
#cnf(i_0_155, negated_conjecture, (~convergent_lines(line_connecting(esk1_0,esk2_0),esk3_0))).
#
#cnf(i_0_202, plain, (convergent_lines(orthogonal_through_point(X1,X2),X3)|convergent_lines(X1,X3))).
#
#cnf(i_0_204, plain, (convergent_lines(X1,orthogonal_through_point(X2,X3))|convergent_lines(X2,X1))).
#
#cnf(i_0_205, plain, (convergent_lines(X2,X1)|distinct_lines(X1,orthogonal_through_point(X2,X3)))).
#
#cnf(i_0_156, negated_conjecture, (~convergent_lines(esk3_0,line_connecting(esk1_0,esk2_0)))).
####
#cnf(i_0_164, plain, (~convergent_lines(X1,parallel_through_point(parallel_through_point(X1,X2),X3)))).
#
#cnf(i_0_215, plain, (distinct_lines(X1,orthogonal_through_point(X2,X3))|distinct_lines(X2,X1))).
############
#cnf(i_0_13, plain, (apart_point_and_line(X1,X3)|distinct_lines(X2,X3)|~apart_point_and_line(X1,X2))).
#
#cnf(i_0_248, plain, (distinct_lines(X3,orthogonal_through_point(X2,X1))|~apart_point_and_line(X1,X3))).
#
#cnf(i_0_250, plain, (distinct_lines(X3,parallel_through_point(X2,X1))|~apart_point_and_line(X1,X3))).
##
#cnf(i_0_257, plain, (convergent_lines(X1,orthogonal_through_point(X2,X3))|~apart_point_and_line(X3,X1))).
#
#cnf(i_0_259, plain, (convergent_lines(orthogonal_through_point(X1,X2),X3)|~apart_point_and_line(X2,X3))).
#
#cnf(i_0_260, plain, (distinct_lines(orthogonal_through_point(X1,X2),X3)|~apart_point_and_line(X2,X3))).
##
#cnf(i_0_262, plain, (convergent_lines(X1,parallel_through_point(X2,X3))|~apart_point_and_line(X3,X1))).
#
#cnf(i_0_289, plain, (distinct_lines(X1,X2)|~apart_point_and_line(X3,X1))).
#
#cnf(i_0_292, plain, (convergent_lines(X1,X2)|~apart_point_and_line(X3,X1))).
#
#cnf(i_0_291, plain, (~apart_point_and_line(X2,X1))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(apart3, axiom, ![X1]:~(convergent_lines(X1,X1)), file('./p/GEO174+3.tptp', apart3)).
fof(ceq3, axiom, ![X1, X2, X3]:(convergent_lines(X1,X2)=>(distinct_lines(X2,X3)|convergent_lines(X1,X3))), file('./p/GEO174+3.tptp', ceq3)).
fof(cp1, axiom, ![X1, X2]:~(convergent_lines(parallel_through_point(X2,X1),X2)), file('./p/GEO174+3.tptp', cp1)).
fof(p1, axiom, ![X1, X2]:(distinct_lines(X1,X2)=>convergent_lines(X1,X2)), file('./p/GEO174+3.tptp', p1)).
fof(cp2, axiom, ![X1, X2]:~(apart_point_and_line(X1,parallel_through_point(X2,X1))), file('./p/GEO174+3.tptp', cp2)).
fof(ceq2, axiom, ![X1, X2, X3]:(apart_point_and_line(X1,X2)=>(distinct_lines(X2,X3)|apart_point_and_line(X1,X3))), file('./p/GEO174+3.tptp', ceq2)).
fof(con, conjecture, ![X1, X2, X4, X5]:(((distinct_points(X1,X2)&convergent_lines(X4,X5))&(apart_point_and_line(X1,X4)|apart_point_and_line(X2,X4)))=>distinct_lines(X4,line_connecting(X1,X2))), file('./p/GEO174+3.tptp', con)).
fof(apart2, axiom, ![X1]:~(distinct_lines(X1,X1)), file('./p/GEO174+3.tptp', apart2)).
fof(c_0_8, plain, ![X13]:~convergent_lines(X13,X13), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[apart3])])).
fof(c_0_9, plain, ![X41, X42, X43]:(~convergent_lines(X41,X42)|(distinct_lines(X42,X43)|convergent_lines(X41,X43))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ceq3])])).
fof(c_0_10, plain, ![X46, X47]:~convergent_lines(parallel_through_point(X47,X46),X47), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[cp1])])).
fof(c_0_11, plain, ![X44, X45]:(~distinct_lines(X44,X45)|convergent_lines(X44,X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[p1])])).
cnf(c_0_12, plain, (~convergent_lines(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (distinct_lines(X2,X3)|convergent_lines(X1,X3)|~convergent_lines(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_14, plain, ![X48, X49]:~apart_point_and_line(X48,parallel_through_point(X49,X48)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[cp2])])).
fof(c_0_15, plain, ![X38, X39, X40]:(~apart_point_and_line(X38,X39)|(distinct_lines(X39,X40)|apart_point_and_line(X38,X40))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ceq2])])).
cnf(c_0_16, plain, (~convergent_lines(parallel_through_point(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_17, plain, (convergent_lines(X1,X2)|~distinct_lines(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, plain, (distinct_lines(X1,X2)|~convergent_lines(X2,X1)), inference(pm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_19, plain, (~apart_point_and_line(X1,parallel_through_point(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (distinct_lines(X2,X3)|apart_point_and_line(X1,X3)|~apart_point_and_line(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_21, plain, (distinct_lines(X1,X2)|~convergent_lines(parallel_through_point(X2,X3),X1)), inference(pm,[status(thm)],[c_0_16, c_0_13])).
cnf(c_0_22, plain, (convergent_lines(X1,X2)|~convergent_lines(X2,X1)), inference(pm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, plain, (distinct_lines(X1,parallel_through_point(X2,X3))|~apart_point_and_line(X3,X1)), inference(pm,[status(thm)],[c_0_19, c_0_20])).
fof(c_0_24, negated_conjecture, ~(![X1, X2, X4, X5]:(((distinct_points(X1,X2)&convergent_lines(X4,X5))&(apart_point_and_line(X1,X4)|apart_point_and_line(X2,X4)))=>distinct_lines(X4,line_connecting(X1,X2)))), inference(assume_negation,[status(cth)],[con])).
fof(c_0_25, plain, ![X12]:~distinct_lines(X12,X12), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[apart2])])).
cnf(c_0_26, plain, (distinct_lines(X1,X2)|~convergent_lines(X1,parallel_through_point(X2,X3))), inference(pm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, plain, (convergent_lines(X1,parallel_through_point(X2,X3))|~apart_point_and_line(X3,X1)), inference(pm,[status(thm)],[c_0_17, c_0_23])).
fof(c_0_28, negated_conjecture, (((distinct_points(esk1_0,esk2_0)&convergent_lines(esk3_0,esk4_0))&(apart_point_and_line(esk1_0,esk3_0)|apart_point_and_line(esk2_0,esk3_0)))&~distinct_lines(esk3_0,line_connecting(esk1_0,esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_24])])])).
cnf(c_0_29, plain, (~distinct_lines(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, plain, (distinct_lines(X1,X2)|~apart_point_and_line(X3,X1)), inference(pm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31, negated_conjecture, (apart_point_and_line(esk1_0,esk3_0)|apart_point_and_line(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_32, plain, (~apart_point_and_line(X1,X2)), inference(pm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_33, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_31, c_0_32]), c_0_32]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 34
# Proof object clause steps            : 17
# Proof object formula steps           : 17
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 2
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 36
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 50
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 50
# Processed clauses                    : 124
# ...of these trivial                  : 1
# ...subsumed                          : 43
# ...remaining for further processing  : 80
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 0
# Generated clauses                    : 261
# ...of the previous two non-trivial   : 222
# Contextual simplify-reflections      : 0
# Paramodulations                      : 255
# Factorizations                       : 4
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 73
#    Positive orientable unit clauses  : 13
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 47
# Current number of unprocessed clauses: 145
# ...number of literals in the above   : 381
# Current number of archived formulas  : 0
# Current number of archived clauses   : 7
# Clause-clause subsumption calls (NU) : 231
# Rec. Clause-clause subsumption calls : 213
# Non-unit clause-clause subsumptions  : 31
# Unit Clause-clause subsumption calls : 163
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4562
