# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_49, negated_conjecture, (distinct_points(esk1_0,esk2_0))).
#
#cnf(i_0_15, plain, (convergent_lines(X1,X2)|~distinct_lines(X1,X2))).
#
#cnf(i_0_37, plain, (equal_points(X1,X2)|distinct_points(X1,X2))).
#
#cnf(i_0_1, plain, (~distinct_points(X1,X1))).
#
#cnf(i_0_39, plain, (equal_lines(X1,X2)|distinct_lines(X1,X2))).
#
#cnf(i_0_19, plain, (unorthogonal_lines(X1,X2)|convergent_lines(X1,X2))).
##
#cnf(i_0_2, plain, (~distinct_lines(X1,X1))).
#
#cnf(i_0_41, plain, (parallel_lines(X1,X2)|convergent_lines(X1,X2))).
#
#cnf(i_0_43, plain, (incident_point_and_line(X1,X2)|apart_point_and_line(X1,X2))).
#
#cnf(i_0_45, plain, (orthogonal_lines(X1,X2)|unorthogonal_lines(X1,X2))).
#
#cnf(i_0_3, plain, (~convergent_lines(X1,X1))).
#
#cnf(i_0_48, negated_conjecture, (apart_point_and_line(esk3_0,line_connecting(esk1_0,esk2_0)))).
#
#cnf(i_0_35, plain, (line(parallel_through_point(X1,X2))|~line(X1)|~point(X2))).
#
#cnf(i_0_36, plain, (line(orthogonal_through_point(X1,X2))|~line(X1)|~point(X2))).
##
#cnf(i_0_4, plain, (distinct_points(X2,X3)|distinct_points(X1,X3)|~distinct_points(X1,X2))).
###
#cnf(i_0_17, plain, (~apart_point_and_line(X1,parallel_through_point(X2,X1)))).
#
#cnf(i_0_53, plain, (distinct_points(X2,X1)|~distinct_points(X1,X2))).
#
#cnf(i_0_60, negated_conjecture, (distinct_points(esk2_0,esk1_0))).
#
#cnf(i_0_57, negated_conjecture, (distinct_points(esk2_0,X1)|distinct_points(esk1_0,X1))).
#
#cnf(i_0_25, plain, (~apart_point_and_line(X1,orthogonal_through_point(X2,X1)))).
#
#cnf(i_0_64, negated_conjecture, (distinct_points(X1,esk1_0)|distinct_points(esk2_0,X1))).
#
#cnf(i_0_66, negated_conjecture, (distinct_points(X1,esk2_0)|distinct_points(X1,esk1_0))).
##
#cnf(i_0_16, plain, (~convergent_lines(parallel_through_point(X1,X2),X1))).
#
#cnf(i_0_72, negated_conjecture, (distinct_points(esk1_0,X1)|distinct_points(X1,esk2_0))).
###
#cnf(i_0_24, plain, (~unorthogonal_lines(orthogonal_through_point(X1,X2),X1))).
#
#cnf(i_0_80, plain, (convergent_lines(orthogonal_through_point(X1,X2),X1))).
#
#cnf(i_0_12, plain, (apart_point_and_line(X3,X2)|distinct_points(X1,X3)|~apart_point_and_line(X1,X2))).
#
#cnf(i_0_84, plain, (distinct_points(X3,X1)|~apart_point_and_line(X3,orthogonal_through_point(X2,X1)))).
#
#cnf(i_0_74, plain, (~distinct_lines(parallel_through_point(X1,X2),X1))).
#
#cnf(i_0_85, plain, (distinct_points(X3,X1)|~apart_point_and_line(X3,parallel_through_point(X2,X1)))).
#
#cnf(i_0_5, plain, (distinct_lines(X2,X3)|distinct_lines(X1,X3)|~distinct_lines(X1,X2))).
##
#cnf(i_0_38, plain, (~equal_points(X1,X2)|~distinct_points(X1,X2))).
##
#cnf(i_0_93, plain, (distinct_lines(X2,X1)|~distinct_lines(X1,X2))).
#
#cnf(i_0_92, plain, (distinct_lines(X3,X1)|~distinct_lines(parallel_through_point(X1,X2),X3))).
#
#cnf(i_0_99, plain, (~distinct_lines(X1,parallel_through_point(X1,X2)))).
#
#cnf(i_0_95, plain, (distinct_lines(X3,X1)|~distinct_lines(X3,parallel_through_point(X1,X2)))).
###
#cnf(i_0_40, plain, (~equal_lines(X1,X2)|~distinct_lines(X1,X2))).
#
#cnf(i_0_107, plain, (distinct_lines(X3,parallel_through_point(X1,X2))|~distinct_lines(X3,X1))).
#
#cnf(i_0_108, plain, (distinct_lines(X3,parallel_through_point(X1,X2))|~distinct_lines(X1,X3))).
###
#cnf(i_0_117, plain, (distinct_lines(parallel_through_point(X1,X2),X3)|~distinct_lines(X3,X1))).
#
#cnf(i_0_124, plain, (distinct_lines(parallel_through_point(X1,X2),X3)|~distinct_lines(X1,X3))).
#####
#cnf(i_0_14, plain, (convergent_lines(X1,X3)|distinct_lines(X2,X3)|~convergent_lines(X1,X2))).
##
#cnf(i_0_149, plain, (distinct_lines(X2,X1)|~convergent_lines(X1,X2))).
#
#cnf(i_0_153, plain, (distinct_lines(X1,orthogonal_through_point(X1,X2)))).
###
#cnf(i_0_141, plain, (convergent_lines(X3,X1)|~convergent_lines(X3,parallel_through_point(X1,X2)))).
#
#cnf(i_0_148, plain, (distinct_lines(X3,X1)|~convergent_lines(parallel_through_point(X1,X2),X3))).
#
#cnf(i_0_156, plain, (distinct_lines(orthogonal_through_point(parallel_through_point(X1,X2),X3),X1))).
##
#cnf(i_0_158, plain, (convergent_lines(orthogonal_through_point(parallel_through_point(X1,X2),X3),X1))).
#
#cnf(i_0_159, plain, (convergent_lines(X1,X2)|~distinct_lines(X1,parallel_through_point(X2,X3)))).
###
#cnf(i_0_177, plain, (convergent_lines(X1,X2)|~distinct_lines(X2,X1))).
##
#cnf(i_0_165, plain, (distinct_lines(X1,orthogonal_through_point(parallel_through_point(X1,X2),X3)))).
###
#cnf(i_0_173, plain, (convergent_lines(X1,X2)|~distinct_lines(parallel_through_point(X2,X3),X1))).
########
#cnf(i_0_143, plain, (convergent_lines(X3,parallel_through_point(X1,X2))|~convergent_lines(X3,X1))).
#
#cnf(i_0_200, plain, (distinct_lines(parallel_through_point(X1,X2),X3)|~convergent_lines(X3,X1))).
##
#cnf(i_0_209, plain, (distinct_lines(X1,X2)|~convergent_lines(X1,X2))).
#
#cnf(i_0_212, plain, (distinct_lines(orthogonal_through_point(X1,X2),X1))).
###
#cnf(i_0_207, plain, (distinct_lines(X1,parallel_through_point(X2,X3))|~convergent_lines(X1,X2))).
#####
#cnf(i_0_151, plain, (convergent_lines(orthogonal_through_point(X1,X2),X3)|distinct_lines(X1,X3))).
#
#cnf(i_0_157, plain, (distinct_lines(orthogonal_through_point(X1,X2),X3)|distinct_lines(X1,X3))).
##
#cnf(i_0_235, plain, (distinct_lines(X1,parallel_through_point(orthogonal_through_point(X1,X2),X3)))).
#
#cnf(i_0_245, plain, (convergent_lines(X1,orthogonal_through_point(X1,X2)))).
####
#cnf(i_0_234, plain, (distinct_lines(X1,orthogonal_through_point(X2,X3))|distinct_lines(X2,X1))).
###
#cnf(i_0_248, plain, (convergent_lines(X1,X2)|distinct_lines(orthogonal_through_point(X1,X3),X2))).
#
#cnf(i_0_256, plain, (convergent_lines(X1,parallel_through_point(orthogonal_through_point(X1,X2),X3)))).
#
#cnf(i_0_265, plain, (distinct_lines(parallel_through_point(orthogonal_through_point(X1,X2),X3),X1))).
#
#cnf(i_0_42, plain, (~parallel_lines(X1,X2)|~convergent_lines(X1,X2))).
##
#cnf(i_0_261, plain, (convergent_lines(X2,X1)|distinct_lines(X1,orthogonal_through_point(X2,X3)))).
##
#cnf(i_0_44, plain, (~incident_point_and_line(X1,X2)|~apart_point_and_line(X1,X2))).
#
#cnf(i_0_13, plain, (apart_point_and_line(X1,X3)|distinct_lines(X2,X3)|~apart_point_and_line(X1,X2))).
#
#cnf(i_0_281, plain, (apart_point_and_line(X3,X1)|~apart_point_and_line(X3,parallel_through_point(X1,X2)))).
#
#cnf(i_0_284, plain, (apart_point_and_line(X3,parallel_through_point(X1,X2))|~apart_point_and_line(X3,X1))).
#
#cnf(i_0_301, plain, (~apart_point_and_line(X1,X2))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(cp1, axiom, ![X1, X2]:~(convergent_lines(parallel_through_point(X2,X1),X2)), file('./p/GEO181+3.tptp', cp1)).
fof(p1, axiom, ![X1, X2]:(distinct_lines(X1,X2)=>convergent_lines(X1,X2)), file('./p/GEO181+3.tptp', p1)).
fof(apart2, axiom, ![X1]:~(distinct_lines(X1,X1)), file('./p/GEO181+3.tptp', apart2)).
fof(apart5, axiom, ![X1, X2, X3]:(distinct_lines(X1,X2)=>(distinct_lines(X1,X3)|distinct_lines(X2,X3))), file('./p/GEO181+3.tptp', apart5)).
fof(ceq2, axiom, ![X1, X2, X3]:(apart_point_and_line(X1,X2)=>(distinct_lines(X2,X3)|apart_point_and_line(X1,X3))), file('./p/GEO181+3.tptp', ceq2)).
fof(con, conjecture, ![X1, X2, X3]:(distinct_points(X1,X2)=>(apart_point_and_line(X3,line_connecting(X1,X2))=>apart_point_and_line(X2,line_connecting(X1,X3)))), file('./p/GEO181+3.tptp', con)).
fof(cp2, axiom, ![X1, X2]:~(apart_point_and_line(X1,parallel_through_point(X2,X1))), file('./p/GEO181+3.tptp', cp2)).
fof(c_0_7, plain, ![X46, X47]:~convergent_lines(parallel_through_point(X47,X46),X47), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[cp1])])).
fof(c_0_8, plain, ![X44, X45]:(~distinct_lines(X44,X45)|convergent_lines(X44,X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[p1])])).
fof(c_0_9, plain, ![X12]:~distinct_lines(X12,X12), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[apart2])])).
fof(c_0_10, plain, ![X17, X18, X19]:(~distinct_lines(X17,X18)|(distinct_lines(X17,X19)|distinct_lines(X18,X19))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[apart5])])).
cnf(c_0_11, plain, (~convergent_lines(parallel_through_point(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (convergent_lines(X1,X2)|~distinct_lines(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_13, plain, (~distinct_lines(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_14, plain, (distinct_lines(X1,X3)|distinct_lines(X2,X3)|~distinct_lines(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (~distinct_lines(parallel_through_point(X1,X2),X1)), inference(pm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, plain, (distinct_lines(X1,X2)|~distinct_lines(X2,X1)), inference(pm,[status(thm)],[c_0_13, c_0_14])).
fof(c_0_17, plain, ![X38, X39, X40]:(~apart_point_and_line(X38,X39)|(distinct_lines(X39,X40)|apart_point_and_line(X38,X40))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ceq2])])).
fof(c_0_18, negated_conjecture, ~(![X1, X2, X3]:(distinct_points(X1,X2)=>(apart_point_and_line(X3,line_connecting(X1,X2))=>apart_point_and_line(X2,line_connecting(X1,X3))))), inference(assume_negation,[status(cth)],[con])).
fof(c_0_19, plain, ![X48, X49]:~apart_point_and_line(X48,parallel_through_point(X49,X48)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[cp2])])).
cnf(c_0_20, plain, (~distinct_lines(X1,parallel_through_point(X1,X2))), inference(pm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_21, plain, (distinct_lines(X2,X3)|apart_point_and_line(X1,X3)|~apart_point_and_line(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_17])).
fof(c_0_22, negated_conjecture, (distinct_points(esk1_0,esk2_0)&(apart_point_and_line(esk3_0,line_connecting(esk1_0,esk2_0))&~apart_point_and_line(esk2_0,line_connecting(esk1_0,esk3_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])).
cnf(c_0_23, plain, (~apart_point_and_line(X1,parallel_through_point(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (apart_point_and_line(X1,parallel_through_point(X2,X3))|~apart_point_and_line(X1,X2)), inference(pm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_25, negated_conjecture, (apart_point_and_line(esk3_0,line_connecting(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_26, plain, (~apart_point_and_line(X1,X2)), inference(pm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_27, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_25, c_0_26]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 28
# Proof object clause steps            : 13
# Proof object formula steps           : 15
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 1
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 36
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 49
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 49
# Processed clauses                    : 120
# ...of these trivial                  : 1
# ...subsumed                          : 47
# ...remaining for further processing  : 72
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 259
# ...of the previous two non-trivial   : 211
# Contextual simplify-reflections      : 0
# Paramodulations                      : 253
# Factorizations                       : 4
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 70
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 48
# Current number of unprocessed clauses: 139
# ...number of literals in the above   : 372
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 241
# Rec. Clause-clause subsumption calls : 199
# Non-unit clause-clause subsumptions  : 27
# Unit Clause-clause subsumption calls : 206
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4499
