# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_51, negated_conjecture, (line(esk1_0))).
#
#cnf(i_0_50, negated_conjecture, (line(esk2_0))).
#
#cnf(i_0_49, negated_conjecture, (convergent_lines(esk1_0,esk2_0))).
#
#cnf(i_0_1, plain, (~distinct_points(X1,X1))).
#
#cnf(i_0_15, plain, (convergent_lines(X1,X2)|~distinct_lines(X1,X2))).
#
#cnf(i_0_37, plain, (equal_points(X1,X2)|distinct_points(X1,X2))).
#
#cnf(i_0_39, plain, (equal_lines(X1,X2)|distinct_lines(X1,X2))).
#
#cnf(i_0_2, plain, (~distinct_lines(X1,X1))).
#
#cnf(i_0_19, plain, (unorthogonal_lines(X1,X2)|convergent_lines(X1,X2))).
##
#cnf(i_0_41, plain, (parallel_lines(X1,X2)|convergent_lines(X1,X2))).
#
#cnf(i_0_3, plain, (~convergent_lines(X1,X1))).
#
#cnf(i_0_43, plain, (incident_point_and_line(X1,X2)|apart_point_and_line(X1,X2))).
#
#cnf(i_0_45, plain, (orthogonal_lines(X1,X2)|unorthogonal_lines(X1,X2))).
#
#cnf(i_0_35, negated_conjecture, (line(parallel_through_point(X1,X2))|~line(X1))).
#
#cnf(i_0_17, plain, (~apart_point_and_line(X1,parallel_through_point(X2,X1)))).
#
#cnf(i_0_36, negated_conjecture, (line(orthogonal_through_point(X1,X2))|~line(X1))).
#
#cnf(i_0_33, negated_conjecture, (line(line_connecting(X1,X2))|~distinct_points(X1,X2))).
#
#cnf(i_0_4, plain, (distinct_points(X2,X3)|distinct_points(X1,X3)|~distinct_points(X1,X2))).
#
#cnf(i_0_25, plain, (~apart_point_and_line(X1,orthogonal_through_point(X2,X1)))).
###
#cnf(i_0_54, plain, (distinct_points(X2,X1)|~distinct_points(X1,X2))).
#
#cnf(i_0_16, plain, (~convergent_lines(parallel_through_point(X1,X2),X1))).
#
#cnf(i_0_12, plain, (apart_point_and_line(X3,X2)|distinct_points(X1,X3)|~apart_point_and_line(X1,X2))).
#
#cnf(i_0_65, plain, (distinct_points(X3,X1)|~apart_point_and_line(X3,orthogonal_through_point(X2,X1)))).
#
#cnf(i_0_66, plain, (distinct_points(X3,X1)|~apart_point_and_line(X3,parallel_through_point(X2,X1)))).
#
#cnf(i_0_24, plain, (~unorthogonal_lines(orthogonal_through_point(X1,X2),X1))).
#
#cnf(i_0_70, plain, (convergent_lines(orthogonal_through_point(X1,X2),X1))).
#
#cnf(i_0_5, plain, (distinct_lines(X2,X3)|distinct_lines(X1,X3)|~distinct_lines(X1,X2))).
##
#cnf(i_0_38, plain, (~equal_points(X1,X2)|~distinct_points(X1,X2))).
##
#cnf(i_0_73, plain, (distinct_lines(X2,X1)|~distinct_lines(X1,X2))).
#
#cnf(i_0_80, plain, (convergent_lines(X1,X2)|~distinct_lines(X2,X1))).
#
#cnf(i_0_40, plain, (~equal_lines(X1,X2)|~distinct_lines(X1,X2))).
##
#cnf(i_0_14, plain, (convergent_lines(X1,X3)|distinct_lines(X2,X3)|~convergent_lines(X1,X2))).
#
#cnf(i_0_95, plain, (distinct_lines(X2,X1)|~convergent_lines(X1,X2))).
##
#cnf(i_0_102, plain, (distinct_lines(X1,X2)|~convergent_lines(X1,X2))).
##
#cnf(i_0_99, plain, (convergent_lines(X1,X2)|~convergent_lines(X2,X1))).
##
#cnf(i_0_111, negated_conjecture, (convergent_lines(esk2_0,esk1_0))).
#
#cnf(i_0_112, plain, (convergent_lines(X1,orthogonal_through_point(X1,X2)))).
##
#cnf(i_0_108, plain, (~convergent_lines(X1,parallel_through_point(X1,X2)))).
#
#cnf(i_0_94, plain, (distinct_lines(X3,X1)|~convergent_lines(parallel_through_point(X1,X2),X3))).
#
#cnf(i_0_116, plain, (distinct_lines(orthogonal_through_point(parallel_through_point(X1,X2),X3),X1))).
#
#cnf(i_0_118, plain, (distinct_lines(X1,X2)|~convergent_lines(X1,parallel_through_point(X2,X3)))).
#
#cnf(i_0_42, plain, (~parallel_lines(X1,X2)|~convergent_lines(X1,X2))).
#
#cnf(i_0_119, plain, (convergent_lines(orthogonal_through_point(parallel_through_point(X1,X2),X3),X1))).
#
#cnf(i_0_121, plain, (convergent_lines(X1,orthogonal_through_point(parallel_through_point(X1,X2),X3)))).
#
#cnf(i_0_122, plain, (distinct_lines(X1,orthogonal_through_point(parallel_through_point(X1,X2),X3)))).
#
#cnf(i_0_44, plain, (~incident_point_and_line(X1,X2)|~apart_point_and_line(X1,X2))).
###
#cnf(i_0_115, plain, (distinct_lines(X3,parallel_through_point(X1,X2))|~convergent_lines(X1,X3))).
##
#cnf(i_0_137, plain, (convergent_lines(X1,parallel_through_point(X2,X3))|~convergent_lines(X2,X1))).
##
#cnf(i_0_139, plain, (convergent_lines(parallel_through_point(X1,X2),X3)|~convergent_lines(X1,X3))).
###
#cnf(i_0_140, plain, (distinct_lines(parallel_through_point(X1,X2),X3)|~convergent_lines(X1,X3))).
########
#cnf(i_0_97, plain, (convergent_lines(orthogonal_through_point(X1,X2),X3)|distinct_lines(X1,X3))).
#
#cnf(i_0_162, plain, (distinct_lines(X1,orthogonal_through_point(X1,X2)))).
##
#cnf(i_0_160, plain, (distinct_lines(X1,parallel_through_point(orthogonal_through_point(X1,X2),X3)))).
#
#cnf(i_0_168, plain, (convergent_lines(X1,parallel_through_point(orthogonal_through_point(X1,X2),X3)))).
#
#cnf(i_0_170, plain, (convergent_lines(parallel_through_point(orthogonal_through_point(X1,X2),X3),X1))).
##
#cnf(i_0_114, plain, (convergent_lines(X1,X2)|distinct_lines(orthogonal_through_point(X1,X3),X2))).
#
#cnf(i_0_164, plain, (convergent_lines(X1,orthogonal_through_point(X2,X3))|distinct_lines(X2,X1))).
#
#cnf(i_0_166, plain, (distinct_lines(orthogonal_through_point(X1,X2),X3)|distinct_lines(X1,X3))).
#
#cnf(i_0_46, plain, (~orthogonal_lines(X1,X2)|~unorthogonal_lines(X1,X2))).
#
#cnf(i_0_179, plain, (convergent_lines(orthogonal_through_point(X1,X2),X3)|convergent_lines(X1,X3))).
#
#cnf(i_0_181, plain, (convergent_lines(X1,orthogonal_through_point(X2,X3))|convergent_lines(X2,X1))).
#
#cnf(i_0_182, plain, (convergent_lines(X2,X1)|distinct_lines(X1,orthogonal_through_point(X2,X3)))).
#
#cnf(i_0_47, negated_conjecture, (~incident_point_and_line(X1,esk1_0)|~incident_point_and_line(X1,esk2_0))).
#
#cnf(i_0_209, negated_conjecture, (apart_point_and_line(X1,esk1_0)|~incident_point_and_line(X1,esk2_0))).
#
#cnf(i_0_210, negated_conjecture, (apart_point_and_line(X1,esk1_0)|apart_point_and_line(X1,esk2_0))).
##
#cnf(i_0_141, plain, (~convergent_lines(X1,parallel_through_point(parallel_through_point(X1,X2),X3)))).
###
#cnf(i_0_192, plain, (distinct_lines(X1,orthogonal_through_point(X2,X3))|distinct_lines(X2,X1))).
#########
#cnf(i_0_8, plain, (~apart_point_and_line(X2,line_connecting(X1,X2))|~distinct_points(X1,X2))).
###
#cnf(i_0_13, plain, (apart_point_and_line(X1,X3)|distinct_lines(X2,X3)|~apart_point_and_line(X1,X2))).
#
#cnf(i_0_7, plain, (~apart_point_and_line(X1,line_connecting(X1,X2))|~distinct_points(X1,X2))).
#
#cnf(i_0_229, plain, (distinct_lines(X3,orthogonal_through_point(X2,X1))|~apart_point_and_line(X1,X3))).
#
#cnf(i_0_230, plain, (distinct_lines(X3,parallel_through_point(X2,X1))|~apart_point_and_line(X1,X3))).
#
#cnf(i_0_240, plain, (convergent_lines(X1,orthogonal_through_point(X2,X3))|~apart_point_and_line(X3,X1))).
##
#cnf(i_0_242, plain, (convergent_lines(orthogonal_through_point(X1,X2),X3)|~apart_point_and_line(X2,X3))).
#
#cnf(i_0_243, plain, (distinct_lines(orthogonal_through_point(X1,X2),X3)|~apart_point_and_line(X2,X3))).
#
#cnf(i_0_245, plain, (convergent_lines(X1,parallel_through_point(X2,X3))|~apart_point_and_line(X3,X1))).
#
#cnf(i_0_270, plain, (~apart_point_and_line(X3,X1))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(apart3, axiom, ![X1]:~(convergent_lines(X1,X1)), file('./p/GEO226+3.tptp', apart3)).
fof(ceq3, axiom, ![X1, X2, X3]:(convergent_lines(X1,X2)=>(distinct_lines(X2,X3)|convergent_lines(X1,X3))), file('./p/GEO226+3.tptp', ceq3)).
fof(p1, axiom, ![X1, X2]:(distinct_lines(X1,X2)=>convergent_lines(X1,X2)), file('./p/GEO226+3.tptp', p1)).
fof(cp1, axiom, ![X1, X2]:~(convergent_lines(parallel_through_point(X2,X1),X2)), file('./p/GEO226+3.tptp', cp1)).
fof(con, conjecture, ![X6, X7]:(((line(X6)&line(X7))&convergent_lines(X6,X7))=>?[X1]:(point(X1)=>(incident_point_and_line(X1,X6)&incident_point_and_line(X1,X7)))), file('./p/GEO226+3.tptp', con)).
fof(cp2, axiom, ![X1, X2]:~(apart_point_and_line(X1,parallel_through_point(X2,X1))), file('./p/GEO226+3.tptp', cp2)).
fof(ceq2, axiom, ![X1, X2, X3]:(apart_point_and_line(X1,X2)=>(distinct_lines(X2,X3)|apart_point_and_line(X1,X3))), file('./p/GEO226+3.tptp', ceq2)).
fof(a4, axiom, ![X1, X2]:(incident_point_and_line(X1,X2)<=>~(apart_point_and_line(X1,X2))), file('./p/GEO226+3.tptp', a4)).
fof(c_0_8, plain, ![X13]:~convergent_lines(X13,X13), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[apart3])])).
fof(c_0_9, plain, ![X41, X42, X43]:(~convergent_lines(X41,X42)|(distinct_lines(X42,X43)|convergent_lines(X41,X43))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ceq3])])).
fof(c_0_10, plain, ![X44, X45]:(~distinct_lines(X44,X45)|convergent_lines(X44,X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[p1])])).
cnf(c_0_11, plain, (~convergent_lines(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (distinct_lines(X2,X3)|convergent_lines(X1,X3)|~convergent_lines(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_13, plain, ![X46, X47]:~convergent_lines(parallel_through_point(X47,X46),X47), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[cp1])])).
cnf(c_0_14, plain, (convergent_lines(X1,X2)|~distinct_lines(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, plain, (distinct_lines(X1,X2)|~convergent_lines(X2,X1)), inference(pm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, plain, (~convergent_lines(parallel_through_point(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_17, plain, (convergent_lines(X1,X2)|~convergent_lines(X2,X1)), inference(pm,[status(thm)],[c_0_14, c_0_15])).
fof(c_0_18, negated_conjecture, ~(![X6, X7]:(((line(X6)&line(X7))&convergent_lines(X6,X7))=>?[X1]:(point(X1)=>(incident_point_and_line(X1,X6)&incident_point_and_line(X1,X7))))), inference(assume_negation,[status(cth)],[con])).
cnf(c_0_19, plain, (~convergent_lines(X1,parallel_through_point(X1,X2))), inference(pm,[status(thm)],[c_0_16, c_0_17])).
fof(c_0_20, plain, ![X48, X49]:~apart_point_and_line(X48,parallel_through_point(X49,X48)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[cp2])])).
fof(c_0_21, plain, ![X38, X39, X40]:(~apart_point_and_line(X38,X39)|(distinct_lines(X39,X40)|apart_point_and_line(X38,X40))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ceq2])])).
fof(c_0_22, negated_conjecture, ![X94]:(((line(esk1_0)&line(esk2_0))&convergent_lines(esk1_0,esk2_0))&(point(X94)&(~incident_point_and_line(X94,esk1_0)|~incident_point_and_line(X94,esk2_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_18])])])])).
fof(c_0_23, plain, ![X88, X89]:((~incident_point_and_line(X88,X89)|~apart_point_and_line(X88,X89))&(apart_point_and_line(X88,X89)|incident_point_and_line(X88,X89))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[a4])])])).
cnf(c_0_24, plain, (distinct_lines(X1,parallel_through_point(X2,X3))|~convergent_lines(X2,X1)), inference(pm,[status(thm)],[c_0_19, c_0_12])).
cnf(c_0_25, plain, (~apart_point_and_line(X1,parallel_through_point(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (distinct_lines(X2,X3)|apart_point_and_line(X1,X3)|~apart_point_and_line(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, negated_conjecture, (~incident_point_and_line(X1,esk1_0)|~incident_point_and_line(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (apart_point_and_line(X1,X2)|incident_point_and_line(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_29, plain, (convergent_lines(X1,parallel_through_point(X2,X3))|~convergent_lines(X2,X1)), inference(pm,[status(thm)],[c_0_14, c_0_24])).
cnf(c_0_30, plain, (distinct_lines(X1,parallel_through_point(X2,X3))|~apart_point_and_line(X3,X1)), inference(pm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_31, negated_conjecture, (apart_point_and_line(X1,esk1_0)|~incident_point_and_line(X1,esk2_0)), inference(pm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32, plain, (~convergent_lines(X1,parallel_through_point(parallel_through_point(X1,X2),X3))), inference(pm,[status(thm)],[c_0_16, c_0_29])).
cnf(c_0_33, plain, (convergent_lines(X1,parallel_through_point(X2,X3))|~apart_point_and_line(X3,X1)), inference(pm,[status(thm)],[c_0_14, c_0_30])).
cnf(c_0_34, negated_conjecture, (apart_point_and_line(X1,esk1_0)|apart_point_and_line(X1,esk2_0)), inference(pm,[status(thm)],[c_0_31, c_0_28])).
cnf(c_0_35, plain, (~apart_point_and_line(X1,X2)), inference(pm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_36, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_34, c_0_35]), c_0_35]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 37
# Proof object clause steps            : 20
# Proof object formula steps           : 17
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 2
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 36
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 51
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 49
# Processed clauses                    : 116
# ...of these trivial                  : 1
# ...subsumed                          : 38
# ...remaining for further processing  : 77
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 225
# ...of the previous two non-trivial   : 189
# Contextual simplify-reflections      : 0
# Paramodulations                      : 219
# Factorizations                       : 4
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 75
#    Positive orientable unit clauses  : 14
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 51
# Current number of unprocessed clauses: 121
# ...number of literals in the above   : 327
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3
# Clause-clause subsumption calls (NU) : 197
# Rec. Clause-clause subsumption calls : 185
# Non-unit clause-clause subsumptions  : 23
# Unit Clause-clause subsumption calls : 142
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 4445
