# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_5, plain, (multiply(identity,X1)=X1)).
#
#cnf(i_0_6, plain, (multiply(left_inverse(X1),X1)=identity)).
#
#cnf(i_0_7, plain, (multiply(multiply(multiply(X1,X2),X1),X3)=multiply(X1,multiply(X2,multiply(X1,X3))))).
#
#cnf(i_0_8, negated_conjecture, (multiply(multiply(multiply(a,b),c),b)!=multiply(a,multiply(b,multiply(c,b))))).
#
#cnf(i_0_12, plain, (multiply(multiply(X1,identity),X2)=multiply(X1,X2))).
#
#cnf(i_0_16, plain, (multiply(left_inverse(identity),X1)=X1)).
#
#cnf(i_0_15, plain, (multiply(multiply(X1,X1),X2)=multiply(X1,multiply(X1,X2)))).
#
#cnf(i_0_19, plain, (multiply(multiply(X1,left_inverse(identity)),X2)=multiply(X1,X2))).
#
#cnf(i_0_35, plain, (multiply(left_inverse(left_inverse(identity)),X1)=X1)).
#
#cnf(i_0_44, plain, (left_inverse(identity)=identity)).
#
#cnf(i_0_20, plain, (multiply(multiply(X1,multiply(X1,identity)),X2)=multiply(X1,multiply(X1,X2)))).
##
#cnf(i_0_24, plain, (multiply(multiply(X1,multiply(X1,X1)),X2)=multiply(X1,multiply(X1,multiply(X1,X2))))).
##
#cnf(i_0_11, plain, (multiply(left_inverse(X1),multiply(X1,multiply(left_inverse(X1),X2)))=multiply(left_inverse(X1),X2))).
#
#cnf(i_0_81, plain, (multiply(left_inverse(X1),multiply(X1,identity))=identity)).
#
#cnf(i_0_95, plain, (multiply(left_inverse(multiply(X1,X1)),multiply(X1,multiply(X1,identity)))=identity)).
#
#cnf(i_0_13, plain, (multiply(multiply(multiply(X1,X2),multiply(X1,identity)),X3)=multiply(X1,multiply(X2,multiply(X1,X3))))).
#
#cnf(i_0_17, plain, (multiply(multiply(X1,multiply(X2,multiply(X1,identity))),X3)=multiply(X1,multiply(X2,multiply(X1,X3))))).
#
#cnf(i_0_185, plain, (multiply(X1,multiply(left_inverse(X1),multiply(X1,X2)))=multiply(X1,X2))).
#
#cnf(i_0_215, plain, (multiply(left_inverse(X1),multiply(left_inverse(left_inverse(X1)),identity))=identity)).
##
#cnf(i_0_184, plain, (multiply(X1,multiply(left_inverse(multiply(X1,identity)),multiply(X1,X2)))=multiply(X1,X2))).
###
#cnf(i_0_246, plain, (multiply(left_inverse(X1),multiply(left_inverse(multiply(left_inverse(X1),identity)),identity))=identity)).
##########
#cnf(i_0_22, plain, (multiply(multiply(X1,multiply(X1,multiply(X1,X1))),X2)=multiply(X1,multiply(X1,multiply(X1,multiply(X1,X2)))))).
##
#cnf(i_0_91, plain, (multiply(left_inverse(multiply(multiply(X1,X2),X1)),multiply(X1,multiply(X2,multiply(X1,identity))))=identity)).
#
#cnf(i_0_348, plain, (multiply(left_inverse(multiply(multiply(X1,left_inverse(X1)),X1)),multiply(X1,identity))=identity)).
##
#cnf(i_0_400, plain, (multiply(X1,multiply(left_inverse(multiply(multiply(X1,left_inverse(X1)),X1)),multiply(X1,X2)))=multiply(X1,X2))).
#
#cnf(i_0_93, plain, (multiply(left_inverse(multiply(X1,multiply(X1,X1))),multiply(X1,multiply(X1,multiply(X1,identity))))=identity)).
#
#cnf(i_0_345, plain, (multiply(left_inverse(multiply(multiply(X1,left_inverse(multiply(X1,identity))),X1)),multiply(X1,identity))=identity)).
####
#cnf(i_0_210, plain, (multiply(X1,multiply(multiply(left_inverse(X1),multiply(X1,X2)),multiply(X1,X3)))=multiply(X1,multiply(X2,multiply(X1,X3))))).
###
#cnf(i_0_46, plain, (multiply(multiply(X1,multiply(X1,multiply(X1,multiply(X1,identity)))),X2)=multiply(X1,multiply(X1,multiply(X1,multiply(X1,X2)))))).
######
#cnf(i_0_175, plain, (multiply(X1,multiply(multiply(X2,multiply(X1,identity)),multiply(X1,X3)))=multiply(multiply(X1,multiply(X2,multiply(X1,X1))),X3))).
#
#cnf(i_0_188, plain, (multiply(multiply(X1,multiply(X2,multiply(X2,multiply(X1,identity)))),X3)=multiply(X1,multiply(X2,multiply(X2,multiply(X1,X3)))))).
#
#cnf(i_0_844, plain, (multiply(X1,multiply(multiply(X1,multiply(left_inverse(X1),identity)),multiply(X1,X2)))=multiply(X1,multiply(X1,X2)))).
##
#cnf(i_0_208, plain, (multiply(X1,multiply(X1,multiply(left_inverse(multiply(X1,X1)),multiply(X1,multiply(X1,X2)))))=multiply(X1,multiply(X1,X2)))).
######
#cnf(i_0_721, plain, (multiply(multiply(X1,multiply(X2,multiply(X1,multiply(X1,identity)))),X3)=multiply(multiply(X1,multiply(X2,multiply(X1,X1))),X3))).
#
#cnf(i_0_1116, plain, (multiply(multiply(X1,multiply(left_inverse(multiply(X1,multiply(X1,identity))),multiply(X1,X1))),X2)=multiply(X1,X2))).
###
#cnf(i_0_861, plain, (multiply(left_inverse(left_inverse(X1)),multiply(left_inverse(X1),multiply(left_inverse(X1),multiply(left_inverse(left_inverse(X1)),X2))))=X2)).
#
#cnf(i_0_1275, plain, (multiply(left_inverse(left_inverse(X1)),multiply(left_inverse(X1),X2))=X2)).
#
#cnf(i_0_1331, plain, (multiply(X1,multiply(left_inverse(X1),X2))=X2)).
#
#cnf(i_0_1377, plain, (multiply(X1,identity)=X1)).
#
#cnf(i_0_1380, plain, (multiply(X1,left_inverse(X1))=identity)).
#
#cnf(i_0_1397, plain, (left_inverse(left_inverse(X1))=X1)).
#
#cnf(i_0_1410, plain, (multiply(left_inverse(X1),multiply(X1,X2))=X2)).
###
#cnf(i_0_1381, plain, (multiply(multiply(X1,X2),X1)=multiply(X1,multiply(X2,X1)))).
#
#cnf(i_0_1401, plain, (multiply(X1,multiply(X1,left_inverse(multiply(X1,X1))))=identity)).
#
#cnf(i_0_1450, plain, (multiply(X1,left_inverse(multiply(X1,X1)))=left_inverse(X1))).
#
#cnf(i_0_1468, plain, (left_inverse(multiply(X1,X1))=multiply(left_inverse(X1),left_inverse(X1)))).
##
#cnf(i_0_1440, plain, (multiply(X1,multiply(multiply(left_inverse(X1),X2),X1))=multiply(X2,X1))).
#
#cnf(i_0_1484, plain, (multiply(multiply(X1,X2),left_inverse(X1))=multiply(X1,multiply(X2,left_inverse(X1))))).
#
#cnf(i_0_1491, plain, (multiply(multiply(left_inverse(X1),X2),X1)=multiply(left_inverse(X1),multiply(X2,X1)))).
######
#cnf(i_0_1430, plain, (multiply(left_inverse(multiply(X1,X2)),multiply(X1,multiply(X2,X1)))=X1)).
#
#cnf(i_0_17, plain, (multiply(multiply(X1,multiply(X2,X1)),X3)=multiply(X1,multiply(X2,multiply(X1,X3))))).
#
#cnf(i_0_1591, plain, (multiply(X1,multiply(multiply(X2,X1),left_inverse(X1)))=multiply(X1,X2))).
#
#cnf(i_0_1611, plain, (multiply(multiply(X2,left_inverse(X1)),X1)=X2)).
#
#cnf(i_0_1620, plain, (multiply(multiply(X2,X1),left_inverse(X1))=X2)).
##
#cnf(i_0_1673, plain, (multiply(left_inverse(multiply(X1,X2)),X1)=left_inverse(X2))).
#
#cnf(i_0_1687, plain, (multiply(X2,left_inverse(multiply(X1,X2)))=left_inverse(X1))).
##
#cnf(i_0_1705, plain, (left_inverse(multiply(X1,X2))=multiply(left_inverse(X2),left_inverse(X1)))).
###
#cnf(i_0_1666, plain, (multiply(multiply(X1,multiply(X1,left_inverse(X2))),X2)=multiply(X1,X1))).
#
#cnf(i_0_1694, plain, (multiply(multiply(X1,multiply(X1,X2)),left_inverse(X2))=multiply(X1,X1))).
##
#cnf(i_0_1844, plain, (multiply(multiply(X1,X2),multiply(left_inverse(X2),X1))=multiply(X1,X1))).
#
#cnf(i_0_1861, plain, (multiply(multiply(X1,left_inverse(X2)),multiply(X2,X1))=multiply(X1,X1))).
#
#cnf(i_0_1883, plain, (multiply(X2,multiply(multiply(left_inverse(X2),X1),X1))=multiply(X1,X1))).
#
#cnf(i_0_1928, plain, (multiply(multiply(X1,X2),X2)=multiply(X1,multiply(X2,X2)))).
#
#cnf(i_0_1887, plain, (multiply(multiply(X2,left_inverse(X1)),multiply(X1,X1))=multiply(X2,X1))).
#
#cnf(i_0_1968, plain, (multiply(multiply(X1,multiply(X2,X2)),left_inverse(X2))=multiply(X1,X2))).
##
#cnf(i_0_1760, plain, (multiply(multiply(left_inverse(X2),left_inverse(X1)),multiply(X1,X2))=identity)).
#
#cnf(i_0_1763, plain, (multiply(multiply(X1,X2),multiply(left_inverse(X2),left_inverse(X1)))=identity)).
#
#cnf(i_0_2107, plain, (multiply(multiply(left_inverse(X1),X2),multiply(left_inverse(X2),X1))=identity)).
#
#cnf(i_0_2110, plain, (multiply(multiply(X1,left_inverse(X2)),multiply(X2,left_inverse(X1)))=identity)).
######
#cnf(i_0_1801, plain, (multiply(multiply(X1,X2),multiply(X1,X2))=multiply(X1,multiply(X2,multiply(X1,X2))))).
##########
#cnf(i_0_1972, plain, (multiply(multiply(X1,multiply(X2,X2)),X2)=multiply(multiply(X1,X2),multiply(X2,X2)))).
#
#cnf(i_0_1989, plain, (multiply(multiply(X1,X2),multiply(X1,X1))=multiply(X1,multiply(X2,multiply(X1,X1))))).
#
#cnf(i_0_1990, plain, (multiply(multiply(X1,multiply(X1,X2)),X2)=multiply(X1,multiply(X1,multiply(X2,X2))))).
####
#cnf(i_0_1430, plain, (multiply(multiply(left_inverse(X2),left_inverse(X1)),multiply(X1,multiply(X2,X1)))=X1)).
#
#cnf(i_0_1762, plain, (multiply(multiply(X1,X2),multiply(multiply(left_inverse(X2),left_inverse(X1)),X3))=X3)).
#
#cnf(i_0_1764, plain, (multiply(multiply(left_inverse(X2),left_inverse(X1)),multiply(multiply(X1,X2),X3))=X3)).
#
#cnf(i_0_1767, plain, (multiply(multiply(X1,multiply(left_inverse(X3),left_inverse(X2))),multiply(X2,X3))=X1)).
#
#cnf(i_0_1768, plain, (multiply(multiply(X1,multiply(X2,X3)),multiply(left_inverse(X3),left_inverse(X2)))=X1)).
##
#cnf(i_0_1967, plain, (multiply(multiply(left_inverse(X2),left_inverse(X1)),multiply(X1,multiply(X2,X2)))=X2)).
####
#cnf(i_0_2566, plain, (multiply(multiply(X1,left_inverse(X2)),multiply(X2,multiply(left_inverse(X1),X2)))=X2)).
#
#cnf(i_0_2597, plain, (multiply(multiply(left_inverse(X1),X2),multiply(multiply(left_inverse(X2),X1),X3))=X3)).
#
#cnf(i_0_2600, plain, (multiply(multiply(X1,left_inverse(X2)),multiply(multiply(X2,left_inverse(X1)),X3))=X3)).
#
#cnf(i_0_2608, plain, (multiply(multiply(X1,X2),multiply(left_inverse(X2),multiply(left_inverse(X1),X2)))=X2)).
##
#cnf(i_0_2728, plain, (multiply(multiply(left_inverse(X1),X2),multiply(left_inverse(X2),multiply(X1,X2)))=X2)).
#####
#cnf(i_0_2807, plain, (multiply(multiply(X1,multiply(left_inverse(X2),X3)),multiply(left_inverse(X3),X2))=X1)).
#
#cnf(i_0_2810, plain, (multiply(multiply(X1,multiply(X2,left_inverse(X3))),multiply(X3,left_inverse(X2)))=X1)).
###
#cnf(i_0_3001, plain, (multiply(multiply(left_inverse(X1),X2),multiply(left_inverse(X2),multiply(X1,X1)))=X1)).
############
#cnf(i_0_1961, plain, (multiply(multiply(X1,multiply(left_inverse(X2),left_inverse(X2))),X2)=multiply(X1,left_inverse(X2)))).
#
#cnf(i_0_1981, plain, (multiply(multiply(X1,X2),multiply(left_inverse(X2),left_inverse(X2)))=multiply(X1,left_inverse(X2)))).
#############
#cnf(i_0_1506, plain, (multiply(multiply(left_inverse(X2),left_inverse(X1)),multiply(X1,multiply(X2,left_inverse(X1))))=left_inverse(X1))).
###
#cnf(i_0_1505, plain, (multiply(multiply(X1,multiply(X2,left_inverse(X1))),multiply(X1,X2))=multiply(X1,multiply(X2,X2)))).
##
#cnf(i_0_1598, plain, (multiply(X1,multiply(multiply(left_inverse(X1),X2),multiply(X1,X3)))=multiply(multiply(X2,X1),X3))).
#
#cnf(i_0_1653, plain, (multiply(X1,multiply(multiply(X2,left_inverse(X1)),multiply(X1,X3)))=multiply(multiply(X1,X2),X3))).
#
#cnf(i_0_1664, plain, (multiply(multiply(X1,multiply(X2,multiply(X1,left_inverse(X3)))),X3)=multiply(X1,multiply(X2,X1)))).
#
#cnf(i_0_1672, plain, (multiply(multiply(X1,X2),multiply(left_inverse(X2),multiply(X1,X2)))=multiply(X1,multiply(X1,X2)))).
#
#cnf(i_0_1692, plain, (multiply(multiply(X1,multiply(X2,multiply(X1,X3))),left_inverse(X3))=multiply(X1,multiply(X2,X1)))).
##
#cnf(i_0_1988, plain, (multiply(multiply(X1,multiply(X1,left_inverse(X2))),multiply(X2,X2))=multiply(X1,multiply(X1,X2)))).
####
#cnf(i_0_2065, plain, (multiply(multiply(X1,multiply(X1,multiply(X2,X2))),left_inverse(X2))=multiply(X1,multiply(X1,X2)))).
#####
#cnf(i_0_4378, plain, (multiply(X1,multiply(multiply(X2,multiply(X3,left_inverse(X1))),X1))=multiply(multiply(X1,X2),X3))).
#
#cnf(i_0_4990, plain, (multiply(multiply(X1,X2),multiply(X3,X1))=multiply(X1,multiply(multiply(X2,X3),X1)))).
#######
#cnf(i_0_4912, plain, (multiply(multiply(multiply(X1,X2),multiply(X1,left_inverse(X2))),X2)=multiply(X1,multiply(X2,X1)))).
#
#cnf(i_0_4980, plain, (multiply(X1,multiply(multiply(multiply(left_inverse(X1),X2),X3),X1))=multiply(X2,multiply(X3,X1)))).
#
#cnf(i_0_5530, plain, (multiply(multiply(X2,X1),multiply(X1,X1))=multiply(X2,multiply(X1,multiply(X1,X1))))).
#
#cnf(i_0_1972, plain, (multiply(multiply(X1,multiply(X2,X2)),X2)=multiply(X1,multiply(X2,multiply(X2,X2))))).
##
#cnf(i_0_5462, plain, (multiply(X2,multiply(multiply(multiply(left_inverse(X2),X1),X3),X1))=multiply(X1,multiply(X3,X1)))).
#
#cnf(i_0_5777, plain, (multiply(multiply(multiply(X1,X2),X3),X2)=multiply(X1,multiply(X2,multiply(X3,X2))))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(moufang3, axiom, (multiply(multiply(multiply(X1,X2),X1),X3)=multiply(X1,multiply(X2,multiply(X1,X3)))), file('./p/GRP203-1.tptp', moufang3)).
cnf(left_identity, axiom, (multiply(identity,X1)=X1), file('./p/GRP203-1.tptp', left_identity)).
cnf(left_inverse, axiom, (multiply(left_inverse(X1),X1)=identity), file('./p/GRP203-1.tptp', left_inverse)).
cnf(prove_moufang2, negated_conjecture, (multiply(multiply(multiply(a,b),c),b)!=multiply(a,multiply(b,multiply(c,b)))), file('./p/GRP203-1.tptp', prove_moufang2)).
cnf(c_0_4, axiom, (multiply(multiply(multiply(X1,X2),X1),X3)=multiply(X1,multiply(X2,multiply(X1,X3)))), moufang3).
cnf(c_0_5, axiom, (multiply(identity,X1)=X1), left_identity).
cnf(c_0_6, axiom, (multiply(left_inverse(X1),X1)=identity), left_inverse).
cnf(c_0_7, plain, (multiply(multiply(X1,identity),X2)=multiply(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_4, c_0_5]), c_0_5]), c_0_5])).
cnf(c_0_8, plain, (multiply(left_inverse(X1),multiply(X1,multiply(left_inverse(X1),X2)))=multiply(left_inverse(X1),X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_4, c_0_6]), c_0_5])).
cnf(c_0_9, plain, (multiply(multiply(X1,X1),X2)=multiply(X1,multiply(X1,X2))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_4, c_0_7]), c_0_5])).
cnf(c_0_10, plain, (multiply(multiply(X1,multiply(X2,multiply(X1,identity))),X3)=multiply(X1,multiply(X2,multiply(X1,X3)))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_4]), c_0_4])).
cnf(c_0_11, plain, (multiply(left_inverse(X1),multiply(X1,identity))=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_8, c_0_6]), c_0_6])).
cnf(c_0_12, plain, (multiply(multiply(X1,multiply(X1,identity)),X2)=multiply(X1,multiply(X1,X2))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_9]), c_0_9])).
cnf(c_0_13, plain, (multiply(X1,multiply(left_inverse(X1),multiply(X1,X2)))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_10, c_0_11]), c_0_7])).
cnf(c_0_14, plain, (multiply(multiply(X1,multiply(X2,multiply(X2,multiply(X1,identity)))),X3)=multiply(X1,multiply(X2,multiply(X2,multiply(X1,X3))))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_10, c_0_12]), c_0_12])).
cnf(c_0_15, plain, (multiply(left_inverse(X1),multiply(left_inverse(left_inverse(X1)),identity))=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_13, c_0_11]), c_0_11])).
cnf(c_0_16, plain, (multiply(left_inverse(left_inverse(X1)),multiply(left_inverse(X1),multiply(left_inverse(X1),multiply(left_inverse(left_inverse(X1)),X2))))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_14, c_0_15]), c_0_11]), c_0_5])).
cnf(c_0_17, plain, (multiply(left_inverse(left_inverse(X1)),multiply(left_inverse(X1),X2))=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_8, c_0_16]), c_0_16])).
cnf(c_0_18, plain, (multiply(X1,multiply(left_inverse(X1),X2))=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_17, c_0_8]), c_0_17])).
cnf(c_0_19, plain, (multiply(X1,identity)=X1), inference(pm,[status(thm)],[c_0_18, c_0_6])).
cnf(c_0_20, plain, (multiply(multiply(X1,X2),X1)=multiply(X1,multiply(X2,X1))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_4, c_0_19]), c_0_19])).
cnf(c_0_21, plain, (multiply(X1,left_inverse(X1))=identity), inference(pm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22, plain, (multiply(X1,multiply(multiply(left_inverse(X1),X2),X1))=multiply(X2,X1)), inference(pm,[status(thm)],[c_0_20, c_0_18])).
cnf(c_0_23, plain, (left_inverse(left_inverse(X1))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_18, c_0_21]), c_0_19])).
cnf(c_0_24, plain, (multiply(multiply(X1,X2),left_inverse(X1))=multiply(X1,multiply(X2,left_inverse(X1)))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_18, c_0_22]), c_0_23])).
cnf(c_0_25, plain, (multiply(multiply(X1,multiply(X2,X1)),X3)=multiply(X1,multiply(X2,multiply(X1,X3)))), inference(rw,[status(thm)],[c_0_10, c_0_19])).
cnf(c_0_26, plain, (multiply(left_inverse(X1),multiply(X1,X2))=X2), inference(pm,[status(thm)],[c_0_18, c_0_23])).
cnf(c_0_27, plain, (multiply(X1,multiply(multiply(X2,X1),left_inverse(X1)))=multiply(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_24, c_0_25]), c_0_21]), c_0_19])).
cnf(c_0_28, plain, (multiply(multiply(X1,X2),left_inverse(X2))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_26, c_0_27]), c_0_26])).
cnf(c_0_29, plain, (multiply(multiply(X1,left_inverse(X2)),X2)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_18, c_0_27]), c_0_18]), c_0_23])).
cnf(c_0_30, plain, (multiply(left_inverse(multiply(X1,X2)),X1)=left_inverse(X2)), inference(pm,[status(thm)],[c_0_26, c_0_28])).
cnf(c_0_31, plain, (multiply(multiply(X1,multiply(X2,multiply(X1,left_inverse(X3)))),X3)=multiply(X1,multiply(X2,X1))), inference(pm,[status(thm)],[c_0_29, c_0_25])).
cnf(c_0_32, plain, (left_inverse(multiply(X1,X2))=multiply(left_inverse(X2),left_inverse(X1))), inference(pm,[status(thm)],[c_0_28, c_0_30])).
cnf(c_0_33, plain, (multiply(X1,multiply(multiply(X2,multiply(X3,left_inverse(X1))),X1))=multiply(multiply(X1,X2),X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_31, c_0_29]), c_0_32]), c_0_23])).
cnf(c_0_34, plain, (multiply(X1,multiply(multiply(multiply(left_inverse(X1),X2),X3),X1))=multiply(X2,multiply(X3,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_33, c_0_33]), c_0_21]), c_0_23]), c_0_5])).
cnf(c_0_35, plain, (multiply(X1,multiply(multiply(multiply(left_inverse(X1),X2),X3),X2))=multiply(X2,multiply(X3,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_34, c_0_18]), c_0_32]), c_0_23])).
cnf(c_0_36, negated_conjecture, (multiply(multiply(multiply(a,b),c),b)!=multiply(a,multiply(b,multiply(c,b)))), prove_moufang2).
cnf(c_0_37, plain, (multiply(multiply(multiply(X1,X2),X3),X2)=multiply(X1,multiply(X2,multiply(X3,X2)))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_18, c_0_35]), c_0_23])).
cnf(c_0_38, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_37])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 39
# Proof object clause steps            : 35
# Proof object formula steps           : 4
# Proof object conjectures             : 3
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 4
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 29
# Proof object simplifying inferences  : 33
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 4
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 4
# Processed clauses                    : 223
# ...of these trivial                  : 117
# ...subsumed                          : 0
# ...remaining for further processing  : 106
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 47
# Generated clauses                    : 6125
# ...of the previous two non-trivial   : 3551
# Contextual simplify-reflections      : 0
# Paramodulations                      : 6125
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 59
#    Positive orientable unit clauses  : 59
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 3259
# ...number of literals in the above   : 3259
# Current number of archived formulas  : 0
# Current number of archived clauses   : 47
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 325
# BW rewrite match successes           : 30
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 103844
