# Initializing proof state
# Scanning for AC axioms
# multiply is associative
#
#cnf(i_0_26, plain, (multiply(identity,X1)=X1)).
#
#cnf(i_0_27, plain, (multiply(inverse(X1),X1)=identity)).
#
#cnf(i_0_33, negated_conjecture, (inverse(sk_c1)=sk_c8|inverse(sk_c5)=sk_c8)).
#
#cnf(i_0_50, negated_conjecture, (multiply(X5,sk_c8)!=sk_c7|multiply(X2,sk_c6)!=sk_c7|multiply(X1,sk_c8)!=sk_c7|multiply(sk_c6,X3)!=sk_c7|inverse(sk_c6)!=sk_c7|multiply(X4,sk_c6)!=X3|inverse(X5)!=sk_c8|inverse(X4)!=sk_c6|inverse(X2)!=sk_c6|inverse(X1)!=sk_c8)).
#
#cnf(i_0_34, negated_conjecture, (inverse(sk_c1)=sk_c8|inverse(sk_c6)=sk_c7)).
#
#cnf(i_0_39, negated_conjecture, (inverse(sk_c5)=sk_c8|inverse(sk_c2)=sk_c6)).
#
#cnf(i_0_48, negated_conjecture, (inverse(sk_c5)=sk_c8|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_52, negated_conjecture, (multiply(sk_c6,multiply(X4,sk_c6))!=sk_c7|multiply(X1,sk_c8)!=sk_c7|multiply(X2,sk_c6)!=sk_c7|multiply(X3,sk_c8)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c8|inverse(X4)!=sk_c6|inverse(X2)!=sk_c6|inverse(X3)!=sk_c8)).
#
#cnf(i_0_40, negated_conjecture, (inverse(sk_c6)=sk_c7|inverse(sk_c2)=sk_c6)).
#
#cnf(i_0_49, negated_conjecture, (inverse(sk_c6)=sk_c7|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_28, plain, (multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3)))).
#
#cnf(i_0_61, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c7|multiply(X1,sk_c8)!=sk_c7|multiply(X2,sk_c6)!=sk_c7|multiply(X3,sk_c8)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(identity)!=sk_c6|inverse(X1)!=sk_c8|inverse(X2)!=sk_c6|inverse(X3)!=sk_c8)).
#
#cnf(i_0_73, plain, (multiply(inverse(X1),multiply(X1,X2))=X2)).
#
#cnf(i_0_85, plain, (multiply(inverse(identity),X1)=X1)).
#
#cnf(i_0_83, plain, (multiply(inverse(inverse(X1)),identity)=X1)).
#
#cnf(i_0_75, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c7|inverse(inverse(sk_c8))!=sk_c8|multiply(X1,sk_c8)!=sk_c7|multiply(X2,sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(identity)!=sk_c6|inverse(X1)!=sk_c8|inverse(X2)!=sk_c6|sk_c7!=identity)).
#
#cnf(i_0_101, plain, (multiply(inverse(inverse(identity)),X1)=X1)).
#
#cnf(i_0_130, plain, (inverse(identity)=identity)).
#
#cnf(i_0_82, plain, (multiply(inverse(inverse(X1)),X2)=multiply(X1,X2))).
#
#cnf(i_0_75, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c7|inverse(inverse(sk_c8))!=sk_c8|multiply(X1,sk_c8)!=sk_c7|multiply(X2,sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c8|inverse(X2)!=sk_c6|sk_c6!=identity|sk_c7!=identity)).
#
#cnf(i_0_83, plain, (multiply(X1,identity)=X1)).
#
#cnf(i_0_160, plain, (inverse(inverse(X1))=X1)).
#
#cnf(i_0_165, plain, (multiply(X1,inverse(X1))=identity)).
#
#cnf(i_0_75, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c7|multiply(X1,sk_c8)!=sk_c7|multiply(X2,sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c8|inverse(X2)!=sk_c6|sk_c6!=identity|sk_c7!=identity)).
#
#cnf(i_0_166, plain, (multiply(X1,multiply(inverse(X1),X2))=X2)).
##
#cnf(i_0_167, negated_conjecture, (inverse(sk_c6)=sk_c3|inverse(sk_c5)=sk_c8)).
#
#cnf(i_0_187, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c7|multiply(X1,sk_c8)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c8|sk_c7!=identity|sk_c6!=identity)).
#
#cnf(i_0_168, negated_conjecture, (inverse(sk_c6)=sk_c3|inverse(sk_c6)=sk_c7)).
#
#cnf(i_0_216, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c3!=sk_c7)).
#
#cnf(i_0_226, negated_conjecture, (inverse(sk_c7)=sk_c6|sk_c3!=sk_c7)).
#
#cnf(i_0_214, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=identity|sk_c6!=identity)).
##
#cnf(i_0_223, negated_conjecture, (multiply(sk_c7,sk_c6)=identity|sk_c3!=sk_c7)).
##
#cnf(i_0_192, negated_conjecture, (multiply(X1,sk_c8)!=sk_c7|multiply(X2,sk_c6)!=sk_c7|multiply(X3,sk_c8)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c8|inverse(X2)!=sk_c6|inverse(X3)!=sk_c8|sk_c7!=sk_c6)).
#
#cnf(i_0_225, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|sk_c3!=sk_c7)).
###
#cnf(i_0_240, negated_conjecture, (multiply(X1,sk_c8)!=sk_c7|multiply(X2,sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c8|inverse(X2)!=sk_c6|sk_c7!=identity|sk_c7!=sk_c6)).
##
#cnf(i_0_169, negated_conjecture, (inverse(sk_c6)=sk_c2|inverse(sk_c5)=sk_c8)).
#
#cnf(i_0_266, negated_conjecture, (inverse(sk_c5)=sk_c8|sk_c3=sk_c2)).
#
#cnf(i_0_254, negated_conjecture, (multiply(X1,sk_c8)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c8|sk_c7!=identity|sk_c7!=sk_c6)).
#
#cnf(i_0_272, negated_conjecture, (inverse(sk_c8)=sk_c5|sk_c3=sk_c2)).
##
#cnf(i_0_170, negated_conjecture, (inverse(sk_c6)=sk_c2|inverse(sk_c6)=sk_c7)).
#
#cnf(i_0_274, negated_conjecture, (inverse(sk_c6)!=sk_c7|sk_c7!=identity|sk_c7!=sk_c6)).
#
#cnf(i_0_282, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c2!=sk_c7)).
#
#cnf(i_0_301, negated_conjecture, (inverse(sk_c7)=sk_c6|sk_c2!=sk_c7)).
##
#cnf(i_0_297, negated_conjecture, (sk_c7!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
#
#cnf(i_0_290, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c3=sk_c2)).
#
#cnf(i_0_316, negated_conjecture, (inverse(sk_c7)=sk_c6|sk_c3=sk_c2)).
##
#cnf(i_0_307, negated_conjecture, (sk_c7!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7)).
#
#cnf(i_0_322, negated_conjecture, (sk_c3=sk_c2|sk_c7!=identity|sk_c7!=sk_c6)).
#
#cnf(i_0_298, negated_conjecture, (multiply(sk_c7,sk_c6)=identity|sk_c2!=sk_c7)).
###
#cnf(i_0_300, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|sk_c2!=sk_c7)).
#####
#cnf(i_0_171, negated_conjecture, (inverse(sk_c7)=sk_c6|inverse(sk_c1)=sk_c8)).
#
#cnf(i_0_172, negated_conjecture, (inverse(sk_c8)=sk_c5|inverse(sk_c1)=sk_c8)).
#
#cnf(i_0_246, negated_conjecture, (multiply(X1,sk_c8)!=sk_c7|multiply(X2,sk_c8)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c8|inverse(X2)!=sk_c8|sk_c7!=sk_c6|sk_c6!=identity)).
#########
#cnf(i_0_269, negated_conjecture, (multiply(sk_c8,sk_c5)=identity|sk_c3=sk_c2)).
##
#cnf(i_0_271, negated_conjecture, (multiply(sk_c5,sk_c8)=identity|sk_c3=sk_c2)).
#
#cnf(i_0_352, negated_conjecture, (multiply(X1,sk_c8)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c8|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c6!=identity)).
####
#cnf(i_0_367, negated_conjecture, (inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c6!=identity)).
#
#cnf(i_0_313, negated_conjecture, (multiply(sk_c7,sk_c6)=identity|sk_c3=sk_c2)).
##
#cnf(i_0_315, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|sk_c3=sk_c2)).
#
#cnf(i_0_373, negated_conjecture, (sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c6!=identity|sk_c3!=sk_c7)).
####
#cnf(i_0_374, negated_conjecture, (sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c6!=identity|sk_c2!=sk_c7)).
#
#cnf(i_0_292, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c7!=identity|sk_c7!=sk_c6)).
#
#cnf(i_0_389, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c7!=identity|sk_c7!=sk_c6)).
###
#cnf(i_0_175, plain, (multiply(X1,multiply(X2,inverse(multiply(X1,X2))))=identity)).
#
#cnf(i_0_404, plain, (multiply(X2,inverse(multiply(X1,X2)))=inverse(X1))).
##########
#cnf(i_0_242, negated_conjecture, (multiply(X1,sk_c8)!=sk_c7|multiply(X2,sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c8|inverse(X2)!=sk_c6|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6)).
##
#cnf(i_0_430, plain, (multiply(inverse(multiply(X1,X2)),X1)=inverse(X2))).
##
#cnf(i_0_455, negated_conjecture, (multiply(X1,sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c6|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6)).
##########
#cnf(i_0_429, plain, (inverse(multiply(X2,X1))=multiply(inverse(X1),inverse(X2)))).
#
#cnf(i_0_32, negated_conjecture, (multiply(sk_c5,sk_c8)=sk_c7|inverse(sk_c1)=sk_c8)).
##
#cnf(i_0_30, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|inverse(sk_c5)=sk_c8)).
#
#cnf(i_0_42, negated_conjecture, (multiply(sk_c6,sk_c4)=sk_c7|inverse(sk_c5)=sk_c8)).
#
#cnf(i_0_36, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c7|inverse(sk_c5)=sk_c8)).
##
#cnf(i_0_45, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c4|inverse(sk_c5)=sk_c8)).
#
#cnf(i_0_31, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|inverse(sk_c6)=sk_c7)).
#
#cnf(i_0_43, negated_conjecture, (multiply(sk_c6,sk_c4)=sk_c7|inverse(sk_c6)=sk_c7)).
#
#cnf(i_0_498, negated_conjecture, (multiply(X1,multiply(X2,sk_c6))!=sk_c7|multiply(inverse(X2),inverse(X1))!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6)).
#
#cnf(i_0_37, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c7|inverse(sk_c6)=sk_c7)).
#
#cnf(i_0_46, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c4|inverse(sk_c6)=sk_c7)).
#
#cnf(i_0_38, negated_conjecture, (multiply(sk_c5,sk_c8)=sk_c7|inverse(sk_c2)=sk_c6)).
##
#cnf(i_0_47, negated_conjecture, (multiply(sk_c5,sk_c8)=sk_c7|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_51, negated_conjecture, (multiply(sk_c8,sk_c5)=identity|inverse(sk_c1)=sk_c8)).
###
#cnf(i_0_57, negated_conjecture, (multiply(sk_c7,sk_c6)=identity|inverse(sk_c1)=sk_c8)).
##
#cnf(i_0_58, negated_conjecture, (multiply(sk_c6,sk_c2)=identity|inverse(sk_c5)=sk_c8)).
###
#cnf(i_0_59, negated_conjecture, (multiply(sk_c6,sk_c3)=identity|inverse(sk_c5)=sk_c8)).
###
#cnf(i_0_62, negated_conjecture, (multiply(sk_c6,sk_c2)=identity|inverse(sk_c6)=sk_c7)).
##
#cnf(i_0_63, negated_conjecture, (multiply(sk_c6,sk_c3)=identity|inverse(sk_c6)=sk_c7)).
###
#cnf(i_0_179, negated_conjecture, (multiply(sk_c3,sk_c6)=identity|inverse(sk_c5)=sk_c8)).
#
#cnf(i_0_654, negated_conjecture, (inverse(sk_c5)=sk_c8|sk_c4=identity)).
#
#cnf(i_0_572, negated_conjecture, (multiply(sk_c6,inverse(X1))!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|X1!=sk_c7)).
#
#cnf(i_0_660, negated_conjecture, (inverse(sk_c8)=sk_c5|sk_c4=identity)).
####
#cnf(i_0_657, negated_conjecture, (multiply(sk_c8,sk_c5)=identity|sk_c4=identity)).
##
#cnf(i_0_659, negated_conjecture, (multiply(sk_c5,sk_c8)=identity|sk_c4=identity)).
#####
#cnf(i_0_685, negated_conjecture, (multiply(sk_c6,X1)!=sk_c6|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6)).
#
#cnf(i_0_180, negated_conjecture, (multiply(sk_c3,sk_c6)=identity|inverse(sk_c6)=sk_c7)).
#
#cnf(i_0_731, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c4=identity)).
#
#cnf(i_0_736, negated_conjecture, (inverse(sk_c7)=sk_c6|sk_c4=identity)).
###
#cnf(i_0_743, negated_conjecture, (sk_c4=identity|sk_c7!=identity|sk_c7!=sk_c6)).
###
#cnf(i_0_733, negated_conjecture, (multiply(sk_c7,sk_c6)=identity|sk_c4=identity)).
##
#cnf(i_0_735, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|sk_c4=identity)).
#
#cnf(i_0_671, negated_conjecture, (multiply(sk_c6,sk_c7)!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
#####
#cnf(i_0_181, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|inverse(sk_c5)=sk_c8)).
#
#cnf(i_0_780, negated_conjecture, (inverse(sk_c5)=sk_c8|sk_c7=identity)).
#
#cnf(i_0_785, negated_conjecture, (inverse(sk_c8)=sk_c5|sk_c7=identity)).
#
#cnf(i_0_672, negated_conjecture, (multiply(sk_c6,sk_c7)!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7)).
###
#cnf(i_0_782, negated_conjecture, (multiply(sk_c8,sk_c5)=identity|sk_c7=identity)).
###
#cnf(i_0_784, negated_conjecture, (multiply(sk_c5,sk_c8)=identity|sk_c7=identity)).
##
#cnf(i_0_678, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
#
#cnf(i_0_809, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c7=identity)).
#
#cnf(i_0_810, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c7=identity)).
#
#cnf(i_0_811, negated_conjecture, (inverse(sk_c3)=sk_c6|sk_c7=identity)).
#
#cnf(i_0_679, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7)).
#
#cnf(i_0_815, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c7=identity)).
#
#cnf(i_0_840, negated_conjecture, (sk_c5=sk_c1|sk_c7=identity)).
#
#cnf(i_0_821, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c7=identity)).
##
#cnf(i_0_827, negated_conjecture, (inverse(sk_c6)=sk_c3|sk_c7=identity)).
#
#cnf(i_0_887, negated_conjecture, (sk_c3=sk_c2|sk_c7=identity)).
#######
#cnf(i_0_843, negated_conjecture, (sk_c6=identity|sk_c5=sk_c1|sk_c3!=sk_c7)).
###
#cnf(i_0_710, negated_conjecture, (multiply(inverse(X1),sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|X1!=sk_c6)).
#
#cnf(i_0_846, negated_conjecture, (sk_c6=identity|sk_c5=sk_c1|sk_c2!=sk_c7)).
####
#cnf(i_0_866, negated_conjecture, (sk_c2=sk_c7|sk_c7=identity|sk_c3!=sk_c7)).
#
#cnf(i_0_882, negated_conjecture, (sk_c3=sk_c7|sk_c7=identity|sk_c2!=sk_c7)).
####
#cnf(i_0_812, negated_conjecture, (multiply(sk_c8,sk_c1)=identity|sk_c7=identity)).
###
#cnf(i_0_814, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|sk_c7=identity)).
###
#cnf(i_0_912, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
#
#cnf(i_0_967, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c7=identity)).
#
#cnf(i_0_987, negated_conjecture, (sk_c2=sk_c7|sk_c7=identity)).
#
#cnf(i_0_988, negated_conjecture, (sk_c3=sk_c7|sk_c7=identity)).
###
#cnf(i_0_977, negated_conjecture, (inverse(sk_c7)=sk_c6|sk_c7=identity)).
##
#cnf(i_0_913, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7)).
###
#cnf(i_0_818, negated_conjecture, (multiply(sk_c6,sk_c2)=identity|sk_c7=identity)).
###
#cnf(i_0_820, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|sk_c7=identity)).
##
#cnf(i_0_923, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
###
#cnf(i_0_824, negated_conjecture, (multiply(sk_c6,sk_c3)=identity|sk_c7=identity)).
#
#cnf(i_0_924, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7)).
##
#cnf(i_0_826, negated_conjecture, (multiply(sk_c3,sk_c6)=identity|sk_c7=identity)).
###########
#cnf(i_0_974, negated_conjecture, (multiply(sk_c7,sk_c6)=identity|sk_c7=identity)).
##
#cnf(i_0_976, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|sk_c7=identity)).
################
#cnf(i_0_849, negated_conjecture, (sk_c6=identity|sk_c3=sk_c2|sk_c5=sk_c1)).
#
#cnf(i_0_569, negated_conjecture, (multiply(inverse(sk_c7),inverse(X1))!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7|X1!=sk_c7)).
#####
#cnf(i_0_854, negated_conjecture, (sk_c6=identity|sk_c4=identity|sk_c5=sk_c1)).
#######
#cnf(i_0_1126, negated_conjecture, (multiply(inverse(sk_c7),sk_c6)!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
#########
#cnf(i_0_182, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|inverse(sk_c6)=sk_c7)).
####
#cnf(i_0_183, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|inverse(sk_c1)=sk_c8)).
##
#cnf(i_0_184, negated_conjecture, (multiply(sk_c5,sk_c8)=identity|inverse(sk_c1)=sk_c8)).
#########
#cnf(i_0_1139, negated_conjecture, (multiply(inverse(sk_c7),X1)!=sk_c6|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
###########
#cnf(i_0_268, negated_conjecture, (inverse(sk_c5)=sk_c8|sk_c2=sk_c7|sk_c3!=sk_c7)).
###
#cnf(i_0_303, negated_conjecture, (inverse(sk_c5)=sk_c8|sk_c3=sk_c7|sk_c2!=sk_c7)).
###
#cnf(i_0_1204, negated_conjecture, (inverse(sk_c8)=sk_c5|sk_c2=sk_c7|sk_c3!=sk_c7)).
#
#cnf(i_0_1212, negated_conjecture, (inverse(sk_c8)=sk_c5|sk_c3=sk_c7|sk_c2!=sk_c7)).
####
#cnf(i_0_224, negated_conjecture, (multiply(sk_c7,multiply(sk_c6,X1))=X1|sk_c3!=sk_c7)).
##
#cnf(i_0_1240, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c7))!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c6|sk_c7!=sk_c8|sk_c8!=identity|sk_c3!=sk_c7)).
######
#cnf(i_0_227, negated_conjecture, (multiply(sk_c6,multiply(sk_c7,X1))=X1|sk_c3!=sk_c7)).
#########
#cnf(i_0_299, negated_conjecture, (multiply(sk_c7,multiply(sk_c6,X1))=X1|sk_c2!=sk_c7)).
#####
#cnf(i_0_1281, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c6))!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c6|sk_c7!=sk_c8|sk_c8!=identity|sk_c3!=sk_c7)).
###
#cnf(i_0_302, negated_conjecture, (multiply(sk_c6,multiply(sk_c7,X1))=X1|sk_c2!=sk_c7)).
##########
#cnf(i_0_1258, negated_conjecture, (multiply(sk_c6,X1)=X1|sk_c5=sk_c1|sk_c3!=sk_c7)).
#####
#cnf(i_0_1365, negated_conjecture, (inverse(sk_c6)=identity|sk_c5=sk_c1|sk_c3!=sk_c7)).
##
#cnf(i_0_1298, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c7))!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c6|sk_c7!=sk_c8|sk_c8!=identity|sk_c2!=sk_c7)).
#######
#cnf(i_0_1316, negated_conjecture, (multiply(sk_c6,X1)=X1|sk_c5=sk_c1|sk_c2!=sk_c7)).
######
#cnf(i_0_1457, negated_conjecture, (inverse(sk_c6)=identity|sk_c5=sk_c1|sk_c2!=sk_c7)).
#######
#cnf(i_0_1339, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c6))!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c6|sk_c7!=sk_c8|sk_c8!=identity|sk_c2!=sk_c7)).
##
#cnf(i_0_1394, negated_conjecture, (multiply(sk_c7,X1)=X1|sk_c5=sk_c1|sk_c3!=sk_c7)).
####
#cnf(i_0_1551, negated_conjecture, (inverse(sk_c7)=identity|sk_c5=sk_c1|sk_c3!=sk_c7)).
###########
#cnf(i_0_1487, negated_conjecture, (multiply(sk_c7,X1)=X1|sk_c5=sk_c1|sk_c2!=sk_c7)).
####
#cnf(i_0_1605, negated_conjecture, (inverse(sk_c7)=identity|sk_c5=sk_c1|sk_c2!=sk_c7)).
###
#cnf(i_0_570, negated_conjecture, (multiply(inverse(sk_c7),inverse(X1))!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7|X1!=sk_c7)).
############
#cnf(i_0_1675, negated_conjecture, (multiply(inverse(sk_c7),sk_c6)!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7)).
########################
#cnf(i_0_1691, negated_conjecture, (multiply(inverse(sk_c7),X1)!=sk_c6|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7)).
##
#cnf(i_0_841, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c6=identity|sk_c5=sk_c1)).
####################
#cnf(i_0_1730, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c6=identity|sk_c5=sk_c1)).
######
#cnf(i_0_1197, negated_conjecture, (multiply(inverse(X1),inverse(sk_c7))!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7|X1!=sk_c6)).
##
#cnf(i_0_270, negated_conjecture, (multiply(sk_c8,multiply(sk_c5,X1))=X1|sk_c3=sk_c2)).
####
#cnf(i_0_273, negated_conjecture, (multiply(sk_c5,multiply(sk_c8,X1))=X1|sk_c3=sk_c2)).
######
#cnf(i_0_1787, negated_conjecture, (multiply(sk_c6,inverse(sk_c7))!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
##
#cnf(i_0_314, negated_conjecture, (multiply(sk_c7,multiply(sk_c6,X1))=X1|sk_c3=sk_c2)).
############
#cnf(i_0_317, negated_conjecture, (multiply(sk_c6,multiply(sk_c7,X1))=X1|sk_c3=sk_c2)).
##########
#cnf(i_0_1806, negated_conjecture, (multiply(X1,inverse(sk_c7))!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c6|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
##
#cnf(i_0_658, negated_conjecture, (multiply(sk_c8,multiply(sk_c5,X1))=X1|sk_c4=identity)).
####
#cnf(i_0_661, negated_conjecture, (multiply(sk_c5,multiply(sk_c8,X1))=X1|sk_c4=identity)).
###########
#cnf(i_0_734, negated_conjecture, (multiply(sk_c7,multiply(sk_c6,X1))=X1|sk_c4=identity)).
############
#cnf(i_0_737, negated_conjecture, (multiply(sk_c6,multiply(sk_c7,X1))=X1|sk_c4=identity)).
###
#cnf(i_0_1341, negated_conjecture, (multiply(inverse(X1),inverse(sk_c7))!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7|X1!=sk_c6)).
######
#cnf(i_0_783, negated_conjecture, (multiply(sk_c8,multiply(sk_c5,X1))=X1|sk_c7=identity)).
####
#cnf(i_0_786, negated_conjecture, (multiply(sk_c5,multiply(sk_c8,X1))=X1|sk_c7=identity)).
##
#cnf(i_0_2021, negated_conjecture, (multiply(sk_c6,inverse(sk_c7))!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7)).
#########
#cnf(i_0_813, negated_conjecture, (multiply(sk_c8,multiply(sk_c1,X1))=X1|sk_c7=identity)).
####
#cnf(i_0_2086, negated_conjecture, (multiply(sk_c5,X1)=multiply(sk_c1,X1)|sk_c7=identity)).
####
#cnf(i_0_2102, negated_conjecture, (multiply(sk_c1,inverse(sk_c5))=identity|sk_c7=identity)).
##
#cnf(i_0_2123, negated_conjecture, (inverse(sk_c5)=inverse(sk_c1)|sk_c7=identity)).
#####
#cnf(i_0_2039, negated_conjecture, (multiply(X1,inverse(sk_c7))!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X1)!=sk_c6|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7)).
##
#cnf(i_0_2124, negated_conjecture, (multiply(sk_c5,inverse(sk_c1))=identity|sk_c7=identity)).
#####
#cnf(i_0_2132, negated_conjecture, (multiply(inverse(sk_c1),sk_c5)=identity|sk_c7=identity)).
########
#cnf(i_0_2180, negated_conjecture, (multiply(inverse(sk_c5),sk_c1)=identity|sk_c7=identity)).
############
#cnf(i_0_816, negated_conjecture, (multiply(sk_c1,multiply(sk_c8,X1))=X1|sk_c7=identity)).
#####
#cnf(i_0_686, negated_conjecture, (multiply(sk_c6,multiply(inverse(X2),inverse(X1)))!=sk_c6|inverse(sk_c6)!=sk_c7|multiply(X1,X2)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6)).
######
#cnf(i_0_819, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,X1))=X1|sk_c7=identity)).
####
#cnf(i_0_822, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,X1))=X1|sk_c7=identity)).
##############
#cnf(i_0_2258, negated_conjecture, (multiply(sk_c6,multiply(inverse(X1),X2))!=sk_c6|multiply(inverse(X2),X1)!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6)).
#
#cnf(i_0_825, negated_conjecture, (multiply(sk_c6,multiply(sk_c3,X1))=X1|sk_c7=identity)).
####
#cnf(i_0_2436, negated_conjecture, (multiply(sk_c3,X1)=multiply(sk_c2,X1)|sk_c7=identity)).
####
#cnf(i_0_2479, negated_conjecture, (multiply(sk_c2,X1)=multiply(sk_c7,X1)|sk_c7=identity)).
####
#cnf(i_0_2466, negated_conjecture, (multiply(sk_c2,inverse(sk_c3))=identity|sk_c7=identity)).
##
#cnf(i_0_2509, negated_conjecture, (inverse(sk_c3)=inverse(sk_c2)|sk_c7=identity)).
####
#cnf(i_0_2539, negated_conjecture, (inverse(sk_c2)=inverse(sk_c7)|sk_c7=identity)).
######
#cnf(i_0_2494, negated_conjecture, (multiply(sk_c7,inverse(sk_c2))=identity|sk_c7=identity)).
######
#cnf(i_0_2510, negated_conjecture, (multiply(sk_c3,inverse(sk_c2))=identity|sk_c7=identity)).
#
#cnf(i_0_2413, negated_conjecture, (multiply(sk_c6,multiply(X1,X2))!=sk_c6|multiply(inverse(X2),inverse(X1))!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6)).
#####
#cnf(i_0_2512, negated_conjecture, (multiply(sk_c7,inverse(sk_c3))=identity|sk_c7=identity)).
#
#cnf(i_0_2681, negated_conjecture, (inverse(sk_c3)=inverse(sk_c7)|sk_c7=identity)).
###########
#cnf(i_0_2516, negated_conjecture, (multiply(sk_c2,inverse(sk_c7))=identity|sk_c7=identity)).
########
#cnf(i_0_2521, negated_conjecture, (multiply(inverse(sk_c2),sk_c3)=identity|sk_c7=identity)).
######
#cnf(i_0_2540, negated_conjecture, (multiply(inverse(sk_c7),sk_c2)=identity|sk_c7=identity)).
#
#cnf(i_0_2610, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c6))!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
##########
#cnf(i_0_2577, negated_conjecture, (multiply(sk_c3,inverse(sk_c7))=identity|sk_c7=identity)).
##########
#cnf(i_0_2611, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c6))!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7)).
###
#cnf(i_0_2699, negated_conjecture, (multiply(inverse(sk_c7),sk_c3)=identity|sk_c7=identity)).
###########
#cnf(i_0_2742, negated_conjecture, (multiply(inverse(sk_c3),sk_c2)=identity|sk_c7=identity)).
#####
#cnf(i_0_2749, negated_conjecture, (multiply(inverse(sk_c2),sk_c7)=identity|sk_c7=identity)).
#
#cnf(i_0_2636, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c7))!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
#############
#cnf(i_0_2842, negated_conjecture, (multiply(inverse(sk_c3),sk_c7)=identity|sk_c7=identity)).
#######
#cnf(i_0_2637, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c7))!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7)).
#####################
#cnf(i_0_828, negated_conjecture, (multiply(sk_c3,multiply(sk_c6,X1))=X1|sk_c7=identity)).
#######
#cnf(i_0_2308, negated_conjecture, (multiply(sk_c6,multiply(X1,inverse(X2)))!=sk_c6|multiply(X2,inverse(X1))!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6)).
###################
#cnf(i_0_975, negated_conjecture, (multiply(sk_c7,multiply(sk_c6,X1))=X1|sk_c7=identity)).
#######
#cnf(i_0_3090, negated_conjecture, (multiply(sk_c3,X1)=multiply(sk_c7,X1)|sk_c7=identity)).
##########
#cnf(i_0_936, negated_conjecture, (multiply(inverse(X2),multiply(inverse(X1),sk_c6))!=sk_c7|inverse(sk_c6)!=sk_c7|multiply(X1,X2)!=sk_c6|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6)).
######
#cnf(i_0_978, negated_conjecture, (multiply(sk_c6,multiply(sk_c7,X1))=X1|sk_c7=identity)).
######################
#cnf(i_0_3181, negated_conjecture, (multiply(inverse(X1),multiply(X2,sk_c6))!=sk_c7|multiply(inverse(X2),X1)!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6)).
#################
#cnf(i_0_1872, negated_conjecture, (multiply(sk_c6,X1)=X1|sk_c3=sk_c2|sk_c5=sk_c1)).
#####
#cnf(i_0_3367, negated_conjecture, (inverse(sk_c6)=identity|sk_c3=sk_c2|sk_c5=sk_c1)).
#########
#cnf(i_0_1962, negated_conjecture, (multiply(sk_c6,X1)=X1|sk_c4=identity|sk_c5=sk_c1)).
######
#cnf(i_0_3476, negated_conjecture, (inverse(sk_c6)=identity|sk_c4=identity|sk_c5=sk_c1)).
###
#cnf(i_0_3235, negated_conjecture, (multiply(X1,multiply(inverse(X2),sk_c6))!=sk_c7|multiply(X2,inverse(X1))!=sk_c6|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6)).
#########
#cnf(i_0_3404, negated_conjecture, (multiply(sk_c7,X1)=X1|sk_c3=sk_c2|sk_c5=sk_c1)).
####
#cnf(i_0_3656, negated_conjecture, (inverse(sk_c7)=identity|sk_c3=sk_c2|sk_c5=sk_c1)).
##########
#cnf(i_0_3514, negated_conjecture, (multiply(sk_c7,X1)=X1|sk_c4=identity|sk_c5=sk_c1)).
####
#cnf(i_0_3752, negated_conjecture, (inverse(sk_c7)=identity|sk_c4=identity|sk_c5=sk_c1)).
######################
#cnf(i_0_1238, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c6,X1)|sk_c3!=sk_c7)).
#############
#cnf(i_0_1242, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c7,X1)|sk_c3!=sk_c7)).
###############
#cnf(i_0_1296, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c6,X1)|sk_c2!=sk_c7)).
#############
#cnf(i_0_1300, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c7,X1)|sk_c2!=sk_c7)).
###############
#cnf(i_0_1811, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c5,X1)|sk_c3=sk_c2)).
##########
#cnf(i_0_1815, negated_conjecture, (multiply(inverse(sk_c5),X1)=multiply(sk_c8,X1)|sk_c3=sk_c2)).
##############
#cnf(i_0_1848, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c6,X1)|sk_c3=sk_c2)).
#############
#cnf(i_0_1854, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c7,X1)|sk_c3=sk_c2)).
###############
#cnf(i_0_1912, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c5,X1)|sk_c4=identity)).
##########
#cnf(i_0_1916, negated_conjecture, (multiply(inverse(sk_c5),X1)=multiply(sk_c8,X1)|sk_c4=identity)).
##############
#cnf(i_0_1938, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c6,X1)|sk_c4=identity)).
#############
#cnf(i_0_1944, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c7,X1)|sk_c4=identity)).
###############
#cnf(i_0_2044, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c5,X1)|sk_c7=identity)).
############
#cnf(i_0_2048, negated_conjecture, (multiply(inverse(sk_c5),X1)=multiply(sk_c8,X1)|sk_c7=identity)).
##############
#cnf(i_0_2081, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c1,X1)|sk_c7=identity)).
##############
#cnf(i_0_2088, negated_conjecture, (multiply(inverse(sk_c1),X1)=multiply(sk_c8,X1)|sk_c7=identity)).
##############
#cnf(i_0_2317, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c2,X1)|sk_c7=identity)).
##############
#cnf(i_0_2331, negated_conjecture, (multiply(inverse(sk_c2),X1)=multiply(sk_c6,X1)|sk_c7=identity)).
################
#cnf(i_0_2427, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c3,X1)|sk_c7=identity)).
###
#cnf(i_0_2377, negated_conjecture, (multiply(sk_c6,multiply(sk_c7,X1))!=sk_c6|multiply(inverse(X1),sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
###########
#cnf(i_0_2445, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c6,X1)|sk_c7=identity)).
################
#cnf(i_0_3067, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c6,X1)|sk_c7=identity)).
#########
#cnf(i_0_2378, negated_conjecture, (multiply(sk_c6,multiply(sk_c7,X1))!=sk_c6|multiply(inverse(X1),sk_c6)!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c2!=sk_c7)).
######
#cnf(i_0_3074, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c7,X1)|sk_c7=identity)).
##############################
#cnf(i_0_2394, negated_conjecture, (multiply(sk_c6,multiply(sk_c6,X1))!=sk_c6|multiply(inverse(X1),sk_c7)!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
#
#cnf(i_0_386, negated_conjecture, (multiply(sk_c8,sk_c1)=identity|sk_c7!=identity|sk_c7!=sk_c6)).
##
#cnf(i_0_388, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|sk_c7!=identity|sk_c7!=sk_c6)).
#####
#cnf(i_0_4944, negated_conjecture, (multiply(sk_c6,sk_c7)!=sk_c7|inverse(sk_c6)!=sk_c7|sk_c7!=sk_c8|sk_c8!=identity|sk_c7!=sk_c6|sk_c3!=sk_c7)).
#
#cnf(i_0_505, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c6))=identity|sk_c3!=sk_c7)).
#########
#cnf(i_0_506, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c6))=identity|sk_c2!=sk_c7)).
#########
#cnf(i_0_509, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c7))=identity|sk_c3!=sk_c7)).
##########
#cnf(i_0_510, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c7))=identity|sk_c2!=sk_c7)).
#########
#cnf(i_0_29, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|multiply(sk_c5,sk_c8)=sk_c7)).
#
#cnf(i_0_5213, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|sk_c7=identity)).
##
#cnf(i_0_5228, negated_conjecture, (sk_c7=identity)).
#
#cnf(i_0_301, negated_conjecture, (sk_c6=identity|sk_c2!=identity)).
##
#cnf(i_0_307, negated_conjecture, (sk_c6!=identity|sk_c2!=identity)).
##
#cnf(i_0_226, negated_conjecture, (sk_c6=identity|sk_c3!=identity)).
##
#cnf(i_0_297, negated_conjecture, (sk_c6!=identity|sk_c3!=identity)).
##
#cnf(i_0_743, negated_conjecture, (sk_c4=identity|sk_c6!=identity)).
#
#cnf(i_0_322, negated_conjecture, (sk_c3=sk_c2|sk_c6!=identity)).
##
#cnf(i_0_733, negated_conjecture, (sk_c6=identity|sk_c4=identity)).
###
#cnf(i_0_274, negated_conjecture, (inverse(sk_c6)!=identity|sk_c6!=identity)).
#
#cnf(i_0_316, negated_conjecture, (sk_c6=identity|sk_c3=sk_c2)).
#####
#cnf(i_0_505, negated_conjecture, (inverse(sk_c6)=identity|sk_c3!=identity)).
#####
#cnf(i_0_506, negated_conjecture, (inverse(sk_c6)=identity|sk_c2!=identity)).
########
#cnf(i_0_389, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c6!=identity)).
##
#cnf(i_0_292, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c6!=identity)).
######
#cnf(i_0_1296, negated_conjecture, (multiply(sk_c6,X1)=X1|sk_c2!=identity)).
#######
#cnf(i_0_1238, negated_conjecture, (multiply(sk_c6,X1)=X1|sk_c3!=identity)).
############
#cnf(i_0_731, negated_conjecture, (inverse(sk_c6)=identity|sk_c4=identity)).
#######
#cnf(i_0_290, negated_conjecture, (inverse(sk_c6)=identity|sk_c3=sk_c2)).
######
#cnf(i_0_183, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c6=identity)).
####
#cnf(i_0_5313, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c6=identity)).
###
#cnf(i_0_1848, negated_conjecture, (multiply(sk_c6,X1)=X1|sk_c3=sk_c2)).
####
#cnf(i_0_734, negated_conjecture, (multiply(sk_c6,X1)=X1|sk_c4=identity)).
###############
#cnf(i_0_388, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|sk_c6!=identity)).
##
#cnf(i_0_386, negated_conjecture, (multiply(sk_c8,sk_c1)=identity|sk_c6!=identity)).
###########
#cnf(i_0_170, negated_conjecture, (inverse(sk_c6)=sk_c2|inverse(sk_c6)=identity)).
###
#cnf(i_0_168, negated_conjecture, (inverse(sk_c6)=sk_c3|inverse(sk_c6)=identity)).
####
#cnf(i_0_49, negated_conjecture, (inverse(sk_c6)=identity|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_40, negated_conjecture, (inverse(sk_c6)=identity|inverse(sk_c2)=sk_c6)).
#
#cnf(i_0_34, negated_conjecture, (inverse(sk_c1)=sk_c8|inverse(sk_c6)=identity)).
##################
#cnf(i_0_1300, negated_conjecture, (multiply(inverse(sk_c6),X1)=X1|sk_c2!=identity)).
##########
#cnf(i_0_1242, negated_conjecture, (multiply(inverse(sk_c6),X1)=X1|sk_c3!=identity)).
################
#cnf(i_0_5310, negated_conjecture, (multiply(sk_c8,sk_c1)=identity|sk_c6=identity)).
###
#cnf(i_0_5312, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|sk_c6=identity)).
#####
#cnf(i_0_5387, negated_conjecture, (multiply(sk_c6,X1)=X1|inverse(sk_c1)=sk_c8)).
###########
#cnf(i_0_1944, negated_conjecture, (multiply(inverse(sk_c6),X1)=X1|sk_c4=identity)).
###########
#cnf(i_0_1854, negated_conjecture, (multiply(inverse(sk_c6),X1)=X1|sk_c3=sk_c2)).
################
#cnf(i_0_5401, negated_conjecture, (multiply(inverse(X1),sk_c6)=inverse(X1)|sk_c2!=identity)).
#####
#cnf(i_0_5489, negated_conjecture, (inverse(sk_c6)=sk_c6|sk_c2!=identity)).
#######
#cnf(i_0_5520, negated_conjecture, (multiply(X1,sk_c6)=X1|sk_c2!=identity)).
##########
#cnf(i_0_5523, negated_conjecture, (multiply(sk_c6,sk_c6)=identity|sk_c2!=identity)).
#
#cnf(i_0_52, negated_conjecture, (multiply(sk_c6,multiply(X1,sk_c6))!=identity|multiply(X2,sk_c8)!=identity|multiply(X3,sk_c6)!=identity|multiply(X4,sk_c8)!=identity|inverse(sk_c6)!=identity|inverse(X2)!=sk_c8|inverse(X1)!=sk_c6|inverse(X3)!=sk_c6|inverse(X4)!=sk_c8)).
############
#cnf(i_0_5589, negated_conjecture, (multiply(sk_c6,sk_c6)!=identity|multiply(X1,sk_c8)!=identity|multiply(X2,sk_c6)!=identity|multiply(X3,sk_c8)!=identity|inverse(sk_c6)!=identity|inverse(X1)!=sk_c8|inverse(X2)!=sk_c6|inverse(X3)!=sk_c8|sk_c2!=identity)).
#
#cnf(i_0_5419, negated_conjecture, (multiply(inverse(X1),sk_c6)=inverse(X1)|sk_c3!=identity)).
###
#cnf(i_0_5613, negated_conjecture, (multiply(sk_c6,sk_c6)!=identity|multiply(X1,sk_c8)!=identity|multiply(X2,sk_c6)!=identity|inverse(sk_c6)!=identity|inverse(X1)!=sk_c8|inverse(X2)!=sk_c6|sk_c2!=identity)).
##
#cnf(i_0_5661, negated_conjecture, (inverse(sk_c6)=sk_c6|sk_c3!=identity)).
######
#cnf(i_0_5712, negated_conjecture, (multiply(sk_c6,sk_c6)!=identity|multiply(X1,sk_c8)!=identity|inverse(sk_c6)!=identity|inverse(X1)!=sk_c8|sk_c2!=identity)).
#
#cnf(i_0_5692, negated_conjecture, (multiply(X1,sk_c6)=X1|sk_c3!=identity)).
###
#cnf(i_0_5761, negated_conjecture, (multiply(sk_c6,sk_c6)!=identity|inverse(sk_c6)!=identity|sk_c2!=identity)).
####
#cnf(i_0_5803, negated_conjecture, (inverse(sk_c6)!=identity|sk_c2!=identity)).
#
#cnf(i_0_5808, negated_conjecture, (sk_c4=identity|sk_c2!=identity)).
#
#cnf(i_0_5811, negated_conjecture, (sk_c3=sk_c2|sk_c2!=identity)).
##
#cnf(i_0_5810, negated_conjecture, (sk_c2!=identity)).
#########
#cnf(i_0_5733, negated_conjecture, (multiply(sk_c6,sk_c6)=identity|sk_c3!=identity)).
###########
#cnf(i_0_5590, negated_conjecture, (multiply(sk_c6,sk_c6)!=identity|multiply(X1,sk_c8)!=identity|multiply(X2,sk_c6)!=identity|multiply(X3,sk_c8)!=identity|inverse(sk_c6)!=identity|inverse(X1)!=sk_c8|inverse(X2)!=sk_c6|inverse(X3)!=sk_c8|sk_c3!=identity)).
####
#cnf(i_0_5841, negated_conjecture, (multiply(sk_c6,sk_c6)!=identity|multiply(X1,sk_c8)!=identity|multiply(X2,sk_c6)!=identity|inverse(sk_c6)!=identity|inverse(X1)!=sk_c8|inverse(X2)!=sk_c6|sk_c3!=identity)).
####
#cnf(i_0_5902, negated_conjecture, (multiply(sk_c6,sk_c6)!=identity|multiply(X1,sk_c8)!=identity|inverse(sk_c6)!=identity|inverse(X1)!=sk_c8|sk_c3!=identity)).
####
#cnf(i_0_5933, negated_conjecture, (multiply(sk_c6,sk_c6)!=identity|inverse(sk_c6)!=identity|sk_c3!=identity)).
##
#cnf(i_0_5453, negated_conjecture, (multiply(inverse(X1),sk_c6)=inverse(X1)|sk_c4=identity)).
##
#cnf(i_0_5951, negated_conjecture, (inverse(sk_c6)!=identity|sk_c3!=identity)).
#
#cnf(i_0_6003, negated_conjecture, (sk_c3=sk_c2|sk_c3!=identity)).
#
#cnf(i_0_6000, negated_conjecture, (sk_c4=identity|sk_c3!=identity)).
##
#cnf(i_0_6001, negated_conjecture, (sk_c3!=identity)).
#########
#cnf(i_0_5958, negated_conjecture, (inverse(sk_c6)=sk_c6|sk_c4=identity)).
########
#cnf(i_0_5993, negated_conjecture, (multiply(X1,sk_c6)=X1|sk_c4=identity)).
######################
#cnf(i_0_6013, negated_conjecture, (multiply(sk_c6,sk_c6)=identity|sk_c4=identity)).
############
#cnf(i_0_6041, negated_conjecture, (multiply(sk_c8,sk_c5)=sk_c6|sk_c4=identity)).
###
#cnf(i_0_6043, negated_conjecture, (multiply(sk_c5,sk_c8)=sk_c6|sk_c4=identity)).
###########
#cnf(i_0_6049, negated_conjecture, (multiply(sk_c8,sk_c6)=inverse(sk_c5)|sk_c4=identity)).
###
#cnf(i_0_50, negated_conjecture, (multiply(X1,sk_c8)!=identity|multiply(X2,sk_c6)!=identity|multiply(X3,sk_c8)!=identity|multiply(sk_c6,X4)!=identity|inverse(sk_c6)!=identity|multiply(X5,sk_c6)!=X4|inverse(X1)!=sk_c8|inverse(X5)!=sk_c6|inverse(X2)!=sk_c6|inverse(X3)!=sk_c8)).
##
#cnf(i_0_6051, negated_conjecture, (multiply(sk_c5,sk_c6)=inverse(sk_c8)|sk_c4=identity)).
##
#cnf(i_0_6152, negated_conjecture, (multiply(X1,sk_c8)!=identity|multiply(X2,sk_c6)!=identity|multiply(sk_c6,X3)!=identity|inverse(sk_c6)!=identity|multiply(X4,sk_c6)!=X3|inverse(X1)!=sk_c8|inverse(X4)!=sk_c6|inverse(X2)!=sk_c6)).
####
#cnf(i_0_6262, negated_conjecture, (multiply(X1,sk_c8)!=identity|multiply(sk_c6,X2)!=identity|inverse(sk_c6)!=identity|multiply(X3,sk_c6)!=X2|inverse(X1)!=sk_c8|inverse(X3)!=sk_c6)).
##
#cnf(i_0_5471, negated_conjecture, (multiply(inverse(X1),sk_c6)=inverse(X1)|sk_c3=sk_c2)).
##
#cnf(i_0_6336, negated_conjecture, (multiply(sk_c6,X1)!=identity|inverse(sk_c6)!=identity|multiply(X2,sk_c6)!=X1|inverse(X2)!=sk_c6)).
###
#cnf(i_0_6345, negated_conjecture, (inverse(sk_c6)=sk_c6|sk_c3=sk_c2)).
#
#cnf(i_0_6403, negated_conjecture, (multiply(sk_c6,X1)!=identity|inverse(sk_c6)!=identity|identity!=X1)).
####
#cnf(i_0_6450, negated_conjecture, (inverse(sk_c6)!=identity)).
#
#cnf(i_0_731, negated_conjecture, (sk_c4=identity)).
#
#cnf(i_0_290, negated_conjecture, (sk_c3=sk_c2)).
#
#cnf(i_0_170, negated_conjecture, (inverse(sk_c6)=sk_c2)).
###
#cnf(i_0_49, negated_conjecture, (inverse(sk_c2)=sk_c6)).
###
#cnf(i_0_34, negated_conjecture, (inverse(sk_c1)=sk_c8)).
##
#cnf(i_0_6486, negated_conjecture, (inverse(sk_c8)=sk_c1)).
####
#cnf(i_0_6471, negated_conjecture, (multiply(sk_c2,sk_c6)=identity)).
##
#cnf(i_0_6473, negated_conjecture, (multiply(sk_c6,sk_c2)=identity)).
####
#cnf(i_0_6483, negated_conjecture, (multiply(sk_c8,sk_c1)=identity)).
#
#cnf(i_0_6485, negated_conjecture, (multiply(sk_c1,sk_c8)=identity)).
#########
#cnf(i_0_6472, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,X1))=X1)).
######
#cnf(i_0_6475, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,X1))=X1)).
######
#cnf(i_0_6484, negated_conjecture, (multiply(sk_c8,multiply(sk_c1,X1))=X1)).
##
#cnf(i_0_6487, negated_conjecture, (multiply(sk_c1,multiply(sk_c8,X1))=X1)).
#########
#cnf(i_0_167, negated_conjecture, (inverse(sk_c6)=sk_c2|inverse(sk_c5)=sk_c8)).
###########
#cnf(i_0_43, negated_conjecture, (sk_c6=identity)).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(associativity, axiom, (multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), file('./p/GRP268-1.tptp', associativity)).
cnf(left_inverse, axiom, (multiply(inverse(X1),X1)=identity), file('./p/GRP268-1.tptp', left_inverse)).
cnf(left_identity, axiom, (multiply(identity,X1)=X1), file('./p/GRP268-1.tptp', left_identity)).
cnf(prove_this_11, negated_conjecture, (inverse(sk_c2)=sk_c6|inverse(sk_c5)=sk_c8), file('./p/GRP268-1.tptp', prove_this_11)).
cnf(prove_this_8, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c7|inverse(sk_c5)=sk_c8), file('./p/GRP268-1.tptp', prove_this_8)).
cnf(prove_this_4, negated_conjecture, (inverse(sk_c1)=sk_c8|multiply(sk_c5,sk_c8)=sk_c7), file('./p/GRP268-1.tptp', prove_this_4)).
cnf(prove_this_12, negated_conjecture, (inverse(sk_c2)=sk_c6|inverse(sk_c6)=sk_c7), file('./p/GRP268-1.tptp', prove_this_12)).
cnf(prove_this_1, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|multiply(sk_c5,sk_c8)=sk_c7), file('./p/GRP268-1.tptp', prove_this_1)).
cnf(prove_this_22, negated_conjecture, (multiply(X1,sk_c8)!=sk_c7|inverse(X1)!=sk_c8|multiply(X2,sk_c6)!=sk_c7|inverse(X2)!=sk_c6|multiply(sk_c6,X3)!=sk_c7|multiply(X4,sk_c6)!=X3|inverse(X4)!=sk_c6|multiply(X5,sk_c8)!=sk_c7|inverse(X5)!=sk_c8|inverse(sk_c6)!=sk_c7), file('./p/GRP268-1.tptp', prove_this_22)).
cnf(prove_this_21, negated_conjecture, (inverse(sk_c3)=sk_c6|inverse(sk_c6)=sk_c7), file('./p/GRP268-1.tptp', prove_this_21)).
cnf(prove_this_18, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c4|inverse(sk_c6)=sk_c7), file('./p/GRP268-1.tptp', prove_this_18)).
cnf(prove_this_15, negated_conjecture, (multiply(sk_c6,sk_c4)=sk_c7|inverse(sk_c6)=sk_c7), file('./p/GRP268-1.tptp', prove_this_15)).
cnf(c_0_12, axiom, (multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), associativity).
cnf(c_0_13, axiom, (multiply(inverse(X1),X1)=identity), left_inverse).
cnf(c_0_14, axiom, (multiply(identity,X1)=X1), left_identity).
cnf(c_0_15, plain, (multiply(inverse(X1),multiply(X1,X2))=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_12, c_0_13]), c_0_14])).
cnf(c_0_16, plain, (multiply(inverse(inverse(X1)),identity)=X1), inference(pm,[status(thm)],[c_0_15, c_0_13])).
cnf(c_0_17, plain, (multiply(inverse(inverse(X1)),X2)=multiply(X1,X2)), inference(pm,[status(thm)],[c_0_15, c_0_15])).
cnf(c_0_18, plain, (multiply(X1,identity)=X1), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_19, plain, (inverse(inverse(X1))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_17, c_0_18]), c_0_18])).
cnf(c_0_20, plain, (multiply(X1,inverse(X1))=identity), inference(pm,[status(thm)],[c_0_13, c_0_19])).
cnf(c_0_21, negated_conjecture, (inverse(sk_c2)=sk_c6|inverse(sk_c5)=sk_c8), prove_this_11).
cnf(c_0_22, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c7|inverse(sk_c5)=sk_c8), prove_this_8).
cnf(c_0_23, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|inverse(sk_c5)=sk_c8), inference(pm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_24, negated_conjecture, (inverse(sk_c5)=sk_c8|sk_c7=identity), inference(pm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_25, negated_conjecture, (inverse(sk_c1)=sk_c8|multiply(sk_c5,sk_c8)=sk_c7), prove_this_4).
cnf(c_0_26, negated_conjecture, (multiply(sk_c5,sk_c8)=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_20, c_0_24])).
cnf(c_0_27, plain, (multiply(X1,multiply(inverse(X1),X2))=X2), inference(pm,[status(thm)],[c_0_15, c_0_19])).
cnf(c_0_28, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c7=identity), inference(pm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_29, negated_conjecture, (inverse(sk_c2)=sk_c6|inverse(sk_c6)=sk_c7), prove_this_12).
cnf(c_0_30, negated_conjecture, (multiply(sk_c5,multiply(sk_c8,X1))=X1|sk_c7=identity), inference(pm,[status(thm)],[c_0_27, c_0_24])).
cnf(c_0_31, negated_conjecture, (multiply(sk_c8,multiply(sk_c1,X1))=X1|sk_c7=identity), inference(pm,[status(thm)],[c_0_15, c_0_28])).
cnf(c_0_32, negated_conjecture, (inverse(sk_c6)=sk_c2|inverse(sk_c6)=sk_c7), inference(pm,[status(thm)],[c_0_19, c_0_29])).
cnf(c_0_33, negated_conjecture, (multiply(sk_c5,X1)=multiply(sk_c1,X1)|sk_c7=identity), inference(pm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|multiply(sk_c5,sk_c8)=sk_c7), prove_this_1).
cnf(c_0_35, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c2!=sk_c7), inference(ef,[status(thm)],[c_0_32])).
cnf(c_0_36, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_20, c_0_28])).
cnf(c_0_37, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|sk_c7=identity), inference(pm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_38, negated_conjecture, (multiply(X1,sk_c8)!=sk_c7|inverse(X1)!=sk_c8|multiply(X2,sk_c6)!=sk_c7|inverse(X2)!=sk_c6|multiply(sk_c6,X3)!=sk_c7|multiply(X4,sk_c6)!=X3|inverse(X4)!=sk_c6|multiply(X5,sk_c8)!=sk_c7|inverse(X5)!=sk_c8|inverse(sk_c6)!=sk_c7), prove_this_22).
cnf(c_0_39, negated_conjecture, (multiply(sk_c7,multiply(sk_c6,X1))=X1|sk_c2!=sk_c7), inference(pm,[status(thm)],[c_0_15, c_0_35])).
cnf(c_0_40, negated_conjecture, (sk_c7=identity), inference(pm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41, negated_conjecture, (multiply(sk_c6,multiply(X1,sk_c6))!=sk_c7|multiply(X2,sk_c8)!=sk_c7|multiply(X3,sk_c6)!=sk_c7|multiply(X4,sk_c8)!=sk_c7|inverse(sk_c6)!=sk_c7|inverse(X2)!=sk_c8|inverse(X1)!=sk_c6|inverse(X3)!=sk_c6|inverse(X4)!=sk_c8), inference(er,[status(thm)],[c_0_38])).
cnf(c_0_42, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c7,X1)|sk_c2!=sk_c7), inference(pm,[status(thm)],[c_0_39, c_0_27])).
cnf(c_0_43, plain, (multiply(X1,multiply(X2,inverse(multiply(X1,X2))))=identity), inference(pm,[status(thm)],[c_0_12, c_0_20])).
cnf(c_0_44, negated_conjecture, (multiply(X1,sk_c8)!=identity|multiply(X2,sk_c6)!=identity|multiply(X3,sk_c8)!=identity|multiply(sk_c6,X4)!=identity|inverse(sk_c6)!=identity|multiply(X5,sk_c6)!=X4|inverse(X1)!=sk_c8|inverse(X5)!=sk_c6|inverse(X2)!=sk_c6|inverse(X3)!=sk_c8), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_40]), c_0_40]), c_0_40]), c_0_40]), c_0_40])).
cnf(c_0_45, negated_conjecture, (multiply(sk_c6,multiply(X1,sk_c6))!=identity|multiply(X2,sk_c8)!=identity|multiply(X3,sk_c6)!=identity|multiply(X4,sk_c8)!=identity|inverse(sk_c6)!=identity|inverse(X2)!=sk_c8|inverse(X1)!=sk_c6|inverse(X3)!=sk_c6|inverse(X4)!=sk_c8), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_41, c_0_40]), c_0_40]), c_0_40]), c_0_40]), c_0_40])).
cnf(c_0_46, negated_conjecture, (multiply(inverse(sk_c6),X1)=X1|sk_c2!=identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_40]), c_0_14]), c_0_40])).
cnf(c_0_47, plain, (multiply(X1,inverse(multiply(X2,X1)))=inverse(X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_15, c_0_43]), c_0_18])).
cnf(c_0_48, negated_conjecture, (multiply(X1,sk_c8)!=identity|multiply(X2,sk_c6)!=identity|multiply(sk_c6,X3)!=identity|inverse(sk_c6)!=identity|multiply(X4,sk_c6)!=X3|inverse(X1)!=sk_c8|inverse(X4)!=sk_c6|inverse(X2)!=sk_c6), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_44, c_0_13]), c_0_19])])).
cnf(c_0_49, negated_conjecture, (multiply(sk_c6,sk_c6)!=identity|multiply(X1,sk_c8)!=identity|multiply(X2,sk_c6)!=identity|multiply(X3,sk_c8)!=identity|inverse(sk_c6)!=identity|inverse(X1)!=sk_c8|inverse(X2)!=sk_c6|inverse(X3)!=sk_c8|sk_c2!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_45, c_0_46]), c_0_19])])).
cnf(c_0_50, plain, (inverse(multiply(X1,X2))=multiply(inverse(X2),inverse(X1))), inference(pm,[status(thm)],[c_0_15, c_0_47])).
cnf(c_0_51, plain, (multiply(inverse(identity),X1)=X1), inference(pm,[status(thm)],[c_0_15, c_0_14])).
cnf(c_0_52, negated_conjecture, (multiply(X1,sk_c8)!=identity|multiply(sk_c6,X2)!=identity|inverse(sk_c6)!=identity|multiply(X3,sk_c6)!=X2|inverse(X1)!=sk_c8|inverse(X3)!=sk_c6), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_48, c_0_13]), c_0_19])])).
cnf(c_0_53, negated_conjecture, (inverse(sk_c3)=sk_c6|inverse(sk_c6)=sk_c7), prove_this_21).
cnf(c_0_54, negated_conjecture, (multiply(sk_c6,sk_c6)!=identity|multiply(X1,sk_c8)!=identity|multiply(X2,sk_c6)!=identity|inverse(sk_c6)!=identity|inverse(X1)!=sk_c8|inverse(X2)!=sk_c6|sk_c2!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_49, c_0_13]), c_0_19])])).
cnf(c_0_55, negated_conjecture, (multiply(inverse(X1),sk_c6)=inverse(X1)|sk_c2!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_50, c_0_46]), c_0_19])).
cnf(c_0_56, plain, (multiply(inverse(inverse(identity)),X1)=X1), inference(pm,[status(thm)],[c_0_15, c_0_51])).
cnf(c_0_57, negated_conjecture, (multiply(sk_c6,X1)!=identity|inverse(sk_c6)!=identity|multiply(X2,sk_c6)!=X1|inverse(X2)!=sk_c6), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_52, c_0_13]), c_0_19])])).
cnf(c_0_58, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c4|inverse(sk_c6)=sk_c7), prove_this_18).
cnf(c_0_59, negated_conjecture, (multiply(sk_c3,sk_c6)=identity|inverse(sk_c6)=sk_c7), inference(pm,[status(thm)],[c_0_20, c_0_53])).
cnf(c_0_60, negated_conjecture, (multiply(sk_c6,sk_c6)!=identity|multiply(X1,sk_c8)!=identity|inverse(sk_c6)!=identity|inverse(X1)!=sk_c8|sk_c2!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_54, c_0_13]), c_0_19])])).
cnf(c_0_61, negated_conjecture, (inverse(sk_c6)=sk_c6|sk_c2!=identity), inference(pm,[status(thm)],[c_0_46, c_0_55])).
cnf(c_0_62, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|sk_c2!=sk_c7), inference(pm,[status(thm)],[c_0_20, c_0_35])).
cnf(c_0_63, plain, (inverse(identity)=identity), inference(pm,[status(thm)],[c_0_13, c_0_56])).
cnf(c_0_64, negated_conjecture, (multiply(sk_c6,X1)!=identity|inverse(sk_c6)!=identity|identity!=X1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_57, c_0_13]), c_0_19])])).
cnf(c_0_65, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c4=identity), inference(pm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_66, negated_conjecture, (multiply(sk_c6,sk_c6)!=identity|inverse(sk_c6)!=identity|sk_c2!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_60, c_0_13]), c_0_19])])).
cnf(c_0_67, negated_conjecture, (multiply(sk_c6,sk_c6)=identity|sk_c2!=identity), inference(pm,[status(thm)],[c_0_13, c_0_61])).
cnf(c_0_68, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c6))=identity|sk_c2!=sk_c7), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_50, c_0_62]), c_0_63])).
cnf(c_0_69, negated_conjecture, (inverse(sk_c6)=sk_c2|inverse(sk_c6)=identity), inference(rw,[status(thm)],[c_0_32, c_0_40])).
cnf(c_0_70, negated_conjecture, (inverse(sk_c6)!=identity), inference(pm,[status(thm)],[c_0_64, c_0_20])).
cnf(c_0_71, negated_conjecture, (inverse(sk_c6)=identity|sk_c4=identity), inference(rw,[status(thm)],[c_0_65, c_0_40])).
cnf(c_0_72, negated_conjecture, (inverse(sk_c6)!=identity|sk_c2!=identity), inference(pm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_73, negated_conjecture, (inverse(sk_c6)=identity|sk_c2!=identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_68, c_0_40]), c_0_63]), c_0_14]), c_0_40])).
cnf(c_0_74, negated_conjecture, (inverse(sk_c6)=sk_c2), inference(sr,[status(thm)],[c_0_69, c_0_70])).
cnf(c_0_75, negated_conjecture, (multiply(sk_c6,sk_c4)=sk_c7|inverse(sk_c6)=sk_c7), prove_this_15).
cnf(c_0_76, negated_conjecture, (sk_c4=identity), inference(sr,[status(thm)],[c_0_71, c_0_70])).
cnf(c_0_77, negated_conjecture, (sk_c2!=identity), inference(pm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_78, negated_conjecture, (multiply(sk_c6,sk_c2)=identity), inference(pm,[status(thm)],[c_0_20, c_0_74])).
cnf(c_0_79, negated_conjecture, (sk_c6=identity), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_40]), c_0_40]), c_0_76]), c_0_18]), c_0_74]), c_0_77])).
cnf(c_0_80, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_79]), c_0_14]), c_0_77]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 81
# Proof object clause steps            : 69
# Proof object formula steps           : 12
# Proof object conjectures             : 62
# Proof object clause conjectures      : 53
# Proof object formula conjectures     : 9
# Proof object initial clauses used    : 12
# Proof object initial formulas used   : 12
# Proof object generating inferences   : 46
# Proof object simplifying inferences  : 52
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 25
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 25
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 25
# Processed clauses                    : 1789
# ...of these trivial                  : 25
# ...subsumed                          : 1399
# ...remaining for further processing  : 365
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 50
# Backward-rewritten                   : 287
# Generated clauses                    : 6505
# ...of the previous two non-trivial   : 6076
# Contextual simplify-reflections      : 0
# Paramodulations                      : 6489
# Factorizations                       : 4
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 21
#    Positive orientable unit clauses  : 20
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 3676
# ...number of literals in the above   : 22346
# Current number of archived formulas  : 0
# Current number of archived clauses   : 344
# Clause-clause subsumption calls (NU) : 6298
# Rec. Clause-clause subsumption calls : 3872
# Non-unit clause-clause subsumptions  : 1353
# Unit Clause-clause subsumption calls : 105
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 20
# BW rewrite match successes           : 19
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 33170
