# Initializing proof state
# Scanning for AC axioms
# multiply is associative
#
#cnf(i_0_41, plain, (multiply(identity,X1)=X1)).
#
#cnf(i_0_42, plain, (multiply(inverse(X1),X1)=identity)).
#
#cnf(i_0_58, negated_conjecture, (inverse(sk_c3)=sk_c7|inverse(sk_c1)=sk_c8)).
#
#cnf(i_0_80, negated_conjecture, (multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X3,sk_c7)!=sk_c8|multiply(X2,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c7|multiply(sk_c8,X4)!=sk_c7|multiply(X5,sk_c8)!=X4|inverse(X5)!=sk_c8|inverse(X3)!=sk_c7|inverse(X2)!=sk_c8|inverse(X1)!=sk_c8)).
#
#cnf(i_0_76, negated_conjecture, (inverse(sk_c3)=sk_c7|inverse(sk_c2)=sk_c8)).
#
#cnf(i_0_61, negated_conjecture, (inverse(sk_c4)=sk_c8|inverse(sk_c1)=sk_c8)).
#
#cnf(i_0_79, negated_conjecture, (inverse(sk_c4)=sk_c8|inverse(sk_c2)=sk_c8)).
#
#cnf(i_0_82, negated_conjecture, (multiply(sk_c8,multiply(X4,sk_c8))!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c7)!=sk_c8|multiply(X2,sk_c8)!=sk_c6|multiply(X3,sk_c8)!=sk_c7|inverse(X4)!=sk_c8|inverse(X1)!=sk_c7|inverse(X2)!=sk_c8|inverse(X3)!=sk_c8)).
#
#cnf(i_0_43, plain, (multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3)))).
#
#cnf(i_0_101, plain, (multiply(inverse(X1),multiply(X1,X2))=X2)).
#
#cnf(i_0_107, plain, (multiply(inverse(identity),X1)=X1)).
#
#cnf(i_0_91, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c7)!=sk_c8|multiply(X2,sk_c8)!=sk_c6|multiply(X3,sk_c8)!=sk_c7|inverse(identity)!=sk_c8|inverse(X1)!=sk_c7|inverse(X2)!=sk_c8|inverse(X3)!=sk_c8)).
#
#cnf(i_0_105, plain, (multiply(inverse(inverse(X1)),identity)=X1)).
#
#cnf(i_0_118, plain, (multiply(inverse(inverse(identity)),X1)=X1)).
#
#cnf(i_0_146, plain, (inverse(identity)=identity)).
#
#cnf(i_0_91, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c7)!=sk_c8|multiply(X2,sk_c8)!=sk_c6|multiply(X3,sk_c8)!=sk_c7|inverse(X1)!=sk_c7|inverse(X2)!=sk_c8|inverse(X3)!=sk_c8|sk_c8!=identity)).
#
#cnf(i_0_104, plain, (multiply(inverse(inverse(X1)),X2)=multiply(X1,X2))).
#
#cnf(i_0_105, plain, (multiply(X1,identity)=X1)).
#
#cnf(i_0_176, plain, (inverse(inverse(X1))=X1)).
#
#cnf(i_0_154, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c6|multiply(X2,sk_c8)!=sk_c7|inverse(X1)!=sk_c8|inverse(X2)!=sk_c8|sk_c8!=identity)).
#
#cnf(i_0_181, plain, (multiply(X1,inverse(X1))=identity)).
#
#cnf(i_0_182, plain, (multiply(X1,multiply(inverse(X1),X2))=X2)).
##
#cnf(i_0_188, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_183, negated_conjecture, (inverse(sk_c8)=sk_c2|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_184, negated_conjecture, (inverse(sk_c8)=sk_c1|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_230, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c2=sk_c1)).
#
#cnf(i_0_218, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_234, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c2=sk_c1)).
##
#cnf(i_0_185, negated_conjecture, (inverse(sk_c8)=sk_c4|inverse(sk_c1)=sk_c8)).
#
#cnf(i_0_190, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c7|inverse(X1)!=sk_c8|sk_c6!=identity|sk_c8!=identity)).
#
#cnf(i_0_186, negated_conjecture, (inverse(sk_c8)=sk_c4|inverse(sk_c2)=sk_c8)).
#####
#cnf(i_0_253, negated_conjecture, (inverse(sk_c8)=sk_c2|inverse(sk_c8)=sk_c4)).
#
#cnf(i_0_255, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c4!=sk_c2)).
#
#cnf(i_0_219, negated_conjecture, (multiply(X1,multiply(X2,sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|inverse(multiply(X1,X2))!=sk_c8|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_266, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c4!=sk_c2)).
##
#cnf(i_0_263, negated_conjecture, (multiply(sk_c2,sk_c8)=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_275, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
##
#cnf(i_0_265, negated_conjecture, (multiply(sk_c8,sk_c2)=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_259, negated_conjecture, (inverse(sk_c4)=sk_c8|inverse(sk_c8)=sk_c2)).
###
#cnf(i_0_231, negated_conjecture, (multiply(sk_c7,sk_c3)=identity|sk_c2=sk_c1)).
###
#cnf(i_0_233, negated_conjecture, (multiply(sk_c3,sk_c7)=identity|sk_c2=sk_c1)).
#####
#cnf(i_0_193, plain, (multiply(X1,multiply(X2,inverse(multiply(X1,X2))))=identity)).
#
#cnf(i_0_312, plain, (multiply(X2,inverse(multiply(X1,X2)))=inverse(X1))).
#####
#cnf(i_0_271, negated_conjecture, (inverse(multiply(X1,inverse(sk_c8)))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6)).
##
#cnf(i_0_334, plain, (multiply(inverse(multiply(X1,X2)),X1)=inverse(X2))).
#######
#cnf(i_0_333, plain, (inverse(multiply(X2,X1))=multiply(inverse(X1),inverse(X2)))).
#
#cnf(i_0_46, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c6|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_64, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c7|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_271, negated_conjecture, (multiply(sk_c8,inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6)).
#
#cnf(i_0_52, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_70, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_49, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c6|inverse(sk_c4)=sk_c8)).
##
#cnf(i_0_67, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c7|inverse(sk_c4)=sk_c8)).
#
#cnf(i_0_55, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|inverse(sk_c4)=sk_c8)).
#
#cnf(i_0_73, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|inverse(sk_c4)=sk_c8)).
##
#cnf(i_0_59, negated_conjecture, (multiply(sk_c8,sk_c5)=sk_c7|inverse(sk_c1)=sk_c8)).
#
#cnf(i_0_56, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c6|inverse(sk_c1)=sk_c8)).
#
#cnf(i_0_57, negated_conjecture, (multiply(sk_c3,sk_c7)=sk_c8|inverse(sk_c1)=sk_c8)).
#
#cnf(i_0_431, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c8,X1)!=sk_c8|inverse(X1)!=sk_c6|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_60, negated_conjecture, (multiply(sk_c4,sk_c8)=sk_c5|inverse(sk_c1)=sk_c8)).
#
#cnf(i_0_77, negated_conjecture, (multiply(sk_c8,sk_c5)=sk_c7|inverse(sk_c2)=sk_c8)).
#
#cnf(i_0_74, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c6|inverse(sk_c2)=sk_c8)).
##
#cnf(i_0_75, negated_conjecture, (multiply(sk_c3,sk_c7)=sk_c8|inverse(sk_c2)=sk_c8)).
#
#cnf(i_0_78, negated_conjecture, (multiply(sk_c4,sk_c8)=sk_c5|inverse(sk_c2)=sk_c8)).
#
#cnf(i_0_81, negated_conjecture, (multiply(sk_c8,sk_c1)=identity|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_493, negated_conjecture, (multiply(inverse(X1),sk_c8)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity|X1!=sk_c8)).
##
#cnf(i_0_87, negated_conjecture, (multiply(sk_c8,sk_c2)=identity|inverse(sk_c3)=sk_c7)).
###
#cnf(i_0_88, negated_conjecture, (multiply(sk_c8,sk_c4)=identity|inverse(sk_c1)=sk_c8)).
##
#cnf(i_0_89, negated_conjecture, (multiply(sk_c8,sk_c4)=identity|inverse(sk_c2)=sk_c8)).
###
#cnf(i_0_197, negated_conjecture, (multiply(sk_c2,sk_c8)=identity|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_575, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c6=identity)).
#
#cnf(i_0_549, negated_conjecture, (multiply(sk_c2,sk_c8)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_580, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c6=identity)).
####
#cnf(i_0_577, negated_conjecture, (multiply(sk_c7,sk_c3)=identity|sk_c6=identity)).
##
#cnf(i_0_579, negated_conjecture, (multiply(sk_c3,sk_c7)=identity|sk_c6=identity)).
#####
#cnf(i_0_419, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity|sk_c6!=sk_c2|sk_c4!=sk_c2)).
#
#cnf(i_0_198, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_609, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c7=identity)).
#
#cnf(i_0_614, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c7=identity)).
####
#cnf(i_0_611, negated_conjecture, (multiply(sk_c7,sk_c3)=identity|sk_c7=identity)).
#
#cnf(i_0_539, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity|sk_c2!=sk_c8|sk_c4!=sk_c2)).
##
#cnf(i_0_613, negated_conjecture, (multiply(sk_c3,sk_c7)=identity|sk_c7=identity)).
##
#cnf(i_0_249, negated_conjecture, (multiply(X1,multiply(X2,sk_c8))!=sk_c7|multiply(inverse(X2),inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c8!=identity)).
###
#cnf(i_0_199, negated_conjecture, (multiply(sk_c4,sk_c8)=identity|inverse(sk_c1)=sk_c8)).
##
#cnf(i_0_666, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c5=identity)).
#
#cnf(i_0_671, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c5=identity)).
##
#cnf(i_0_650, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c8!=identity|sk_c6!=identity)).
##
#cnf(i_0_684, negated_conjecture, (sk_c2=sk_c1|sk_c5=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_668, negated_conjecture, (multiply(sk_c8,sk_c1)=identity|sk_c5=identity)).
###
#cnf(i_0_670, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|sk_c5=identity)).
##
#cnf(i_0_646, negated_conjecture, (multiply(sk_c8,inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c8!=identity|X1!=sk_c7)).
###
#cnf(i_0_200, negated_conjecture, (multiply(sk_c4,sk_c8)=identity|inverse(sk_c2)=sk_c8)).
##
#cnf(i_0_737, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c5=identity)).
#
#cnf(i_0_741, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c5=identity)).
#
#cnf(i_0_757, negated_conjecture, (sk_c2=sk_c1|sk_c5=identity)).
####
#cnf(i_0_738, negated_conjecture, (multiply(sk_c8,sk_c2)=identity|sk_c5=identity)).
###
#cnf(i_0_740, negated_conjecture, (multiply(sk_c2,sk_c8)=identity|sk_c5=identity)).
######
#cnf(i_0_720, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c8,X1)!=sk_c8|inverse(X1)!=sk_c7|sk_c6!=identity|sk_c8!=identity)).
#######
#cnf(i_0_250, negated_conjecture, (multiply(sk_c8,sk_c2)=identity|inverse(sk_c8)=sk_c4)).
#
#cnf(i_0_779, negated_conjecture, (multiply(inverse(X1),sk_c8)!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c8!=identity|X1!=sk_c8)).
##
#cnf(i_0_252, negated_conjecture, (multiply(sk_c2,sk_c8)=identity|inverse(sk_c8)=sk_c4)).
###
#cnf(i_0_256, negated_conjecture, (multiply(sk_c4,sk_c8)=identity|inverse(sk_c8)=sk_c2)).
##
#cnf(i_0_258, negated_conjecture, (multiply(sk_c8,sk_c4)=identity|inverse(sk_c8)=sk_c2)).
#####
#cnf(i_0_791, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c8!=identity|sk_c2!=sk_c8|sk_c4!=sk_c2)).
###
#cnf(i_0_479, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c6=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_807, negated_conjecture, (multiply(sk_c2,sk_c8)!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_852, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c6=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_264, negated_conjecture, (multiply(sk_c2,multiply(sk_c8,X1))=X1|sk_c4!=sk_c2)).
#######
#cnf(i_0_267, negated_conjecture, (multiply(sk_c8,multiply(sk_c2,X1))=X1|sk_c4!=sk_c2)).
##
#cnf(i_0_702, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c8!=identity|sk_c2!=sk_c7|sk_c4!=sk_c2)).
#######
#cnf(i_0_491, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c8=identity|sk_c2=sk_c1)).
#
#cnf(i_0_638, negated_conjecture, (multiply(inverse(sk_c2),inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c8!=identity|sk_c4!=sk_c2|X1!=sk_c7)).
#
#cnf(i_0_522, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c8=identity|sk_c2=sk_c1)).
#
#cnf(i_0_601, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c8=identity|sk_c6=identity)).
#
#cnf(i_0_602, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c8=identity|sk_c6=identity)).
##
#cnf(i_0_633, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c8=identity|sk_c7=identity)).
#
#cnf(i_0_634, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c8=identity|sk_c7=identity)).
##
#cnf(i_0_960, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|inverse(sk_c2)!=sk_c7|sk_c8!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_699, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c7=identity|sk_c5=identity)).
#####
#cnf(i_0_758, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c8=sk_c7|sk_c2=sk_c1)).
#
#cnf(i_0_759, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c8=sk_c7|sk_c2=sk_c1)).
#
#cnf(i_0_953, negated_conjecture, (multiply(inverse(sk_c2),X1)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c7|sk_c6!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_776, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c6=identity|sk_c5=identity)).
#
#cnf(i_0_922, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c8=identity|sk_c2=sk_c1)).
##
#cnf(i_0_1082, negated_conjecture, (sk_c2=sk_c1|sk_c8=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_964, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c8=identity|sk_c2=sk_c1)).
##
#cnf(i_0_1099, negated_conjecture, (sk_c2=sk_c1|sk_c8=identity)).
#
#cnf(i_0_1113, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c8=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_1128, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c8=identity|sk_c4!=sk_c2)).
####
#cnf(i_0_975, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c8=identity|sk_c6=identity)).
##
#cnf(i_0_985, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c8=identity|sk_c6=identity)).
##
#cnf(i_0_996, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c8=identity|sk_c7=identity)).
#
#cnf(i_0_930, negated_conjecture, (multiply(inverse(sk_c2),sk_c8)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c8!=identity|sk_c4!=sk_c2|sk_c2!=sk_c7)).
#
#cnf(i_0_1006, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c8=identity|sk_c7=identity)).
##
#cnf(i_0_1019, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c7=identity|sk_c5=identity)).
##
#cnf(i_0_1240, negated_conjecture, (sk_c4=sk_c1|sk_c5=identity|sk_c7=identity)).
#
#cnf(i_0_1241, negated_conjecture, (sk_c4=sk_c2|sk_c5=identity|sk_c7=identity)).
#
#cnf(i_0_1029, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c8=sk_c7|sk_c2=sk_c1)).
#
#cnf(i_0_905, negated_conjecture, (multiply(inverse(X1),inverse(sk_c2))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2|X1!=sk_c8)).
#
#cnf(i_0_1263, negated_conjecture, (sk_c2=sk_c1|sk_c8=sk_c7|sk_c4!=sk_c2)).
##
#cnf(i_0_1039, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c8=sk_c7|sk_c2=sk_c1)).
##
#cnf(i_0_1340, negated_conjecture, (sk_c2=sk_c1|sk_c8=sk_c7)).
#
#cnf(i_0_1354, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c8=sk_c7|sk_c4!=sk_c2)).
#
#cnf(i_0_1370, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c8=sk_c7|sk_c4!=sk_c2)).
#
#cnf(i_0_1320, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|inverse(sk_c2)!=sk_c6|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_1063, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c6=identity|sk_c5=identity)).
##
#cnf(i_0_1414, negated_conjecture, (sk_c4=sk_c1|sk_c5=identity|sk_c6=identity)).
#
#cnf(i_0_1415, negated_conjecture, (sk_c4=sk_c2|sk_c5=identity|sk_c6=identity)).
######
#cnf(i_0_1280, negated_conjecture, (multiply(sk_c8,inverse(sk_c2))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2|sk_c2!=sk_c8)).
########
#cnf(i_0_906, negated_conjecture, (multiply(inverse(X1),inverse(sk_c2))!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c8!=identity|sk_c4!=sk_c2|X1!=sk_c8)).
#
#cnf(i_0_232, negated_conjecture, (multiply(sk_c7,multiply(sk_c3,X1))=X1|sk_c2=sk_c1)).
####
#cnf(i_0_235, negated_conjecture, (multiply(sk_c3,multiply(sk_c7,X1))=X1|sk_c2=sk_c1)).
#########
#cnf(i_0_578, negated_conjecture, (multiply(sk_c7,multiply(sk_c3,X1))=X1|sk_c6=identity)).
####
#cnf(i_0_581, negated_conjecture, (multiply(sk_c3,multiply(sk_c7,X1))=X1|sk_c6=identity)).
#########
#cnf(i_0_612, negated_conjecture, (multiply(sk_c7,multiply(sk_c3,X1))=X1|sk_c7=identity)).
####
#cnf(i_0_615, negated_conjecture, (multiply(sk_c3,multiply(sk_c7,X1))=X1|sk_c7=identity)).
#
#cnf(i_0_1318, negated_conjecture, (multiply(X1,inverse(sk_c2))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#########
#cnf(i_0_669, negated_conjecture, (multiply(sk_c8,multiply(sk_c1,X1))=X1|sk_c5=identity)).
####
#cnf(i_0_672, negated_conjecture, (multiply(sk_c1,multiply(sk_c8,X1))=X1|sk_c5=identity)).
###
#cnf(i_0_1482, negated_conjecture, (multiply(X1,inverse(sk_c2))!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###############
#cnf(i_0_739, negated_conjecture, (multiply(sk_c8,multiply(sk_c2,X1))=X1|sk_c5=identity)).
#
#cnf(i_0_219, negated_conjecture, (multiply(X1,multiply(X2,sk_c8))!=sk_c6|multiply(inverse(X2),inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity)).
###
#cnf(i_0_1629, negated_conjecture, (multiply(sk_c2,X1)=multiply(sk_c1,X1)|sk_c5=identity)).
####
#cnf(i_0_1692, negated_conjecture, (multiply(sk_c1,inverse(sk_c2))=identity|sk_c5=identity)).
###
#cnf(i_0_1712, negated_conjecture, (inverse(sk_c2)=inverse(sk_c1)|sk_c5=identity)).
#######
#cnf(i_0_1713, negated_conjecture, (multiply(sk_c2,inverse(sk_c1))=identity|sk_c5=identity)).
######
#cnf(i_0_1723, negated_conjecture, (multiply(inverse(sk_c1),sk_c2)=identity|sk_c5=identity)).
#####
#cnf(i_0_1647, negated_conjecture, (multiply(inverse(sk_c2),inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2|X1!=sk_c6)).
###
#cnf(i_0_1766, negated_conjecture, (multiply(inverse(sk_c2),sk_c1)=identity|sk_c5=identity)).
#######
#cnf(i_0_742, negated_conjecture, (multiply(sk_c2,multiply(sk_c8,X1))=X1|sk_c5=identity)).
##########
#cnf(i_0_1817, negated_conjecture, (multiply(inverse(sk_c2),X1)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c6|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##########
#cnf(i_0_879, negated_conjecture, (multiply(inverse(sk_c2),X1)=multiply(sk_c8,X1)|sk_c4!=sk_c2)).
###########
#cnf(i_0_883, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c2,X1)|sk_c4!=sk_c2)).
#############
#cnf(i_0_1487, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c3,X1)|sk_c2=sk_c1)).
###########
#cnf(i_0_1491, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c7,X1)|sk_c2=sk_c1)).
#############
#cnf(i_0_1511, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c3,X1)|sk_c6=identity)).
###########
#cnf(i_0_1515, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c7,X1)|sk_c6=identity)).
###
#cnf(i_0_1892, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c2!=sk_c8|sk_c8!=identity|sk_c4!=sk_c2)).
########
#cnf(i_0_2086, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c2!=sk_c8|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_1535, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c3,X1)|sk_c7=identity)).
##
#cnf(i_0_2091, negated_conjecture, (multiply(sk_c2,sk_c8)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c2!=sk_c8|sk_c8!=identity|sk_c4!=sk_c2)).
#########
#cnf(i_0_1539, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c7,X1)|sk_c7=identity)).
#############
#cnf(i_0_1572, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c1,X1)|sk_c5=identity)).
######
#cnf(i_0_1778, negated_conjecture, (multiply(inverse(sk_c2),sk_c8)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2|sk_c6!=sk_c2)).
######
#cnf(i_0_1581, negated_conjecture, (multiply(inverse(sk_c1),X1)=multiply(sk_c8,X1)|sk_c5=identity)).
###############
#cnf(i_0_1624, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c2,X1)|sk_c5=identity)).
#############
#cnf(i_0_1635, negated_conjecture, (multiply(inverse(sk_c2),X1)=multiply(sk_c8,X1)|sk_c5=identity)).
##
#cnf(i_0_432, negated_conjecture, (multiply(sk_c8,multiply(inverse(X2),inverse(X1)))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,X2)!=sk_c6|sk_c7!=identity|sk_c8!=identity)).
##############
#cnf(i_0_393, negated_conjecture, (multiply(inverse(sk_c8),inverse(sk_c2))=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_396, negated_conjecture, (multiply(inverse(sk_c2),inverse(sk_c8))=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_47, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c6|multiply(sk_c8,sk_c5)=sk_c7)).
#
#cnf(i_0_44, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c6|multiply(sk_c7,sk_c8)=sk_c6)).
#
#cnf(i_0_45, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c6|multiply(sk_c3,sk_c7)=sk_c8)).
##
#cnf(i_0_48, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c6|multiply(sk_c4,sk_c8)=sk_c5)).
#
#cnf(i_0_65, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c7|multiply(sk_c8,sk_c5)=sk_c7)).
#
#cnf(i_0_62, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c7|multiply(sk_c7,sk_c8)=sk_c6)).
##
#cnf(i_0_63, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c7|multiply(sk_c3,sk_c7)=sk_c8)).
#
#cnf(i_0_66, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c7|multiply(sk_c4,sk_c8)=sk_c5)).
#
#cnf(i_0_53, negated_conjecture, (multiply(sk_c8,sk_c5)=sk_c7|multiply(sk_c1,sk_c8)=sk_c7)).
##
#cnf(i_0_71, negated_conjecture, (multiply(sk_c8,sk_c5)=sk_c7|multiply(sk_c2,sk_c8)=sk_c6)).
#
#cnf(i_0_50, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c6|multiply(sk_c1,sk_c8)=sk_c7)).
#
#cnf(i_0_68, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c6|multiply(sk_c2,sk_c8)=sk_c6)).
##
#cnf(i_0_51, negated_conjecture, (multiply(sk_c3,sk_c7)=sk_c8|multiply(sk_c1,sk_c8)=sk_c7)).
#
#cnf(i_0_69, negated_conjecture, (multiply(sk_c3,sk_c7)=sk_c8|multiply(sk_c2,sk_c8)=sk_c6)).
#
#cnf(i_0_54, negated_conjecture, (multiply(sk_c4,sk_c8)=sk_c5|multiply(sk_c1,sk_c8)=sk_c7)).
##
#cnf(i_0_72, negated_conjecture, (multiply(sk_c4,sk_c8)=sk_c5|multiply(sk_c2,sk_c8)=sk_c6)).
#
#cnf(i_0_405, negated_conjecture, (multiply(inverse(sk_c8),sk_c6)=sk_c7|inverse(sk_c3)=sk_c7)).
####
#cnf(i_0_2780, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c7|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_2781, negated_conjecture, (multiply(sk_c1,sk_c6)=sk_c7|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_2473, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c8,sk_c2)!=sk_c6|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_408, negated_conjecture, (multiply(inverse(sk_c8),sk_c7)=sk_c6|inverse(sk_c3)=sk_c7)).
####
#cnf(i_0_2820, negated_conjecture, (multiply(sk_c2,sk_c7)=sk_c6|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_2821, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_436, negated_conjecture, (multiply(inverse(sk_c1),sk_c7)=sk_c8|inverse(sk_c3)=sk_c7)).
####
#cnf(i_0_2855, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c8|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_2366, negated_conjecture, (multiply(sk_c8,multiply(inverse(X1),X2))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(inverse(X2),X1)!=sk_c6|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_2867, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c6=sk_c8)).
#
#cnf(i_0_2943, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c6=sk_c8)).
###
#cnf(i_0_2940, negated_conjecture, (multiply(sk_c7,sk_c3)=identity|sk_c6=sk_c8)).
###
#cnf(i_0_2942, negated_conjecture, (multiply(sk_c3,sk_c7)=identity|sk_c6=sk_c8)).
#####
#cnf(i_0_2984, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c8=identity|sk_c6=sk_c8)).
#
#cnf(i_0_2985, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c8=identity|sk_c6=sk_c8)).
##
#cnf(i_0_2989, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c8=identity|sk_c6=sk_c8)).
#
#cnf(i_0_3008, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c8=identity|sk_c6=sk_c8)).
#######
#cnf(i_0_2941, negated_conjecture, (multiply(sk_c7,multiply(sk_c3,X1))=X1|sk_c6=sk_c8)).
####
#cnf(i_0_2944, negated_conjecture, (multiply(sk_c3,multiply(sk_c7,X1))=X1|sk_c6=sk_c8)).
#########
#cnf(i_0_3105, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c3,X1)|sk_c6=sk_c8)).
##
#cnf(i_0_2912, negated_conjecture, (multiply(sk_c8,multiply(X1,X2))!=sk_c8|multiply(inverse(X2),inverse(X1))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity)).
#########
#cnf(i_0_3110, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c7,X1)|sk_c6=sk_c8)).
#############
#cnf(i_0_440, negated_conjecture, (multiply(inverse(sk_c2),sk_c6)=sk_c8|inverse(sk_c3)=sk_c7)).
####
#cnf(i_0_3318, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c8|inverse(sk_c3)=sk_c7)).
###
#cnf(i_0_3336, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c8=sk_c7)).
#
#cnf(i_0_3362, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c8=sk_c7)).
###
#cnf(i_0_3359, negated_conjecture, (multiply(sk_c7,sk_c3)=identity|sk_c8=sk_c7)).
##
#cnf(i_0_3361, negated_conjecture, (multiply(sk_c3,sk_c7)=identity|sk_c8=sk_c7)).
#####
#cnf(i_0_3172, negated_conjecture, (multiply(inverse(sk_c8),inverse(sk_c2))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_3407, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c8=identity|sk_c8=sk_c7)).
#
#cnf(i_0_3408, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c8=identity|sk_c8=sk_c7)).
#
#cnf(i_0_3444, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c8=identity|sk_c8=sk_c7)).
##
#cnf(i_0_3463, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c8=identity|sk_c8=sk_c7)).
########
#cnf(i_0_3360, negated_conjecture, (multiply(sk_c7,multiply(sk_c3,X1))=X1|sk_c8=sk_c7)).
####
#cnf(i_0_3363, negated_conjecture, (multiply(sk_c3,multiply(sk_c7,X1))=X1|sk_c8=sk_c7)).
###
#cnf(i_0_3219, negated_conjecture, (multiply(inverse(sk_c2),inverse(sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_3566, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c3,X1)|sk_c8=sk_c7)).
###########
#cnf(i_0_3572, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c7,X1)|sk_c8=sk_c7)).
#######
#cnf(i_0_2410, negated_conjecture, (multiply(sk_c8,multiply(X1,inverse(X2)))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X2,inverse(X1))!=sk_c6|sk_c7!=identity|sk_c8!=identity)).
######
#cnf(i_0_452, negated_conjecture, (multiply(inverse(sk_c8),sk_c6)=sk_c7|inverse(sk_c4)=sk_c8)).
####
#cnf(i_0_455, negated_conjecture, (multiply(inverse(sk_c8),sk_c7)=sk_c6|inverse(sk_c4)=sk_c8)).
####
#cnf(i_0_465, negated_conjecture, (multiply(inverse(sk_c1),sk_c7)=sk_c8|inverse(sk_c4)=sk_c8)).
####
#cnf(i_0_469, negated_conjecture, (multiply(inverse(sk_c2),sk_c6)=sk_c8|inverse(sk_c4)=sk_c8)).
####
#cnf(i_0_481, negated_conjecture, (multiply(inverse(sk_c8),sk_c7)=sk_c5|inverse(sk_c1)=sk_c8)).
####
#cnf(i_0_3890, negated_conjecture, (multiply(sk_c4,sk_c7)=sk_c5|inverse(sk_c1)=sk_c8)).
#
#cnf(i_0_485, negated_conjecture, (multiply(inverse(sk_c7),sk_c6)=sk_c8|inverse(sk_c1)=sk_c8)).
####
#cnf(i_0_489, negated_conjecture, (multiply(inverse(sk_c3),sk_c8)=sk_c7|inverse(sk_c1)=sk_c8)).
####
#cnf(i_0_500, negated_conjecture, (multiply(inverse(sk_c4),sk_c5)=sk_c8|inverse(sk_c1)=sk_c8)).
####
#cnf(i_0_3969, negated_conjecture, (multiply(sk_c8,sk_c5)=sk_c8|inverse(sk_c1)=sk_c8)).
#
#cnf(i_0_552, negated_conjecture, (multiply(inverse(X2),multiply(inverse(X1),sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,X2)!=sk_c8|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_3985, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c8=sk_c7)).
#
#cnf(i_0_4147, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c8=sk_c7)).
####
#cnf(i_0_4144, negated_conjecture, (multiply(sk_c8,sk_c1)=identity|sk_c8=sk_c7)).
####
#cnf(i_0_4146, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|sk_c8=sk_c7)).
############
#cnf(i_0_4190, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c4=sk_c1|sk_c8=sk_c7)).
####
#cnf(i_0_4225, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c7=identity|sk_c8=sk_c7)).
#
#cnf(i_0_4048, negated_conjecture, (multiply(inverse(X1),multiply(X2,sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(inverse(X2),X1)!=sk_c8|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_4231, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c4=sk_c1|sk_c8=sk_c7)).
#
#cnf(i_0_4272, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c7=identity|sk_c8=sk_c7)).
#
#cnf(i_0_4500, negated_conjecture, (sk_c4=sk_c1|sk_c8=sk_c7|sk_c7=identity)).
########
#cnf(i_0_4145, negated_conjecture, (multiply(sk_c8,multiply(sk_c1,X1))=X1|sk_c8=sk_c7)).
######
#cnf(i_0_4148, negated_conjecture, (multiply(sk_c1,multiply(sk_c8,X1))=X1|sk_c8=sk_c7)).
##################
#cnf(i_0_4523, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c1,X1)|sk_c8=sk_c7)).
#
#cnf(i_0_4122, negated_conjecture, (multiply(X1,multiply(inverse(X2),sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X2,inverse(X1))!=sk_c8|sk_c7!=identity|sk_c8!=identity)).
##########
#cnf(i_0_4537, negated_conjecture, (multiply(inverse(sk_c1),X1)=multiply(sk_c8,X1)|sk_c8=sk_c7)).
#############
#cnf(i_0_511, negated_conjecture, (multiply(inverse(sk_c8),sk_c7)=sk_c5|inverse(sk_c2)=sk_c8)).
####
#cnf(i_0_516, negated_conjecture, (multiply(inverse(sk_c7),sk_c6)=sk_c8|inverse(sk_c2)=sk_c8)).
####
#cnf(i_0_520, negated_conjecture, (multiply(inverse(sk_c3),sk_c8)=sk_c7|inverse(sk_c2)=sk_c8)).
#
#cnf(i_0_721, negated_conjecture, (multiply(sk_c8,multiply(inverse(X2),inverse(X1)))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,X2)!=sk_c7|sk_c6!=identity|sk_c8!=identity)).
###
#cnf(i_0_524, negated_conjecture, (multiply(inverse(sk_c4),sk_c5)=sk_c8|inverse(sk_c2)=sk_c8)).
####
#cnf(i_0_5022, negated_conjecture, (multiply(sk_c8,sk_c5)=sk_c8|inverse(sk_c2)=sk_c8)).
###
#cnf(i_0_5040, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c8=sk_c7)).
#
#cnf(i_0_5052, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c8=sk_c7)).
#
#cnf(i_0_5006, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c8,sk_c2)!=sk_c7|sk_c6!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_5049, negated_conjecture, (multiply(sk_c8,sk_c2)=identity|sk_c8=sk_c7)).
####
#cnf(i_0_5051, negated_conjecture, (multiply(sk_c2,sk_c8)=identity|sk_c8=sk_c7)).
#########
#cnf(i_0_5140, negated_conjecture, (sk_c4=sk_c2|sk_c7=identity|sk_c8=sk_c7)).
##
#cnf(i_0_4925, negated_conjecture, (multiply(sk_c8,multiply(inverse(X1),X2))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(inverse(X2),X1)!=sk_c7|sk_c6!=identity|sk_c8!=identity)).
##
#cnf(i_0_5141, negated_conjecture, (multiply(sk_c2,X1)=multiply(sk_c1,X1)|sk_c8=sk_c7)).
####
#cnf(i_0_5324, negated_conjecture, (multiply(sk_c1,inverse(sk_c2))=identity|sk_c8=sk_c7)).
###
#cnf(i_0_5356, negated_conjecture, (inverse(sk_c2)=inverse(sk_c1)|sk_c8=sk_c7)).
######
#cnf(i_0_5357, negated_conjecture, (multiply(sk_c2,inverse(sk_c1))=identity|sk_c8=sk_c7)).
######
#cnf(i_0_5373, negated_conjecture, (multiply(inverse(sk_c1),sk_c2)=identity|sk_c8=sk_c7)).
#########
#cnf(i_0_5447, negated_conjecture, (multiply(inverse(sk_c2),sk_c1)=identity|sk_c8=sk_c7)).
##########
#cnf(i_0_5262, negated_conjecture, (multiply(sk_c8,multiply(X1,X2))!=sk_c8|multiply(inverse(X2),inverse(X1))!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c8!=identity)).
#####
#cnf(i_0_5177, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c6=identity|sk_c8=sk_c7)).
########
#cnf(i_0_5638, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c6=identity|sk_c8=sk_c7)).
#
#cnf(i_0_5712, negated_conjecture, (sk_c4=sk_c1|sk_c8=sk_c7|sk_c6=identity)).
#
#cnf(i_0_5714, negated_conjecture, (sk_c4=sk_c2|sk_c8=sk_c7|sk_c6=identity)).
####
#cnf(i_0_5716, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c8=sk_c7|sk_c4=sk_c1)).
###
#cnf(i_0_5749, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c8=sk_c7|sk_c4=sk_c2)).
#
#cnf(i_0_5784, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c8=sk_c7|sk_c4=sk_c1)).
##
#cnf(i_0_5886, negated_conjecture, (sk_c4=sk_c1|sk_c8=sk_c7)).
####
#cnf(i_0_5812, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c8=sk_c7|sk_c4=sk_c2)).
#
#cnf(i_0_5965, negated_conjecture, (sk_c4=sk_c2|sk_c8=sk_c7)).
#########
#cnf(i_0_5050, negated_conjecture, (multiply(sk_c8,multiply(sk_c2,X1))=X1|sk_c8=sk_c7)).
#
#cnf(i_0_5495, negated_conjecture, (multiply(inverse(sk_c8),inverse(sk_c2))!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_5053, negated_conjecture, (multiply(sk_c2,multiply(sk_c8,X1))=X1|sk_c8=sk_c7)).
#############
#cnf(i_0_5553, negated_conjecture, (multiply(inverse(sk_c2),inverse(sk_c8))!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###########
#cnf(i_0_5991, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c2,X1)|sk_c8=sk_c7)).
#############
#cnf(i_0_4983, negated_conjecture, (multiply(sk_c8,multiply(X1,inverse(X2)))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X2,inverse(X1))!=sk_c7|sk_c6!=identity|sk_c8!=identity)).
#
#cnf(i_0_6009, negated_conjecture, (multiply(inverse(sk_c2),X1)=multiply(sk_c8,X1)|sk_c8=sk_c7)).
#################
#cnf(i_0_2797, negated_conjecture, (multiply(inverse(sk_c2),sk_c7)=sk_c6|inverse(sk_c3)=sk_c7)).
#####
#cnf(i_0_2805, negated_conjecture, (multiply(inverse(sk_c1),sk_c7)=sk_c6|inverse(sk_c3)=sk_c7)).
#######
#cnf(i_0_2837, negated_conjecture, (multiply(inverse(sk_c2),sk_c6)=sk_c7|inverse(sk_c3)=sk_c7)).
#######
#cnf(i_0_2846, negated_conjecture, (multiply(inverse(sk_c1),sk_c6)=sk_c7|inverse(sk_c3)=sk_c7)).
###
#cnf(i_0_812, negated_conjecture, (multiply(inverse(X2),multiply(inverse(X1),sk_c8))!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,X2)!=sk_c8|sk_c6!=identity|sk_c8!=identity)).
##
#cnf(i_0_3909, negated_conjecture, (multiply(inverse(sk_c4),sk_c5)=sk_c7|inverse(sk_c1)=sk_c8)).
#######
#cnf(i_0_849, negated_conjecture, (multiply(sk_c8,sk_c4)=identity|sk_c6=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_851, negated_conjecture, (multiply(sk_c4,sk_c8)=identity|sk_c6=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_1110, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|sk_c8=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_1112, negated_conjecture, (multiply(sk_c8,sk_c1)=identity|sk_c8=identity|sk_c4!=sk_c2)).
#########
#cnf(i_0_6514, negated_conjecture, (multiply(inverse(X1),multiply(X2,sk_c8))!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(inverse(X2),X1)!=sk_c8|sk_c6!=identity|sk_c8!=identity)).
#########
#cnf(i_0_2662, negated_conjecture, (multiply(sk_c8,sk_c5)=sk_c7|sk_c6=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_2697, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c6|sk_c6=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_2732, negated_conjecture, (multiply(sk_c3,sk_c7)=sk_c8|sk_c6=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_6894, negated_conjecture, (sk_c8=identity|sk_c6=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_6854, negated_conjecture, (multiply(sk_c2,sk_c7)=sk_c5|sk_c6=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_6879, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c8|sk_c6=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_6889, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c7|sk_c6=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_6975, negated_conjecture, (sk_c6=sk_c7|sk_c6=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_6976, negated_conjecture, (sk_c6=identity|sk_c7!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_7023, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
####
#cnf(i_0_7025, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_394, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c3))=identity|sk_c2=sk_c1)).
#####
#cnf(i_0_395, negated_conjecture, (multiply(inverse(sk_c3),inverse(sk_c7))=identity|sk_c2=sk_c1)).
#
#cnf(i_0_7039, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|inverse(sk_c2)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
####
#cnf(i_0_7042, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c8,sk_c2)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_597, negated_conjecture, (multiply(inverse(sk_c3),inverse(sk_c7))=identity|sk_c6=identity)).
#####
#cnf(i_0_600, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c3))=identity|sk_c6=identity)).
#####
#cnf(i_0_627, negated_conjecture, (multiply(inverse(sk_c3),inverse(sk_c7))=identity|sk_c7=identity)).
#####
#cnf(i_0_7026, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c6!=sk_c2|sk_c4!=sk_c2)).
#
#cnf(i_0_632, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c3))=identity|sk_c7=identity)).
###
#cnf(i_0_7041, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c2!=sk_c8|sk_c8!=identity|sk_c4!=sk_c2|sk_c7!=identity)).
##
#cnf(i_0_692, negated_conjecture, (multiply(inverse(sk_c1),inverse(sk_c8))=identity|sk_c5=identity)).
######
#cnf(i_0_7021, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_695, negated_conjecture, (multiply(inverse(sk_c8),inverse(sk_c1))=identity|sk_c5=identity)).
#######
#cnf(i_0_7300, negated_conjecture, (multiply(sk_c2,sk_c8)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_762, negated_conjecture, (multiply(inverse(sk_c2),inverse(sk_c8))=identity|sk_c5=identity)).
######
#cnf(i_0_765, negated_conjecture, (multiply(inverse(sk_c8),inverse(sk_c2))=identity|sk_c5=identity)).
########
#cnf(i_0_2974, negated_conjecture, (multiply(inverse(sk_c3),inverse(sk_c7))=identity|sk_c6=sk_c8)).
#
#cnf(i_0_7027, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c8,X1)!=sk_c8|inverse(X1)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_2980, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c3))=identity|sk_c6=sk_c8)).
#####
#cnf(i_0_3393, negated_conjecture, (multiply(inverse(sk_c3),inverse(sk_c7))=identity|sk_c8=sk_c7)).
##
#cnf(i_0_7505, negated_conjecture, (multiply(inverse(X1),sk_c8)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2|X1!=sk_c8)).
###
#cnf(i_0_3401, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c3))=identity|sk_c8=sk_c7)).
######
#cnf(i_0_4211, negated_conjecture, (multiply(inverse(sk_c1),inverse(sk_c8))=identity|sk_c8=sk_c7)).
######
#cnf(i_0_4218, negated_conjecture, (multiply(inverse(sk_c8),inverse(sk_c1))=identity|sk_c8=sk_c7)).
########
#cnf(i_0_5149, negated_conjecture, (multiply(inverse(sk_c2),inverse(sk_c8))=identity|sk_c8=sk_c7)).
#######
#cnf(i_0_5159, negated_conjecture, (multiply(inverse(sk_c8),inverse(sk_c2))=identity|sk_c8=sk_c7)).
##
#cnf(i_0_7038, negated_conjecture, (multiply(X1,inverse(sk_c2))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
########################
#cnf(i_0_7907, negated_conjecture, (multiply(sk_c8,inverse(sk_c2))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c2!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_828, negated_conjecture, (inverse(sk_c4)=sk_c8|inverse(sk_c8)=sk_c4|sk_c6=identity)).
#
#cnf(i_0_7948, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c6=identity)).
##
#cnf(i_0_7982, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c6=identity)).
##
#cnf(i_0_7979, negated_conjecture, (multiply(sk_c4,sk_c8)=identity|sk_c6=identity)).
#
#cnf(i_0_7528, negated_conjecture, (multiply(inverse(X1),inverse(sk_c2))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2|X1!=sk_c8)).
##
#cnf(i_0_7981, negated_conjecture, (multiply(sk_c8,sk_c4)=identity|sk_c6=identity)).
########
#cnf(i_0_8027, negated_conjecture, (sk_c4=sk_c2|sk_c8=identity|sk_c6=identity)).
#
#cnf(i_0_8028, negated_conjecture, (sk_c4=sk_c1|sk_c8=identity|sk_c6=identity)).
##################
#cnf(i_0_7980, negated_conjecture, (multiply(sk_c4,multiply(sk_c8,X1))=X1|sk_c6=identity)).
#############
#cnf(i_0_7983, negated_conjecture, (multiply(sk_c8,multiply(sk_c4,X1))=X1|sk_c6=identity)).
#########
#cnf(i_0_8287, negated_conjecture, (multiply(inverse(sk_c4),X1)=multiply(sk_c8,X1)|sk_c6=identity)).
###########
#cnf(i_0_8294, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c4,X1)|sk_c6=identity)).
######
#cnf(i_0_6596, negated_conjecture, (multiply(X1,multiply(inverse(X2),sk_c8))!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X2,inverse(X1))!=sk_c8|sk_c6!=identity|sk_c8!=identity)).
#######
#cnf(i_0_8314, negated_conjecture, (multiply(sk_c4,sk_c7)=sk_c5|sk_c6=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_8083, negated_conjecture, (multiply(inverse(sk_c8),inverse(sk_c4))=identity|sk_c6=identity)).
#####
#cnf(i_0_8201, negated_conjecture, (multiply(inverse(sk_c4),inverse(sk_c8))=identity|sk_c6=identity)).
#################
#cnf(i_0_7307, negated_conjecture, (multiply(X1,multiply(X2,sk_c8))!=sk_c6|multiply(inverse(X2),inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
################
#cnf(i_0_8769, negated_conjecture, (multiply(sk_c8,inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2|X1!=sk_c6)).
#############
#cnf(i_0_6632, negated_conjecture, (inverse(sk_c4)=sk_c7|inverse(sk_c1)=sk_c8|sk_c2=sk_c1)).
####
#cnf(i_0_8884, negated_conjecture, (inverse(sk_c7)=sk_c4|inverse(sk_c1)=sk_c8|sk_c2=sk_c1)).
#
#cnf(i_0_8930, negated_conjecture, (inverse(sk_c8)=sk_c1|inverse(sk_c7)=sk_c4|sk_c2=sk_c1)).
#
#cnf(i_0_9006, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c2=sk_c1|sk_c4!=sk_c2)).
##
#cnf(i_0_9073, negated_conjecture, (sk_c4=sk_c3|sk_c2=sk_c1|sk_c4!=sk_c2)).
##
#cnf(i_0_9044, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c2=sk_c1|sk_c4!=sk_c2)).
######
#cnf(i_0_9041, negated_conjecture, (multiply(sk_c4,sk_c7)=identity|sk_c2=sk_c1|sk_c4!=sk_c2)).
##
#cnf(i_0_9043, negated_conjecture, (multiply(sk_c7,sk_c4)=identity|sk_c2=sk_c1|sk_c4!=sk_c2)).
#
#cnf(i_0_8722, negated_conjecture, (multiply(inverse(sk_c2),inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2|X1!=sk_c6)).
############################
#cnf(i_0_9223, negated_conjecture, (multiply(inverse(sk_c2),X1)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#####################
#cnf(i_0_108, negated_conjecture, (multiply(sk_c8,multiply(sk_c2,X1))=X1|inverse(sk_c3)=sk_c7)).
#######
#cnf(i_0_9168, negated_conjecture, (multiply(inverse(sk_c2),sk_c8)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2|sk_c6!=sk_c2)).
###
#cnf(i_0_109, negated_conjecture, (multiply(sk_c8,multiply(sk_c1,X1))=X1|inverse(sk_c3)=sk_c7)).
###########
#cnf(i_0_110, negated_conjecture, (multiply(sk_c8,multiply(sk_c4,X1))=X1|inverse(sk_c1)=sk_c8)).
#######
#cnf(i_0_111, negated_conjecture, (multiply(sk_c8,multiply(sk_c4,X1))=X1|inverse(sk_c2)=sk_c8)).
#####
#cnf(i_0_209, negated_conjecture, (multiply(sk_c2,multiply(sk_c8,X1))=X1|inverse(sk_c3)=sk_c7)).
##
#cnf(i_0_8747, negated_conjecture, (multiply(X1,multiply(sk_c8,sk_c8))!=sk_c6|multiply(sk_c2,inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_9490, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c7|inverse(sk_c3)=sk_c7)).
##
#cnf(i_0_9560, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c6=sk_c7)).
##
#cnf(i_0_9579, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c6=sk_c7)).
###
#cnf(i_0_9576, negated_conjecture, (multiply(sk_c7,sk_c3)=identity|sk_c6=sk_c7)).
###
#cnf(i_0_9578, negated_conjecture, (multiply(sk_c3,sk_c7)=identity|sk_c6=sk_c7)).
###
#cnf(i_0_8764, negated_conjecture, (multiply(X1,multiply(sk_c2,sk_c8))!=sk_c6|multiply(sk_c8,inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###########
#cnf(i_0_9687, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c8=identity|sk_c6=sk_c7)).
##
#cnf(i_0_9688, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c8=identity|sk_c6=sk_c7)).
#
#cnf(i_0_9740, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c8=identity|sk_c6=sk_c7)).
#
#cnf(i_0_9785, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c8=identity|sk_c6=sk_c7)).
#######
#cnf(i_0_9577, negated_conjecture, (multiply(sk_c7,multiply(sk_c3,X1))=X1|sk_c6=sk_c7)).
##
#cnf(i_0_8798, negated_conjecture, (multiply(sk_c8,multiply(X1,sk_c8))!=sk_c6|multiply(inverse(X1),sk_c2)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_9580, negated_conjecture, (multiply(sk_c3,multiply(sk_c7,X1))=X1|sk_c6=sk_c7)).
#########
#cnf(i_0_9491, negated_conjecture, (multiply(sk_c2,X1)=multiply(sk_c1,X1)|inverse(sk_c3)=sk_c7)).
######
#cnf(i_0_10153, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c6|inverse(sk_c3)=sk_c7)).
#######
#cnf(i_0_8815, negated_conjecture, (multiply(sk_c2,multiply(X1,sk_c8))!=sk_c6|multiply(inverse(X1),sk_c8)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_9998, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c3,X1)|sk_c6=sk_c7)).
###########
#cnf(i_0_10006, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c7,X1)|sk_c6=sk_c7)).
#############
#cnf(i_0_9535, negated_conjecture, (multiply(inverse(sk_c2),sk_c7)=sk_c8|inverse(sk_c3)=sk_c7)).
#######
#cnf(i_0_10133, negated_conjecture, (multiply(sk_c1,inverse(sk_c2))=identity|inverse(sk_c3)=sk_c7)).
###
#cnf(i_0_10439, negated_conjecture, (inverse(sk_c2)=inverse(sk_c1)|inverse(sk_c3)=sk_c7)).
###
#cnf(i_0_9508, negated_conjecture, (multiply(sk_c8,multiply(sk_c8,sk_c8))!=sk_c6|multiply(sk_c2,sk_c2)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_10180, negated_conjecture, (multiply(inverse(sk_c1),sk_c6)=sk_c8|inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_9525, negated_conjecture, (multiply(sk_c2,multiply(sk_c8,sk_c8))!=sk_c6|multiply(sk_c2,sk_c8)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_10440, negated_conjecture, (multiply(sk_c2,inverse(sk_c1))=identity|inverse(sk_c3)=sk_c7)).
##
#cnf(i_0_9711, negated_conjecture, (multiply(sk_c8,multiply(sk_c2,sk_c8))!=sk_c6|multiply(sk_c8,sk_c2)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_10461, negated_conjecture, (multiply(inverse(sk_c1),sk_c2)=identity|inverse(sk_c3)=sk_c7)).
#######
#cnf(i_0_9728, negated_conjecture, (multiply(sk_c2,multiply(sk_c2,sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_10624, negated_conjecture, (multiply(inverse(sk_c2),sk_c1)=identity|inverse(sk_c3)=sk_c7)).
#######
#cnf(i_0_9669, negated_conjecture, (multiply(inverse(sk_c3),inverse(sk_c7))=identity|sk_c6=sk_c7)).
#####
#cnf(i_0_9679, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c3))=identity|sk_c6=sk_c7)).
######
#cnf(i_0_7644, negated_conjecture, (multiply(inverse(X2),multiply(inverse(X1),sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,X2)!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
######################
#cnf(i_0_210, negated_conjecture, (multiply(sk_c1,multiply(sk_c8,X1))=X1|inverse(sk_c3)=sk_c7)).
######
#cnf(i_0_10818, negated_conjecture, (multiply(inverse(X1),multiply(X2,sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(inverse(X2),X1)!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#################
#cnf(i_0_211, negated_conjecture, (multiply(sk_c4,multiply(sk_c8,X1))=X1|inverse(sk_c1)=sk_c8)).
################
#cnf(i_0_212, negated_conjecture, (multiply(sk_c4,multiply(sk_c8,X1))=X1|inverse(sk_c2)=sk_c8)).
####
#cnf(i_0_11230, negated_conjecture, (multiply(sk_c4,sk_c7)=sk_c5|inverse(sk_c2)=sk_c8)).
#########
#cnf(i_0_11257, negated_conjecture, (multiply(inverse(sk_c4),sk_c5)=sk_c7|inverse(sk_c2)=sk_c8)).
######
#cnf(i_0_10912, negated_conjecture, (multiply(X1,multiply(inverse(X2),sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X2,inverse(X1))!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_11279, negated_conjecture, (inverse(sk_c4)=sk_c7|inverse(sk_c2)=sk_c8|sk_c2=sk_c1)).
###
#cnf(i_0_11423, negated_conjecture, (inverse(sk_c7)=sk_c4|inverse(sk_c2)=sk_c8|sk_c2=sk_c1)).
#
#cnf(i_0_11475, negated_conjecture, (inverse(sk_c8)=sk_c2|inverse(sk_c7)=sk_c4|sk_c2=sk_c1)).
#
#cnf(i_0_11626, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c2=sk_c1)).
##
#cnf(i_0_11665, negated_conjecture, (sk_c4=sk_c3|sk_c2=sk_c1)).
#
#cnf(i_0_11631, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c2=sk_c1)).
####
#cnf(i_0_11628, negated_conjecture, (multiply(sk_c4,sk_c7)=identity|sk_c2=sk_c1)).
###
#cnf(i_0_11630, negated_conjecture, (multiply(sk_c7,sk_c4)=identity|sk_c2=sk_c1)).
###########################
#cnf(i_0_10756, negated_conjecture, (multiply(inverse(X1),multiply(sk_c8,sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c2,X1)!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_11678, negated_conjecture, (multiply(sk_c4,X1)=multiply(sk_c3,X1)|sk_c2=sk_c1)).
####
#cnf(i_0_11667, negated_conjecture, (multiply(inverse(sk_c3),sk_c4)=identity|sk_c2=sk_c1)).
#######
#cnf(i_0_11668, negated_conjecture, (multiply(sk_c4,inverse(sk_c3))=identity|sk_c2=sk_c1)).
#
#cnf(i_0_11992, negated_conjecture, (inverse(sk_c4)=inverse(sk_c3)|sk_c2=sk_c1)).
#####
#cnf(i_0_11928, negated_conjecture, (multiply(sk_c3,inverse(sk_c4))=identity|sk_c2=sk_c1)).
######
#cnf(i_0_11970, negated_conjecture, (multiply(inverse(sk_c4),sk_c3)=identity|sk_c2=sk_c1)).
###########
#cnf(i_0_10799, negated_conjecture, (multiply(inverse(X1),multiply(sk_c2,sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c8,X1)!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_11723, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c6=identity|sk_c2=sk_c1)).
##########################
#cnf(i_0_12219, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c6=identity|sk_c2=sk_c1)).
####
#cnf(i_0_11629, negated_conjecture, (multiply(sk_c4,multiply(sk_c7,X1))=X1|sk_c2=sk_c1)).
##############
#cnf(i_0_11632, negated_conjecture, (multiply(sk_c7,multiply(sk_c4,X1))=X1|sk_c2=sk_c1)).
###
#cnf(i_0_10850, negated_conjecture, (multiply(sk_c8,multiply(inverse(X1),sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c2)!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###################
#cnf(i_0_12372, negated_conjecture, (multiply(inverse(sk_c4),X1)=multiply(sk_c7,X1)|sk_c2=sk_c1)).
############
#cnf(i_0_12384, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c4,X1)|sk_c2=sk_c1)).
#####
#cnf(i_0_10893, negated_conjecture, (multiply(sk_c2,multiply(inverse(X1),sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##################
#cnf(i_0_11789, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c4))=identity|sk_c2=sk_c1)).
########
#cnf(i_0_11804, negated_conjecture, (multiply(inverse(sk_c4),inverse(sk_c7))=identity|sk_c2=sk_c1)).
############
#cnf(i_0_11691, negated_conjecture, (inverse(sk_c3)=sk_c8|inverse(sk_c1)=sk_c8|sk_c2=sk_c1)).
#
#cnf(i_0_11692, negated_conjecture, (inverse(sk_c3)=sk_c8|inverse(sk_c2)=sk_c8|sk_c2=sk_c1)).
##
#cnf(i_0_11709, negated_conjecture, (inverse(sk_c3)=sk_c8|inverse(sk_c8)=sk_c2|sk_c2=sk_c1)).
##########
#cnf(i_0_12887, negated_conjecture, (inverse(sk_c8)=sk_c1|inverse(sk_c3)=sk_c8|sk_c2=sk_c1)).
##
#cnf(i_0_13184, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c2=sk_c1)).
####
#cnf(i_0_13189, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c2=sk_c1)).
##
#cnf(i_0_13300, negated_conjecture, (sk_c2=sk_c3|sk_c2=sk_c1|sk_c4!=sk_c2)).
#
#cnf(i_0_8878, negated_conjecture, (multiply(sk_c8,multiply(inverse(X2),inverse(X1)))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,X2)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_13186, negated_conjecture, (multiply(sk_c8,sk_c3)=identity|sk_c2=sk_c1)).
####
#cnf(i_0_13188, negated_conjecture, (multiply(sk_c3,sk_c8)=identity|sk_c2=sk_c1)).
#####
#cnf(i_0_13253, negated_conjecture, (multiply(sk_c8,sk_c4)=identity|sk_c2=sk_c1)).
#
#cnf(i_0_13573, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c2=sk_c1)).
###
#cnf(i_0_13594, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c2=sk_c1)).
#######
#cnf(i_0_13254, negated_conjecture, (multiply(sk_c4,sk_c8)=identity|sk_c2=sk_c1)).
###########
#cnf(i_0_13412, negated_conjecture, (multiply(sk_c8,multiply(inverse(X1),X2))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(inverse(X2),X1)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
########################################
#cnf(i_0_13840, negated_conjecture, (multiply(sk_c8,multiply(X1,X2))!=sk_c8|multiply(inverse(X2),inverse(X1))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_13240, negated_conjecture, (multiply(sk_c8,X1)=multiply(sk_c7,X1)|sk_c2=sk_c1)).
#######
#cnf(i_0_13228, negated_conjecture, (multiply(sk_c8,inverse(sk_c7))=identity|sk_c2=sk_c1)).
#
#cnf(i_0_14321, negated_conjecture, (inverse(sk_c8)=inverse(sk_c7)|sk_c2=sk_c1)).
########
#cnf(i_0_14333, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c2=sk_c1|sk_c4!=sk_c2)).
#
#cnf(i_0_13943, negated_conjecture, (multiply(inverse(sk_c8),inverse(sk_c2))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_14451, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c2=sk_c1|sk_c4!=sk_c2)).
######
#cnf(i_0_13229, negated_conjecture, (multiply(inverse(sk_c7),sk_c8)=identity|sk_c2=sk_c1)).
###########
#cnf(i_0_14033, negated_conjecture, (multiply(inverse(sk_c2),inverse(sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_14190, negated_conjecture, (multiply(sk_c7,inverse(sk_c8))=identity|sk_c2=sk_c1)).
############
#cnf(i_0_14672, negated_conjecture, (multiply(inverse(sk_c8),sk_c7)=identity|sk_c2=sk_c1)).
###############################
#cnf(i_0_13776, negated_conjecture, (multiply(sk_c8,multiply(sk_c8,X1))!=sk_c8|multiply(inverse(X1),sk_c2)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
####################
#cnf(i_0_13820, negated_conjecture, (multiply(sk_c8,multiply(sk_c2,X1))!=sk_c8|multiply(inverse(X1),sk_c8)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##########
#cnf(i_0_13187, negated_conjecture, (multiply(sk_c8,multiply(sk_c3,X1))=X1|sk_c2=sk_c1)).
###############
#cnf(i_0_13190, negated_conjecture, (multiply(sk_c3,multiply(sk_c8,X1))=X1|sk_c2=sk_c1)).
###############
#cnf(i_0_13479, negated_conjecture, (multiply(sk_c8,multiply(X1,inverse(X2)))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X2,inverse(X1))!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##############
#cnf(i_0_13572, negated_conjecture, (multiply(sk_c8,multiply(sk_c4,X1))=X1|sk_c2=sk_c1)).
#########
#cnf(i_0_13592, negated_conjecture, (multiply(sk_c4,multiply(sk_c8,X1))=X1|sk_c2=sk_c1)).
#################
#cnf(i_0_15055, negated_conjecture, (multiply(sk_c8,multiply(X1,sk_c8))!=sk_c8|multiply(sk_c2,inverse(X1))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##################
#cnf(i_0_14937, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c3,X1)|sk_c2=sk_c1)).
######
#cnf(i_0_15101, negated_conjecture, (multiply(sk_c8,multiply(X1,sk_c2))!=sk_c8|multiply(sk_c8,inverse(X1))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_14969, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c8,X1)|sk_c2=sk_c1)).
#################
#cnf(i_0_15366, negated_conjecture, (multiply(sk_c8,multiply(sk_c8,sk_c8))!=sk_c8|multiply(sk_c2,sk_c2)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_15210, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c4,X1)|sk_c2=sk_c1)).
##########
#cnf(i_0_15564, negated_conjecture, (multiply(sk_c8,multiply(sk_c2,sk_c2))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_15244, negated_conjecture, (multiply(inverse(sk_c4),X1)=multiply(sk_c8,X1)|sk_c2=sk_c1)).
###
#cnf(i_0_13350, negated_conjecture, (multiply(sk_c8,multiply(inverse(X1),sk_c8))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c2,X1)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#####################
#cnf(i_0_14211, negated_conjecture, (multiply(sk_c7,sk_c2)=identity|sk_c2=sk_c1|sk_c4!=sk_c2)).
####
#cnf(i_0_14448, negated_conjecture, (multiply(sk_c2,sk_c7)=identity|sk_c2=sk_c1|sk_c4!=sk_c2)).
###########
#cnf(i_0_13393, negated_conjecture, (multiply(sk_c8,multiply(inverse(X1),sk_c2))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c8,X1)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_13539, negated_conjecture, (multiply(inverse(sk_c3),inverse(sk_c8))=identity|sk_c2=sk_c1)).
#########
#cnf(i_0_13554, negated_conjecture, (multiply(inverse(sk_c8),inverse(sk_c3))=identity|sk_c2=sk_c1)).
#########
#cnf(i_0_13574, negated_conjecture, (multiply(inverse(sk_c4),inverse(sk_c8))=identity|sk_c2=sk_c1)).
##########
#cnf(i_0_13739, negated_conjecture, (multiply(inverse(sk_c8),inverse(sk_c4))=identity|sk_c2=sk_c1)).
###
#cnf(i_0_13417, negated_conjecture, (multiply(sk_c8,multiply(sk_c8,inverse(X1)))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c2)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
########################
#cnf(i_0_13460, negated_conjecture, (multiply(sk_c8,multiply(sk_c2,inverse(X1)))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###################################
#cnf(i_0_251, negated_conjecture, (multiply(sk_c8,multiply(sk_c2,X1))=X1|inverse(sk_c8)=sk_c4)).
#
#cnf(i_0_14854, negated_conjecture, (multiply(inverse(X1),multiply(sk_c8,sk_c2))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c8,X1)!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_254, negated_conjecture, (multiply(sk_c2,multiply(sk_c8,X1))=X1|inverse(sk_c8)=sk_c4)).
##############
#cnf(i_0_16683, negated_conjecture, (multiply(X1,multiply(sk_c8,sk_c2))!=sk_c6|multiply(sk_c8,inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_257, negated_conjecture, (multiply(sk_c4,multiply(sk_c8,X1))=X1|inverse(sk_c8)=sk_c2)).
################
#cnf(i_0_260, negated_conjecture, (multiply(sk_c8,multiply(sk_c4,X1))=X1|inverse(sk_c8)=sk_c2)).
#####
#cnf(i_0_16616, negated_conjecture, (multiply(sk_c8,multiply(sk_c8,sk_c2))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c8,sk_c2)!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
########
#cnf(i_0_16662, negated_conjecture, (multiply(sk_c2,multiply(sk_c8,sk_c2))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###############
#cnf(i_0_972, negated_conjecture, (multiply(sk_c8,sk_c1)=identity|sk_c8=identity|sk_c6=identity)).
######
#cnf(i_0_974, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|sk_c8=identity|sk_c6=identity)).
####
#cnf(i_0_982, negated_conjecture, (multiply(sk_c8,sk_c2)=identity|sk_c8=identity|sk_c6=identity)).
####
#cnf(i_0_984, negated_conjecture, (multiply(sk_c2,sk_c8)=identity|sk_c8=identity|sk_c6=identity)).
#####
#cnf(i_0_993, negated_conjecture, (multiply(sk_c8,sk_c1)=identity|sk_c8=identity|sk_c7=identity)).
####
#cnf(i_0_995, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|sk_c8=identity|sk_c7=identity)).
##
#cnf(i_0_1925, negated_conjecture, (multiply(sk_c2,sk_c8)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c7|inverse(X1)!=sk_c8|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_17159, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c7=identity|sk_c8=identity)).
#
#cnf(i_0_17230, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c7=identity|sk_c8=identity)).
##
#cnf(i_0_17366, negated_conjecture, (sk_c4=sk_c2|sk_c8=identity|sk_c7=identity)).
#
#cnf(i_0_17367, negated_conjecture, (sk_c4=sk_c1|sk_c8=identity|sk_c7=identity)).
######
#cnf(i_0_17198, negated_conjecture, (multiply(sk_c2,sk_c8)!=sk_c7|multiply(sk_c2,sk_c8)!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_1003, negated_conjecture, (multiply(sk_c8,sk_c2)=identity|sk_c8=identity|sk_c7=identity)).
##
#cnf(i_0_1005, negated_conjecture, (multiply(sk_c2,sk_c8)=identity|sk_c8=identity|sk_c7=identity)).
####
#cnf(i_0_1016, negated_conjecture, (multiply(sk_c8,sk_c4)=identity|sk_c7=identity|sk_c5=identity)).
#
#cnf(i_0_1929, negated_conjecture, (multiply(sk_c2,sk_c8)!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c8!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_1018, negated_conjecture, (multiply(sk_c4,sk_c8)=identity|sk_c7=identity|sk_c5=identity)).
###################
#cnf(i_0_7931, negated_conjecture, (multiply(X1,multiply(X2,inverse(sk_c2)))!=sk_c6|multiply(inverse(X2),inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#########
#cnf(i_0_2664, negated_conjecture, (multiply(sk_c8,sk_c5)=sk_c7|sk_c6=identity|sk_c5=identity)).
#
#cnf(i_0_2677, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c6|sk_c7=identity|sk_c5=identity)).
#
#cnf(i_0_2699, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c6|sk_c6=identity|sk_c5=identity)).
###
#cnf(i_0_2712, negated_conjecture, (multiply(sk_c3,sk_c7)=sk_c8|sk_c7=identity|sk_c5=identity)).
#
#cnf(i_0_18020, negated_conjecture, (sk_c8=identity|sk_c7=identity|sk_c5=identity)).
##
#cnf(i_0_2734, negated_conjecture, (multiply(sk_c3,sk_c7)=sk_c8|sk_c6=identity|sk_c5=identity)).
#
#cnf(i_0_18051, negated_conjecture, (sk_c8=identity|sk_c6=identity|sk_c5=identity)).
###
#cnf(i_0_2986, negated_conjecture, (multiply(sk_c8,sk_c1)=identity|sk_c8=identity|sk_c6=sk_c8)).
####
#cnf(i_0_2988, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|sk_c8=identity|sk_c6=sk_c8)).
##
#cnf(i_0_3005, negated_conjecture, (multiply(sk_c8,sk_c2)=identity|sk_c8=identity|sk_c6=sk_c8)).
###
#cnf(i_0_3007, negated_conjecture, (multiply(sk_c2,sk_c8)=identity|sk_c8=identity|sk_c6=sk_c8)).
###########
#cnf(i_0_4226, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c6|sk_c7=identity|sk_c8=sk_c7)).
#
#cnf(i_0_4227, negated_conjecture, (multiply(sk_c3,sk_c7)=sk_c8|sk_c7=identity|sk_c8=sk_c7)).
#
#cnf(i_0_18239, negated_conjecture, (sk_c8=identity|sk_c7=identity|sk_c8=sk_c7)).
###
#cnf(i_0_18384, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c8=identity|sk_c7=identity)).
#
#cnf(i_0_18623, negated_conjecture, (sk_c2=sk_c3|sk_c7=identity|sk_c8=identity)).
##
#cnf(i_0_18694, negated_conjecture, (sk_c1=sk_c3|sk_c8=identity|sk_c7=identity)).
#
#cnf(i_0_19083, negated_conjecture, (sk_c7=identity|sk_c2=sk_c1|sk_c1=sk_c3)).
##
#cnf(i_0_17805, negated_conjecture, (multiply(X1,multiply(sk_c8,inverse(sk_c2)))!=sk_c6|multiply(sk_c2,inverse(X1))!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_18385, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c8=identity|sk_c7=identity)).
####
#cnf(i_0_18562, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c7=identity|sk_c8=identity)).
##
#cnf(i_0_19442, negated_conjecture, (sk_c4=sk_c3|sk_c7=identity|sk_c8=identity)).
######
#cnf(i_0_18567, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c8=identity|sk_c7=identity)).
####
#cnf(i_0_18673, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c8=identity|sk_c7=identity)).
####
#cnf(i_0_19302, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c8=identity|sk_c7=identity)).
#
#cnf(i_0_17889, negated_conjecture, (multiply(sk_c8,multiply(X1,inverse(sk_c2)))!=sk_c6|multiply(inverse(X1),sk_c2)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_19386, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c7=identity|sk_c8=identity)).
##########
#cnf(i_0_19646, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c8=identity|sk_c7=identity)).
##############
#cnf(i_0_18849, negated_conjecture, (sk_c2=identity|sk_c1=sk_c3|sk_c7=identity|sk_c4!=sk_c2)).
#########
#cnf(i_0_17912, negated_conjecture, (multiply(sk_c2,multiply(X1,inverse(sk_c2)))!=sk_c6|multiply(inverse(X1),sk_c8)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_4269, negated_conjecture, (multiply(sk_c8,sk_c4)=identity|sk_c7=identity|sk_c8=sk_c7)).
#####
#cnf(i_0_4271, negated_conjecture, (multiply(sk_c4,sk_c8)=identity|sk_c7=identity|sk_c8=sk_c7)).
##############
#cnf(i_0_5178, negated_conjecture, (multiply(sk_c8,sk_c5)=sk_c7|sk_c6=identity|sk_c8=sk_c7)).
#####
#cnf(i_0_20266, negated_conjecture, (sk_c8=sk_c7|sk_c6=identity|sk_c8=identity)).
#####
#cnf(i_0_20270, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c8=sk_c7|sk_c8=identity)).
####
#cnf(i_0_20310, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c8=sk_c7|sk_c8=identity)).
#
#cnf(i_0_19265, negated_conjecture, (multiply(sk_c8,multiply(sk_c8,inverse(sk_c2)))!=sk_c6|multiply(sk_c2,sk_c2)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
########
#cnf(i_0_19288, negated_conjecture, (multiply(sk_c2,multiply(sk_c8,inverse(sk_c2)))!=sk_c6|multiply(sk_c2,sk_c8)!=sk_c8|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_5179, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c6|sk_c6=identity|sk_c8=sk_c7)).
##
#cnf(i_0_5180, negated_conjecture, (multiply(sk_c3,sk_c7)=sk_c8|sk_c6=identity|sk_c8=sk_c7)).
################
#cnf(i_0_8102, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c6|sk_c5=identity|sk_c6=identity)).
#
#cnf(i_0_8173, negated_conjecture, (multiply(inverse(X2),multiply(inverse(X1),inverse(sk_c2)))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,X2)!=sk_c8|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_8103, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c7|sk_c5=identity|sk_c6=identity)).
#
#cnf(i_0_8104, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|sk_c5=identity|sk_c6=identity)).
#
#cnf(i_0_20941, negated_conjecture, (sk_c7=identity|sk_c5=identity|sk_c6=identity)).
##
#cnf(i_0_20962, negated_conjecture, (sk_c8=sk_c7|sk_c6=identity|sk_c7=identity)).
######
#cnf(i_0_8105, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|sk_c5=identity|sk_c6=identity)).
##
#cnf(i_0_21042, negated_conjecture, (sk_c6=identity|sk_c5=identity)).
#
#cnf(i_0_21075, negated_conjecture, (sk_c8=sk_c7|sk_c6=identity)).
###
#cnf(i_0_21078, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c8=sk_c7)).
##
#cnf(i_0_21118, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c8=sk_c7)).
##########
#cnf(i_0_21079, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c6|sk_c8=sk_c7)).
#
#cnf(i_0_21080, negated_conjecture, (multiply(sk_c3,sk_c7)=sk_c8|sk_c8=sk_c7)).
#
#cnf(i_0_21341, negated_conjecture, (sk_c8=identity|sk_c8=sk_c7)).
##
#cnf(i_0_21343, negated_conjecture, (sk_c8=identity|sk_c7!=identity)).
#
#cnf(i_0_21742, negated_conjecture, (sk_c2=identity|sk_c4!=sk_c2|sk_c7!=identity)).
##
#cnf(i_0_21884, negated_conjecture, (multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_21773, negated_conjecture, (sk_c1=identity|sk_c5=identity|sk_c7!=identity)).
##
#cnf(i_0_22008, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_21781, negated_conjecture, (sk_c2=identity|sk_c5=identity|sk_c7!=identity)).
##
#cnf(i_0_22029, negated_conjecture, (sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_21854, negated_conjecture, (sk_c1=identity|sk_c8=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_22049, negated_conjecture, (sk_c7!=identity|sk_c8!=sk_c7|sk_c6!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_21863, negated_conjecture, (sk_c2=identity|sk_c8=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_22382, negated_conjecture, (sk_c7!=identity|sk_c8!=sk_c7|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_21905, negated_conjecture, (sk_c4=identity|sk_c6=identity|sk_c7!=identity)).
##
#cnf(i_0_21739, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_21943, negated_conjecture, (sk_c3=identity|sk_c2=sk_c1|sk_c7!=identity)).
##
#cnf(i_0_22576, negated_conjecture, (multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_21950, negated_conjecture, (sk_c4=identity|sk_c2=sk_c1|sk_c7!=identity)).
##
#cnf(i_0_22682, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_21997, negated_conjecture, (sk_c6=sk_c7|sk_c8=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_21998, negated_conjecture, (sk_c4=identity|sk_c8=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_22768, negated_conjecture, (sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
############
#cnf(i_0_21754, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
####
#cnf(i_0_22905, negated_conjecture, (multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
####
#cnf(i_0_22914, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
##
#cnf(i_0_21115, negated_conjecture, (multiply(sk_c8,sk_c4)=identity|sk_c8=sk_c7)).
##
#cnf(i_0_22924, negated_conjecture, (sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
######
#cnf(i_0_21117, negated_conjecture, (multiply(sk_c4,sk_c8)=identity|sk_c8=sk_c7)).
#########
#cnf(i_0_21318, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c8|sk_c8=sk_c7)).
#
#cnf(i_0_21823, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|inverse(sk_c2)!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_23034, negated_conjecture, (sk_c8=sk_c3|sk_c8=sk_c7)).
#
#cnf(i_0_23046, negated_conjecture, (sk_c8=sk_c3|sk_c7!=sk_c3)).
###
#cnf(i_0_21334, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c7|sk_c8=sk_c7)).
##
#cnf(i_0_23581, negated_conjecture, (sk_c6=sk_c7|sk_c8=sk_c7)).
##
#cnf(i_0_23629, negated_conjecture, (sk_c7=identity|sk_c8=sk_c7)).
#
#cnf(i_0_23875, negated_conjecture, (sk_c7=identity|sk_c1=sk_c3)).
##
#cnf(i_0_23552, negated_conjecture, (multiply(sk_c3,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_23950, negated_conjecture, (multiply(sk_c3,sk_c8)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_23962, negated_conjecture, (multiply(sk_c3,sk_c8)!=sk_c7|sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_23975, negated_conjecture, (sk_c7!=sk_c3|sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
####
#cnf(i_0_23982, negated_conjecture, (sk_c8!=identity|sk_c7!=sk_c3|sk_c6!=sk_c7|sk_c7!=identity|sk_c4!=sk_c2)).
####
#cnf(i_0_23990, negated_conjecture, (sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3|sk_c4!=sk_c2)).
#####
#cnf(i_0_21273, negated_conjecture, (multiply(sk_c4,X1)=multiply(sk_c1,X1)|sk_c8=sk_c7)).
###
#cnf(i_0_23554, negated_conjecture, (multiply(sk_c3,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_21279, negated_conjecture, (multiply(sk_c4,X1)=multiply(sk_c2,X1)|sk_c8=sk_c7)).
###
#cnf(i_0_24075, negated_conjecture, (multiply(sk_c3,sk_c8)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_24157, negated_conjecture, (multiply(sk_c3,sk_c8)!=sk_c7|sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_21757, negated_conjecture, (inverse(sk_c2)=identity|sk_c4!=sk_c2|sk_c7!=identity)).
###
#cnf(i_0_24170, negated_conjecture, (sk_c7!=sk_c3|sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
#####
#cnf(i_0_23372, negated_conjecture, (inverse(sk_c3)=sk_c2|sk_c4!=sk_c2|sk_c7!=sk_c3)).
#
#cnf(i_0_24209, negated_conjecture, (inverse(sk_c8)=identity|sk_c4!=sk_c2|sk_c7!=identity)).
######
#cnf(i_0_21771, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c2|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_24270, negated_conjecture, (inverse(sk_c2)=sk_c3|sk_c4!=sk_c2|sk_c7!=sk_c3)).
######
#cnf(i_0_24435, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=sk_c2|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_21751, negated_conjecture, (multiply(sk_c2,X1)=X1|sk_c4!=sk_c2|sk_c7!=identity)).
##########
#cnf(i_0_24220, negated_conjecture, (multiply(sk_c8,X1)=X1|sk_c4!=sk_c2|sk_c7!=identity)).
###############
#cnf(i_0_21415, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c8=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_21493, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c8=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_21673, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c1=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_21674, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c6=sk_c2|sk_c7!=identity)).
##
#cnf(i_0_21675, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c6=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_21677, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c1=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_21678, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c6=sk_c2|sk_c7!=identity)).
#
#cnf(i_0_21679, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c5=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_21680, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c6=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_21834, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(sk_c8,sk_c2)!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_21681, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c5=sk_c4|sk_c7!=identity)).
#
#cnf(i_0_21682, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c5=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_21683, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c6=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_21684, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c5=sk_c4|sk_c7!=identity)).
#
#cnf(i_0_21706, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c1=identity|sk_c7!=identity)).
#
#cnf(i_0_21711, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c2=identity|sk_c7!=identity)).
##
#cnf(i_0_21713, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c4=identity|sk_c7!=identity)).
#
#cnf(i_0_21714, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c4=identity|sk_c7!=identity)).
#
#cnf(i_0_21722, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c4=identity|sk_c7!=identity)).
##
#cnf(i_0_25822, negated_conjecture, (sk_c2=identity|sk_c4=identity|sk_c7!=identity)).
###
#cnf(i_0_23369, negated_conjecture, (multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_25909, negated_conjecture, (multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
#######
#cnf(i_0_21744, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c2=identity|sk_c7!=identity)).
#####
#cnf(i_0_25925, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
#######
#cnf(i_0_21783, negated_conjecture, (inverse(sk_c1)=identity|sk_c5=identity|sk_c7!=identity)).
#
#cnf(i_0_26028, negated_conjecture, (multiply(sk_c7,sk_c3)!=sk_c6|multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
########
#cnf(i_0_23384, negated_conjecture, (multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c7!=sk_c3)).
##
#cnf(i_0_21790, negated_conjecture, (inverse(sk_c2)=identity|sk_c5=identity|sk_c7!=identity)).
##
#cnf(i_0_26105, negated_conjecture, (multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c7!=sk_c3)).
###########
#cnf(i_0_21850, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c5=sk_c8|sk_c7!=identity)).
#
#cnf(i_0_26208, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_21851, negated_conjecture, (inverse(sk_c1)=identity|sk_c8=sk_c7|sk_c7!=identity)).
#######
#cnf(i_0_26268, negated_conjecture, (multiply(sk_c7,sk_c3)!=sk_c6|multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|sk_c8!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c7!=sk_c3)).
###
#cnf(i_0_21869, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c5=sk_c8|sk_c7!=identity)).
##
#cnf(i_0_21872, negated_conjecture, (inverse(sk_c2)=identity|sk_c8=sk_c7|sk_c7!=identity)).
###
#cnf(i_0_23406, negated_conjecture, (multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c8!=identity|sk_c6!=identity|sk_c7!=sk_c3)).
#######
#cnf(i_0_21910, negated_conjecture, (inverse(sk_c4)=identity|sk_c6=identity|sk_c7!=identity)).
########
#cnf(i_0_21923, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c2=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_24276, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=identity|sk_c6!=identity|sk_c4!=sk_c2|sk_c7!=sk_c3)).
#
#cnf(i_0_21933, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c6=sk_c1|sk_c7!=identity)).
#
#cnf(i_0_21947, negated_conjecture, (inverse(sk_c7)=identity|sk_c2=sk_c1|sk_c7!=identity)).
########
#cnf(i_0_21961, negated_conjecture, (inverse(sk_c4)=identity|sk_c2=sk_c1|sk_c7!=identity)).
#########
#cnf(i_0_21964, negated_conjecture, (inverse(sk_c3)=identity|sk_c2=sk_c1|sk_c7!=identity)).
#
#cnf(i_0_25911, negated_conjecture, (multiply(sk_c8,sk_c3)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
########
#cnf(i_0_25917, negated_conjecture, (multiply(sk_c3,sk_c8)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
#####
#cnf(i_0_22468, negated_conjecture, (inverse(sk_c8)=identity|sk_c8=sk_c7|sk_c7!=identity)).
###
#cnf(i_0_26107, negated_conjecture, (multiply(sk_c8,sk_c3)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c7!=sk_c3)).
########
#cnf(i_0_26113, negated_conjecture, (multiply(sk_c3,sk_c8)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c7!=sk_c3)).
########
#cnf(i_0_26509, negated_conjecture, (multiply(sk_c8,sk_c3)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c8!=identity|sk_c6!=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_26515, negated_conjecture, (multiply(sk_c3,sk_c8)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c8!=identity|sk_c6!=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_22752, negated_conjecture, (inverse(sk_c8)=identity|sk_c2=sk_c1|sk_c7!=identity)).
###
#cnf(i_0_26649, negated_conjecture, (multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=identity|sk_c6!=identity|sk_c4!=sk_c2|sk_c7!=sk_c3)).
####
#cnf(i_0_26916, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c8,sk_c3)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
##
#cnf(i_0_23118, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c8=sk_c3|sk_c4!=sk_c2)).
#
#cnf(i_0_23397, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c7=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_27204, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=sk_c3|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_23400, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c7=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_23409, negated_conjecture, (inverse(sk_c3)=sk_c1|sk_c5=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_23423, negated_conjecture, (inverse(sk_c3)=sk_c2|sk_c5=identity|sk_c7!=sk_c3)).
###
#cnf(i_0_23483, negated_conjecture, (inverse(sk_c3)=sk_c1|sk_c8=sk_c7|sk_c7!=sk_c3)).
#
#cnf(i_0_27600, negated_conjecture, (sk_c1=sk_c7|sk_c8=sk_c7|sk_c7!=sk_c3)).
#
#cnf(i_0_27328, negated_conjecture, (multiply(sk_c7,sk_c3)!=sk_c6|multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|sk_c7!=sk_c3|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_23493, negated_conjecture, (inverse(sk_c3)=sk_c2|sk_c8=sk_c7|sk_c7!=sk_c3)).
#
#cnf(i_0_27959, negated_conjecture, (sk_c2=sk_c7|sk_c8=sk_c7|sk_c7!=sk_c3)).
####
#cnf(i_0_23512, negated_conjecture, (inverse(sk_c3)=sk_c4|sk_c6=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_28133, negated_conjecture, (sk_c4=sk_c7|sk_c6=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_27041, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c8,sk_c3)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_23521, negated_conjecture, (inverse(sk_c3)=sk_c3|sk_c2=sk_c1|sk_c7!=sk_c3)).
##
#cnf(i_0_23531, negated_conjecture, (inverse(sk_c3)=sk_c4|sk_c2=sk_c1|sk_c7!=sk_c3)).
#
#cnf(i_0_28219, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c7!=sk_c3|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
#
#cnf(i_0_28332, negated_conjecture, (sk_c4=sk_c7|sk_c2=sk_c1|sk_c7!=sk_c3)).
#
#cnf(i_0_28360, negated_conjecture, (sk_c4=sk_c8|sk_c2=sk_c1|sk_c7!=sk_c3)).
#####
#cnf(i_0_23548, negated_conjecture, (inverse(sk_c3)=sk_c4|sk_c8=sk_c7|sk_c7!=sk_c3)).
#
#cnf(i_0_28381, negated_conjecture, (multiply(sk_c7,sk_c3)!=sk_c6|multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|sk_c7!=sk_c3|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
#
#cnf(i_0_28613, negated_conjecture, (sk_c4=sk_c7|sk_c8=sk_c7|sk_c7!=sk_c3)).
#######
#cnf(i_0_27067, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c8,sk_c3)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c8!=identity|sk_c6!=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_23705, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c7=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_28781, negated_conjecture, (sk_c2=sk_c3|sk_c7=identity|sk_c4!=sk_c2)).
####
#cnf(i_0_24759, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c8=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_27080, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c8)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c8!=identity|sk_c6!=identity|sk_c7!=sk_c3)).
##
#cnf(i_0_24822, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c8=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_27192, negated_conjecture, (multiply(sk_c8,sk_c3)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=identity|sk_c6!=identity|sk_c4!=sk_c2|sk_c7!=sk_c3)).
##
#cnf(i_0_24887, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c1=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_24938, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c6=sk_c2|sk_c7!=identity)).
#
#cnf(i_0_27198, negated_conjecture, (multiply(sk_c3,sk_c8)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=identity|sk_c6!=identity|sk_c4!=sk_c2|sk_c7!=sk_c3)).
#
#cnf(i_0_24989, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c6=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_29219, negated_conjecture, (sk_c4=identity|sk_c6=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_25042, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c1=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_28736, negated_conjecture, (multiply(sk_c3,sk_c3)!=sk_c7|multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c8!=identity|sk_c6!=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_29362, negated_conjecture, (sk_c4=identity|sk_c1=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_25095, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c6=sk_c2|sk_c7!=identity)).
#
#cnf(i_0_28914, negated_conjecture, (multiply(sk_c7,sk_c3)!=sk_c6|multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c8)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|sk_c8!=sk_c7|sk_c8!=identity|sk_c6!=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_29556, negated_conjecture, (sk_c4=identity|sk_c6=sk_c2|sk_c7!=identity)).
##
#cnf(i_0_29620, negated_conjecture, (sk_c2=sk_c7|sk_c4=identity|sk_c7!=identity)).
#
#cnf(i_0_28991, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c8,sk_c3)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=identity|sk_c6!=identity|sk_c4!=sk_c2|sk_c7!=sk_c3)).
#
#cnf(i_0_25148, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c5=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_29808, negated_conjecture, (sk_c1=identity|sk_c5=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_25194, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c6=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_29919, negated_conjecture, (sk_c4=sk_c1|sk_c6=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_29923, negated_conjecture, (sk_c1=identity|sk_c6=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_25261, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c5=sk_c4|sk_c7!=identity)).
##
#cnf(i_0_30121, negated_conjecture, (sk_c1=identity|sk_c5=sk_c4|sk_c7!=identity)).
#
#cnf(i_0_30145, negated_conjecture, (sk_c4=identity|sk_c1=identity|sk_c7!=identity)).
#
#cnf(i_0_30146, negated_conjecture, (sk_c4=sk_c7|sk_c1=identity|sk_c7!=identity)).
#
#cnf(i_0_29121, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c8)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=identity|sk_c6!=identity|sk_c4!=sk_c2|sk_c7!=sk_c3)).
##
#cnf(i_0_25307, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c5=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_30426, negated_conjecture, (sk_c2=sk_c1|sk_c5=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_29371, negated_conjecture, (multiply(sk_c7,sk_c3)!=sk_c6|multiply(sk_c3,sk_c3)!=sk_c7|multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|sk_c8!=sk_c7|sk_c8!=identity|sk_c6!=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_30430, negated_conjecture, (sk_c2=identity|sk_c5=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_25379, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c6=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_30566, negated_conjecture, (sk_c4=sk_c2|sk_c6=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_29711, negated_conjecture, (multiply(sk_c3,sk_c3)!=sk_c7|multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=identity|sk_c6!=identity|sk_c4!=sk_c2|sk_c7!=sk_c3)).
#
#cnf(i_0_30570, negated_conjecture, (sk_c2=sk_c1|sk_c6=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_30573, negated_conjecture, (sk_c2=identity|sk_c6=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_25451, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c5=sk_c4|sk_c7!=identity)).
#
#cnf(i_0_30328, negated_conjecture, (multiply(sk_c7,sk_c3)!=sk_c6|multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c8)!=sk_c7|multiply(sk_c3,sk_c6)!=sk_c7|sk_c8!=identity|sk_c6!=identity|sk_c4!=sk_c2|sk_c7!=sk_c3)).
#
#cnf(i_0_30835, negated_conjecture, (sk_c2=sk_c1|sk_c5=sk_c4|sk_c7!=identity)).
#
#cnf(i_0_30837, negated_conjecture, (sk_c2=identity|sk_c5=sk_c4|sk_c7!=identity)).
##
#cnf(i_0_30633, negated_conjecture, (multiply(sk_c7,sk_c3)!=sk_c6|multiply(sk_c3,sk_c3)!=sk_c7|multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|sk_c8!=identity|sk_c6!=identity|sk_c4!=sk_c2|sk_c7!=sk_c3)).
#
#cnf(i_0_30869, negated_conjecture, (sk_c4=sk_c7|sk_c2=sk_c1|sk_c7!=identity)).
##
#cnf(i_0_30892, negated_conjecture, (sk_c4=sk_c7|sk_c2=identity|sk_c7!=identity)).
###
#cnf(i_0_30947, negated_conjecture, (sk_c7=sk_c3|sk_c2=sk_c1|sk_c7!=identity)).
#
#cnf(i_0_25523, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c1=identity|sk_c7!=identity)).
##
#cnf(i_0_25574, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c2=identity|sk_c7!=identity)).
#
#cnf(i_0_25625, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c4=identity|sk_c7!=identity)).
#######
#cnf(i_0_25938, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c2=identity|sk_c7!=identity)).
##
#cnf(i_0_26073, negated_conjecture, (inverse(sk_c8)=identity|sk_c5=identity|sk_c7!=identity)).
##############
#cnf(i_0_26221, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c5=sk_c8|sk_c7!=identity)).
#
#cnf(i_0_31675, negated_conjecture, (sk_c1=identity|sk_c5=sk_c8|sk_c7!=identity)).
####
#cnf(i_0_31701, negated_conjecture, (sk_c4=sk_c8|sk_c1=identity|sk_c7!=identity)).
#########
#cnf(i_0_26348, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c5=sk_c8|sk_c7!=identity)).
##
#cnf(i_0_31894, negated_conjecture, (sk_c2=sk_c1|sk_c5=sk_c8|sk_c7!=identity)).
#
#cnf(i_0_31896, negated_conjecture, (sk_c2=identity|sk_c5=sk_c8|sk_c7!=identity)).
####
#cnf(i_0_31925, negated_conjecture, (sk_c4=sk_c8|sk_c2=sk_c1|sk_c7!=identity)).
####
#cnf(i_0_31951, negated_conjecture, (sk_c4=sk_c8|sk_c2=identity|sk_c7!=identity)).
#
#cnf(i_0_32124, negated_conjecture, (sk_c2=identity|sk_c2!=sk_c8|sk_c7!=identity)).
##
#cnf(i_0_32000, negated_conjecture, (sk_c8=sk_c3|sk_c2=sk_c1|sk_c7!=identity)).
#############
#cnf(i_0_26557, negated_conjecture, (inverse(sk_c8)=identity|sk_c6=identity|sk_c7!=identity)).
########
#cnf(i_0_26590, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c2=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_26662, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c6=sk_c1|sk_c7!=identity)).
################################
#cnf(i_0_27215, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c8=sk_c3|sk_c4!=sk_c2)).
####
#cnf(i_0_27280, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c7=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_21726, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_27324, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c7=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_27341, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c7=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_32948, negated_conjecture, (sk_c2=sk_c1|sk_c7=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_27415, negated_conjecture, (inverse(sk_c1)=sk_c3|sk_c5=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_32764, negated_conjecture, (multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c8!=sk_c7|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_27485, negated_conjecture, (inverse(sk_c2)=sk_c3|sk_c5=identity|sk_c7!=sk_c3)).
###
#cnf(i_0_27555, negated_conjecture, (inverse(sk_c1)=sk_c3|sk_c8=sk_c7|sk_c7!=sk_c3)).
##
#cnf(i_0_27914, negated_conjecture, (inverse(sk_c2)=sk_c3|sk_c8=sk_c7|sk_c7!=sk_c3)).
#
#cnf(i_0_33212, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_28059, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c8=sk_c7|sk_c7!=sk_c3)).
##
#cnf(i_0_33582, negated_conjecture, (multiply(X1,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_28102, negated_conjecture, (inverse(sk_c4)=sk_c3|sk_c6=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_28306, negated_conjecture, (inverse(sk_c4)=sk_c3|sk_c2=sk_c1|sk_c7!=sk_c3)).
##
#cnf(i_0_28436, negated_conjecture, (inverse(sk_c7)=sk_c7|sk_c2=sk_c1|sk_c7!=sk_c3)).
####
#cnf(i_0_28454, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c2=sk_c1|sk_c7!=sk_c3)).
#
#cnf(i_0_33672, negated_conjecture, (inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6)).
#####
#cnf(i_0_28522, negated_conjecture, (inverse(sk_c8)=sk_c7|sk_c2=sk_c1|sk_c7!=sk_c3)).
##
#cnf(i_0_28540, negated_conjecture, (inverse(sk_c8)=sk_c8|sk_c2=sk_c1|sk_c7!=sk_c3)).
######
#cnf(i_0_28568, negated_conjecture, (inverse(sk_c4)=sk_c3|sk_c8=sk_c7|sk_c7!=sk_c3)).
####
#cnf(i_0_28746, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c7=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_34252, negated_conjecture, (inverse(X1)!=sk_c6|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c8)).
###################
#cnf(i_0_30943, negated_conjecture, (inverse(sk_c7)=sk_c7|sk_c2=sk_c1|sk_c7!=identity)).
####
#cnf(i_0_30961, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c2=sk_c1|sk_c7!=identity)).
##############
#cnf(i_0_31470, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c2=identity|sk_c7!=identity)).
#
#cnf(i_0_35082, negated_conjecture, (sk_c8=sk_c3|sk_c2=identity|sk_c7!=identity)).
#
#cnf(i_0_35210, negated_conjecture, (sk_c1=identity|sk_c8=sk_c3|sk_c7!=identity)).
##
#cnf(i_0_35461, negated_conjecture, (sk_c3=identity|sk_c1=identity|sk_c7!=identity)).
#
#cnf(i_0_35476, negated_conjecture, (sk_c7=sk_c3|sk_c1=identity|sk_c7!=identity)).
######
#cnf(i_0_33674, negated_conjecture, (multiply(X1,sk_c3)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
###
#cnf(i_0_31996, negated_conjecture, (inverse(sk_c8)=sk_c7|sk_c2=sk_c1|sk_c7!=identity)).
#
#cnf(i_0_35722, negated_conjecture, (sk_c6!=sk_c3|sk_c8!=identity|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c7!=sk_c3)).
###
#cnf(i_0_32014, negated_conjecture, (inverse(sk_c8)=sk_c8|sk_c2=sk_c1|sk_c7!=identity)).
##########
#cnf(i_0_32125, negated_conjecture, (inverse(sk_c8)=sk_c8|sk_c2=identity|sk_c7!=identity)).
#######
#cnf(i_0_34253, negated_conjecture, (multiply(inverse(X2),inverse(X1))!=sk_c8|multiply(X1,X2)!=sk_c6|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
#####
#cnf(i_0_32681, negated_conjecture, (inverse(sk_c3)=sk_c1|sk_c7=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_36613, negated_conjecture, (sk_c1=sk_c7|sk_c7=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_36652, negated_conjecture, (sk_c8=sk_c1|sk_c7=identity|sk_c7!=sk_c3)).
######
#cnf(i_0_32683, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c7=identity|sk_c7!=sk_c3)).
##
#cnf(i_0_32782, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c7=identity|sk_c7!=sk_c3)).
###
#cnf(i_0_37245, negated_conjecture, (sk_c2=sk_c3|sk_c7=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_32952, negated_conjecture, (inverse(sk_c3)=sk_c2|sk_c7=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_37428, negated_conjecture, (sk_c2=sk_c7|sk_c7=identity|sk_c7!=sk_c3)).
##
#cnf(i_0_37466, negated_conjecture, (sk_c2=sk_c8|sk_c7=identity|sk_c7!=sk_c3)).
#########
#cnf(i_0_32954, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c7=identity|sk_c7!=sk_c3)).
##############
#cnf(i_0_34845, negated_conjecture, (multiply(inverse(X2),inverse(X1))!=sk_c6|multiply(X1,X2)!=sk_c8|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_33598, negated_conjecture, (inverse(sk_c8)=sk_c7|sk_c8=sk_c7|sk_c7!=sk_c3)).
###################################
#cnf(i_0_35010, negated_conjecture, (inverse(sk_c8)=sk_c7|sk_c2=identity|sk_c7!=identity)).
##########
#cnf(i_0_36575, negated_conjecture, (inverse(sk_c1)=sk_c3|sk_c7=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_36330, negated_conjecture, (multiply(inverse(X1),X2)!=sk_c8|multiply(inverse(X2),X1)!=sk_c6|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_36719, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c7=identity|sk_c7!=sk_c3)).
#####
#cnf(i_0_37034, negated_conjecture, (inverse(sk_c1)=sk_c1|sk_c7=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_37035, negated_conjecture, (inverse(sk_c1)=sk_c2|sk_c7=identity|sk_c7!=sk_c3)).
######
#cnf(i_0_37052, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c7=identity|sk_c7!=sk_c3)).
#######
#cnf(i_0_36507, negated_conjecture, (multiply(X1,inverse(X2))!=sk_c8|multiply(X2,inverse(X1))!=sk_c6|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_37254, negated_conjecture, (inverse(sk_c3)=sk_c3|sk_c7=identity|sk_c7!=sk_c3)).
#########
#cnf(i_0_37390, negated_conjecture, (inverse(sk_c2)=sk_c3|sk_c7=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_37711, negated_conjecture, (inverse(sk_c8)=sk_c8|sk_c7=identity|sk_c7!=sk_c3)).
#######
#cnf(i_0_37729, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c7=identity|sk_c7!=sk_c3)).
############
#cnf(i_0_39080, negated_conjecture, (inverse(sk_c8)=sk_c7|sk_c7=identity|sk_c7!=sk_c3)).
#########
#cnf(i_0_39247, negated_conjecture, (inverse(sk_c2)=sk_c1|sk_c7=identity|sk_c7!=sk_c3)).
###########
#cnf(i_0_39402, negated_conjecture, (inverse(sk_c7)=sk_c7|sk_c7=identity|sk_c7!=sk_c3)).
########################################
#cnf(i_0_21270, negated_conjecture, (multiply(inverse(sk_c1),sk_c4)=identity|sk_c8=sk_c7)).
#########
#cnf(i_0_21272, negated_conjecture, (multiply(sk_c4,inverse(sk_c1))=identity|sk_c8=sk_c7)).
###
#cnf(i_0_40791, negated_conjecture, (inverse(sk_c4)=inverse(sk_c1)|sk_c8=sk_c7)).
#############
#cnf(i_0_21275, negated_conjecture, (multiply(inverse(sk_c2),sk_c4)=identity|sk_c8=sk_c7)).
###########
#cnf(i_0_21276, negated_conjecture, (multiply(sk_c4,inverse(sk_c2))=identity|sk_c8=sk_c7)).
##
#cnf(i_0_41030, negated_conjecture, (inverse(sk_c4)=inverse(sk_c2)|sk_c8=sk_c7)).
#############
#cnf(i_0_33768, negated_conjecture, (multiply(X1,multiply(X2,sk_c8))!=sk_c6|multiply(inverse(X2),inverse(X1))!=sk_c8|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
###
#cnf(i_0_21299, negated_conjecture, (multiply(inverse(sk_c7),sk_c6)=sk_c8|sk_c8=sk_c7)).
###
#cnf(i_0_41416, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c8=sk_c7)).
###
#cnf(i_0_41462, negated_conjecture, (inverse(sk_c8)=sk_c7|sk_c8=sk_c7)).
#
#cnf(i_0_41623, negated_conjecture, (sk_c1=sk_c7|sk_c8=sk_c7)).
#
#cnf(i_0_41626, negated_conjecture, (sk_c2=sk_c7|sk_c8=sk_c7)).
##
#cnf(i_0_41651, negated_conjecture, (sk_c4=sk_c7|sk_c8=sk_c7)).
#########
#cnf(i_0_42091, negated_conjecture, (inverse(sk_c1)=inverse(sk_c7)|sk_c8=sk_c7)).
##
#cnf(i_0_41304, negated_conjecture, (multiply(sk_c8,inverse(X1))!=sk_c8|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6)).
###
#cnf(i_0_42247, negated_conjecture, (inverse(sk_c2)=inverse(sk_c7)|sk_c8=sk_c7)).
#####################
#cnf(i_0_42529, negated_conjecture, (multiply(sk_c8,X1)!=sk_c8|inverse(X1)!=sk_c6|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
##########
#cnf(i_0_41459, negated_conjecture, (multiply(sk_c8,sk_c7)=identity|sk_c8=sk_c7)).
#######
#cnf(i_0_41461, negated_conjecture, (multiply(sk_c7,sk_c8)=identity|sk_c8=sk_c7)).
###
#cnf(i_0_42674, negated_conjecture, (multiply(inverse(X1),sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c8)).
#####
#cnf(i_0_41534, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c8|sk_c8=sk_c7)).
#####
#cnf(i_0_43131, negated_conjecture, (sk_c6=sk_c8|sk_c8=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_43188, negated_conjecture, (sk_c8=sk_c7|sk_c7!=identity)).
############
#cnf(i_0_43081, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c6|sk_c8=sk_c7)).
#
#cnf(i_0_43515, negated_conjecture, (multiply(X1,sk_c7)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c8!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
###
#cnf(i_0_43082, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|sk_c8=sk_c7)).
#
#cnf(i_0_43644, negated_conjecture, (sk_c6!=sk_c7|sk_c8!=identity|sk_c8!=sk_c7|sk_c7!=identity)).
#######
#cnf(i_0_43121, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c8|sk_c8=sk_c7)).
#
#cnf(i_0_43449, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
####
#cnf(i_0_43774, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
####
#cnf(i_0_43792, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
####
#cnf(i_0_43809, negated_conjecture, (sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
#
#cnf(i_0_43742, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|sk_c8=sk_c7)).
###
#cnf(i_0_43826, negated_conjecture, (sk_c8!=identity|sk_c6!=sk_c7|sk_c7!=identity|sk_c4!=sk_c2)).
####
#cnf(i_0_43877, negated_conjecture, (sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_43743, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c7|sk_c8=sk_c7)).
##
#cnf(i_0_43452, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
####
#cnf(i_0_43982, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
####
#cnf(i_0_44000, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
####
#cnf(i_0_44017, negated_conjecture, (sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
########
#cnf(i_0_43499, negated_conjecture, (multiply(sk_c7,sk_c7)!=sk_c6|multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|sk_c7!=sk_c3|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
####
#cnf(i_0_43504, negated_conjecture, (multiply(sk_c7,sk_c7)!=sk_c6|multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|sk_c7!=sk_c3|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
####
#cnf(i_0_43269, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
####
#cnf(i_0_44050, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
############
#cnf(i_0_44077, negated_conjecture, (multiply(sk_c7,sk_c7)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
############
#cnf(i_0_43284, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
#
#cnf(i_0_41519, negated_conjecture, (multiply(sk_c8,X1)=multiply(sk_c3,X1)|sk_c8=sk_c7)).
##
#cnf(i_0_44152, negated_conjecture, (multiply(sk_c3,sk_c2)=identity|sk_c8=sk_c7)).
#
#cnf(i_0_44104, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
##
#cnf(i_0_44279, negated_conjecture, (inverse(sk_c3)=sk_c2|sk_c8=sk_c7)).
###
#cnf(i_0_44324, negated_conjecture, (inverse(sk_c2)=sk_c3|sk_c8=sk_c7)).
####
#cnf(i_0_44453, negated_conjecture, (inverse(sk_c1)=sk_c3|sk_c8=sk_c7)).
###
#cnf(i_0_44310, negated_conjecture, (multiply(sk_c7,sk_c7)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
###
#cnf(i_0_44496, negated_conjecture, (inverse(sk_c3)=sk_c1|sk_c8=sk_c7)).
#########
#cnf(i_0_44057, negated_conjecture, (multiply(sk_c8,sk_c7)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
######
#cnf(i_0_44159, negated_conjecture, (multiply(sk_c3,sk_c1)=identity|sk_c8=sk_c7)).
##
#cnf(i_0_44066, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
###
#cnf(i_0_44166, negated_conjecture, (multiply(sk_c3,sk_c4)=identity|sk_c8=sk_c7)).
###
#cnf(i_0_44712, negated_conjecture, (inverse(sk_c3)=sk_c4|sk_c8=sk_c7)).
##
#cnf(i_0_44111, negated_conjecture, (multiply(sk_c8,sk_c7)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
###
#cnf(i_0_44746, negated_conjecture, (inverse(sk_c4)=sk_c3|sk_c8=sk_c7)).
#####
#cnf(i_0_44120, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
########
#cnf(i_0_44665, negated_conjecture, (multiply(sk_c7,sk_c7)!=sk_c6|multiply(sk_c8,sk_c7)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
####
#cnf(i_0_44940, negated_conjecture, (multiply(sk_c7,sk_c7)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_44321, negated_conjecture, (multiply(sk_c2,sk_c3)=identity|sk_c8=sk_c7)).
#########
#cnf(i_0_44454, negated_conjecture, (multiply(sk_c1,sk_c3)=identity|sk_c8=sk_c7)).
##
#cnf(i_0_44828, negated_conjecture, (multiply(sk_c7,sk_c7)!=sk_c6|multiply(sk_c8,sk_c7)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
####
#cnf(i_0_45013, negated_conjecture, (multiply(sk_c7,sk_c7)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
##
#cnf(i_0_44466, negated_conjecture, (multiply(sk_c4,sk_c3)=identity|sk_c8=sk_c7)).
##########
#cnf(i_0_43488, negated_conjecture, (multiply(sk_c7,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_45089, negated_conjecture, (multiply(sk_c7,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_43489, negated_conjecture, (multiply(sk_c7,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c7!=sk_c3)).
####
#cnf(i_0_45106, negated_conjecture, (multiply(sk_c7,sk_c7)!=sk_c6|multiply(sk_c3,sk_c6)!=sk_c7|sk_c8!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c7!=sk_c3)).
########
#cnf(i_0_44071, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity|sk_c7!=sk_c3)).
############
#cnf(i_0_44304, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity|sk_c7!=sk_c3)).
############################
#cnf(i_0_32686, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6)).
############
#cnf(i_0_45147, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c8)).
############
#cnf(i_0_45157, negated_conjecture, (multiply(sk_c8,sk_c7)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6)).
#
#cnf(i_0_41625, negated_conjecture, (multiply(sk_c1,X1)=multiply(sk_c7,X1)|sk_c8=sk_c7)).
###
#cnf(i_0_45196, negated_conjecture, (multiply(sk_c8,sk_c7)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6)).
####
#cnf(i_0_45265, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6)).
#
#cnf(i_0_41631, negated_conjecture, (multiply(sk_c2,X1)=multiply(sk_c7,X1)|sk_c8=sk_c7)).
###
#cnf(i_0_45286, negated_conjecture, (inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6)).
###############
#cnf(i_0_44791, negated_conjecture, (multiply(sk_c4,X1)=multiply(sk_c7,X1)|sk_c8=sk_c7)).
#####
#cnf(i_0_45580, negated_conjecture, (inverse(X1)!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c8)).
#########
#cnf(i_0_43272, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c4!=sk_c2|sk_c7!=identity)).
#
#cnf(i_0_43473, negated_conjecture, (inverse(sk_c7)=identity|sk_c4!=sk_c2|sk_c7!=identity)).
###
#cnf(i_0_45866, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c4!=sk_c2|sk_c7!=identity)).
#####
#cnf(i_0_43474, negated_conjecture, (multiply(sk_c7,X1)=X1|sk_c4!=sk_c2|sk_c7!=identity)).
########
#cnf(i_0_41758, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c1=sk_c7|sk_c4!=sk_c2)).
#
#cnf(i_0_41964, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c1=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_46347, negated_conjecture, (sk_c4=sk_c3|sk_c1=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_46455, negated_conjecture, (sk_c3=identity|sk_c1=sk_c7|sk_c7!=identity)).
#########
#cnf(i_0_43308, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c5=identity|sk_c7!=identity)).
##
#cnf(i_0_43322, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c5=identity|sk_c7!=identity)).
#####
#cnf(i_0_43410, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c6=identity|sk_c7!=identity)).
#
#cnf(i_0_46761, negated_conjecture, (sk_c4=sk_c3|sk_c6=identity|sk_c7!=identity)).
######
#cnf(i_0_43451, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c2=identity|sk_c7!=identity)).
#
#cnf(i_0_46939, negated_conjecture, (sk_c4=sk_c3|sk_c2=identity|sk_c7!=identity)).
#
#cnf(i_0_47032, negated_conjecture, (sk_c2=identity|sk_c2!=sk_c3|sk_c7!=identity)).
#
#cnf(i_0_45581, negated_conjecture, (multiply(inverse(X2),inverse(X1))!=sk_c8|multiply(X1,X2)!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_47040, negated_conjecture, (sk_c3=identity|sk_c2=identity|sk_c7!=identity)).
#
#cnf(i_0_47051, negated_conjecture, (sk_c7=sk_c3|sk_c2=identity|sk_c7!=identity)).
############
#cnf(i_0_43457, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c4=identity|sk_c7!=identity)).
###
#cnf(i_0_43476, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c6=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_48000, negated_conjecture, (sk_c4=sk_c3|sk_c6=sk_c7|sk_c7!=identity)).
###
#cnf(i_0_43478, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c6=sk_c2|sk_c7!=identity)).
#
#cnf(i_0_48163, negated_conjecture, (sk_c4=sk_c3|sk_c6=sk_c2|sk_c7!=identity)).
###
#cnf(i_0_48244, negated_conjecture, (sk_c2=sk_c7|sk_c4=sk_c3|sk_c7!=identity)).
#
#cnf(i_0_48335, negated_conjecture, (sk_c3=identity|sk_c2=sk_c7|sk_c7!=identity)).
###
#cnf(i_0_43479, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c5=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_43480, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c6=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_48631, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_43481, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c5=sk_c4|sk_c7!=identity)).
##
#cnf(i_0_43482, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c5=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_43483, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c6=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_48970, negated_conjecture, (sk_c2=sk_c3|sk_c6=sk_c7|sk_c7!=identity)).
####
#cnf(i_0_43484, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c5=sk_c4|sk_c7!=identity)).
##
#cnf(i_0_43485, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c4=identity|sk_c7!=identity)).
###
#cnf(i_0_43490, negated_conjecture, (inverse(sk_c7)=identity|sk_c5=identity|sk_c7!=identity)).
####
#cnf(i_0_43492, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c5=sk_c8|sk_c7!=identity)).
#
#cnf(i_0_43493, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c5=sk_c8|sk_c7!=identity)).
#
#cnf(i_0_45861, negated_conjecture, (multiply(inverse(X2),inverse(X1))!=sk_c6|multiply(X1,X2)!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_43496, negated_conjecture, (inverse(sk_c7)=identity|sk_c6=identity|sk_c7!=identity)).
#
#cnf(i_0_50142, negated_conjecture, (sk_c3=identity|sk_c6=identity|sk_c7!=identity)).
#######
#cnf(i_0_43517, negated_conjecture, (inverse(sk_c7)=sk_c7|sk_c2=identity|sk_c7!=identity)).
########
#cnf(i_0_46200, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c1=sk_c7|sk_c4!=sk_c2)).
####
#cnf(i_0_46282, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c1=sk_c7|sk_c7!=identity)).
####
#cnf(i_0_46444, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c1=sk_c7|sk_c7!=identity)).
####
#cnf(i_0_46557, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c5=identity|sk_c7!=identity)).
####
#cnf(i_0_46645, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c5=identity|sk_c7!=identity)).
####
#cnf(i_0_46729, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c6=identity|sk_c7!=identity)).
####
#cnf(i_0_46871, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c2=identity|sk_c7!=identity)).
####
#cnf(i_0_47033, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c2=identity|sk_c7!=identity)).
####
#cnf(i_0_47862, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c4=identity|sk_c7!=identity)).
####
#cnf(i_0_47950, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c6=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_47324, negated_conjecture, (multiply(inverse(X1),X2)!=sk_c8|multiply(inverse(X2),X1)!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_48095, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c6=sk_c2|sk_c7!=identity)).
####
#cnf(i_0_48485, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c5=sk_c7|sk_c7!=identity)).
####
#cnf(i_0_48581, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c6=sk_c7|sk_c7!=identity)).
####
#cnf(i_0_48731, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c5=sk_c4|sk_c7!=identity)).
####
#cnf(i_0_48827, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c5=sk_c7|sk_c7!=identity)).
####
#cnf(i_0_48920, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c6=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_47530, negated_conjecture, (multiply(X1,inverse(X2))!=sk_c8|multiply(X2,inverse(X1))!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_49067, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c5=sk_c4|sk_c7!=identity)).
####
#cnf(i_0_49160, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c4=identity|sk_c7!=identity)).
####
#cnf(i_0_49364, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c5=sk_c8|sk_c7!=identity)).
####
#cnf(i_0_49464, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c5=sk_c8|sk_c7!=identity)).
####
#cnf(i_0_50143, negated_conjecture, (inverse(sk_c3)=identity|sk_c6=identity|sk_c7!=identity)).
############
#cnf(i_0_50575, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c1=sk_c7|sk_c7!=identity)).
###############
#cnf(i_0_51018, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c2=identity|sk_c7!=identity)).
################################
#cnf(i_0_21325, negated_conjecture, (multiply(inverse(sk_c3),sk_c8)=sk_c7|sk_c8=sk_c7)).
#############
#cnf(i_0_53834, negated_conjecture, (multiply(sk_c4,sk_c8)=sk_c7|sk_c8=sk_c7)).
########################
#cnf(i_0_23018, negated_conjecture, (multiply(inverse(sk_c3),sk_c8)=sk_c6|sk_c8=sk_c7)).
############
#cnf(i_0_53990, negated_conjecture, (multiply(sk_c4,sk_c8)=sk_c6|sk_c8=sk_c7)).
#################
#cnf(i_0_24023, negated_conjecture, (multiply(sk_c1,inverse(sk_c4))=identity|sk_c8=sk_c7)).
########
#cnf(i_0_24104, negated_conjecture, (multiply(sk_c2,inverse(sk_c4))=identity|sk_c8=sk_c7)).
##########
#cnf(i_0_40733, negated_conjecture, (multiply(inverse(sk_c4),sk_c1)=identity|sk_c8=sk_c7)).
################
#cnf(i_0_40952, negated_conjecture, (multiply(inverse(sk_c4),sk_c2)=identity|sk_c8=sk_c7)).
#####################
#cnf(i_0_41516, negated_conjecture, (multiply(sk_c8,inverse(sk_c3))=identity|sk_c8=sk_c7)).
####
#cnf(i_0_54370, negated_conjecture, (inverse(sk_c8)=inverse(sk_c3)|sk_c8=sk_c7)).
###################
#cnf(i_0_41518, negated_conjecture, (multiply(inverse(sk_c3),sk_c8)=identity|sk_c8=sk_c7)).
##############
#cnf(i_0_41622, negated_conjecture, (multiply(inverse(sk_c1),sk_c7)=identity|sk_c8=sk_c7)).
########
#cnf(i_0_41624, negated_conjecture, (multiply(sk_c7,inverse(sk_c1))=identity|sk_c8=sk_c7)).
#######
#cnf(i_0_41627, negated_conjecture, (multiply(inverse(sk_c2),sk_c7)=identity|sk_c8=sk_c7)).
########
#cnf(i_0_41628, negated_conjecture, (multiply(sk_c7,inverse(sk_c2))=identity|sk_c8=sk_c7)).
#######
#cnf(i_0_42092, negated_conjecture, (multiply(sk_c1,inverse(sk_c7))=identity|sk_c8=sk_c7)).
######
#cnf(i_0_42093, negated_conjecture, (multiply(inverse(sk_c7),sk_c1)=identity|sk_c8=sk_c7)).
###########
#cnf(i_0_42113, negated_conjecture, (multiply(inverse(sk_c7),sk_c4)=identity|sk_c8=sk_c7)).
########
#cnf(i_0_42114, negated_conjecture, (multiply(sk_c4,inverse(sk_c7))=identity|sk_c8=sk_c7)).
#
#cnf(i_0_55146, negated_conjecture, (inverse(sk_c4)=inverse(sk_c7)|sk_c8=sk_c7)).
######################
#cnf(i_0_42306, negated_conjecture, (multiply(sk_c2,inverse(sk_c7))=identity|sk_c8=sk_c7)).
########
#cnf(i_0_42307, negated_conjecture, (multiply(inverse(sk_c7),sk_c2)=identity|sk_c8=sk_c7)).
##################
#cnf(i_0_43530, negated_conjecture, (multiply(inverse(sk_c1),sk_c6)=sk_c8|sk_c8=sk_c7)).
###########
#cnf(i_0_43646, negated_conjecture, (multiply(inverse(sk_c2),sk_c6)=sk_c8|sk_c8=sk_c7)).
###########
#cnf(i_0_43842, negated_conjecture, (multiply(inverse(sk_c1),sk_c7)=sk_c8|sk_c8=sk_c7)).
##############
#cnf(i_0_43913, negated_conjecture, (multiply(inverse(sk_c2),sk_c7)=sk_c8|sk_c8=sk_c7)).
###################
#cnf(i_0_44131, negated_conjecture, (multiply(sk_c3,inverse(sk_c8))=identity|sk_c8=sk_c7)).
##
#cnf(i_0_47119, negated_conjecture, (multiply(inverse(X1),multiply(inverse(X3),inverse(X2)))!=sk_c8|multiply(X2,multiply(X3,X1))!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
#############
#cnf(i_0_44452, negated_conjecture, (multiply(inverse(sk_c8),sk_c3)=identity|sk_c8=sk_c7)).
##########################
#cnf(i_0_45601, negated_conjecture, (multiply(sk_c7,inverse(sk_c4))=identity|sk_c8=sk_c7)).
##########
#cnf(i_0_53839, negated_conjecture, (multiply(inverse(sk_c4),sk_c7)=sk_c8|sk_c8=sk_c7)).
##################
#cnf(i_0_53995, negated_conjecture, (multiply(inverse(sk_c4),sk_c6)=sk_c8|sk_c8=sk_c7)).
##################################
#cnf(i_0_54315, negated_conjecture, (multiply(inverse(sk_c4),sk_c7)=identity|sk_c8=sk_c7)).
###############################
#cnf(i_0_49601, negated_conjecture, (multiply(inverse(X1),multiply(inverse(X3),inverse(X2)))!=sk_c6|multiply(X2,multiply(X3,X1))!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
###################################################################
#cnf(i_0_21774, negated_conjecture, (multiply(sk_c1,X1)=X1|sk_c5=identity|sk_c7!=identity)).
###########
#cnf(i_0_21791, negated_conjecture, (multiply(sk_c2,X1)=X1|sk_c5=identity|sk_c7!=identity)).
###################
#cnf(i_0_21906, negated_conjecture, (multiply(sk_c4,X1)=X1|sk_c6=identity|sk_c7!=identity)).
#########
#cnf(i_0_21942, negated_conjecture, (multiply(sk_c3,X1)=X1|sk_c2=sk_c1|sk_c7!=identity)).
#########
#cnf(i_0_21949, negated_conjecture, (multiply(sk_c7,X1)=X1|sk_c2=sk_c1|sk_c7!=identity)).
########
#cnf(i_0_21958, negated_conjecture, (multiply(sk_c4,X1)=X1|sk_c2=sk_c1|sk_c7!=identity)).
#########
#cnf(i_0_51519, negated_conjecture, (multiply(inverse(X3),multiply(inverse(X2),X1))!=sk_c6|multiply(inverse(X1),multiply(X2,X3))!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
#########
#cnf(i_0_22751, negated_conjecture, (multiply(sk_c8,X1)=X1|sk_c2=sk_c1|sk_c7!=identity)).
##############
#cnf(i_0_26076, negated_conjecture, (multiply(sk_c8,X1)=X1|sk_c5=identity|sk_c7!=identity)).
###############
#cnf(i_0_26558, negated_conjecture, (multiply(sk_c8,X1)=X1|sk_c6=identity|sk_c7!=identity)).
#######################################
#cnf(i_0_49257, negated_conjecture, (multiply(sk_c7,X1)=X1|sk_c5=identity|sk_c7!=identity)).
#####
#cnf(i_0_50104, negated_conjecture, (multiply(sk_c7,X1)=X1|sk_c6=identity|sk_c7!=identity)).
#######
#cnf(i_0_50137, negated_conjecture, (multiply(sk_c3,X1)=X1|sk_c6=identity|sk_c7!=identity)).
#######################################################
#cnf(i_0_51796, negated_conjecture, (multiply(inverse(X2),multiply(inverse(X1),X3))!=sk_c8|multiply(inverse(X3),multiply(X1,X2))!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
##########################
#cnf(i_0_21452, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c5=identity|sk_c8=identity)).
#
#cnf(i_0_21466, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c5=identity|sk_c8=identity)).
###
#cnf(i_0_21476, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c8=identity|sk_c6=identity)).
#
#cnf(i_0_62701, negated_conjecture, (sk_c2=sk_c3|sk_c6=identity|sk_c8=identity)).
##
#cnf(i_0_21478, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c8=identity|sk_c6=identity)).
#
#cnf(i_0_62884, negated_conjecture, (sk_c1=sk_c3|sk_c6=identity|sk_c8=identity)).
#####
#cnf(i_0_21516, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c8=identity|sk_c6=sk_c8)).
##
#cnf(i_0_63081, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c8|sk_c8=identity)).
#
#cnf(i_0_63217, negated_conjecture, (sk_c8=identity|sk_c1=sk_c3)).
###
#cnf(i_0_63546, negated_conjecture, (sk_c8=sk_c7|sk_c1=sk_c3)).
##
#cnf(i_0_63289, negated_conjecture, (sk_c2=identity|sk_c1=sk_c3|sk_c4!=sk_c2)).
####
#cnf(i_0_63467, negated_conjecture, (sk_c3=identity|sk_c1=sk_c3|sk_c7!=sk_c3)).
####
#cnf(i_0_63318, negated_conjecture, (sk_c1=identity|sk_c5=identity|sk_c1=sk_c3)).
####
#cnf(i_0_63326, negated_conjecture, (sk_c2=identity|sk_c5=identity|sk_c1=sk_c3)).
############
#cnf(i_0_63416, negated_conjecture, (sk_c4=identity|sk_c6=identity|sk_c1=sk_c3)).
####
#cnf(i_0_63429, negated_conjecture, (sk_c3=identity|sk_c2=sk_c1|sk_c1=sk_c3)).
####
#cnf(i_0_63436, negated_conjecture, (sk_c4=identity|sk_c2=sk_c1|sk_c1=sk_c3)).
################################################
#cnf(i_0_63303, negated_conjecture, (inverse(sk_c2)=identity|sk_c1=sk_c3|sk_c4!=sk_c2)).
###########
#cnf(i_0_63631, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c1=sk_c3|sk_c4!=sk_c2)).
#
#cnf(i_0_64444, negated_conjecture, (inverse(sk_c8)=identity|sk_c1=sk_c3|sk_c4!=sk_c2)).
#####
#cnf(i_0_52620, negated_conjecture, (multiply(X3,multiply(inverse(X2),inverse(X1)))!=sk_c6|multiply(X1,multiply(X2,inverse(X3)))!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
###
#cnf(i_0_64530, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c1=sk_c3|sk_c4!=sk_c2)).
########
#cnf(i_0_64802, negated_conjecture, (inverse(sk_c7)=identity|sk_c1=sk_c3|sk_c4!=sk_c2)).
########
#cnf(i_0_63298, negated_conjecture, (multiply(sk_c2,X1)=X1|sk_c1=sk_c3|sk_c4!=sk_c2)).
############
#cnf(i_0_64455, negated_conjecture, (multiply(sk_c8,X1)=X1|sk_c1=sk_c3|sk_c4!=sk_c2)).
###########
#cnf(i_0_66032, negated_conjecture, (multiply(sk_c7,X1)=X1|sk_c1=sk_c3|sk_c4!=sk_c2)).
#############
#cnf(i_0_21517, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c8=identity|sk_c6=sk_c8)).
##
#cnf(i_0_66595, negated_conjecture, (sk_c2=sk_c3|sk_c6=sk_c8|sk_c8=identity)).
#
#cnf(i_0_66727, negated_conjecture, (sk_c8=identity|sk_c2=sk_c3)).
###########
#cnf(i_0_66748, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c8=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_66838, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c8=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_66886, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c8=identity|sk_c4!=sk_c2)).
####
#cnf(i_0_66992, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c8=identity|sk_c4!=sk_c2)).
###########
#cnf(i_0_21582, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c6=identity|sk_c8=identity)).
#
#cnf(i_0_67424, negated_conjecture, (sk_c4=sk_c3|sk_c6=identity|sk_c8=identity)).
####
#cnf(i_0_21595, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c8=identity|sk_c6=sk_c7)).
###
#cnf(i_0_21605, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c8=identity|sk_c6=sk_c7)).
######################
#cnf(i_0_23155, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c5=identity|sk_c8=sk_c3)).
##
#cnf(i_0_23169, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c5=identity|sk_c8=sk_c3)).
#####
#cnf(i_0_23258, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c6=identity|sk_c8=sk_c3)).
#
#cnf(i_0_68164, negated_conjecture, (sk_c4=sk_c3|sk_c6=identity|sk_c8=sk_c3)).
#
#cnf(i_0_52704, negated_conjecture, (multiply(X1,multiply(inverse(X3),inverse(X2)))!=sk_c8|multiply(X2,multiply(X3,inverse(X1)))!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
#################
#cnf(i_0_23742, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c5=identity|sk_c7=identity)).
##
#cnf(i_0_23756, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c5=identity|sk_c7=identity)).
##
#cnf(i_0_69665, negated_conjecture, (sk_c2=sk_c3|sk_c7=identity|sk_c5=identity)).
#####
#cnf(i_0_23779, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c7=identity|sk_c5=identity)).
#
#cnf(i_0_69824, negated_conjecture, (sk_c4=sk_c3|sk_c7=identity|sk_c5=identity)).
#######
#cnf(i_0_23845, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c6=identity|sk_c7=identity)).
#
#cnf(i_0_69991, negated_conjecture, (sk_c4=sk_c3|sk_c6=identity|sk_c7=identity)).
############
#cnf(i_0_23904, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c5=identity|sk_c7=identity)).
#######################################
#cnf(i_0_41794, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c5=identity|sk_c1=sk_c7)).
#
#cnf(i_0_41808, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c5=identity|sk_c1=sk_c7)).
######
#cnf(i_0_41896, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c6=identity|sk_c1=sk_c7)).
#
#cnf(i_0_70597, negated_conjecture, (sk_c4=sk_c3|sk_c6=identity|sk_c1=sk_c7)).
##############################################
#cnf(i_0_56736, negated_conjecture, (multiply(inverse(X1),multiply(X2,inverse(X3)))!=sk_c8|multiply(X3,multiply(inverse(X2),X1))!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
#########################################################################################################################
#cnf(i_0_62399, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c5=identity|sk_c8=identity)).
####
#cnf(i_0_62526, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c5=identity|sk_c8=identity)).
####
#cnf(i_0_62649, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c8=identity|sk_c6=identity)).
###
#cnf(i_0_58316, negated_conjecture, (multiply(inverse(X1),multiply(X2,inverse(X3)))!=sk_c6|multiply(X3,multiply(inverse(X2),X1))!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_62832, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c8=identity|sk_c6=identity)).
####
#cnf(i_0_63021, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c8=identity|sk_c6=sk_c8)).
######
#cnf(i_0_63220, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c1=sk_c7|sk_c1=sk_c3)).
###
#cnf(i_0_63221, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c6=sk_c2|sk_c1=sk_c3)).
#
#cnf(i_0_63222, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c6=sk_c7|sk_c1=sk_c3)).
###
#cnf(i_0_63224, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c1=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_63225, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c6=sk_c2|sk_c1=sk_c3)).
##
#cnf(i_0_63226, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c5=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_63227, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c6=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_63228, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c5=sk_c4|sk_c1=sk_c3)).
##
#cnf(i_0_63229, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c5=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_63230, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c6=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_63231, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c5=sk_c4|sk_c1=sk_c3)).
##
#cnf(i_0_63253, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c1=identity|sk_c1=sk_c3)).
#
#cnf(i_0_63258, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c2=identity|sk_c1=sk_c3)).
#
#cnf(i_0_63260, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c4=identity|sk_c1=sk_c3)).
##
#cnf(i_0_63261, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c4=identity|sk_c1=sk_c3)).
#
#cnf(i_0_63269, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c4=identity|sk_c1=sk_c3)).
#
#cnf(i_0_75516, negated_conjecture, (sk_c2=identity|sk_c4=identity|sk_c1=sk_c3)).
####################
#cnf(i_0_63291, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c2=identity|sk_c1=sk_c3)).
############
#cnf(i_0_63328, negated_conjecture, (inverse(sk_c1)=identity|sk_c5=identity|sk_c1=sk_c3)).
############
#cnf(i_0_63335, negated_conjecture, (inverse(sk_c2)=identity|sk_c5=identity|sk_c1=sk_c3)).
###############
#cnf(i_0_63387, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c5=sk_c8|sk_c1=sk_c3)).
####
#cnf(i_0_63404, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c5=sk_c8|sk_c1=sk_c3)).
####
#cnf(i_0_63421, negated_conjecture, (inverse(sk_c4)=identity|sk_c6=identity|sk_c1=sk_c3)).
#########
#cnf(i_0_63425, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c2=sk_c7|sk_c1=sk_c3)).
###
#cnf(i_0_63427, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c6=sk_c1|sk_c1=sk_c3)).
#
#cnf(i_0_63433, negated_conjecture, (inverse(sk_c7)=identity|sk_c2=sk_c1|sk_c1=sk_c3)).
############
#cnf(i_0_63440, negated_conjecture, (inverse(sk_c4)=identity|sk_c2=sk_c1|sk_c1=sk_c3)).
############
#cnf(i_0_63442, negated_conjecture, (inverse(sk_c3)=identity|sk_c2=sk_c1|sk_c1=sk_c3)).
##############################
#cnf(i_0_63666, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c5=identity|sk_c1=sk_c3)).
#####
#cnf(i_0_77187, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c8|sk_c5=identity)).
###
#cnf(i_0_77195, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c7|sk_c5=identity)).
#
#cnf(i_0_63680, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c5=identity|sk_c1=sk_c3)).
#####
#cnf(i_0_63762, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c6=identity|sk_c1=sk_c3)).
##
#cnf(i_0_77526, negated_conjecture, (sk_c4=sk_c3|sk_c6=identity|sk_c1=sk_c3)).
#############
#cnf(i_0_64322, negated_conjecture, (inverse(sk_c8)=identity|sk_c2=sk_c1|sk_c1=sk_c3)).
############
#cnf(i_0_66532, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c8=identity|sk_c6=sk_c8)).
####
#cnf(i_0_66754, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c8=identity|sk_c6=identity)).
#######
#cnf(i_0_66759, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c5=identity|sk_c8=identity)).
##
#cnf(i_0_66791, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c8=identity|sk_c6=sk_c8)).
######
#cnf(i_0_66813, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c8=identity|sk_c6=sk_c7)).
######
#cnf(i_0_67369, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c6=identity|sk_c8=identity)).
####
#cnf(i_0_67569, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c8=identity|sk_c6=sk_c7)).
####
#cnf(i_0_67706, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c8=identity|sk_c6=sk_c7)).
####
#cnf(i_0_67838, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c5=identity|sk_c8=sk_c3)).
####
#cnf(i_0_67976, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c5=identity|sk_c8=sk_c3)).
####
#cnf(i_0_68109, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c6=identity|sk_c8=sk_c3)).
####
#cnf(i_0_69461, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c5=identity|sk_c7=identity)).
######
#cnf(i_0_69604, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c5=identity|sk_c7=identity)).
#####
#cnf(i_0_69763, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c7=identity|sk_c5=identity)).
####
#cnf(i_0_69933, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c6=identity|sk_c7=identity)).
#####
#cnf(i_0_70142, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c5=identity|sk_c7=identity)).
######
#cnf(i_0_70252, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c5=identity|sk_c1=sk_c7)).
####
#cnf(i_0_70398, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c5=identity|sk_c1=sk_c7)).
####
#cnf(i_0_70539, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c6=identity|sk_c1=sk_c7)).
####
#cnf(i_0_70713, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c1=sk_c7|sk_c4=sk_c3)).
########
#cnf(i_0_72115, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c5=identity|sk_c8=identity)).
#########
#cnf(i_0_73611, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c1=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_80693, negated_conjecture, (sk_c1=sk_c3|sk_c5=identity|sk_c1=sk_c7)).
###
#cnf(i_0_73716, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c6=sk_c2|sk_c1=sk_c3)).
#
#cnf(i_0_80902, negated_conjecture, (sk_c1=sk_c3|sk_c5=identity|sk_c6=sk_c2)).
##
#cnf(i_0_73821, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c6=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_81084, negated_conjecture, (sk_c4=identity|sk_c6=sk_c7|sk_c1=sk_c3)).
###
#cnf(i_0_73943, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c1=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_81307, negated_conjecture, (sk_c4=identity|sk_c1=sk_c7|sk_c1=sk_c3)).
#####
#cnf(i_0_74065, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c6=sk_c2|sk_c1=sk_c3)).
##
#cnf(i_0_81605, negated_conjecture, (sk_c4=identity|sk_c6=sk_c2|sk_c1=sk_c3)).
##
#cnf(i_0_81680, negated_conjecture, (sk_c2=sk_c7|sk_c4=identity|sk_c1=sk_c3)).
####
#cnf(i_0_74187, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c5=sk_c7|sk_c1=sk_c3)).
##
#cnf(i_0_81917, negated_conjecture, (sk_c1=sk_c3|sk_c2=sk_c1|sk_c5=sk_c7)).
#
#cnf(i_0_81978, negated_conjecture, (sk_c1=identity|sk_c5=sk_c7|sk_c1=sk_c3)).
####
#cnf(i_0_74294, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c6=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_82143, negated_conjecture, (sk_c1=sk_c3|sk_c2=sk_c1|sk_c6=sk_c7)).
##
#cnf(i_0_82193, negated_conjecture, (sk_c4=sk_c1|sk_c6=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_82204, negated_conjecture, (sk_c1=identity|sk_c6=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_74401, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c5=sk_c4|sk_c1=sk_c3)).
##
#cnf(i_0_82498, negated_conjecture, (sk_c1=sk_c3|sk_c2=sk_c1|sk_c5=sk_c4)).
#
#cnf(i_0_82559, negated_conjecture, (sk_c1=identity|sk_c5=sk_c4|sk_c1=sk_c3)).
###
#cnf(i_0_82594, negated_conjecture, (sk_c4=sk_c7|sk_c1=sk_c3|sk_c2=sk_c1)).
#
#cnf(i_0_82621, negated_conjecture, (sk_c4=identity|sk_c1=identity|sk_c1=sk_c3)).
#
#cnf(i_0_82630, negated_conjecture, (sk_c4=sk_c7|sk_c1=identity|sk_c1=sk_c3)).
##
#cnf(i_0_82679, negated_conjecture, (sk_c7=sk_c3|sk_c2=sk_c1|sk_c1=sk_c3)).
############
#cnf(i_0_74508, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c5=sk_c7|sk_c1=sk_c3)).
##
#cnf(i_0_83369, negated_conjecture, (sk_c2=identity|sk_c5=sk_c7|sk_c1=sk_c3)).
###
#cnf(i_0_74644, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c6=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_83551, negated_conjecture, (sk_c4=sk_c2|sk_c6=sk_c7|sk_c1=sk_c3)).
###
#cnf(i_0_83565, negated_conjecture, (sk_c2=identity|sk_c6=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_74780, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c5=sk_c4|sk_c1=sk_c3)).
###
#cnf(i_0_83856, negated_conjecture, (sk_c2=identity|sk_c5=sk_c4|sk_c1=sk_c3)).
###
#cnf(i_0_83893, negated_conjecture, (sk_c4=sk_c7|sk_c2=identity|sk_c1=sk_c3)).
#
#cnf(i_0_84023, negated_conjecture, (sk_c2=identity|sk_c1=sk_c3|sk_c2!=sk_c7)).
####
#cnf(i_0_74916, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c1=identity|sk_c1=sk_c3)).
###
#cnf(i_0_75021, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c2=identity|sk_c1=sk_c3)).
##
#cnf(i_0_75126, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c4=identity|sk_c1=sk_c3)).
########
#cnf(i_0_75517, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c4=identity|sk_c1=sk_c3)).
##
#cnf(i_0_75653, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c2=identity|sk_c1=sk_c3)).
##
#cnf(i_0_75806, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c2=identity|sk_c1=sk_c3)).
##
#cnf(i_0_84980, negated_conjecture, (sk_c4=sk_c3|sk_c2=identity|sk_c1=sk_c3)).
#
#cnf(i_0_85112, negated_conjecture, (sk_c2=identity|sk_c1=sk_c3|sk_c2!=sk_c3)).
###
#cnf(i_0_85121, negated_conjecture, (sk_c3=identity|sk_c2=identity|sk_c1=sk_c3)).
#
#cnf(i_0_85127, negated_conjecture, (sk_c7=sk_c3|sk_c2=identity|sk_c1=sk_c3)).
#########
#cnf(i_0_85303, negated_conjecture, (sk_c1=identity|sk_c3=identity|sk_c1=sk_c3)).
#######
#cnf(i_0_85480, negated_conjecture, (sk_c1=identity|sk_c7=sk_c3|sk_c1=sk_c3)).
#
#cnf(i_0_85597, negated_conjecture, (sk_c1=identity|sk_c7=sk_c3|sk_c3!=identity)).
##
#cnf(i_0_85801, negated_conjecture, (sk_c8=sk_c7|sk_c7=sk_c3|sk_c3!=identity)).
##
#cnf(i_0_75874, negated_conjecture, (inverse(sk_c8)=identity|sk_c5=identity|sk_c1=sk_c3)).
##############
#cnf(i_0_76091, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c5=sk_c8|sk_c1=sk_c3)).
#
#cnf(i_0_86455, negated_conjecture, (sk_c1=sk_c3|sk_c2=sk_c1|sk_c5=sk_c8)).
#
#cnf(i_0_86520, negated_conjecture, (sk_c1=identity|sk_c5=sk_c8|sk_c1=sk_c3)).
####
#cnf(i_0_86558, negated_conjecture, (sk_c4=sk_c8|sk_c1=sk_c3|sk_c2=sk_c1)).
####
#cnf(i_0_86598, negated_conjecture, (sk_c4=sk_c8|sk_c1=identity|sk_c1=sk_c3)).
##
#cnf(i_0_86649, negated_conjecture, (sk_c8=sk_c3|sk_c2=sk_c1|sk_c1=sk_c3)).
##########
#cnf(i_0_76199, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c5=sk_c8|sk_c1=sk_c3)).
###
#cnf(i_0_87223, negated_conjecture, (sk_c2=identity|sk_c5=sk_c8|sk_c1=sk_c3)).
####
#cnf(i_0_87264, negated_conjecture, (sk_c4=sk_c8|sk_c2=identity|sk_c1=sk_c3)).
#
#cnf(i_0_87398, negated_conjecture, (sk_c2=identity|sk_c1=sk_c3|sk_c2!=sk_c8)).
######
#cnf(i_0_87414, negated_conjecture, (sk_c8=sk_c3|sk_c2=identity|sk_c1=sk_c3)).
####
#cnf(i_0_87604, negated_conjecture, (sk_c1=identity|sk_c8=sk_c3|sk_c1=sk_c3)).
####
#cnf(i_0_76402, negated_conjecture, (inverse(sk_c8)=identity|sk_c6=identity|sk_c1=sk_c3)).
########
#cnf(i_0_76474, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c2=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_88300, negated_conjecture, (sk_c1=sk_c3|sk_c5=identity|sk_c2=sk_c7)).
#
#cnf(i_0_76579, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c6=sk_c1|sk_c1=sk_c3)).
##
#cnf(i_0_88487, negated_conjecture, (sk_c1=sk_c3|sk_c5=identity|sk_c6=sk_c1)).
####################
#cnf(i_0_77115, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c5=identity|sk_c1=sk_c3)).
#####
#cnf(i_0_77312, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c5=identity|sk_c1=sk_c3)).
#####
#cnf(i_0_77462, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c6=identity|sk_c1=sk_c3)).
##################
#cnf(i_0_77991, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c8=identity|sk_c6=identity)).
#####
#cnf(i_0_78114, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c5=identity|sk_c8=identity)).
####
#cnf(i_0_78237, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c8=identity|sk_c6=sk_c8)).
####
#cnf(i_0_78360, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c8=identity|sk_c6=sk_c7)).
#################################
#cnf(i_0_80358, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c1=sk_c7|sk_c4=sk_c3)).
######
#cnf(i_0_80468, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c5=identity|sk_c8=identity)).
#########
#cnf(i_0_81085, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c6=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_90057, negated_conjecture, (sk_c4=sk_c3|sk_c6=sk_c7|sk_c1=sk_c3)).
######
#cnf(i_0_81305, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c1=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_90369, negated_conjecture, (sk_c4=sk_c3|sk_c1=sk_c7|sk_c1=sk_c3)).
####
#cnf(i_0_90518, negated_conjecture, (sk_c3=identity|sk_c1=sk_c7|sk_c1=sk_c3)).
#####
#cnf(i_0_81606, negated_conjecture, (inverse(sk_c7)=sk_c4|sk_c6=sk_c2|sk_c1=sk_c3)).
##
#cnf(i_0_90803, negated_conjecture, (sk_c4=sk_c3|sk_c6=sk_c2|sk_c1=sk_c3)).
##
#cnf(i_0_90891, negated_conjecture, (sk_c2=sk_c7|sk_c4=sk_c3|sk_c1=sk_c3)).
###
#cnf(i_0_91034, negated_conjecture, (sk_c3=identity|sk_c2=sk_c7|sk_c1=sk_c3)).
#######
#cnf(i_0_81979, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c5=sk_c7|sk_c1=sk_c3)).
##
#cnf(i_0_91282, negated_conjecture, (sk_c1=sk_c3|sk_c6=identity|sk_c5=sk_c7)).
###
#cnf(i_0_91290, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c8|sk_c5=sk_c7)).
###
#cnf(i_0_91298, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c7|sk_c5=sk_c7)).
#
#cnf(i_0_91372, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c2|sk_c5=sk_c7)).
#
#cnf(i_0_91373, negated_conjecture, (sk_c1=sk_c3|sk_c1=sk_c7|sk_c5=sk_c7)).
####
#cnf(i_0_91379, negated_conjecture, (sk_c1=sk_c3|sk_c2=sk_c7|sk_c5=sk_c7)).
##
#cnf(i_0_91380, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c1|sk_c5=sk_c7)).
############
#cnf(i_0_82205, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c6=sk_c7|sk_c1=sk_c3)).
#
#cnf(i_0_91755, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c7)).
############
#cnf(i_0_82560, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c5=sk_c4|sk_c1=sk_c3)).
###
#cnf(i_0_91975, negated_conjecture, (sk_c1=sk_c3|sk_c6=identity|sk_c5=sk_c4)).
##
#cnf(i_0_91983, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c8|sk_c5=sk_c4)).
####
#cnf(i_0_92065, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c2|sk_c5=sk_c4)).
##
#cnf(i_0_92066, negated_conjecture, (sk_c1=sk_c3|sk_c1=sk_c7|sk_c5=sk_c4)).
####
#cnf(i_0_92072, negated_conjecture, (sk_c1=sk_c3|sk_c2=sk_c7|sk_c5=sk_c4)).
#
#cnf(i_0_92073, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c1|sk_c5=sk_c4)).
###
#cnf(i_0_92117, negated_conjecture, (sk_c4=sk_c7|sk_c1=sk_c3|sk_c6=identity)).
#
#cnf(i_0_92152, negated_conjecture, (sk_c4=identity|sk_c1=sk_c3|sk_c6=sk_c8)).
#
#cnf(i_0_92164, negated_conjecture, (sk_c4=sk_c7|sk_c1=sk_c3|sk_c6=sk_c8)).
###
#cnf(i_0_92211, negated_conjecture, (sk_c4=sk_c7|sk_c1=sk_c3|sk_c6=sk_c2)).
###
#cnf(i_0_92258, negated_conjecture, (sk_c4=sk_c7|sk_c1=sk_c3|sk_c1=sk_c7)).
##
#cnf(i_0_92305, negated_conjecture, (sk_c4=sk_c7|sk_c1=sk_c3|sk_c2=sk_c7)).
##
#cnf(i_0_92346, negated_conjecture, (sk_c4=identity|sk_c1=sk_c3|sk_c6=sk_c1)).
#
#cnf(i_0_92352, negated_conjecture, (sk_c4=sk_c7|sk_c1=sk_c3|sk_c6=sk_c1)).
#####
#cnf(i_0_92489, negated_conjecture, (sk_c2=sk_c8|sk_c4=identity|sk_c1=sk_c3)).
##########
#cnf(i_0_92646, negated_conjecture, (sk_c2=sk_c8|sk_c4=sk_c7|sk_c1=sk_c3)).
####
#cnf(i_0_92794, negated_conjecture, (sk_c7=sk_c3|sk_c1=sk_c7|sk_c1=sk_c3)).
####
#cnf(i_0_92944, negated_conjecture, (sk_c7=sk_c3|sk_c2=sk_c7|sk_c1=sk_c3)).
#####
#cnf(i_0_93023, negated_conjecture, (sk_c8=sk_c1|sk_c4=identity|sk_c1=sk_c3)).
####
#cnf(i_0_93105, negated_conjecture, (sk_c8=sk_c1|sk_c4=sk_c7|sk_c1=sk_c3)).
###################
#cnf(i_0_82675, negated_conjecture, (inverse(sk_c7)=sk_c7|sk_c2=sk_c1|sk_c1=sk_c3)).
######
#cnf(i_0_82693, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c2=sk_c1|sk_c1=sk_c3)).
##########
#cnf(i_0_83370, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c5=sk_c7|sk_c1=sk_c3)).
####
#cnf(i_0_83857, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c5=sk_c4|sk_c1=sk_c3)).
######
#cnf(i_0_84525, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c4=identity|sk_c1=sk_c3)).
###################
#cnf(i_0_84529, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c4=identity|sk_c1=sk_c3)).
#####
#cnf(i_0_84816, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c2=identity|sk_c1=sk_c3)).
####
#cnf(i_0_84820, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c2=identity|sk_c1=sk_c3)).
######
#cnf(i_0_85111, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c2=identity|sk_c1=sk_c3)).
#########
#cnf(i_0_86356, negated_conjecture, (inverse(sk_c7)=identity|sk_c5=identity|sk_c1=sk_c3)).
#########
#cnf(i_0_86521, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c5=sk_c8|sk_c1=sk_c3)).
###
#cnf(i_0_95739, negated_conjecture, (sk_c1=sk_c3|sk_c6=identity|sk_c5=sk_c8)).
###
#cnf(i_0_95747, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c8|sk_c5=sk_c8)).
####
#cnf(i_0_95829, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c2|sk_c5=sk_c8)).
#
#cnf(i_0_95830, negated_conjecture, (sk_c1=sk_c3|sk_c1=sk_c7|sk_c5=sk_c8)).
####
#cnf(i_0_95836, negated_conjecture, (sk_c1=sk_c3|sk_c2=sk_c7|sk_c5=sk_c8)).
#
#cnf(i_0_95837, negated_conjecture, (sk_c1=sk_c3|sk_c6=sk_c1|sk_c5=sk_c8)).
####
#cnf(i_0_95893, negated_conjecture, (sk_c4=sk_c8|sk_c1=sk_c3|sk_c6=identity)).
####
#cnf(i_0_95946, negated_conjecture, (sk_c4=sk_c8|sk_c1=sk_c3|sk_c6=sk_c8)).
####
#cnf(i_0_95999, negated_conjecture, (sk_c4=sk_c8|sk_c1=sk_c3|sk_c6=sk_c2)).
####
#cnf(i_0_96052, negated_conjecture, (sk_c4=sk_c8|sk_c1=sk_c3|sk_c1=sk_c7)).
####
#cnf(i_0_96105, negated_conjecture, (sk_c4=sk_c8|sk_c1=sk_c3|sk_c2=sk_c7)).
####
#cnf(i_0_96158, negated_conjecture, (sk_c4=sk_c8|sk_c1=sk_c3|sk_c6=sk_c1)).
##########
#cnf(i_0_96387, negated_conjecture, (sk_c2=sk_c8|sk_c4=sk_c8|sk_c1=sk_c3)).
####
#cnf(i_0_96539, negated_conjecture, (sk_c8=sk_c3|sk_c1=sk_c7|sk_c1=sk_c3)).
#####
#cnf(i_0_96699, negated_conjecture, (sk_c8=sk_c3|sk_c2=sk_c7|sk_c1=sk_c3)).
#####
#cnf(i_0_96788, negated_conjecture, (sk_c8=sk_c1|sk_c4=sk_c8|sk_c1=sk_c3)).
######################
#cnf(i_0_86645, negated_conjecture, (inverse(sk_c8)=sk_c7|sk_c2=sk_c1|sk_c1=sk_c3)).
######
#cnf(i_0_86663, negated_conjecture, (inverse(sk_c8)=sk_c8|sk_c2=sk_c1|sk_c1=sk_c3)).
############
#cnf(i_0_87224, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c5=sk_c8|sk_c1=sk_c3)).
###
#cnf(i_0_87397, negated_conjecture, (inverse(sk_c8)=sk_c8|sk_c2=identity|sk_c1=sk_c3)).
########
#cnf(i_0_88146, negated_conjecture, (inverse(sk_c7)=identity|sk_c6=identity|sk_c1=sk_c3)).
#####
#cnf(i_0_98444, negated_conjecture, (sk_c3=identity|sk_c6=identity|sk_c1=sk_c3)).
##########################
#cnf(i_0_90211, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c1=sk_c7|sk_c1=sk_c3)).
#####
#cnf(i_0_90504, negated_conjecture, (inverse(sk_c3)=sk_c8|sk_c1=sk_c7|sk_c1=sk_c3)).
####
#cnf(i_0_90646, negated_conjecture, (inverse(sk_c4)=sk_c7|sk_c6=sk_c2|sk_c1=sk_c3)).
####
#cnf(i_0_91218, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c5=sk_c7|sk_c1=sk_c3)).
#######
#cnf(i_0_91911, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c5=sk_c4|sk_c1=sk_c3)).
####
#cnf(i_0_92774, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c1=sk_c7|sk_c1=sk_c3)).
######
#cnf(i_0_94369, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c5=sk_c7|sk_c1=sk_c3)).
####
#cnf(i_0_94535, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c5=sk_c4|sk_c1=sk_c3)).
####
#cnf(i_0_94701, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c4=identity|sk_c1=sk_c3)).
######
#cnf(i_0_95023, negated_conjecture, (inverse(sk_c8)=sk_c7|sk_c2=identity|sk_c1=sk_c3)).
######
#cnf(i_0_95346, negated_conjecture, (inverse(sk_c7)=sk_c7|sk_c2=identity|sk_c1=sk_c3)).
##########
#cnf(i_0_95362, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c2=identity|sk_c1=sk_c3)).
##############
#cnf(i_0_95675, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c5=sk_c8|sk_c1=sk_c3)).
####
#cnf(i_0_96518, negated_conjecture, (inverse(sk_c8)=sk_c8|sk_c1=sk_c7|sk_c1=sk_c3)).
######
#cnf(i_0_97300, negated_conjecture, (inverse(sk_c3)=sk_c4|sk_c1=sk_c7|sk_c1=sk_c3)).
############
#cnf(i_0_98022, negated_conjecture, (inverse(sk_c2)=sk_c7|sk_c5=sk_c8|sk_c1=sk_c3)).
######
#cnf(i_0_98445, negated_conjecture, (inverse(sk_c3)=identity|sk_c6=identity|sk_c1=sk_c3)).
###############
#cnf(i_0_98783, negated_conjecture, (inverse(sk_c7)=sk_c7|sk_c1=sk_c7|sk_c1=sk_c3)).
###############
#cnf(i_0_98789, negated_conjecture, (inverse(sk_c8)=sk_c7|sk_c1=sk_c7|sk_c1=sk_c3)).
#############
#cnf(i_0_98796, negated_conjecture, (inverse(sk_c8)=sk_c3|sk_c1=sk_c7|sk_c1=sk_c3)).
##############
#cnf(i_0_56208, negated_conjecture, (multiply(inverse(X1),multiply(inverse(X3),multiply(inverse(X2),inverse(X4))))!=sk_c8|multiply(X4,multiply(X2,multiply(X3,X1)))!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
#########################
#cnf(i_0_100807, negated_conjecture, (inverse(sk_c4)=sk_c3|sk_c1=sk_c7|sk_c1=sk_c3)).
######################
#cnf(i_0_101795, negated_conjecture, (inverse(sk_c3)=sk_c3|sk_c1=sk_c7|sk_c1=sk_c3)).
##############
#cnf(i_0_21116, negated_conjecture, (multiply(sk_c8,multiply(sk_c4,X1))=X1|sk_c8=sk_c7)).
##################
#cnf(i_0_21119, negated_conjecture, (multiply(sk_c4,multiply(sk_c8,X1))=X1|sk_c8=sk_c7)).
###############################
#cnf(i_0_41460, negated_conjecture, (multiply(sk_c8,multiply(sk_c7,X1))=X1|sk_c8=sk_c7)).
############################
#cnf(i_0_41463, negated_conjecture, (multiply(sk_c7,multiply(sk_c8,X1))=X1|sk_c8=sk_c7)).
############################################
#cnf(i_0_44233, negated_conjecture, (multiply(sk_c3,multiply(sk_c2,X1))=X1|sk_c8=sk_c7)).
##################
#cnf(i_0_57788, negated_conjecture, (multiply(inverse(X1),multiply(inverse(X3),multiply(inverse(X2),inverse(X4))))!=sk_c6|multiply(X4,multiply(X2,multiply(X3,X1)))!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
###
#cnf(i_0_44236, negated_conjecture, (multiply(sk_c3,multiply(sk_c1,X1))=X1|sk_c8=sk_c7)).
####################
#cnf(i_0_44264, negated_conjecture, (multiply(sk_c1,multiply(sk_c3,X1))=X1|sk_c8=sk_c7)).
##########################
#cnf(i_0_44265, negated_conjecture, (multiply(sk_c2,multiply(sk_c3,X1))=X1|sk_c8=sk_c7)).
######################################
#cnf(i_0_44711, negated_conjecture, (multiply(sk_c3,multiply(sk_c4,X1))=X1|sk_c8=sk_c7)).
###########################
#cnf(i_0_44744, negated_conjecture, (multiply(sk_c4,multiply(sk_c3,X1))=X1|sk_c8=sk_c7)).
######################################################################################
#cnf(i_0_59479, negated_conjecture, (multiply(inverse(X1),multiply(inverse(X3),multiply(inverse(X2),X4)))!=sk_c6|multiply(inverse(X4),multiply(X2,multiply(X3,X1)))!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
#####
#cnf(i_0_63319, negated_conjecture, (multiply(sk_c1,X1)=X1|sk_c5=identity|sk_c1=sk_c3)).
##########
#cnf(i_0_63336, negated_conjecture, (multiply(sk_c2,X1)=X1|sk_c5=identity|sk_c1=sk_c3)).
###################
#cnf(i_0_63417, negated_conjecture, (multiply(sk_c4,X1)=X1|sk_c6=identity|sk_c1=sk_c3)).
#########
#cnf(i_0_63428, negated_conjecture, (multiply(sk_c3,X1)=X1|sk_c2=sk_c1|sk_c1=sk_c3)).
############
#cnf(i_0_63435, negated_conjecture, (multiply(sk_c7,X1)=X1|sk_c2=sk_c1|sk_c1=sk_c3)).
##########
#cnf(i_0_63438, negated_conjecture, (multiply(sk_c4,X1)=X1|sk_c2=sk_c1|sk_c1=sk_c3)).
####################
#cnf(i_0_64321, negated_conjecture, (multiply(sk_c8,X1)=X1|sk_c2=sk_c1|sk_c1=sk_c3)).
##############
#cnf(i_0_75877, negated_conjecture, (multiply(sk_c8,X1)=X1|sk_c5=identity|sk_c1=sk_c3)).
###############
#cnf(i_0_76403, negated_conjecture, (multiply(sk_c8,X1)=X1|sk_c6=identity|sk_c1=sk_c3)).
#####################################
#cnf(i_0_95485, negated_conjecture, (multiply(sk_c7,X1)=X1|sk_c5=identity|sk_c1=sk_c3)).
#######
#cnf(i_0_98378, negated_conjecture, (multiply(sk_c7,X1)=X1|sk_c6=identity|sk_c1=sk_c3)).
########
#cnf(i_0_98439, negated_conjecture, (multiply(sk_c3,X1)=X1|sk_c6=identity|sk_c1=sk_c3)).
##############################
#cnf(i_0_59728, negated_conjecture, (multiply(inverse(X4),multiply(inverse(X3),multiply(X2,X1)))!=sk_c8|multiply(inverse(X1),multiply(inverse(X2),multiply(X3,X4)))!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
##########################################
#cnf(i_0_42098, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c8,X1)|sk_c8=sk_c7)).
#######################
#cnf(i_0_54510, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c1,X1)|sk_c8=sk_c7)).
#####################
#cnf(i_0_54516, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c2,X1)|sk_c8=sk_c7)).
#######################
#cnf(i_0_104324, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c4,X1)|sk_c8=sk_c7)).
########################
#cnf(i_0_104410, negated_conjecture, (multiply(inverse(sk_c4),X1)=multiply(sk_c8,X1)|sk_c8=sk_c7)).
#####################
#cnf(i_0_104606, negated_conjecture, (multiply(inverse(sk_c8),X1)=multiply(sk_c7,X1)|sk_c8=sk_c7)).
###############################
#cnf(i_0_104860, negated_conjecture, (multiply(inverse(sk_c2),X1)=multiply(sk_c3,X1)|sk_c8=sk_c7)).
###################
#cnf(i_0_61467, negated_conjecture, (multiply(inverse(X1),multiply(inverse(X3),multiply(inverse(X2),X4)))!=sk_c8|multiply(inverse(X4),multiply(X2,multiply(X3,X1)))!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
###
#cnf(i_0_107166, negated_conjecture, (multiply(inverse(sk_c1),X1)=multiply(sk_c3,X1)|sk_c8=sk_c7)).
###########################
#cnf(i_0_107346, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c4,X1)|sk_c8=sk_c7)).
########################
#cnf(i_0_107421, negated_conjecture, (multiply(inverse(sk_c4),X1)=multiply(sk_c3,X1)|sk_c8=sk_c7)).
###################################################
#cnf(i_0_28532, negated_conjecture, (inverse(sk_c8)=inverse(sk_c3)|sk_c2=sk_c1|sk_c7!=sk_c3)).
########
#cnf(i_0_30953, negated_conjecture, (inverse(sk_c7)=inverse(sk_c3)|sk_c2=sk_c1|sk_c7!=identity)).
##
#cnf(i_0_118587, negated_conjecture, (inverse(sk_c3)=sk_c3|sk_c2=sk_c1|sk_c7!=identity)).
######
#cnf(i_0_118614, negated_conjecture, (inverse(sk_c3)=sk_c4|sk_c2=sk_c1|sk_c7!=identity)).
############
#cnf(i_0_118893, negated_conjecture, (inverse(sk_c4)=sk_c3|sk_c2=sk_c1|sk_c7!=identity)).
#############
#cnf(i_0_32006, negated_conjecture, (inverse(sk_c8)=inverse(sk_c3)|sk_c2=sk_c1|sk_c7!=identity)).
#################
#cnf(i_0_26637, negated_conjecture, (sk_c2=sk_c7|sk_c4!=sk_c2|sk_c7!=sk_c3|sk_c7!=identity)).
##
#cnf(i_0_43724, negated_conjecture, (sk_c1=identity|sk_c8!=identity|sk_c8!=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_119506, negated_conjecture, (sk_c1=identity|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_119520, negated_conjecture, (sk_c1=identity|sk_c7!=identity)).
#########################################
#cnf(i_0_65055, negated_conjecture, (multiply(X4,multiply(X3,multiply(inverse(X2),inverse(X1))))!=sk_c8|multiply(X1,multiply(X2,multiply(inverse(X3),inverse(X4))))!=sk_c6|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
#############################################################
#cnf(i_0_43831, negated_conjecture, (sk_c4=identity|sk_c8!=identity|sk_c6!=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_121872, negated_conjecture, (sk_c4=identity|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_122008, negated_conjecture, (sk_c2!=identity|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_122081, negated_conjecture, (sk_c8=sk_c3|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_122100, negated_conjecture, (sk_c8=sk_c3|sk_c7!=identity)).
##
#cnf(i_0_122355, negated_conjecture, (sk_c3=identity|sk_c7!=identity)).
#
#cnf(i_0_122442, negated_conjecture, (sk_c7=sk_c3|sk_c7!=identity)).
############
#cnf(i_0_122183, negated_conjecture, (inverse(sk_c3)=sk_c2|sk_c4!=sk_c2|sk_c7!=identity)).
#
#cnf(i_0_122616, negated_conjecture, (sk_c2=sk_c7|sk_c7!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_122677, negated_conjecture, (sk_c2=sk_c7|sk_c2!=sk_c3|sk_c7!=identity)).
####
#cnf(i_0_122378, negated_conjecture, (inverse(sk_c3)=identity|sk_c4!=sk_c2|sk_c7!=identity)).
#########
#cnf(i_0_122496, negated_conjecture, (inverse(sk_c2)=sk_c3|sk_c4!=sk_c2|sk_c7!=identity)).
#######
#cnf(i_0_122379, negated_conjecture, (multiply(sk_c3,X1)=X1|sk_c4!=sk_c2|sk_c7!=identity)).
########
#cnf(i_0_122218, negated_conjecture, (inverse(sk_c3)=sk_c1|sk_c5=identity|sk_c7!=identity)).
#
#cnf(i_0_123285, negated_conjecture, (sk_c1=sk_c7|sk_c5=identity|sk_c7!=identity)).
##
#cnf(i_0_122232, negated_conjecture, (inverse(sk_c3)=sk_c2|sk_c5=identity|sk_c7!=identity)).
#
#cnf(i_0_123499, negated_conjecture, (sk_c2=sk_c7|sk_c5=identity|sk_c7!=identity)).
#####
#cnf(i_0_122314, negated_conjecture, (inverse(sk_c3)=sk_c4|sk_c6=identity|sk_c7!=identity)).
##
#cnf(i_0_123678, negated_conjecture, (sk_c4=sk_c7|sk_c6=identity|sk_c7!=identity)).
######
#cnf(i_0_122359, negated_conjecture, (inverse(sk_c3)=sk_c4|sk_c2=identity|sk_c7!=identity)).
####
#cnf(i_0_122363, negated_conjecture, (inverse(sk_c3)=sk_c2|sk_c4=identity|sk_c7!=identity)).
####
#cnf(i_0_122381, negated_conjecture, (inverse(sk_c3)=sk_c4|sk_c6=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_124269, negated_conjecture, (sk_c4=sk_c7|sk_c6=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_122382, negated_conjecture, (inverse(sk_c3)=sk_c4|sk_c1=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_124530, negated_conjecture, (sk_c4=sk_c7|sk_c1=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_124547, negated_conjecture, (sk_c4=sk_c8|sk_c1=sk_c7|sk_c7!=identity)).
#########
#cnf(i_0_122383, negated_conjecture, (inverse(sk_c3)=sk_c4|sk_c6=sk_c2|sk_c7!=identity)).
#
#cnf(i_0_125038, negated_conjecture, (sk_c4=sk_c7|sk_c6=sk_c2|sk_c7!=identity)).
###
#cnf(i_0_125167, negated_conjecture, (sk_c2=sk_c7|sk_c4=sk_c7|sk_c7!=identity)).
####
#cnf(i_0_122384, negated_conjecture, (inverse(sk_c3)=sk_c1|sk_c5=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_125459, negated_conjecture, (sk_c1=sk_c7|sk_c5=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_122385, negated_conjecture, (inverse(sk_c3)=sk_c1|sk_c6=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_125670, negated_conjecture, (sk_c1=sk_c7|sk_c6=sk_c7|sk_c7!=identity)).
####
#cnf(i_0_122386, negated_conjecture, (inverse(sk_c3)=sk_c1|sk_c5=sk_c4|sk_c7!=identity)).
#
#cnf(i_0_125938, negated_conjecture, (sk_c1=sk_c7|sk_c5=sk_c4|sk_c7!=identity)).
####
#cnf(i_0_122387, negated_conjecture, (inverse(sk_c3)=sk_c2|sk_c5=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_126166, negated_conjecture, (sk_c2=sk_c7|sk_c5=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_122388, negated_conjecture, (inverse(sk_c3)=sk_c2|sk_c6=sk_c7|sk_c7!=identity)).
##
#cnf(i_0_126377, negated_conjecture, (sk_c2=sk_c7|sk_c6=sk_c7|sk_c7!=identity)).
#########
#cnf(i_0_122389, negated_conjecture, (inverse(sk_c3)=sk_c2|sk_c5=sk_c4|sk_c7!=identity)).
#
#cnf(i_0_126646, negated_conjecture, (sk_c2=sk_c7|sk_c5=sk_c4|sk_c7!=identity)).
######
#cnf(i_0_122390, negated_conjecture, (inverse(sk_c3)=sk_c1|sk_c4=identity|sk_c7!=identity)).
####
#cnf(i_0_122395, negated_conjecture, (inverse(sk_c3)=identity|sk_c5=identity|sk_c7!=identity)).
##########
#cnf(i_0_122398, negated_conjecture, (inverse(sk_c3)=sk_c1|sk_c5=sk_c8|sk_c7!=identity)).
##
#cnf(i_0_127233, negated_conjecture, (sk_c1=sk_c7|sk_c5=sk_c8|sk_c7!=identity)).
#####
#cnf(i_0_122399, negated_conjecture, (inverse(sk_c3)=sk_c2|sk_c5=sk_c8|sk_c7!=identity)).
#
#cnf(i_0_127468, negated_conjecture, (sk_c2=sk_c7|sk_c5=sk_c8|sk_c7!=identity)).
######
#cnf(i_0_127577, negated_conjecture, (sk_c4=sk_c8|sk_c2=sk_c7|sk_c7!=identity)).
#
#cnf(i_0_127748, negated_conjecture, (sk_c2=sk_c7|sk_c2!=sk_c8|sk_c7!=identity)).
###############
#cnf(i_0_122454, negated_conjecture, (inverse(sk_c3)=sk_c3|sk_c1=sk_c7|sk_c7!=identity)).
#####
#cnf(i_0_122455, negated_conjecture, (inverse(sk_c3)=sk_c3|sk_c2=identity|sk_c7!=identity)).
#####
#cnf(i_0_123166, negated_conjecture, (inverse(sk_c1)=sk_c3|sk_c5=identity|sk_c7!=identity)).
#######
#cnf(i_0_123379, negated_conjecture, (inverse(sk_c2)=sk_c3|sk_c5=identity|sk_c7!=identity)).
#####
#cnf(i_0_123591, negated_conjecture, (inverse(sk_c4)=sk_c3|sk_c6=identity|sk_c7!=identity)).
######
#cnf(i_0_123828, negated_conjecture, (inverse(sk_c4)=sk_c3|sk_c2=identity|sk_c7!=identity)).
#####
#cnf(i_0_123996, negated_conjecture, (inverse(sk_c2)=sk_c3|sk_c4=identity|sk_c7!=identity)).
####
#cnf(i_0_124165, negated_conjecture, (inverse(sk_c4)=sk_c3|sk_c6=sk_c7|sk_c7!=identity)).
####
#cnf(i_0_124411, negated_conjecture, (inverse(sk_c4)=sk_c3|sk_c1=sk_c7|sk_c7!=identity)).
#####
#cnf(i_0_124661, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c1=sk_c7|sk_c7!=identity)).
####
#cnf(i_0_124693, negated_conjecture, (inverse(sk_c7)=sk_c7|sk_c1=sk_c7|sk_c7!=identity)).
################
#cnf(i_0_124829, negated_conjecture, (inverse(sk_c8)=sk_c8|sk_c1=sk_c7|sk_c7!=identity)).
#######
#cnf(i_0_124863, negated_conjecture, (inverse(sk_c8)=sk_c7|sk_c1=sk_c7|sk_c7!=identity)).
###############
#cnf(i_0_124920, negated_conjecture, (inverse(sk_c4)=sk_c3|sk_c6=sk_c2|sk_c7!=identity)).
####
#cnf(i_0_125340, negated_conjecture, (inverse(sk_c1)=sk_c3|sk_c5=sk_c7|sk_c7!=identity)).
#####
#cnf(i_0_125566, negated_conjecture, (inverse(sk_c1)=sk_c3|sk_c6=sk_c7|sk_c7!=identity)).
#####
#cnf(i_0_125819, negated_conjecture, (inverse(sk_c1)=sk_c3|sk_c5=sk_c4|sk_c7!=identity)).
######
#cnf(i_0_126046, negated_conjecture, (inverse(sk_c2)=sk_c3|sk_c5=sk_c7|sk_c7!=identity)).
####
#cnf(i_0_126273, negated_conjecture, (inverse(sk_c2)=sk_c3|sk_c6=sk_c7|sk_c7!=identity)).
####
#cnf(i_0_126526, negated_conjecture, (inverse(sk_c2)=sk_c3|sk_c5=sk_c4|sk_c7!=identity)).
####
#cnf(i_0_126754, negated_conjecture, (inverse(sk_c1)=sk_c3|sk_c4=identity|sk_c7!=identity)).
#####
#cnf(i_0_127114, negated_conjecture, (inverse(sk_c1)=sk_c3|sk_c5=sk_c8|sk_c7!=identity)).
#####
#cnf(i_0_127348, negated_conjecture, (inverse(sk_c2)=sk_c3|sk_c5=sk_c8|sk_c7!=identity)).
#######################################################
#cnf(i_0_122396, negated_conjecture, (multiply(sk_c3,X1)=X1|sk_c5=identity|sk_c7!=identity)).
############
#cnf(i_0_43894, negated_conjecture, (sk_c2=sk_c1|sk_c2!=sk_c3|sk_c7!=identity|sk_c8!=identity)).
#####
#cnf(i_0_43906, negated_conjecture, (sk_c2=sk_c1|sk_c2!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
###
#cnf(i_0_43909, negated_conjecture, (sk_c2=sk_c1|sk_c2!=sk_c8|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_122425, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity)).
#######
#cnf(i_0_44044, negated_conjecture, (sk_c4=sk_c1|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_122430, negated_conjecture, (multiply(sk_c3,sk_c7)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c8!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_132307, negated_conjecture, (sk_c4=sk_c1|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_132573, negated_conjecture, (sk_c1=sk_c7|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_132494, negated_conjecture, (sk_c2!=sk_c1|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_132579, negated_conjecture, (sk_c1=sk_c7|sk_c7!=identity)).
#####
#cnf(i_0_132392, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_132739, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_132755, negated_conjecture, (sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|sk_c6!=sk_c1)).
##
#cnf(i_0_132887, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_133155, negated_conjecture, (sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity)).
###
#cnf(i_0_133378, negated_conjecture, (sk_c1=sk_c3|sk_c7!=identity|sk_c8!=identity)).
##
#cnf(i_0_133395, negated_conjecture, (sk_c1=sk_c3|sk_c7!=identity)).
########
#cnf(i_0_133363, negated_conjecture, (sk_c4=sk_c2|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_133364, negated_conjecture, (sk_c2=sk_c1|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_133365, negated_conjecture, (sk_c2=identity|sk_c7!=identity|sk_c8!=identity)).
#
#cnf(i_0_133638, negated_conjecture, (sk_c7!=identity|sk_c8!=identity)).
####
#cnf(i_0_134089, negated_conjecture, (sk_c7!=identity)).
#
#cnf(i_0_23629, negated_conjecture, (sk_c8=sk_c7)).
#
#cnf(i_0_23875, negated_conjecture, (sk_c1=sk_c3)).
#
#cnf(i_0_66727, negated_conjecture, (sk_c2=sk_c3)).
###
#cnf(i_0_609, negated_conjecture, (inverse(sk_c3)=sk_c7)).
#
#cnf(i_0_614, negated_conjecture, (inverse(sk_c7)=sk_c3)).
############
#cnf(i_0_6894, negated_conjecture, (sk_c6=identity|sk_c4!=sk_c3)).
###
#cnf(i_0_611, negated_conjecture, (multiply(sk_c7,sk_c3)=identity)).
#
#cnf(i_0_613, negated_conjecture, (multiply(sk_c3,sk_c7)=identity)).
##
#cnf(i_0_44, negated_conjecture, (multiply(sk_c7,sk_c7)=sk_c6)).
######
#cnf(i_0_134153, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7)).
##
#cnf(i_0_134173, negated_conjecture, (sk_c7=sk_c3|sk_c4!=sk_c3)).
#
#cnf(i_0_1240, negated_conjecture, (sk_c4=sk_c3|sk_c5=identity)).
###
#cnf(i_0_17366, negated_conjecture, (sk_c4=sk_c3)).
#
#cnf(i_0_6894, negated_conjecture, (sk_c6=identity)).
#
#cnf(i_0_134173, negated_conjecture, (sk_c7=sk_c3)).
#
#cnf(i_0_134089, negated_conjecture, (sk_c3!=identity)).
##
#cnf(i_0_23629, negated_conjecture, (sk_c8=sk_c3)).
#
#cnf(i_0_609, negated_conjecture, (inverse(sk_c3)=sk_c3)).
#######
#cnf(i_0_44, negated_conjecture, (multiply(sk_c3,sk_c3)=identity)).
########
#cnf(i_0_18020, negated_conjecture, (sk_c5=identity)).
######################################################################################################################################################
#cnf(i_0_82, negated_conjecture, (multiply(sk_c3,multiply(X1,sk_c3))!=sk_c3|multiply(X2,sk_c3)!=sk_c3|multiply(X3,sk_c3)!=identity|multiply(X4,sk_c3)!=sk_c3|inverse(X1)!=sk_c3|inverse(X2)!=sk_c3|inverse(X3)!=sk_c3|inverse(X4)!=sk_c3)).
####
#cnf(i_0_134185, negated_conjecture, (multiply(X1,sk_c3)!=sk_c3|multiply(X2,sk_c3)!=identity|multiply(X3,sk_c3)!=sk_c3|inverse(X1)!=sk_c3|inverse(X2)!=sk_c3|inverse(X3)!=sk_c3)).
####
#cnf(i_0_134194, negated_conjecture, (multiply(X1,sk_c3)!=sk_c3|multiply(X2,sk_c3)!=sk_c3|inverse(X1)!=sk_c3|inverse(X2)!=sk_c3)).
#################
#cnf(i_0_612, negated_conjecture, (multiply(sk_c3,multiply(sk_c3,X1))=X1)).
###########
#cnf(i_0_134204, negated_conjecture, (multiply(X2,multiply(X3,sk_c3))!=sk_c3|multiply(inverse(X3),inverse(X2))!=sk_c3|multiply(X1,sk_c3)!=sk_c3|inverse(X1)!=sk_c3)).
############################
#cnf(i_0_134218, negated_conjecture, (multiply(sk_c3,inverse(X1))!=sk_c3|multiply(X2,sk_c3)!=sk_c3|inverse(X2)!=sk_c3|X1!=sk_c3)).
####################
#cnf(i_0_134231, negated_conjecture, (multiply(sk_c3,X1)!=sk_c3|multiply(X2,sk_c3)!=sk_c3|inverse(X2)!=sk_c3|inverse(X1)!=sk_c3)).
############################
#cnf(i_0_134239, negated_conjecture, (multiply(inverse(X1),sk_c3)!=sk_c3|multiply(X2,sk_c3)!=sk_c3|inverse(X2)!=sk_c3|X1!=sk_c3)).
################################
#cnf(i_0_134237, negated_conjecture, (multiply(X2,multiply(X3,sk_c3))!=sk_c3|multiply(inverse(X3),inverse(X2))!=sk_c3|multiply(sk_c3,X1)!=sk_c3|inverse(X1)!=sk_c3)).
############################
#cnf(i_0_134252, negated_conjecture, (multiply(sk_c3,inverse(X1))!=sk_c3|multiply(sk_c3,X2)!=sk_c3|inverse(X2)!=sk_c3|X1!=sk_c3)).
####################
#cnf(i_0_134266, negated_conjecture, (multiply(sk_c3,X1)!=sk_c3|multiply(sk_c3,X2)!=sk_c3|inverse(X2)!=sk_c3|inverse(X1)!=sk_c3)).
################################
#cnf(i_0_134271, negated_conjecture, (multiply(inverse(X2),sk_c3)!=sk_c3|multiply(sk_c3,X1)!=sk_c3|inverse(X1)!=sk_c3|X2!=sk_c3)).
################################
#cnf(i_0_134263, negated_conjecture, (multiply(sk_c3,inverse(X1))!=sk_c3|multiply(inverse(X2),sk_c3)!=sk_c3|X2!=sk_c3|X1!=sk_c3)).
############################
#cnf(i_0_134280, negated_conjecture, (multiply(inverse(X1),sk_c3)!=sk_c3|multiply(inverse(X2),sk_c3)!=sk_c3|X2!=sk_c3|X1!=sk_c3)).
########################################
#cnf(i_0_134232, negated_conjecture, (multiply(sk_c3,multiply(inverse(X2),inverse(X1)))!=sk_c3|multiply(X3,sk_c3)!=sk_c3|multiply(X1,X2)!=sk_c3|inverse(X3)!=sk_c3)).
############################
#cnf(i_0_134306, negated_conjecture, (multiply(sk_c3,multiply(inverse(X1),X2))!=sk_c3|multiply(inverse(X2),X1)!=sk_c3|multiply(X3,sk_c3)!=sk_c3|inverse(X3)!=sk_c3)).
####################################
#cnf(i_0_134318, negated_conjecture, (multiply(sk_c3,multiply(X1,X2))!=sk_c3|multiply(inverse(X2),inverse(X1))!=sk_c3|multiply(X3,sk_c3)!=sk_c3|inverse(X3)!=sk_c3)).
############################################
#cnf(i_0_134310, negated_conjecture, (multiply(sk_c3,multiply(X1,inverse(X2)))!=sk_c3|multiply(X2,inverse(X1))!=sk_c3|multiply(X3,sk_c3)!=sk_c3|inverse(X3)!=sk_c3)).
################################
#cnf(i_0_134340, negated_conjecture, (multiply(sk_c3,multiply(X1,sk_c3))!=sk_c3|multiply(sk_c3,inverse(X1))!=sk_c3|multiply(X2,sk_c3)!=sk_c3|inverse(X2)!=sk_c3)).
########################
#cnf(i_0_134305, negated_conjecture, (multiply(sk_c3,multiply(inverse(X1),sk_c3))!=sk_c3|multiply(X2,sk_c3)!=sk_c3|multiply(sk_c3,X1)!=sk_c3|inverse(X2)!=sk_c3)).
########################
#cnf(i_0_134247, negated_conjecture, (multiply(inverse(X2),multiply(inverse(X1),sk_c3))!=sk_c3|multiply(X3,sk_c3)!=sk_c3|multiply(X1,X2)!=sk_c3|inverse(X3)!=sk_c3)).
################################
#cnf(i_0_134366, negated_conjecture, (multiply(inverse(X1),multiply(X2,sk_c3))!=sk_c3|multiply(inverse(X2),X1)!=sk_c3|multiply(X3,sk_c3)!=sk_c3|inverse(X3)!=sk_c3)).
####################################
#cnf(i_0_134382, negated_conjecture, (multiply(sk_c3,multiply(X1,sk_c3))!=sk_c3|multiply(inverse(X1),sk_c3)!=sk_c3|multiply(X2,sk_c3)!=sk_c3|inverse(X2)!=sk_c3)).
########################
#cnf(i_0_134371, negated_conjecture, (multiply(X1,multiply(inverse(X2),sk_c3))!=sk_c3|multiply(X2,inverse(X1))!=sk_c3|multiply(X3,sk_c3)!=sk_c3|inverse(X3)!=sk_c3)).
###########################
# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(associativity, axiom, (multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), file('./p/GRP293-1.tptp', associativity)).
cnf(left_inverse, axiom, (multiply(inverse(X1),X1)=identity), file('./p/GRP293-1.tptp', left_inverse)).
cnf(left_identity, axiom, (multiply(identity,X1)=X1), file('./p/GRP293-1.tptp', left_identity)).
cnf(prove_this_36, negated_conjecture, (inverse(sk_c2)=sk_c8|inverse(sk_c4)=sk_c8), file('./p/GRP293-1.tptp', prove_this_36)).
cnf(prove_this_35, negated_conjecture, (inverse(sk_c2)=sk_c8|multiply(sk_c4,sk_c8)=sk_c5), file('./p/GRP293-1.tptp', prove_this_35)).
cnf(prove_this_30, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|inverse(sk_c4)=sk_c8), file('./p/GRP293-1.tptp', prove_this_30)).
cnf(prove_this_34, negated_conjecture, (inverse(sk_c2)=sk_c8|multiply(sk_c8,sk_c5)=sk_c7), file('./p/GRP293-1.tptp', prove_this_34)).
cnf(prove_this_29, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|multiply(sk_c4,sk_c8)=sk_c5), file('./p/GRP293-1.tptp', prove_this_29)).
cnf(prove_this_37, negated_conjecture, (multiply(sk_c8,sk_c7)!=sk_c6|multiply(X1,sk_c8)!=sk_c7|inverse(X1)!=sk_c8|multiply(sk_c8,sk_c6)!=sk_c7|multiply(X2,sk_c8)!=sk_c6|inverse(X2)!=sk_c8|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X3,sk_c7)!=sk_c8|inverse(X3)!=sk_c7|multiply(sk_c8,X4)!=sk_c7|multiply(X5,sk_c8)!=X4|inverse(X5)!=sk_c8), file('./p/GRP293-1.tptp', prove_this_37)).
cnf(prove_this_28, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|multiply(sk_c8,sk_c5)=sk_c7), file('./p/GRP293-1.tptp', prove_this_28)).
cnf(prove_this_27, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|inverse(sk_c3)=sk_c7), file('./p/GRP293-1.tptp', prove_this_27)).
cnf(prove_this_33, negated_conjecture, (inverse(sk_c2)=sk_c8|inverse(sk_c3)=sk_c7), file('./p/GRP293-1.tptp', prove_this_33)).
cnf(prove_this_19, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c7|multiply(sk_c7,sk_c8)=sk_c6), file('./p/GRP293-1.tptp', prove_this_19)).
cnf(prove_this_21, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c7|inverse(sk_c3)=sk_c7), file('./p/GRP293-1.tptp', prove_this_21)).
cnf(prove_this_20, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c7|multiply(sk_c3,sk_c7)=sk_c8), file('./p/GRP293-1.tptp', prove_this_20)).
cnf(prove_this_26, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|multiply(sk_c3,sk_c7)=sk_c8), file('./p/GRP293-1.tptp', prove_this_26)).
cnf(prove_this_25, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|multiply(sk_c7,sk_c8)=sk_c6), file('./p/GRP293-1.tptp', prove_this_25)).
cnf(prove_this_6, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c6|inverse(sk_c4)=sk_c8), file('./p/GRP293-1.tptp', prove_this_6)).
cnf(prove_this_13, negated_conjecture, (inverse(sk_c1)=sk_c8|multiply(sk_c7,sk_c8)=sk_c6), file('./p/GRP293-1.tptp', prove_this_13)).
cnf(prove_this_12, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|inverse(sk_c4)=sk_c8), file('./p/GRP293-1.tptp', prove_this_12)).
cnf(prove_this_9, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|inverse(sk_c3)=sk_c7), file('./p/GRP293-1.tptp', prove_this_9)).
cnf(prove_this_15, negated_conjecture, (inverse(sk_c1)=sk_c8|inverse(sk_c3)=sk_c7), file('./p/GRP293-1.tptp', prove_this_15)).
cnf(prove_this_18, negated_conjecture, (inverse(sk_c1)=sk_c8|inverse(sk_c4)=sk_c8), file('./p/GRP293-1.tptp', prove_this_18)).
cnf(prove_this_31, negated_conjecture, (inverse(sk_c2)=sk_c8|multiply(sk_c7,sk_c8)=sk_c6), file('./p/GRP293-1.tptp', prove_this_31)).
cnf(prove_this_3, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c6|inverse(sk_c3)=sk_c7), file('./p/GRP293-1.tptp', prove_this_3)).
cnf(prove_this_17, negated_conjecture, (inverse(sk_c1)=sk_c8|multiply(sk_c4,sk_c8)=sk_c5), file('./p/GRP293-1.tptp', prove_this_17)).
cnf(prove_this_32, negated_conjecture, (inverse(sk_c2)=sk_c8|multiply(sk_c3,sk_c7)=sk_c8), file('./p/GRP293-1.tptp', prove_this_32)).
cnf(prove_this_16, negated_conjecture, (inverse(sk_c1)=sk_c8|multiply(sk_c8,sk_c5)=sk_c7), file('./p/GRP293-1.tptp', prove_this_16)).
cnf(prove_this_14, negated_conjecture, (inverse(sk_c1)=sk_c8|multiply(sk_c3,sk_c7)=sk_c8), file('./p/GRP293-1.tptp', prove_this_14)).
cnf(prove_this_8, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|multiply(sk_c3,sk_c7)=sk_c8), file('./p/GRP293-1.tptp', prove_this_8)).
cnf(prove_this_1, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c6|multiply(sk_c7,sk_c8)=sk_c6), file('./p/GRP293-1.tptp', prove_this_1)).
cnf(c_0_31, axiom, (multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), associativity).
cnf(c_0_32, axiom, (multiply(inverse(X1),X1)=identity), left_inverse).
cnf(c_0_33, axiom, (multiply(identity,X1)=X1), left_identity).
cnf(c_0_34, plain, (multiply(inverse(X1),multiply(X1,X2))=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_31, c_0_32]), c_0_33])).
cnf(c_0_35, plain, (multiply(inverse(inverse(X1)),identity)=X1), inference(pm,[status(thm)],[c_0_34, c_0_32])).
cnf(c_0_36, plain, (multiply(inverse(inverse(X1)),X2)=multiply(X1,X2)), inference(pm,[status(thm)],[c_0_34, c_0_34])).
cnf(c_0_37, plain, (multiply(X1,identity)=X1), inference(rw,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_38, plain, (inverse(inverse(X1))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_36, c_0_37]), c_0_37])).
cnf(c_0_39, negated_conjecture, (inverse(sk_c2)=sk_c8|inverse(sk_c4)=sk_c8), prove_this_36).
cnf(c_0_40, plain, (multiply(X1,inverse(X1))=identity), inference(pm,[status(thm)],[c_0_32, c_0_38])).
cnf(c_0_41, negated_conjecture, (inverse(sk_c8)=sk_c4|inverse(sk_c2)=sk_c8), inference(pm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_42, negated_conjecture, (inverse(sk_c2)=sk_c8|multiply(sk_c4,sk_c8)=sk_c5), prove_this_35).
cnf(c_0_43, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|inverse(sk_c4)=sk_c8), prove_this_30).
cnf(c_0_44, negated_conjecture, (multiply(sk_c2,sk_c8)=identity|inverse(sk_c8)=sk_c4), inference(pm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_45, negated_conjecture, (multiply(inverse(sk_c4),sk_c5)=sk_c8|inverse(sk_c2)=sk_c8), inference(pm,[status(thm)],[c_0_34, c_0_42])).
cnf(c_0_46, negated_conjecture, (inverse(sk_c4)=sk_c8|inverse(sk_c8)=sk_c4|sk_c6=identity), inference(pm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_47, negated_conjecture, (inverse(sk_c2)=sk_c8|multiply(sk_c8,sk_c5)=sk_c7), prove_this_34).
cnf(c_0_48, negated_conjecture, (multiply(sk_c8,sk_c5)=sk_c8|inverse(sk_c2)=sk_c8), inference(pm,[status(thm)],[c_0_45, c_0_39])).
cnf(c_0_49, negated_conjecture, (multiply(sk_c4,sk_c8)=identity|inverse(sk_c2)=sk_c8), inference(pm,[status(thm)],[c_0_40, c_0_39])).
cnf(c_0_50, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c6=identity), inference(pm,[status(thm)],[c_0_38, c_0_46])).
cnf(c_0_51, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_52, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c5=identity), inference(pm,[status(thm)],[c_0_42, c_0_49])).
cnf(c_0_53, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|multiply(sk_c4,sk_c8)=sk_c5), prove_this_29).
cnf(c_0_54, negated_conjecture, (multiply(sk_c4,sk_c8)=identity|sk_c6=identity), inference(pm,[status(thm)],[c_0_32, c_0_50])).
cnf(c_0_55, negated_conjecture, (multiply(sk_c8,sk_c7)!=sk_c6|multiply(X1,sk_c8)!=sk_c7|inverse(X1)!=sk_c8|multiply(sk_c8,sk_c6)!=sk_c7|multiply(X2,sk_c8)!=sk_c6|inverse(X2)!=sk_c8|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X3,sk_c7)!=sk_c8|inverse(X3)!=sk_c7|multiply(sk_c8,X4)!=sk_c7|multiply(X5,sk_c8)!=X4|inverse(X5)!=sk_c8), prove_this_37).
cnf(c_0_56, plain, (multiply(inverse(identity),X1)=X1), inference(pm,[status(thm)],[c_0_34, c_0_33])).
cnf(c_0_57, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|multiply(sk_c8,sk_c5)=sk_c7), prove_this_28).
cnf(c_0_58, negated_conjecture, (multiply(sk_c2,sk_c8)=identity|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_40, c_0_51])).
cnf(c_0_59, negated_conjecture, (multiply(sk_c2,sk_c8)=identity|sk_c5=identity), inference(pm,[status(thm)],[c_0_40, c_0_52])).
cnf(c_0_60, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|sk_c5=identity|sk_c6=identity), inference(pm,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_61, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|inverse(sk_c3)=sk_c7), prove_this_27).
cnf(c_0_62, negated_conjecture, (multiply(sk_c8,multiply(X1,sk_c8))!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X2,sk_c7)!=sk_c8|multiply(X3,sk_c8)!=sk_c6|multiply(X4,sk_c8)!=sk_c7|inverse(X1)!=sk_c8|inverse(X2)!=sk_c7|inverse(X3)!=sk_c8|inverse(X4)!=sk_c8), inference(er,[status(thm)],[c_0_55])).
cnf(c_0_63, plain, (multiply(inverse(inverse(identity)),X1)=X1), inference(pm,[status(thm)],[c_0_34, c_0_56])).
cnf(c_0_64, negated_conjecture, (multiply(sk_c8,sk_c5)=sk_c7|sk_c6=identity|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_65, negated_conjecture, (sk_c6=identity|sk_c5=identity), inference(pm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_66, negated_conjecture, (multiply(inverse(sk_c2),sk_c6)=sk_c8|inverse(sk_c3)=sk_c7), inference(pm,[status(thm)],[c_0_34, c_0_61])).
cnf(c_0_67, negated_conjecture, (inverse(sk_c2)=sk_c8|inverse(sk_c3)=sk_c7), prove_this_33).
cnf(c_0_68, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c7)!=sk_c8|multiply(X2,sk_c8)!=sk_c6|multiply(X3,sk_c8)!=sk_c7|inverse(identity)!=sk_c8|inverse(X1)!=sk_c7|inverse(X2)!=sk_c8|inverse(X3)!=sk_c8), inference(pm,[status(thm)],[c_0_62, c_0_33])).
cnf(c_0_69, plain, (inverse(identity)=identity), inference(pm,[status(thm)],[c_0_32, c_0_63])).
cnf(c_0_70, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c7|multiply(sk_c7,sk_c8)=sk_c6), prove_this_19).
cnf(c_0_71, negated_conjecture, (sk_c8=sk_c7|sk_c6=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_64, c_0_65]), c_0_37])])).
cnf(c_0_72, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c7|inverse(sk_c3)=sk_c7), prove_this_21).
cnf(c_0_73, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c8|inverse(sk_c3)=sk_c7), inference(pm,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_74, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c7)!=sk_c8|multiply(X2,sk_c8)!=sk_c6|multiply(X3,sk_c8)!=sk_c7|inverse(X1)!=sk_c7|inverse(X2)!=sk_c8|inverse(X3)!=sk_c8|sk_c8!=identity), inference(rw,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_75, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c6|sk_c8=sk_c7), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_70, c_0_71]), c_0_37])])).
cnf(c_0_76, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_77, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c7|multiply(sk_c3,sk_c7)=sk_c8), prove_this_20).
cnf(c_0_78, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c6|multiply(X2,sk_c8)!=sk_c7|inverse(X1)!=sk_c8|inverse(X2)!=sk_c8|sk_c8!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_74, c_0_32]), c_0_38])])).
cnf(c_0_79, negated_conjecture, (multiply(inverse(sk_c7),sk_c6)=sk_c8|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_34, c_0_75])).
cnf(c_0_80, negated_conjecture, (multiply(sk_c3,sk_c7)=identity|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_40, c_0_76])).
cnf(c_0_81, negated_conjecture, (multiply(sk_c3,sk_c7)=sk_c8|sk_c8=sk_c7), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_77, c_0_71]), c_0_37])])).
cnf(c_0_82, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c7!=identity|sk_c8!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_78, c_0_32]), c_0_38])])).
cnf(c_0_83, negated_conjecture, (multiply(sk_c7,multiply(sk_c3,X1))=X1|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_34, c_0_76])).
cnf(c_0_84, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c8=sk_c7), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_79, c_0_71]), c_0_37])).
cnf(c_0_85, negated_conjecture, (sk_c8=identity|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_86, negated_conjecture, (multiply(X1,multiply(X2,sk_c8))!=sk_c6|multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|inverse(multiply(X1,X2))!=sk_c8|sk_c7!=identity|sk_c8!=identity), inference(pm,[status(thm)],[c_0_82, c_0_31])).
cnf(c_0_87, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c7|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_83, c_0_81])).
cnf(c_0_88, negated_conjecture, (multiply(sk_c8,sk_c6)=sk_c8|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_79, c_0_84])).
cnf(c_0_89, negated_conjecture, (sk_c8=identity|sk_c7!=identity), inference(ef,[status(thm)],[c_0_85])).
cnf(c_0_90, negated_conjecture, (inverse(sk_c8)=sk_c2|inverse(sk_c8)=sk_c4), inference(pm,[status(thm)],[c_0_38, c_0_41])).
cnf(c_0_91, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_86, c_0_34]), c_0_32]), c_0_69])])).
cnf(c_0_92, negated_conjecture, (sk_c6=sk_c7|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_75, c_0_87])).
cnf(c_0_93, negated_conjecture, (sk_c6=sk_c8|sk_c8=sk_c7|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_88, c_0_89]), c_0_33])).
cnf(c_0_94, negated_conjecture, (multiply(sk_c2,sk_c8)=identity|inverse(sk_c3)=sk_c7), inference(pm,[status(thm)],[c_0_40, c_0_67])).
cnf(c_0_95, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c4!=sk_c2), inference(ef,[status(thm)],[c_0_90])).
cnf(c_0_96, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_91, c_0_89]), c_0_33])).
cnf(c_0_97, negated_conjecture, (sk_c8=sk_c7|sk_c7!=identity), inference(pm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_98, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c6=identity), inference(pm,[status(thm)],[c_0_61, c_0_94])).
cnf(c_0_99, negated_conjecture, (multiply(sk_c2,sk_c8)=identity|sk_c4!=sk_c2), inference(pm,[status(thm)],[c_0_32, c_0_95])).
cnf(c_0_100, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|multiply(sk_c3,sk_c7)=sk_c8), prove_this_26).
cnf(c_0_101, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity), inference(pm,[status(thm)],[c_0_96, c_0_97])).
cnf(c_0_102, negated_conjecture, (multiply(sk_c2,sk_c8)=sk_c6|multiply(sk_c7,sk_c8)=sk_c6), prove_this_25).
cnf(c_0_103, negated_conjecture, (multiply(sk_c7,multiply(sk_c3,X1))=X1|sk_c6=identity), inference(pm,[status(thm)],[c_0_34, c_0_98])).
cnf(c_0_104, negated_conjecture, (multiply(sk_c3,sk_c7)=sk_c8|sk_c6=identity|sk_c4!=sk_c2), inference(pm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_105, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_101, c_0_89]), c_0_33])])).
cnf(c_0_106, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c6|sk_c6=identity|sk_c4!=sk_c2), inference(pm,[status(thm)],[c_0_99, c_0_102])).
cnf(c_0_107, negated_conjecture, (multiply(sk_c7,sk_c8)=sk_c7|sk_c6=identity|sk_c4!=sk_c2), inference(pm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_108, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c5=sk_c7|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_47, c_0_89]), c_0_33])).
cnf(c_0_109, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c5=sk_c4|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_42, c_0_89]), c_0_37])).
cnf(c_0_110, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_105, c_0_89]), c_0_37])])).
cnf(c_0_111, negated_conjecture, (sk_c6=sk_c7|sk_c6=identity|sk_c4!=sk_c2), inference(pm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_112, negated_conjecture, (multiply(sk_c8,sk_c2)=identity|inverse(sk_c8)=sk_c4), inference(pm,[status(thm)],[c_0_32, c_0_41])).
cnf(c_0_113, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c5=sk_c7|sk_c7!=identity), inference(pm,[status(thm)],[c_0_38, c_0_108])).
cnf(c_0_114, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c5=sk_c4|sk_c7!=identity), inference(pm,[status(thm)],[c_0_38, c_0_109])).
cnf(c_0_115, negated_conjecture, (sk_c6!=sk_c7|sk_c6!=sk_c8|sk_c8!=identity|sk_c7!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_110, c_0_89]), c_0_37])])).
cnf(c_0_116, negated_conjecture, (sk_c6=identity|sk_c7!=identity|sk_c4!=sk_c2), inference(ef,[status(thm)],[c_0_111])).
cnf(c_0_117, negated_conjecture, (sk_c4=identity|sk_c6=identity|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_54, c_0_89]), c_0_37])).
cnf(c_0_118, negated_conjecture, (multiply(sk_c8,sk_c2)=identity|inverse(sk_c3)=sk_c7), inference(pm,[status(thm)],[c_0_32, c_0_67])).
cnf(c_0_119, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c2=identity|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_112, c_0_89]), c_0_33])).
cnf(c_0_120, negated_conjecture, (sk_c2=identity|sk_c5=sk_c7|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_113, c_0_89]), c_0_69])).
cnf(c_0_121, negated_conjecture, (sk_c2=identity|sk_c5=sk_c4|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_114, c_0_89]), c_0_69])).
cnf(c_0_122, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c8,sk_c6)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_82, c_0_32]), c_0_38])])).
cnf(c_0_123, negated_conjecture, (sk_c8!=identity|sk_c6!=sk_c7|sk_c7!=identity|sk_c4!=sk_c2), inference(pm,[status(thm)],[c_0_115, c_0_116])).
cnf(c_0_124, negated_conjecture, (sk_c4=identity|sk_c8!=identity|sk_c6!=sk_c7|sk_c7!=identity), inference(pm,[status(thm)],[c_0_115, c_0_117])).
cnf(c_0_125, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c2=identity|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_118, c_0_89]), c_0_33])).
cnf(c_0_126, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c2=identity|sk_c7!=identity), inference(pm,[status(thm)],[c_0_38, c_0_119])).
cnf(c_0_127, negated_conjecture, (sk_c4=sk_c7|sk_c2=identity|sk_c7!=identity), inference(pm,[status(thm)],[c_0_120, c_0_121])).
cnf(c_0_128, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_122, c_0_89]), c_0_33])).
cnf(c_0_129, negated_conjecture, (sk_c7!=identity|sk_c8!=identity|sk_c4!=sk_c2), inference(pm,[status(thm)],[c_0_123, c_0_116])).
cnf(c_0_130, negated_conjecture, (sk_c4=identity|sk_c7!=identity|sk_c8!=identity), inference(pm,[status(thm)],[c_0_124, c_0_117])).
cnf(c_0_131, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c2=identity|sk_c7!=identity), inference(pm,[status(thm)],[c_0_38, c_0_125])).
cnf(c_0_132, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c2=identity|sk_c7!=identity), inference(pm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_133, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity), inference(pm,[status(thm)],[c_0_128, c_0_97])).
cnf(c_0_134, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c6|inverse(sk_c4)=sk_c8), prove_this_6).
cnf(c_0_135, negated_conjecture, (inverse(sk_c1)=sk_c8|multiply(sk_c7,sk_c8)=sk_c6), prove_this_13).
cnf(c_0_136, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|inverse(sk_c4)=sk_c8), prove_this_12).
cnf(c_0_137, negated_conjecture, (sk_c2!=identity|sk_c7!=identity|sk_c8!=identity), inference(pm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_138, negated_conjecture, (sk_c8=sk_c3|sk_c2=identity|sk_c7!=identity), inference(pm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_139, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_133, c_0_89]), c_0_33])])).
cnf(c_0_140, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c6=sk_c7|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_134, c_0_89]), c_0_33])).
cnf(c_0_141, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c6=sk_c7|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_135, c_0_89]), c_0_37])).
cnf(c_0_142, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c1=sk_c7|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_136, c_0_89]), c_0_37])).
cnf(c_0_143, negated_conjecture, (sk_c8=sk_c3|sk_c7!=identity|sk_c8!=identity), inference(pm,[status(thm)],[c_0_137, c_0_138])).
cnf(c_0_144, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c7|sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_139, c_0_89]), c_0_37])])).
cnf(c_0_145, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c6=sk_c7|sk_c7!=identity), inference(pm,[status(thm)],[c_0_38, c_0_140])).
cnf(c_0_146, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c6=sk_c7|sk_c7!=identity), inference(pm,[status(thm)],[c_0_38, c_0_141])).
cnf(c_0_147, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|multiply(X1,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_82, c_0_89]), c_0_33])).
cnf(c_0_148, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|inverse(sk_c3)=sk_c7), prove_this_9).
cnf(c_0_149, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c1=sk_c7|sk_c7!=identity), inference(pm,[status(thm)],[c_0_38, c_0_142])).
cnf(c_0_150, negated_conjecture, (sk_c8=sk_c3|sk_c7!=identity), inference(pm,[status(thm)],[c_0_143, c_0_89])).
cnf(c_0_151, negated_conjecture, (sk_c6!=sk_c7|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_144, c_0_89]), c_0_37])])).
cnf(c_0_152, negated_conjecture, (sk_c4=sk_c1|sk_c6=sk_c7|sk_c7!=identity), inference(pm,[status(thm)],[c_0_145, c_0_146])).
cnf(c_0_153, negated_conjecture, (multiply(sk_c8,sk_c8)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_147, c_0_89]), c_0_37])).
cnf(c_0_154, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c1=sk_c7|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_148, c_0_89]), c_0_37])).
cnf(c_0_155, negated_conjecture, (inverse(sk_c3)=sk_c4|sk_c1=sk_c7|sk_c7!=identity), inference(pm,[status(thm)],[c_0_149, c_0_150])).
cnf(c_0_156, negated_conjecture, (sk_c4=sk_c1|sk_c6!=identity|sk_c7!=identity|sk_c8!=identity), inference(pm,[status(thm)],[c_0_151, c_0_152])).
cnf(c_0_157, negated_conjecture, (multiply(sk_c8,sk_c7)!=sk_c7|multiply(sk_c8,sk_c7)!=sk_c6|multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6), inference(pm,[status(thm)],[c_0_153, c_0_97])).
cnf(c_0_158, negated_conjecture, (sk_c4=sk_c7|sk_c1=sk_c7|sk_c7!=identity), inference(pm,[status(thm)],[c_0_154, c_0_155])).
cnf(c_0_159, negated_conjecture, (sk_c4=sk_c1|sk_c7!=identity|sk_c8!=identity), inference(pm,[status(thm)],[c_0_156, c_0_152])).
cnf(c_0_160, negated_conjecture, (multiply(inverse(sk_c1),sk_c7)=sk_c8|inverse(sk_c3)=sk_c7), inference(pm,[status(thm)],[c_0_34, c_0_148])).
cnf(c_0_161, negated_conjecture, (inverse(sk_c1)=sk_c8|inverse(sk_c3)=sk_c7), prove_this_15).
cnf(c_0_162, negated_conjecture, (multiply(sk_c8,sk_c7)!=sk_c7|multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_157, c_0_89]), c_0_33])])).
cnf(c_0_163, negated_conjecture, (inverse(sk_c1)=sk_c8|inverse(sk_c4)=sk_c8), prove_this_18).
cnf(c_0_164, negated_conjecture, (sk_c1=sk_c7|sk_c7!=identity|sk_c8!=identity), inference(pm,[status(thm)],[c_0_158, c_0_159])).
cnf(c_0_165, negated_conjecture, (inverse(sk_c2)=sk_c8|multiply(sk_c7,sk_c8)=sk_c6), prove_this_31).
cnf(c_0_166, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c6|inverse(sk_c3)=sk_c7), prove_this_3).
cnf(c_0_167, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c8|inverse(sk_c3)=sk_c7), inference(pm,[status(thm)],[c_0_160, c_0_161])).
cnf(c_0_168, negated_conjecture, (multiply(sk_c7,sk_c8)!=sk_c6|inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_162, c_0_89]), c_0_33])])).
cnf(c_0_169, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c7!=identity|sk_c8!=identity), inference(pm,[status(thm)],[c_0_163, c_0_159])).
cnf(c_0_170, negated_conjecture, (sk_c1=sk_c7|sk_c7!=identity), inference(pm,[status(thm)],[c_0_164, c_0_89])).
cnf(c_0_171, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c6=sk_c7|sk_c7!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_165, c_0_89]), c_0_37])).
cnf(c_0_172, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c6=sk_c8), inference(pm,[status(thm)],[c_0_166, c_0_167])).
cnf(c_0_173, negated_conjecture, (inverse(sk_c1)=sk_c8|multiply(sk_c4,sk_c8)=sk_c5), prove_this_17).
cnf(c_0_174, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|inverse(sk_c3)=sk_c7), inference(pm,[status(thm)],[c_0_40, c_0_161])).
cnf(c_0_175, negated_conjecture, (inverse(X1)!=sk_c8|sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity|X1!=sk_c6), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_168, c_0_89]), c_0_37])])).
cnf(c_0_176, negated_conjecture, (inverse(sk_c7)=sk_c8|sk_c7!=identity|sk_c8!=identity), inference(pm,[status(thm)],[c_0_169, c_0_170])).
cnf(c_0_177, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c6=sk_c7|sk_c7!=identity), inference(pm,[status(thm)],[c_0_38, c_0_171])).
cnf(c_0_178, negated_conjecture, (inverse(sk_c2)=sk_c8|multiply(sk_c3,sk_c7)=sk_c8), prove_this_32).
cnf(c_0_179, negated_conjecture, (multiply(sk_c3,sk_c7)=identity|sk_c6=identity), inference(pm,[status(thm)],[c_0_40, c_0_98])).
cnf(c_0_180, negated_conjecture, (multiply(sk_c3,sk_c7)=identity|sk_c6=sk_c8), inference(pm,[status(thm)],[c_0_40, c_0_172])).
cnf(c_0_181, negated_conjecture, (multiply(inverse(sk_c4),sk_c5)=sk_c8|inverse(sk_c1)=sk_c8), inference(pm,[status(thm)],[c_0_34, c_0_173])).
cnf(c_0_182, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c7=identity), inference(pm,[status(thm)],[c_0_148, c_0_174])).
cnf(c_0_183, negated_conjecture, (sk_c6!=sk_c7|sk_c7!=identity|sk_c8!=identity), inference(pm,[status(thm)],[c_0_175, c_0_176])).
cnf(c_0_184, negated_conjecture, (sk_c4=sk_c2|sk_c6=sk_c7|sk_c7!=identity), inference(pm,[status(thm)],[c_0_145, c_0_177])).
cnf(c_0_185, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c8=identity|sk_c6=identity), inference(pm,[status(thm)],[c_0_178, c_0_179])).
cnf(c_0_186, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c8=identity|sk_c6=sk_c8), inference(pm,[status(thm)],[c_0_178, c_0_180])).
cnf(c_0_187, negated_conjecture, (inverse(sk_c1)=sk_c8|multiply(sk_c8,sk_c5)=sk_c7), prove_this_16).
cnf(c_0_188, negated_conjecture, (multiply(sk_c8,sk_c5)=sk_c8|inverse(sk_c1)=sk_c8), inference(pm,[status(thm)],[c_0_181, c_0_163])).
cnf(c_0_189, negated_conjecture, (inverse(sk_c1)=sk_c8|multiply(sk_c3,sk_c7)=sk_c8), prove_this_14).
cnf(c_0_190, negated_conjecture, (multiply(sk_c3,sk_c7)=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_40, c_0_182])).
cnf(c_0_191, negated_conjecture, (sk_c4=sk_c2|sk_c7!=identity|sk_c8!=identity), inference(pm,[status(thm)],[c_0_183, c_0_184])).
cnf(c_0_192, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c8=identity|sk_c6=identity), inference(pm,[status(thm)],[c_0_38, c_0_185])).
cnf(c_0_193, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c8=identity|sk_c6=sk_c8), inference(pm,[status(thm)],[c_0_38, c_0_186])).
cnf(c_0_194, negated_conjecture, (inverse(sk_c8)=sk_c2|inverse(sk_c3)=sk_c7), inference(pm,[status(thm)],[c_0_38, c_0_67])).
cnf(c_0_195, negated_conjecture, (inverse(sk_c8)=sk_c1|inverse(sk_c3)=sk_c7), inference(pm,[status(thm)],[c_0_38, c_0_161])).
cnf(c_0_196, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_187, c_0_188])).
cnf(c_0_197, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c8=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_189, c_0_190])).
cnf(c_0_198, negated_conjecture, (sk_c7!=identity|sk_c8!=identity), inference(pm,[status(thm)],[c_0_129, c_0_191])).
cnf(c_0_199, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c6=identity), inference(pm,[status(thm)],[c_0_38, c_0_98])).
cnf(c_0_200, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c8=identity|sk_c6=identity), inference(pm,[status(thm)],[c_0_192, c_0_85])).
cnf(c_0_201, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c6=sk_c8), inference(pm,[status(thm)],[c_0_38, c_0_172])).
cnf(c_0_202, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c8=identity|sk_c6=sk_c8), inference(pm,[status(thm)],[c_0_193, c_0_85])).
cnf(c_0_203, negated_conjecture, (inverse(sk_c3)=sk_c7|sk_c2=sk_c1), inference(pm,[status(thm)],[c_0_194, c_0_195])).
cnf(c_0_204, negated_conjecture, (multiply(sk_c1,sk_c8)=sk_c7|multiply(sk_c3,sk_c7)=sk_c8), prove_this_8).
cnf(c_0_205, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_40, c_0_196])).
cnf(c_0_206, negated_conjecture, (multiply(sk_c1,sk_c8)=identity|sk_c8=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_40, c_0_197])).
cnf(c_0_207, negated_conjecture, (sk_c7=identity|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_71, c_0_92])).
cnf(c_0_208, negated_conjecture, (sk_c7!=identity), inference(pm,[status(thm)],[c_0_198, c_0_89])).
cnf(c_0_209, negated_conjecture, (sk_c2=sk_c3|sk_c6=identity|sk_c8=identity), inference(pm,[status(thm)],[c_0_199, c_0_200])).
cnf(c_0_210, negated_conjecture, (sk_c2=sk_c3|sk_c6=sk_c8|sk_c8=identity), inference(pm,[status(thm)],[c_0_201, c_0_202])).
cnf(c_0_211, negated_conjecture, (multiply(sk_c3,sk_c7)=identity|sk_c2=sk_c1), inference(pm,[status(thm)],[c_0_40, c_0_203])).
cnf(c_0_212, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c8=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_178, c_0_190])).
cnf(c_0_213, negated_conjecture, (multiply(sk_c3,sk_c7)=sk_c8|sk_c7=identity|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_204, c_0_205])).
cnf(c_0_214, negated_conjecture, (inverse(sk_c4)=sk_c8|sk_c7=identity|sk_c8=identity), inference(pm,[status(thm)],[c_0_136, c_0_206])).
cnf(c_0_215, negated_conjecture, (multiply(sk_c8,sk_c7)=sk_c6|multiply(sk_c7,sk_c8)=sk_c6), prove_this_1).
cnf(c_0_216, negated_conjecture, (sk_c8=sk_c7), inference(sr,[status(thm)],[c_0_207, c_0_208])).
cnf(c_0_217, negated_conjecture, (inverse(sk_c7)=sk_c3|sk_c7=identity), inference(pm,[status(thm)],[c_0_38, c_0_182])).
cnf(c_0_218, negated_conjecture, (sk_c8=identity|sk_c2=sk_c3), inference(pm,[status(thm)],[c_0_209, c_0_210])).
cnf(c_0_219, negated_conjecture, (inverse(sk_c1)=sk_c8|sk_c8=identity|sk_c2=sk_c1), inference(pm,[status(thm)],[c_0_211, c_0_189])).
cnf(c_0_220, negated_conjecture, (inverse(sk_c2)=sk_c8|sk_c8=identity|sk_c2=sk_c1), inference(pm,[status(thm)],[c_0_211, c_0_178])).
cnf(c_0_221, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c8=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_38, c_0_212])).
cnf(c_0_222, negated_conjecture, (sk_c8=identity|sk_c7=identity|sk_c8=sk_c7), inference(pm,[status(thm)],[c_0_190, c_0_213])).
cnf(c_0_223, negated_conjecture, (inverse(sk_c8)=sk_c4|sk_c7=identity|sk_c8=identity), inference(pm,[status(thm)],[c_0_38, c_0_214])).
cnf(c_0_224, negated_conjecture, (multiply(sk_c7,sk_c7)=sk_c6), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_215, c_0_216]), c_0_216])])).
cnf(c_0_225, negated_conjecture, (inverse(sk_c7)=sk_c3), inference(sr,[status(thm)],[c_0_217, c_0_208])).
cnf(c_0_226, negated_conjecture, (sk_c8=identity|sk_c6=identity|sk_c4!=sk_c2), inference(pm,[status(thm)],[c_0_179, c_0_104])).
cnf(c_0_227, negated_conjecture, (sk_c2=sk_c3), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_218, c_0_216]), c_0_208])).
cnf(c_0_228, negated_conjecture, (inverse(sk_c8)=sk_c1|sk_c8=identity|sk_c2=sk_c1), inference(pm,[status(thm)],[c_0_38, c_0_219])).
cnf(c_0_229, negated_conjecture, (inverse(sk_c8)=sk_c2|sk_c8=identity|sk_c2=sk_c1), inference(pm,[status(thm)],[c_0_38, c_0_220])).
cnf(c_0_230, negated_conjecture, (inverse(sk_c7)=sk_c2|sk_c8=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_221, c_0_222])).
cnf(c_0_231, negated_conjecture, (sk_c4=sk_c2|sk_c8=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_221, c_0_223])).
cnf(c_0_232, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_34, c_0_224]), c_0_225])).
cnf(c_0_233, negated_conjecture, (sk_c6=identity|sk_c4!=sk_c3), inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_226, c_0_216]), c_0_208]), c_0_227])).
cnf(c_0_234, negated_conjecture, (sk_c2=sk_c1|sk_c8=identity), inference(pm,[status(thm)],[c_0_228, c_0_229])).
cnf(c_0_235, negated_conjecture, (sk_c2=sk_c3|sk_c7=identity|sk_c8=identity), inference(pm,[status(thm)],[c_0_217, c_0_230])).
cnf(c_0_236, negated_conjecture, (sk_c4=sk_c3), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[c_0_231, c_0_208]), c_0_227]), c_0_216]), c_0_208])).
cnf(c_0_237, negated_conjecture, (sk_c7=sk_c3|sk_c4!=sk_c3), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_232, c_0_233]), c_0_37])).
cnf(c_0_238, negated_conjecture, (sk_c1=sk_c3|sk_c8=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_234, c_0_235])).
cnf(c_0_239, negated_conjecture, (sk_c6=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_233, c_0_236])])).
cnf(c_0_240, negated_conjecture, (sk_c7=sk_c3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_237, c_0_236])])).
cnf(c_0_241, negated_conjecture, (sk_c7=identity|sk_c1=sk_c3), inference(pm,[status(thm)],[c_0_238, c_0_207])).
cnf(c_0_242, negated_conjecture, (multiply(sk_c3,sk_c3)=identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_224, c_0_239]), c_0_240]), c_0_240])).
cnf(c_0_243, negated_conjecture, (sk_c1=sk_c3), inference(sr,[status(thm)],[c_0_241, c_0_208])).
cnf(c_0_244, negated_conjecture, (sk_c3!=identity), inference(rw,[status(thm)],[c_0_208, c_0_240])).
cnf(c_0_245, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_204, c_0_216]), c_0_216]), c_0_240]), c_0_242]), c_0_240]), c_0_243]), c_0_240]), c_0_242]), c_0_240])]), c_0_244]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 246
# Proof object clause steps            : 215
# Proof object formula steps           : 31
# Proof object conjectures             : 231
# Proof object clause conjectures      : 203
# Proof object formula conjectures     : 28
# Proof object initial clauses used    : 31
# Proof object initial formulas used   : 31
# Proof object generating inferences   : 170
# Proof object simplifying inferences  : 91
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 40
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 40
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 40
# Processed clauses                    : 9155
# ...of these trivial                  : 623
# ...subsumed                          : 7010
# ...remaining for further processing  : 1521
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 419
# Backward-rewritten                   : 1005
# Generated clauses                    : 134381
# ...of the previous two non-trivial   : 111610
# Contextual simplify-reflections      : 0
# Paramodulations                      : 134316
# Factorizations                       : 9
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 42
#    Positive orientable unit clauses  : 20
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 21
# Current number of unprocessed clauses: 100071
# ...number of literals in the above   : 583705
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1479
# Clause-clause subsumption calls (NU) : 48370
# Rec. Clause-clause subsumption calls : 23126
# Non-unit clause-clause subsumptions  : 7270
# Unit Clause-clause subsumption calls : 358
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 18
# BW rewrite match successes           : 17
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 629264
