# Initializing proof state
# Scanning for AC axioms
# multiply is associative
#
#cnf(i_0_30, plain, (multiply(identity,X1)=X1)).
#
#cnf(i_0_31, plain, (multiply(inverse(X1),X1)=identity)).
#
#cnf(i_0_45, negated_conjecture, (inverse(sk_c3)=sk_c6|inverse(sk_c1)=sk_c7)).
#
#cnf(i_0_58, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X4,sk_c5)!=sk_c6|multiply(X3,sk_c6)!=sk_c7|multiply(X2,sk_c6)!=sk_c5|multiply(X1,sk_c7)!=sk_c6|inverse(X4)!=sk_c6|inverse(X3)!=sk_c6|inverse(X2)!=sk_c6|inverse(X1)!=sk_c7)).
#
#cnf(i_0_55, negated_conjecture, (inverse(sk_c3)=sk_c6|inverse(sk_c2)=sk_c6)).
#
#cnf(i_0_46, negated_conjecture, (inverse(sk_c4)=sk_c6|inverse(sk_c1)=sk_c7)).
#
#cnf(i_0_56, negated_conjecture, (inverse(sk_c4)=sk_c6|inverse(sk_c2)=sk_c6)).
#
#cnf(i_0_60, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|inverse(inverse(sk_c7))!=sk_c7|multiply(X1,sk_c5)!=sk_c6|multiply(X2,sk_c6)!=sk_c7|multiply(X3,sk_c6)!=sk_c5|inverse(X1)!=sk_c6|inverse(X2)!=sk_c6|inverse(X3)!=sk_c6|sk_c6!=identity)).
#
#cnf(i_0_32, plain, (multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3)))).
#
#cnf(i_0_78, plain, (multiply(inverse(X1),multiply(X1,X2))=X2)).
#
#cnf(i_0_84, plain, (multiply(inverse(identity),X1)=X1)).
#
#cnf(i_0_69, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|inverse(inverse(sk_c7))!=sk_c7|inverse(inverse(sk_c5))!=sk_c6|multiply(X1,sk_c6)!=sk_c7|multiply(X2,sk_c6)!=sk_c5|inverse(X1)!=sk_c6|inverse(X2)!=sk_c6|sk_c6!=identity)).
#
#cnf(i_0_82, plain, (multiply(inverse(inverse(X1)),identity)=X1)).
#
#cnf(i_0_97, plain, (multiply(inverse(inverse(identity)),X1)=X1)).
#
#cnf(i_0_123, plain, (inverse(identity)=identity)).
#
#cnf(i_0_100, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|inverse(inverse(sk_c7))!=sk_c7|inverse(inverse(sk_c5))!=sk_c6|inverse(inverse(sk_c6))!=sk_c6|multiply(X1,sk_c6)!=sk_c7|inverse(X1)!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
#
#cnf(i_0_81, plain, (multiply(inverse(inverse(X1)),X2)=multiply(X1,X2))).
#
#cnf(i_0_82, plain, (multiply(X1,identity)=X1)).
#
#cnf(i_0_148, plain, (inverse(inverse(X1))=X1)).
#
#cnf(i_0_100, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c6)!=sk_c7|inverse(X1)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
#
#cnf(i_0_153, plain, (multiply(X1,inverse(X1))=identity)).
#
#cnf(i_0_154, plain, (multiply(X1,multiply(inverse(X1),X2))=X2)).
##
#cnf(i_0_160, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c7!=identity|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
#
#cnf(i_0_155, negated_conjecture, (inverse(sk_c6)=sk_c2|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_156, negated_conjecture, (inverse(sk_c7)=sk_c1|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_157, negated_conjecture, (inverse(sk_c6)=sk_c4|inverse(sk_c1)=sk_c7)).
#
#cnf(i_0_161, negated_conjecture, (multiply(X1,multiply(X2,sk_c6))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|inverse(multiply(X1,X2))!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
#
#cnf(i_0_158, negated_conjecture, (inverse(sk_c6)=sk_c4|inverse(sk_c2)=sk_c6)).
###
#cnf(i_0_204, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c7!=sk_c6|sk_c6!=identity|sk_c5!=sk_c6|sk_c5!=identity)).
#
#cnf(i_0_197, negated_conjecture, (inverse(sk_c7)=sk_c1|inverse(sk_c6)=sk_c4)).
#
#cnf(i_0_210, negated_conjecture, (inverse(sk_c6)=sk_c2|inverse(sk_c6)=sk_c4)).
#
#cnf(i_0_218, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c4!=sk_c2)).
##
#cnf(i_0_227, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c4!=sk_c2)).
##
#cnf(i_0_225, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_228, negated_conjecture, (multiply(sk_c6,sk_c2)=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_221, negated_conjecture, (inverse(sk_c4)=sk_c6|inverse(sk_c6)=sk_c2)).
##
#cnf(i_0_162, plain, (multiply(X1,multiply(X2,inverse(multiply(X1,X2))))=identity)).
#
#cnf(i_0_200, negated_conjecture, (inverse(multiply(X1,inverse(sk_c6)))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|X1!=sk_c7)).
#
#cnf(i_0_254, plain, (multiply(X2,inverse(multiply(X1,X2)))=inverse(X1))).
####
#cnf(i_0_280, plain, (multiply(inverse(multiply(X1,X2)),X1)=inverse(X2))).
####
#cnf(i_0_279, plain, (inverse(multiply(X2,X1))=multiply(inverse(X1),inverse(X2)))).
#
#cnf(i_0_35, negated_conjecture, (multiply(sk_c7,sk_c6)=sk_c5|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_40, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_200, negated_conjecture, (multiply(sk_c6,inverse(X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|X1!=sk_c7)).
#
#cnf(i_0_50, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c5|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_36, negated_conjecture, (multiply(sk_c7,sk_c6)=sk_c5|inverse(sk_c4)=sk_c6)).
#
#cnf(i_0_41, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|inverse(sk_c4)=sk_c6)).
#
#cnf(i_0_340, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c6,X1)!=sk_c6|inverse(X1)!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
#
#cnf(i_0_51, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c5|inverse(sk_c4)=sk_c6)).
#
#cnf(i_0_43, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c5|inverse(sk_c1)=sk_c7)).
#
#cnf(i_0_44, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7|inverse(sk_c1)=sk_c7)).
#
#cnf(i_0_361, negated_conjecture, (multiply(inverse(X1),sk_c6)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|X1!=sk_c6)).
#
#cnf(i_0_47, negated_conjecture, (multiply(sk_c4,sk_c5)=sk_c6|inverse(sk_c1)=sk_c7)).
#
#cnf(i_0_53, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c5|inverse(sk_c2)=sk_c6)).
#
#cnf(i_0_54, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7|inverse(sk_c2)=sk_c6)).
##
#cnf(i_0_57, negated_conjecture, (multiply(sk_c4,sk_c5)=sk_c6|inverse(sk_c2)=sk_c6)).
#
#cnf(i_0_59, negated_conjecture, (multiply(sk_c7,sk_c1)=identity|inverse(sk_c3)=sk_c6)).
###
#cnf(i_0_64, negated_conjecture, (multiply(sk_c6,sk_c2)=identity|inverse(sk_c3)=sk_c6)).
##
#cnf(i_0_65, negated_conjecture, (multiply(sk_c6,sk_c4)=identity|inverse(sk_c1)=sk_c7)).
#
#cnf(i_0_387, negated_conjecture, (multiply(sk_c2,sk_c6)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_66, negated_conjecture, (multiply(sk_c6,sk_c4)=identity|inverse(sk_c2)=sk_c6)).
###
#cnf(i_0_166, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_429, negated_conjecture, (inverse(sk_c3)=sk_c6|sk_c5=identity)).
#
#cnf(i_0_435, negated_conjecture, (inverse(sk_c6)=sk_c3|sk_c5=identity)).
####
#cnf(i_0_446, negated_conjecture, (sk_c2=sk_c3|sk_c5=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_379, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c2!=sk_c6|sk_c4!=sk_c2)).
#
#cnf(i_0_432, negated_conjecture, (multiply(sk_c6,sk_c3)=identity|sk_c5=identity)).
##
#cnf(i_0_436, negated_conjecture, (multiply(sk_c3,sk_c6)=identity|sk_c5=identity)).
#
#cnf(i_0_161, negated_conjecture, (multiply(X1,multiply(X2,sk_c6))!=sk_c7|multiply(inverse(X2),inverse(X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
#####
#cnf(i_0_167, negated_conjecture, (multiply(sk_c1,sk_c7)=identity|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_480, negated_conjecture, (inverse(sk_c3)=sk_c6|sk_c6=identity)).
#
#cnf(i_0_485, negated_conjecture, (inverse(sk_c6)=sk_c3|sk_c6=identity)).
####
#cnf(i_0_496, negated_conjecture, (sk_c2=sk_c3|sk_c6=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_482, negated_conjecture, (multiply(sk_c6,sk_c3)=identity|sk_c6=identity)).
##
#cnf(i_0_486, negated_conjecture, (multiply(sk_c3,sk_c6)=identity|sk_c6=identity)).
######
#cnf(i_0_168, negated_conjecture, (multiply(sk_c4,sk_c6)=identity|inverse(sk_c1)=sk_c7)).
##
#cnf(i_0_169, negated_conjecture, (multiply(sk_c4,sk_c6)=identity|inverse(sk_c2)=sk_c6)).
#####
#cnf(i_0_463, negated_conjecture, (multiply(inverse(sk_c2),inverse(X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2|X1!=sk_c7)).
###
#cnf(i_0_195, negated_conjecture, (multiply(sk_c7,sk_c1)=identity|inverse(sk_c6)=sk_c4)).
###
#cnf(i_0_198, negated_conjecture, (multiply(sk_c1,sk_c7)=identity|inverse(sk_c6)=sk_c4)).
##
#cnf(i_0_543, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|inverse(sk_c2)!=sk_c7|sk_c6!=identity|sk_c5!=sk_c6|sk_c5!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_208, negated_conjecture, (multiply(sk_c6,sk_c2)=identity|inverse(sk_c6)=sk_c4)).
##
#cnf(i_0_211, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|inverse(sk_c6)=sk_c4)).
######
#cnf(i_0_219, negated_conjecture, (multiply(sk_c4,sk_c6)=identity|inverse(sk_c6)=sk_c2)).
##
#cnf(i_0_222, negated_conjecture, (multiply(sk_c6,sk_c4)=identity|inverse(sk_c6)=sk_c2)).
#
#cnf(i_0_537, negated_conjecture, (multiply(inverse(sk_c2),X1)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|inverse(X1)!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_370, negated_conjecture, (inverse(sk_c4)=sk_c6|sk_c5=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_588, negated_conjecture, (inverse(sk_c6)=sk_c4|sk_c5=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_601, negated_conjecture, (sk_c4=sk_c3|sk_c5=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_226, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,X1))=X1|sk_c4!=sk_c2)).
#######
#cnf(i_0_229, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,X1))=X1|sk_c4!=sk_c2)).
##
#cnf(i_0_522, negated_conjecture, (multiply(inverse(sk_c2),sk_c6)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2|sk_c2!=sk_c7)).
####
#cnf(i_0_638, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2|sk_c2!=sk_c7)).
####
#cnf(i_0_581, negated_conjecture, (multiply(inverse(X1),inverse(sk_c2))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2|X1!=sk_c6)).
#
#cnf(i_0_445, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c4=sk_c3|sk_c5=identity)).
#
#cnf(i_0_676, negated_conjecture, (sk_c2=sk_c3|sk_c5=identity|sk_c4=sk_c3)).
#
#cnf(i_0_458, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c7=identity|sk_c5=identity)).
##
#cnf(i_0_459, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c7=identity|sk_c5=identity)).
##
#cnf(i_0_495, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c4=sk_c3|sk_c6=identity)).
##
#cnf(i_0_712, negated_conjecture, (sk_c2=sk_c3|sk_c6=identity|sk_c4=sk_c3)).
#
#cnf(i_0_504, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c7=identity|sk_c6=identity)).
#
#cnf(i_0_505, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c7=identity|sk_c6=identity)).
##
#cnf(i_0_670, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c4=sk_c3|sk_c5=identity)).
#
#cnf(i_0_685, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c7=identity|sk_c5=identity)).
#
#cnf(i_0_693, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c7=identity|sk_c5=identity)).
##
#cnf(i_0_776, negated_conjecture, (sk_c2=sk_c3|sk_c5=identity|sk_c7=identity)).
#
#cnf(i_0_705, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c4=sk_c3|sk_c6=identity)).
#
#cnf(i_0_730, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c7=identity|sk_c6=identity)).
#
#cnf(i_0_646, negated_conjecture, (multiply(sk_c6,inverse(sk_c2))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2|sk_c2!=sk_c6)).
#
#cnf(i_0_738, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c7=identity|sk_c6=identity)).
#
#cnf(i_0_814, negated_conjecture, (sk_c2=sk_c3|sk_c6=identity|sk_c7=identity)).
##########
#cnf(i_0_663, negated_conjecture, (multiply(X1,inverse(sk_c2))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|inverse(X1)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_433, negated_conjecture, (multiply(sk_c6,multiply(sk_c3,X1))=X1|sk_c5=identity)).
####
#cnf(i_0_437, negated_conjecture, (multiply(sk_c3,multiply(sk_c6,X1))=X1|sk_c5=identity)).
#######
#cnf(i_0_872, negated_conjecture, (inverse(sk_c6)=sk_c4|sk_c2=sk_c3|sk_c5=identity)).
##
#cnf(i_0_874, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c4=sk_c3|sk_c5=identity)).
#
#cnf(i_0_341, negated_conjecture, (multiply(sk_c6,multiply(inverse(X2),inverse(X1)))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,X2)!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
###
#cnf(i_0_885, negated_conjecture, (inverse(sk_c4)=sk_c6|sk_c2=sk_c3|sk_c5=identity)).
##
#cnf(i_0_901, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c4=sk_c3|sk_c5=identity)).
#########
#cnf(i_0_483, negated_conjecture, (multiply(sk_c6,multiply(sk_c3,X1))=X1|sk_c6=identity)).
####
#cnf(i_0_487, negated_conjecture, (multiply(sk_c3,multiply(sk_c6,X1))=X1|sk_c6=identity)).
######
#cnf(i_0_939, negated_conjecture, (multiply(sk_c6,multiply(inverse(X1),X2))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(inverse(X2),X1)!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
#
#cnf(i_0_1023, negated_conjecture, (inverse(sk_c6)=sk_c4|sk_c2=sk_c3|sk_c6=identity)).
##
#cnf(i_0_1025, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c4=sk_c3|sk_c6=identity)).
####
#cnf(i_0_1080, negated_conjecture, (inverse(sk_c4)=sk_c6|sk_c2=sk_c3|sk_c6=identity)).
##
#cnf(i_0_1099, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c4=sk_c3|sk_c6=identity)).
##########
#cnf(i_0_607, negated_conjecture, (multiply(inverse(sk_c2),X1)=multiply(sk_c6,X1)|sk_c4!=sk_c2)).
###########
#cnf(i_0_611, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c2,X1)|sk_c4!=sk_c2)).
#############
#cnf(i_0_852, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c3,X1)|sk_c5=identity)).
###########
#cnf(i_0_858, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c6,X1)|sk_c5=identity)).
#############
#cnf(i_0_1003, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c3,X1)|sk_c6=identity)).
#########
#cnf(i_0_1068, negated_conjecture, (multiply(sk_c6,multiply(X1,X2))!=sk_c6|multiply(inverse(X2),inverse(X1))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
##
#cnf(i_0_1009, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c6,X1)|sk_c6=identity)).
#############
#cnf(i_0_319, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c2))=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_320, negated_conjecture, (multiply(inverse(sk_c2),inverse(sk_c6))=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_33, negated_conjecture, (multiply(sk_c7,sk_c6)=sk_c5|multiply(sk_c6,sk_c7)=sk_c5)).
#
#cnf(i_0_34, negated_conjecture, (multiply(sk_c7,sk_c6)=sk_c5|multiply(sk_c3,sk_c6)=sk_c7)).
##
#cnf(i_0_37, negated_conjecture, (multiply(sk_c7,sk_c6)=sk_c5|multiply(sk_c4,sk_c5)=sk_c6)).
#
#cnf(i_0_38, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c5|multiply(sk_c1,sk_c7)=sk_c6)).
#
#cnf(i_0_48, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c5|multiply(sk_c2,sk_c6)=sk_c5)).
##
#cnf(i_0_39, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7|multiply(sk_c1,sk_c7)=sk_c6)).
#
#cnf(i_0_49, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7|multiply(sk_c2,sk_c6)=sk_c5)).
#
#cnf(i_0_42, negated_conjecture, (multiply(sk_c4,sk_c5)=sk_c6|multiply(sk_c1,sk_c7)=sk_c6)).
##
#cnf(i_0_52, negated_conjecture, (multiply(sk_c4,sk_c5)=sk_c6|multiply(sk_c2,sk_c6)=sk_c5)).
#
#cnf(i_0_329, negated_conjecture, (multiply(inverse(sk_c7),sk_c5)=sk_c6|inverse(sk_c3)=sk_c6)).
####
#cnf(i_0_1520, negated_conjecture, (multiply(sk_c1,sk_c5)=sk_c6|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_337, negated_conjecture, (multiply(inverse(sk_c1),sk_c6)=sk_c7|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_1378, negated_conjecture, (multiply(sk_c6,sk_c2)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_1545, negated_conjecture, (multiply(sk_c7,sk_c6)=sk_c7|inverse(sk_c3)=sk_c6)).
###
#cnf(i_0_1561, negated_conjecture, (inverse(sk_c3)=sk_c6|sk_c5=sk_c7)).
#
#cnf(i_0_1580, negated_conjecture, (inverse(sk_c6)=sk_c3|sk_c5=sk_c7)).
###
#cnf(i_0_1601, negated_conjecture, (sk_c2=sk_c3|sk_c5=sk_c7|sk_c4!=sk_c2)).
#
#cnf(i_0_1619, negated_conjecture, (sk_c7=identity|sk_c2=sk_c3|sk_c4!=sk_c2)).
###
#cnf(i_0_1577, negated_conjecture, (multiply(sk_c6,sk_c3)=identity|sk_c5=sk_c7)).
####
#cnf(i_0_1581, negated_conjecture, (multiply(sk_c3,sk_c6)=identity|sk_c5=sk_c7)).
#######
#cnf(i_0_1600, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c4=sk_c3|sk_c5=sk_c7)).
#
#cnf(i_0_1662, negated_conjecture, (sk_c2=sk_c3|sk_c5=sk_c7|sk_c4=sk_c3)).
#
#cnf(i_0_1672, negated_conjecture, (sk_c7=identity|sk_c2=sk_c3|sk_c4=sk_c3)).
#
#cnf(i_0_1271, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c2))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_1639, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c7=identity|sk_c5=sk_c7)).
#
#cnf(i_0_1640, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c7=identity|sk_c5=sk_c7)).
##
#cnf(i_0_1644, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c4=sk_c3|sk_c5=sk_c7)).
#
#cnf(i_0_1723, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c7=identity|sk_c5=sk_c7)).
#
#cnf(i_0_1735, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c7=identity|sk_c5=sk_c7)).
##
#cnf(i_0_1806, negated_conjecture, (sk_c2=sk_c3|sk_c5=sk_c7|sk_c7=identity)).
#
#cnf(i_0_1817, negated_conjecture, (sk_c7=identity|sk_c2=sk_c3)).
#
#cnf(i_0_1819, negated_conjecture, (inverse(sk_c3)=sk_c6|sk_c7=identity)).
##
#cnf(i_0_1853, negated_conjecture, (inverse(sk_c6)=sk_c3|sk_c7=identity)).
####
#cnf(i_0_1850, negated_conjecture, (multiply(sk_c6,sk_c3)=identity|sk_c7=identity)).
###
#cnf(i_0_1290, negated_conjecture, (multiply(inverse(sk_c2),inverse(sk_c6))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_1854, negated_conjecture, (multiply(sk_c3,sk_c6)=identity|sk_c7=identity)).
##
#cnf(i_0_1930, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c7=identity)).
##
#cnf(i_0_1931, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c7=identity)).
#
#cnf(i_0_1935, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c7=identity)).
#
#cnf(i_0_1947, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c7=identity)).
##########
#cnf(i_0_1932, negated_conjecture, (multiply(sk_c7,sk_c1)=identity|sk_c7=identity)).
##
#cnf(i_0_1936, negated_conjecture, (multiply(sk_c1,sk_c7)=identity|sk_c7=identity)).
###
#cnf(i_0_1944, negated_conjecture, (multiply(sk_c6,sk_c2)=identity|sk_c7=identity)).
##
#cnf(i_0_975, negated_conjecture, (multiply(sk_c6,multiply(X1,inverse(X2)))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X2,inverse(X1))!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
#
#cnf(i_0_1948, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|sk_c7=identity)).
#########################
#cnf(i_0_2011, negated_conjecture, (inverse(sk_c4)=sk_c6|sk_c6=identity|sk_c7=identity)).
###
#cnf(i_0_2092, negated_conjecture, (inverse(sk_c4)=sk_c6|sk_c5=identity|sk_c7=identity)).
#
#cnf(i_0_2100, negated_conjecture, (inverse(sk_c6)=sk_c4|sk_c6=identity|sk_c7=identity)).
#
#cnf(i_0_2142, negated_conjecture, (sk_c4=sk_c3|sk_c6=identity|sk_c7=identity)).
###
#cnf(i_0_2151, negated_conjecture, (sk_c4=sk_c2|sk_c7=identity|sk_c6=identity)).
#####
#cnf(i_0_2113, negated_conjecture, (inverse(sk_c6)=sk_c4|sk_c5=identity|sk_c7=identity)).
#
#cnf(i_0_391, negated_conjecture, (multiply(inverse(X2),multiply(inverse(X1),sk_c6))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,X2)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
#
#cnf(i_0_2210, negated_conjecture, (sk_c4=sk_c3|sk_c5=identity|sk_c7=identity)).
##
#cnf(i_0_2222, negated_conjecture, (sk_c4=sk_c2|sk_c7=identity|sk_c5=identity)).
##########
#cnf(i_0_1578, negated_conjecture, (multiply(sk_c6,multiply(sk_c3,X1))=X1|sk_c5=sk_c7)).
####
#cnf(i_0_1582, negated_conjecture, (multiply(sk_c3,multiply(sk_c6,X1))=X1|sk_c5=sk_c7)).
########
#cnf(i_0_2383, negated_conjecture, (inverse(sk_c6)=sk_c4|sk_c2=sk_c3|sk_c5=sk_c7)).
##
#cnf(i_0_2386, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c4=sk_c3|sk_c5=sk_c7)).
#
#cnf(i_0_2268, negated_conjecture, (multiply(inverse(X1),multiply(X2,sk_c6))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(inverse(X2),X1)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
###
#cnf(i_0_2399, negated_conjecture, (inverse(sk_c4)=sk_c6|sk_c2=sk_c3|sk_c5=sk_c7)).
##
#cnf(i_0_2430, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c4=sk_c3|sk_c5=sk_c7)).
#########
#cnf(i_0_1851, negated_conjecture, (multiply(sk_c6,multiply(sk_c3,X1))=X1|sk_c7=identity)).
####
#cnf(i_0_1855, negated_conjecture, (multiply(sk_c3,multiply(sk_c6,X1))=X1|sk_c7=identity)).
##################
#cnf(i_0_2324, negated_conjecture, (multiply(X1,multiply(inverse(X2),sk_c6))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X2,inverse(X1))!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
#
#cnf(i_0_1933, negated_conjecture, (multiply(sk_c7,multiply(sk_c1,X1))=X1|sk_c7=identity)).
####
#cnf(i_0_1937, negated_conjecture, (multiply(sk_c1,multiply(sk_c7,X1))=X1|sk_c7=identity)).
####
#cnf(i_0_1945, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,X1))=X1|sk_c7=identity)).
####
#cnf(i_0_2719, negated_conjecture, (multiply(sk_c2,X1)=multiply(sk_c3,X1)|sk_c7=identity)).
####
#cnf(i_0_2742, negated_conjecture, (multiply(sk_c3,inverse(sk_c2))=identity|sk_c7=identity)).
##
#cnf(i_0_2765, negated_conjecture, (inverse(sk_c2)=inverse(sk_c3)|sk_c7=identity)).
########
#cnf(i_0_2766, negated_conjecture, (multiply(sk_c2,inverse(sk_c3))=identity|sk_c7=identity)).
#######
#cnf(i_0_2781, negated_conjecture, (multiply(inverse(sk_c3),sk_c2)=identity|sk_c7=identity)).
######
#cnf(i_0_69, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c6)!=sk_c7|multiply(X2,sk_c6)!=sk_c5|inverse(X1)!=sk_c6|inverse(X2)!=sk_c6|sk_c5!=sk_c6|sk_c6!=identity)).
##
#cnf(i_0_2830, negated_conjecture, (multiply(inverse(sk_c2),sk_c3)=identity|sk_c7=identity)).
######
#cnf(i_0_2885, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c6)!=sk_c5|inverse(X1)!=sk_c6|sk_c7!=identity|sk_c5!=sk_c6|sk_c6!=identity)).
#####
#cnf(i_0_1949, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,X1))=X1|sk_c7=identity)).
###
#cnf(i_0_2931, negated_conjecture, (multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c7!=identity|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
############
#cnf(i_0_2927, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c2!=sk_c6|sk_c7!=identity|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
####
#cnf(i_0_2935, negated_conjecture, (multiply(X1,multiply(X2,sk_c6))!=sk_c5|multiply(inverse(X2),inverse(X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c7!=identity|sk_c5!=sk_c6|sk_c6!=identity)).
####
#cnf(i_0_3038, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c6!=identity|sk_c7!=identity)).
####
#cnf(i_0_2862, negated_conjecture, (multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c6)!=sk_c7|inverse(X1)!=sk_c6|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_2357, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c3,X1)|sk_c5=sk_c7)).
#####
#cnf(i_0_3054, negated_conjecture, (multiply(sk_c2,sk_c6)!=sk_c7|multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_2366, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c6,X1)|sk_c5=sk_c7)).
######
#cnf(i_0_3050, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c7|multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c2!=sk_c6|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_2568, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c3,X1)|sk_c7=identity)).
#
#cnf(i_0_2883, negated_conjecture, (multiply(sk_c2,sk_c6)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c6)!=sk_c5|inverse(X1)!=sk_c6|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
###########
#cnf(i_0_2577, negated_conjecture, (multiply(inverse(sk_c3),X1)=multiply(sk_c6,X1)|sk_c7=identity)).
#########
#cnf(i_0_3230, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c5|multiply(sk_c2,sk_c6)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c2!=sk_c6|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_2679, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c1,X1)|sk_c7=identity)).
##########
#cnf(i_0_1035, negated_conjecture, (multiply(sk_c6,multiply(sk_c6,X1))!=sk_c6|multiply(inverse(X1),sk_c2)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_2684, negated_conjecture, (multiply(inverse(sk_c1),X1)=multiply(sk_c7,X1)|sk_c7=identity)).
#############
#cnf(i_0_2711, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c2,X1)|sk_c7=identity)).
######
#cnf(i_0_1054, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,X1))!=sk_c6|multiply(inverse(X1),sk_c6)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_2724, negated_conjecture, (multiply(inverse(sk_c2),X1)=multiply(sk_c6,X1)|sk_c7=identity)).
#################
#cnf(i_0_1132, negated_conjecture, (multiply(X1,multiply(sk_c6,sk_c6))!=sk_c7|multiply(sk_c2,inverse(X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_343, negated_conjecture, (multiply(inverse(sk_c2),sk_c5)=sk_c6|inverse(sk_c3)=sk_c6)).
####
#cnf(i_0_3538, negated_conjecture, (multiply(sk_c6,sk_c5)=sk_c6|inverse(sk_c3)=sk_c6)).
##
#cnf(i_0_350, negated_conjecture, (multiply(inverse(sk_c7),sk_c5)=sk_c6|inverse(sk_c4)=sk_c6)).
####
#cnf(i_0_359, negated_conjecture, (multiply(inverse(sk_c1),sk_c6)=sk_c7|inverse(sk_c4)=sk_c6)).
####
#cnf(i_0_365, negated_conjecture, (multiply(inverse(sk_c2),sk_c5)=sk_c6|inverse(sk_c4)=sk_c6)).
####
#cnf(i_0_372, negated_conjecture, (multiply(inverse(sk_c6),sk_c5)=sk_c7|inverse(sk_c1)=sk_c7)).
#
#cnf(i_0_1139, negated_conjecture, (multiply(sk_c6,multiply(X1,sk_c6))!=sk_c7|multiply(inverse(X1),sk_c2)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_375, negated_conjecture, (multiply(inverse(sk_c3),sk_c7)=sk_c6|inverse(sk_c1)=sk_c7)).
####
#cnf(i_0_395, negated_conjecture, (multiply(inverse(sk_c4),sk_c6)=sk_c5|inverse(sk_c1)=sk_c7)).
####
#cnf(i_0_3759, negated_conjecture, (multiply(sk_c6,sk_c6)=sk_c5|inverse(sk_c1)=sk_c7)).
##
#cnf(i_0_399, negated_conjecture, (multiply(inverse(sk_c6),sk_c5)=sk_c7|inverse(sk_c2)=sk_c6)).
####
#cnf(i_0_402, negated_conjecture, (multiply(inverse(sk_c3),sk_c7)=sk_c6|inverse(sk_c2)=sk_c6)).
####
#cnf(i_0_405, negated_conjecture, (multiply(inverse(sk_c4),sk_c6)=sk_c5|inverse(sk_c2)=sk_c6)).
####
#cnf(i_0_3858, negated_conjecture, (multiply(sk_c6,sk_c6)=sk_c5|inverse(sk_c2)=sk_c6)).
#
#cnf(i_0_1527, negated_conjecture, (multiply(inverse(sk_c1),sk_c6)=sk_c5|inverse(sk_c3)=sk_c6)).
##
#cnf(i_0_1157, negated_conjecture, (multiply(X1,multiply(sk_c2,sk_c6))!=sk_c7|multiply(sk_c6,inverse(X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_3769, negated_conjecture, (multiply(inverse(sk_c6),sk_c5)=sk_c6|inverse(sk_c1)=sk_c7)).
#
#cnf(i_0_3953, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c7=sk_c6)).
#
#cnf(i_0_3982, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c7=sk_c6)).
###
#cnf(i_0_3979, negated_conjecture, (multiply(sk_c7,sk_c1)=identity|sk_c7=sk_c6)).
###
#cnf(i_0_3983, negated_conjecture, (multiply(sk_c1,sk_c7)=identity|sk_c7=sk_c6)).
#########
#cnf(i_0_4027, negated_conjecture, (inverse(sk_c4)=sk_c6|sk_c6=identity|sk_c7=sk_c6)).
#
#cnf(i_0_4033, negated_conjecture, (inverse(sk_c6)=sk_c4|sk_c6=identity|sk_c7=sk_c6)).
#
#cnf(i_0_1164, negated_conjecture, (multiply(sk_c2,multiply(X1,sk_c6))!=sk_c7|multiply(inverse(X1),sk_c6)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_4076, negated_conjecture, (sk_c4=sk_c3|sk_c6=identity|sk_c7=sk_c6)).
####
#cnf(i_0_3980, negated_conjecture, (multiply(sk_c7,multiply(sk_c1,X1))=X1|sk_c7=sk_c6)).
####
#cnf(i_0_3984, negated_conjecture, (multiply(sk_c1,multiply(sk_c7,X1))=X1|sk_c7=sk_c6)).
#########
#cnf(i_0_4168, negated_conjecture, (multiply(inverse(sk_c7),X1)=multiply(sk_c1,X1)|sk_c7=sk_c6)).
###########
#cnf(i_0_4173, negated_conjecture, (multiply(inverse(sk_c1),X1)=multiply(sk_c7,X1)|sk_c7=sk_c6)).
#######
#cnf(i_0_2024, negated_conjecture, (multiply(sk_c6,multiply(X1,sk_c6))!=sk_c6|multiply(sk_c2,inverse(X1))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_3867, negated_conjecture, (multiply(inverse(sk_c6),sk_c5)=sk_c6|inverse(sk_c2)=sk_c6)).
#
#cnf(i_0_4350, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c7=sk_c6)).
##
#cnf(i_0_4379, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c7=sk_c6)).
##
#cnf(i_0_4376, negated_conjecture, (multiply(sk_c6,sk_c2)=identity|sk_c7=sk_c6)).
###
#cnf(i_0_4380, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|sk_c7=sk_c6)).
#####
#cnf(i_0_4428, negated_conjecture, (sk_c2=sk_c3|sk_c5=identity|sk_c7=sk_c6)).
##
#cnf(i_0_4430, negated_conjecture, (sk_c2=sk_c3|sk_c6=identity|sk_c7=sk_c6)).
#
#cnf(i_0_4438, negated_conjecture, (sk_c2=sk_c3|sk_c5=sk_c7|sk_c7=sk_c6)).
##
#cnf(i_0_2046, negated_conjecture, (multiply(sk_c6,multiply(X1,sk_c2))!=sk_c6|multiply(sk_c6,inverse(X1))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_4446, negated_conjecture, (sk_c4=sk_c2|sk_c6=identity|sk_c7=sk_c6)).
####################
#cnf(i_0_4479, negated_conjecture, (inverse(sk_c4)=sk_c6|sk_c5=identity|sk_c7=sk_c6)).
######
#cnf(i_0_4647, negated_conjecture, (inverse(sk_c6)=sk_c4|sk_c5=identity|sk_c7=sk_c6)).
##
#cnf(i_0_4689, negated_conjecture, (sk_c4=sk_c3|sk_c5=identity|sk_c7=sk_c6)).
#
#cnf(i_0_4707, negated_conjecture, (sk_c4=sk_c2|sk_c7=sk_c6|sk_c5=identity)).
###
#cnf(i_0_4377, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,X1))=X1|sk_c7=sk_c6)).
#####
#cnf(i_0_4381, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,X1))=X1|sk_c7=sk_c6)).
##################
#cnf(i_0_3502, negated_conjecture, (multiply(sk_c6,multiply(sk_c6,sk_c6))!=sk_c7|multiply(sk_c2,sk_c2)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_4747, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c2,X1)|sk_c7=sk_c6)).
###
#cnf(i_0_3510, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,sk_c6))!=sk_c7|multiply(sk_c2,sk_c6)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_4767, negated_conjecture, (multiply(inverse(sk_c2),X1)=multiply(sk_c6,X1)|sk_c7=sk_c6)).
#####
#cnf(i_0_3704, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,sk_c6))!=sk_c7|multiply(sk_c6,sk_c2)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#########
#cnf(i_0_585, negated_conjecture, (multiply(sk_c6,sk_c4)=identity|sk_c5=identity|sk_c4!=sk_c2)).
####
#cnf(i_0_589, negated_conjecture, (multiply(sk_c4,sk_c6)=identity|sk_c5=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_3930, negated_conjecture, (multiply(sk_c2,multiply(sk_c2,sk_c6))!=sk_c7|multiply(sk_c6,sk_c6)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_1466, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c5|sk_c5=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_1486, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7|sk_c5=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_5030, negated_conjecture, (sk_c7=identity|sk_c5=identity|sk_c4!=sk_c2)).
###########
#cnf(i_0_4305, negated_conjecture, (multiply(sk_c6,multiply(sk_c6,sk_c6))!=sk_c6|multiply(sk_c2,sk_c2)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_5009, negated_conjecture, (multiply(sk_c2,sk_c5)=sk_c7|sk_c5=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_5010, negated_conjecture, (multiply(sk_c3,sk_c5)=sk_c7|sk_c5=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_5026, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c6|sk_c5=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_5113, negated_conjecture, (sk_c5=sk_c6|sk_c5=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_5114, negated_conjecture, (sk_c5=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_4599, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,sk_c2))!=sk_c6|multiply(sk_c6,sk_c6)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_5103, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c7|sk_c5=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_908, negated_conjecture, (multiply(sk_c6,multiply(inverse(X1),sk_c6))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c2,X1)!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_454, negated_conjecture, (multiply(inverse(sk_c3),inverse(sk_c6))=identity|sk_c5=identity)).
#####
#cnf(i_0_457, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c3))=identity|sk_c5=identity)).
#####
#cnf(i_0_499, negated_conjecture, (multiply(inverse(sk_c3),inverse(sk_c6))=identity|sk_c6=identity)).
######
#cnf(i_0_502, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c3))=identity|sk_c6=identity)).
#####
#cnf(i_0_1627, negated_conjecture, (multiply(inverse(sk_c3),inverse(sk_c6))=identity|sk_c5=sk_c7)).
#####
#cnf(i_0_1634, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c3))=identity|sk_c5=sk_c7)).
#####
#cnf(i_0_926, negated_conjecture, (multiply(sk_c6,multiply(inverse(X1),sk_c2))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c6,X1)!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_1890, negated_conjecture, (multiply(inverse(sk_c3),inverse(sk_c6))=identity|sk_c7=identity)).
######
#cnf(i_0_1925, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c3))=identity|sk_c7=identity)).
########
#cnf(i_0_2004, negated_conjecture, (multiply(inverse(sk_c1),inverse(sk_c7))=identity|sk_c7=identity)).
######
#cnf(i_0_2008, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c1))=identity|sk_c7=identity)).
#####
#cnf(i_0_2016, negated_conjecture, (multiply(inverse(sk_c2),inverse(sk_c6))=identity|sk_c7=identity)).
#######
#cnf(i_0_2084, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c2))=identity|sk_c7=identity)).
###
#cnf(i_0_944, negated_conjecture, (multiply(sk_c6,multiply(sk_c6,inverse(X1)))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c2)!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_4018, negated_conjecture, (multiply(inverse(sk_c1),inverse(sk_c7))=identity|sk_c7=sk_c6)).
#####
#cnf(i_0_4023, negated_conjecture, (multiply(inverse(sk_c7),inverse(sk_c1))=identity|sk_c7=sk_c6)).
#####
#cnf(i_0_4449, negated_conjecture, (multiply(inverse(sk_c2),inverse(sk_c6))=identity|sk_c7=sk_c6)).
######
#cnf(i_0_4461, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c2))=identity|sk_c7=sk_c6)).
###
#cnf(i_0_962, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,inverse(X1)))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c6)!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#############
#cnf(i_0_551, negated_conjecture, (inverse(sk_c4)=sk_c6|inverse(sk_c6)=sk_c4|sk_c6=identity)).
#
#cnf(i_0_6103, negated_conjecture, (inverse(sk_c6)=sk_c4|sk_c6=identity)).
#
#cnf(i_0_6160, negated_conjecture, (sk_c4=sk_c3|sk_c6=identity)).
##
#cnf(i_0_6130, negated_conjecture, (inverse(sk_c4)=sk_c6|sk_c6=identity)).
####
#cnf(i_0_6127, negated_conjecture, (multiply(sk_c4,sk_c6)=identity|sk_c6=identity)).
##
#cnf(i_0_6131, negated_conjecture, (multiply(sk_c6,sk_c4)=identity|sk_c6=identity)).
#####################
#cnf(i_0_2225, negated_conjecture, (multiply(inverse(X1),multiply(sk_c6,sk_c6))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c2,X1)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_6155, negated_conjecture, (multiply(sk_c4,X1)=multiply(sk_c3,X1)|sk_c6=identity)).
####
#cnf(i_0_6158, negated_conjecture, (multiply(inverse(sk_c3),sk_c4)=identity|sk_c6=identity)).
######
#cnf(i_0_6159, negated_conjecture, (multiply(sk_c4,inverse(sk_c3))=identity|sk_c6=identity)).
##
#cnf(i_0_6406, negated_conjecture, (inverse(sk_c4)=inverse(sk_c3)|sk_c6=identity)).
#####
#cnf(i_0_6364, negated_conjecture, (multiply(sk_c3,inverse(sk_c4))=identity|sk_c6=identity)).
#####
#cnf(i_0_6390, negated_conjecture, (multiply(inverse(sk_c4),sk_c3)=identity|sk_c6=identity)).
#########
#cnf(i_0_2249, negated_conjecture, (multiply(inverse(X1),multiply(sk_c2,sk_c6))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c6,X1)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
###########
#cnf(i_0_6128, negated_conjecture, (multiply(sk_c4,multiply(sk_c6,X1))=X1|sk_c6=identity)).
################
#cnf(i_0_6132, negated_conjecture, (multiply(sk_c6,multiply(sk_c4,X1))=X1|sk_c6=identity)).
##################
#cnf(i_0_6581, negated_conjecture, (multiply(inverse(sk_c4),X1)=multiply(sk_c6,X1)|sk_c6=identity)).
###
#cnf(i_0_2281, negated_conjecture, (multiply(sk_c6,multiply(inverse(X1),sk_c6))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c2)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#########
#cnf(i_0_6588, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c4,X1)|sk_c6=identity)).
#####################
#cnf(i_0_6451, negated_conjecture, (inverse(sk_c2)=inverse(sk_c3)|sk_c6=identity|sk_c7=sk_c6)).
#####
#cnf(i_0_6241, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c4))=identity|sk_c6=identity)).
#
#cnf(i_0_2305, negated_conjecture, (multiply(sk_c2,multiply(inverse(X1),sk_c6))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c6)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_6260, negated_conjecture, (multiply(inverse(sk_c4),inverse(sk_c6))=identity|sk_c6=identity)).
##############
#cnf(i_0_565, negated_conjecture, (inverse(sk_c4)=sk_c6|inverse(sk_c6)=sk_c4|sk_c5=identity)).
#
#cnf(i_0_7084, negated_conjecture, (inverse(sk_c6)=sk_c4|sk_c5=identity)).
#
#cnf(i_0_7145, negated_conjecture, (sk_c4=sk_c3|sk_c5=identity)).
##
#cnf(i_0_7118, negated_conjecture, (inverse(sk_c4)=sk_c6|sk_c5=identity)).
##
#cnf(i_0_7115, negated_conjecture, (multiply(sk_c4,sk_c6)=identity|sk_c5=identity)).
###
#cnf(i_0_7119, negated_conjecture, (multiply(sk_c6,sk_c4)=identity|sk_c5=identity)).
#####################
#cnf(i_0_7154, negated_conjecture, (multiply(sk_c4,X1)=multiply(sk_c3,X1)|sk_c5=identity)).
####
#cnf(i_0_7148, negated_conjecture, (multiply(inverse(sk_c3),sk_c4)=identity|sk_c5=identity)).
#####
#cnf(i_0_3331, negated_conjecture, (multiply(inverse(X1),multiply(sk_c6,sk_c2))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c6,X1)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_7149, negated_conjecture, (multiply(sk_c4,inverse(sk_c3))=identity|sk_c5=identity)).
#
#cnf(i_0_7406, negated_conjecture, (inverse(sk_c4)=inverse(sk_c3)|sk_c5=identity)).
######
#cnf(i_0_7287, negated_conjecture, (multiply(sk_c3,inverse(sk_c4))=identity|sk_c5=identity)).
#####
#cnf(i_0_7316, negated_conjecture, (multiply(inverse(sk_c4),sk_c3)=identity|sk_c5=identity)).
######
#cnf(i_0_7380, negated_conjecture, (multiply(X1,multiply(sk_c6,sk_c2))!=sk_c7|multiply(sk_c6,inverse(X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_7178, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c4=sk_c6|sk_c4=sk_c3)).
##
#cnf(i_0_7179, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c4=sk_c6|sk_c4=sk_c3)).
#
#cnf(i_0_7190, negated_conjecture, (inverse(sk_c3)=sk_c6|sk_c1=sk_c6|sk_c4=sk_c3)).
###########
#cnf(i_0_7543, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c4=sk_c6|sk_c4=sk_c3)).
#
#cnf(i_0_7575, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c4=sk_c6|sk_c4=sk_c3)).
#
#cnf(i_0_7621, negated_conjecture, (inverse(sk_c6)=sk_c3|sk_c1=sk_c6|sk_c4=sk_c3)).
#
#cnf(i_0_7340, negated_conjecture, (multiply(sk_c6,multiply(sk_c6,sk_c2))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c6,sk_c2)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_7776, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c4=sk_c3|sk_c1=sk_c6)).
#
#cnf(i_0_7364, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,sk_c2))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c6,sk_c6)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_7878, negated_conjecture, (sk_c2=sk_c3|sk_c1=sk_c6|sk_c4=sk_c3)).
#
#cnf(i_0_7821, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c4=sk_c3|sk_c1=sk_c6)).
###
#cnf(i_0_7116, negated_conjecture, (multiply(sk_c4,multiply(sk_c6,X1))=X1|sk_c5=identity)).
################
#cnf(i_0_7120, negated_conjecture, (multiply(sk_c6,multiply(sk_c4,X1))=X1|sk_c5=identity)).
###########
#cnf(i_0_2858, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c6)!=sk_c7|inverse(X1)!=sk_c6|sk_c2!=sk_c6|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_7978, negated_conjecture, (multiply(inverse(sk_c4),X1)=multiply(sk_c6,X1)|sk_c5=identity)).
###
#cnf(i_0_8079, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c7|multiply(sk_c6,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c2!=sk_c6|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
#########
#cnf(i_0_7987, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c4,X1)|sk_c5=identity)).
###
#cnf(i_0_2879, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c6)!=sk_c5|inverse(X1)!=sk_c6|sk_c2!=sk_c6|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
###############
#cnf(i_0_8008, negated_conjecture, (multiply(sk_c4,sk_c5)=sk_c7|sk_c5=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_8009, negated_conjecture, (multiply(sk_c4,sk_c6)=sk_c7|sk_c5=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_849, negated_conjecture, (multiply(X1,multiply(X2,inverse(sk_c2)))!=sk_c7|multiply(inverse(X2),inverse(X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_7236, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c4))=identity|sk_c5=identity)).
########
#cnf(i_0_7257, negated_conjecture, (multiply(inverse(sk_c4),inverse(sk_c6))=identity|sk_c5=identity)).
################################
#cnf(i_0_7184, negated_conjecture, (inverse(sk_c7)=sk_c6|inverse(sk_c3)=sk_c6|sk_c4=sk_c3)).
####
#cnf(i_0_7186, negated_conjecture, (inverse(sk_c7)=sk_c6|inverse(sk_c4)=sk_c6|sk_c4=sk_c3)).
#
#cnf(i_0_8388, negated_conjecture, (multiply(X1,multiply(sk_c6,inverse(sk_c2)))!=sk_c7|multiply(sk_c2,inverse(X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_7188, negated_conjecture, (inverse(sk_c6)=sk_c7|inverse(sk_c1)=sk_c7|sk_c4=sk_c3)).
#
#cnf(i_0_7189, negated_conjecture, (inverse(sk_c6)=sk_c7|inverse(sk_c2)=sk_c6|sk_c4=sk_c3)).
###
#cnf(i_0_7194, negated_conjecture, (inverse(sk_c6)=sk_c6|inverse(sk_c1)=sk_c7|sk_c4=sk_c3)).
#
#cnf(i_0_7195, negated_conjecture, (inverse(sk_c6)=sk_c6|inverse(sk_c2)=sk_c6|sk_c4=sk_c3)).
#######
#cnf(i_0_8555, negated_conjecture, (inverse(sk_c6)=sk_c7|inverse(sk_c3)=sk_c6|sk_c4=sk_c3)).
#
#cnf(i_0_8867, negated_conjecture, (inverse(sk_c3)=sk_c6|sk_c2=sk_c7|sk_c4=sk_c3)).
##
#cnf(i_0_8919, negated_conjecture, (inverse(sk_c6)=sk_c3|sk_c2=sk_c7|sk_c4=sk_c3)).
#####
#cnf(i_0_8983, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c4=sk_c3|sk_c2=sk_c7)).
#
#cnf(i_0_9088, negated_conjecture, (sk_c2=sk_c3|sk_c2=sk_c7|sk_c4=sk_c3)).
#
#cnf(i_0_8442, negated_conjecture, (multiply(sk_c6,multiply(X1,inverse(sk_c2)))!=sk_c7|multiply(inverse(X1),sk_c2)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_9026, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c4=sk_c3|sk_c2=sk_c7)).
####
#cnf(i_0_8594, negated_conjecture, (inverse(sk_c6)=sk_c4|inverse(sk_c7)=sk_c6|sk_c4=sk_c3)).
#
#cnf(i_0_9268, negated_conjecture, (inverse(sk_c6)=sk_c4|sk_c1=sk_c6|sk_c4=sk_c3)).
#
#cnf(i_0_9348, negated_conjecture, (sk_c4=sk_c3|sk_c1=sk_c6)).
####
#cnf(i_0_8685, negated_conjecture, (inverse(sk_c7)=sk_c1|inverse(sk_c6)=sk_c7|sk_c4=sk_c3)).
##
#cnf(i_0_8717, negated_conjecture, (inverse(sk_c6)=sk_c2|inverse(sk_c6)=sk_c7|sk_c4=sk_c3)).
#
#cnf(i_0_9432, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c4=sk_c3|sk_c2!=sk_c7)).
#
#cnf(i_0_9501, negated_conjecture, (inverse(sk_c7)=sk_c6|sk_c4=sk_c3|sk_c2!=sk_c7)).
###
#cnf(i_0_9498, negated_conjecture, (multiply(sk_c7,sk_c6)=identity|sk_c4=sk_c3|sk_c2!=sk_c7)).
###
#cnf(i_0_9502, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|sk_c4=sk_c3|sk_c2!=sk_c7)).
######
#cnf(i_0_8765, negated_conjecture, (inverse(sk_c7)=sk_c1|inverse(sk_c6)=sk_c6|sk_c4=sk_c3)).
##
#cnf(i_0_8797, negated_conjecture, (inverse(sk_c6)=sk_c2|inverse(sk_c6)=sk_c6|sk_c4=sk_c3)).
#
#cnf(i_0_9700, negated_conjecture, (inverse(sk_c6)=sk_c6|sk_c4=sk_c3|sk_c2!=sk_c6)).
####
#cnf(i_0_9767, negated_conjecture, (multiply(sk_c6,sk_c6)=identity|sk_c4=sk_c3|sk_c2!=sk_c6)).
##
#cnf(i_0_8454, negated_conjecture, (multiply(sk_c2,multiply(X1,inverse(sk_c2)))!=sk_c7|multiply(inverse(X1),sk_c6)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_9245, negated_conjecture, (inverse(sk_c6)=sk_c7|inverse(sk_c6)=sk_c4|sk_c4=sk_c3)).
#
#cnf(i_0_9925, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c4=sk_c3|sk_c4!=sk_c7)).
#
#cnf(i_0_10074, negated_conjecture, (sk_c2=sk_c7|sk_c4=sk_c3|sk_c4!=sk_c7)).
##
#cnf(i_0_10002, negated_conjecture, (inverse(sk_c7)=sk_c6|sk_c4=sk_c3|sk_c4!=sk_c7)).
#####
#cnf(i_0_9969, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c4=sk_c3|sk_c5=sk_c7)).
#
#cnf(i_0_10167, negated_conjecture, (sk_c7=sk_c3|sk_c5=sk_c7|sk_c4=sk_c3)).
##
#cnf(i_0_10169, negated_conjecture, (sk_c2=sk_c7|sk_c4=sk_c3|sk_c5=sk_c7)).
##
#cnf(i_0_10226, negated_conjecture, (sk_c7=identity|sk_c4=sk_c3|sk_c7=sk_c3)).
##
#cnf(i_0_10253, negated_conjecture, (sk_c7=identity|sk_c4=sk_c3|sk_c2=sk_c7)).
#
#cnf(i_0_9973, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c4=sk_c3|sk_c7=identity)).
#####
#cnf(i_0_9990, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c4=sk_c2|sk_c4=sk_c3)).
##
#cnf(i_0_10504, negated_conjecture, (sk_c2=sk_c7|sk_c4=sk_c3|sk_c4=sk_c2)).
#
#cnf(i_0_10507, negated_conjecture, (sk_c2=sk_c7|sk_c4=sk_c3|sk_c2!=sk_c3)).
###
#cnf(i_0_9995, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c4=sk_c3|sk_c2=sk_c7)).
#
#cnf(i_0_10687, negated_conjecture, (sk_c2=sk_c7|sk_c4=sk_c3)).
#
#cnf(i_0_10126, negated_conjecture, (inverse(sk_c7)=sk_c6|sk_c4=sk_c3|sk_c5=sk_c7)).
#
#cnf(i_0_8656, negated_conjecture, (multiply(sk_c6,multiply(sk_c6,inverse(sk_c2)))!=sk_c7|multiply(sk_c2,sk_c2)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_10354, negated_conjecture, (inverse(sk_c7)=sk_c6|sk_c4=sk_c3|sk_c7=identity)).
###
#cnf(i_0_10432, negated_conjecture, (inverse(sk_c7)=sk_c6|sk_c4=sk_c2|sk_c4=sk_c3)).
###
#cnf(i_0_8668, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,inverse(sk_c2)))!=sk_c7|multiply(sk_c2,sk_c6)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_10076, negated_conjecture, (sk_c7=sk_c6|sk_c4=sk_c3|sk_c2!=sk_c6|sk_c4!=sk_c7)).
#
#cnf(i_0_9999, negated_conjecture, (multiply(sk_c7,sk_c6)=identity|sk_c4=sk_c3|sk_c4!=sk_c7)).
###
#cnf(i_0_10003, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|sk_c4=sk_c3|sk_c4!=sk_c7)).
##########
#cnf(i_0_538, negated_conjecture, (multiply(inverse(sk_c2),multiply(inverse(X2),inverse(X1)))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,X2)!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_9929, negated_conjecture, (inverse(sk_c4)=sk_c6|inverse(sk_c6)=sk_c7|sk_c4=sk_c3)).
#########################################
#cnf(i_0_11022, negated_conjecture, (multiply(inverse(sk_c2),multiply(inverse(X1),X2))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(inverse(X2),X1)!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#########################################
#cnf(i_0_9811, negated_conjecture, (sk_c6=sk_c3|sk_c5=sk_c7|sk_c4=sk_c3|sk_c2!=sk_c6)).
#
#cnf(i_0_9815, negated_conjecture, (sk_c6=sk_c3|sk_c7=identity|sk_c4=sk_c3|sk_c2!=sk_c6)).
##############
#cnf(i_0_11279, negated_conjecture, (multiply(inverse(sk_c2),multiply(X1,X2))!=sk_c6|multiply(inverse(X2),inverse(X1))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_10200, negated_conjecture, (sk_c7=sk_c6|sk_c4=sk_c3|sk_c5=sk_c7|sk_c2!=sk_c6)).
#
#cnf(i_0_10428, negated_conjecture, (sk_c7=sk_c6|sk_c4=sk_c3|sk_c7=identity|sk_c2!=sk_c6)).
#
#cnf(i_0_10506, negated_conjecture, (sk_c7=sk_c6|sk_c4=sk_c3|sk_c4=sk_c2|sk_c2!=sk_c6)).
##
#cnf(i_0_11717, negated_conjecture, (sk_c7=sk_c6|sk_c4=sk_c3|sk_c2!=sk_c3|sk_c2!=sk_c6)).
########
#cnf(i_0_85, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,X1))=X1|inverse(sk_c3)=sk_c6)).
######
#cnf(i_0_86, negated_conjecture, (multiply(sk_c7,multiply(sk_c1,X1))=X1|inverse(sk_c3)=sk_c6)).
#######
#cnf(i_0_87, negated_conjecture, (multiply(sk_c6,multiply(sk_c4,X1))=X1|inverse(sk_c1)=sk_c7)).
########
#cnf(i_0_88, negated_conjecture, (multiply(sk_c6,multiply(sk_c4,X1))=X1|inverse(sk_c2)=sk_c6)).
########
#cnf(i_0_176, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,X1))=X1|inverse(sk_c3)=sk_c6)).
#################
#cnf(i_0_177, negated_conjecture, (multiply(sk_c1,multiply(sk_c7,X1))=X1|inverse(sk_c3)=sk_c6)).
#
#cnf(i_0_11209, negated_conjecture, (multiply(inverse(sk_c2),multiply(sk_c6,X1))!=sk_c6|multiply(inverse(X1),sk_c2)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_178, negated_conjecture, (multiply(sk_c4,multiply(sk_c6,X1))=X1|inverse(sk_c1)=sk_c7)).
#####
#cnf(i_0_12153, negated_conjecture, (multiply(sk_c4,sk_c5)=sk_c7|inverse(sk_c1)=sk_c7)).
############
#cnf(i_0_12204, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c4=sk_c7|sk_c4=sk_c3)).
##
#cnf(i_0_12221, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c4=sk_c7|sk_c4=sk_c3)).
##
#cnf(i_0_12178, negated_conjecture, (multiply(inverse(sk_c4),sk_c7)=sk_c5|inverse(sk_c1)=sk_c7)).
#########
#cnf(i_0_12066, negated_conjecture, (multiply(inverse(X1),multiply(sk_c6,sk_c2))!=sk_c7|multiply(inverse(sk_c2),X1)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_179, negated_conjecture, (multiply(sk_c4,multiply(sk_c6,X1))=X1|inverse(sk_c2)=sk_c6)).
##
#cnf(i_0_12334, negated_conjecture, (multiply(sk_c6,multiply(sk_c6,sk_c2))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c6!=identity|sk_c5!=sk_c6|sk_c5!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_12467, negated_conjecture, (multiply(sk_c4,sk_c5)=sk_c7|inverse(sk_c2)=sk_c6)).
############
#cnf(i_0_12524, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c4=sk_c7|sk_c4=sk_c3)).
##
#cnf(i_0_12541, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c4=sk_c7|sk_c4=sk_c3)).
##
#cnf(i_0_12498, negated_conjecture, (multiply(inverse(sk_c4),sk_c7)=sk_c5|inverse(sk_c2)=sk_c6)).
#########
#cnf(i_0_12406, negated_conjecture, (multiply(X1,multiply(sk_c6,sk_c2))!=sk_c7|multiply(inverse(sk_c2),inverse(X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_196, negated_conjecture, (multiply(sk_c7,multiply(sk_c1,X1))=X1|inverse(sk_c6)=sk_c4)).
####
#cnf(i_0_199, negated_conjecture, (multiply(sk_c1,multiply(sk_c7,X1))=X1|inverse(sk_c6)=sk_c4)).
####
#cnf(i_0_209, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,X1))=X1|inverse(sk_c6)=sk_c4)).
#####
#cnf(i_0_212, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,X1))=X1|inverse(sk_c6)=sk_c4)).
#########
#cnf(i_0_12368, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,sk_c2))!=sk_c7|multiply(inverse(sk_c2),sk_c6)!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
##########
#cnf(i_0_220, negated_conjecture, (multiply(sk_c4,multiply(sk_c6,X1))=X1|inverse(sk_c6)=sk_c2)).
##########
#cnf(i_0_11096, negated_conjecture, (multiply(inverse(sk_c2),multiply(X1,inverse(X2)))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X2,inverse(X1))!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#####
#cnf(i_0_223, negated_conjecture, (multiply(sk_c6,multiply(sk_c4,X1))=X1|inverse(sk_c6)=sk_c2)).
############################
#cnf(i_0_1641, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|sk_c4=sk_c3|sk_c5=sk_c7)).
##
#cnf(i_0_1645, negated_conjecture, (multiply(sk_c6,sk_c2)=identity|sk_c4=sk_c3|sk_c5=sk_c7)).
#######
#cnf(i_0_2012, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c5|sk_c6=identity|sk_c7=identity)).
#
#cnf(i_0_2013, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7|sk_c6=identity|sk_c7=identity)).
##
#cnf(i_0_13291, negated_conjecture, (sk_c7=identity|sk_c6=identity)).
#######
#cnf(i_0_12970, negated_conjecture, (multiply(inverse(sk_c2),multiply(X1,sk_c6))!=sk_c6|multiply(sk_c2,inverse(X1))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_13345, negated_conjecture, (sk_c1=identity|sk_c7=sk_c6|sk_c6=identity)).
#
#cnf(i_0_13536, negated_conjecture, (sk_c6=identity|sk_c1=identity)).
######################
#cnf(i_0_13298, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c5=sk_c6|sk_c6=identity)).
##
#cnf(i_0_13299, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c5=sk_c6|sk_c6=identity)).
#####
#cnf(i_0_13005, negated_conjecture, (multiply(inverse(sk_c2),multiply(X1,sk_c2))!=sk_c6|multiply(sk_c6,inverse(X1))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
###########
#cnf(i_0_13350, negated_conjecture, (inverse(sk_c1)=identity|sk_c7=sk_c6|sk_c6=identity)).
#########################
#cnf(i_0_13409, negated_conjecture, (multiply(inverse(sk_c2),multiply(sk_c6,sk_c6))!=sk_c6|multiply(sk_c2,sk_c2)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
######
#cnf(i_0_13569, negated_conjecture, (inverse(sk_c7)=identity|sk_c7=sk_c6|sk_c6=identity)).
########
#cnf(i_0_13579, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c5=sk_c6|sk_c6=identity)).
##
#cnf(i_0_10953, negated_conjecture, (multiply(inverse(sk_c2),multiply(inverse(X1),sk_c6))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c2,X1)!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_13617, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c5=sk_c6|sk_c6=identity)).
#
#cnf(i_0_14078, negated_conjecture, (sk_c2=sk_c3|sk_c6=identity|sk_c5=sk_c6)).
###
#cnf(i_0_14098, negated_conjecture, (sk_c4=sk_c2|sk_c6=identity|sk_c5=sk_c6)).
####################
#cnf(i_0_14131, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c6=identity|sk_c2=sk_c3)).
##
#cnf(i_0_14132, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c6=identity|sk_c2=sk_c3)).
##
#cnf(i_0_14165, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c6=identity|sk_c4=sk_c2)).
##
#cnf(i_0_14166, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c6=identity|sk_c4=sk_c2)).
##
#cnf(i_0_14184, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c6=identity|sk_c2=sk_c3)).
##
#cnf(i_0_14224, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c6=identity|sk_c2=sk_c3)).
##
#cnf(i_0_14483, negated_conjecture, (sk_c2=sk_c3|sk_c6=identity)).
####
#cnf(i_0_14291, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c6=identity|sk_c4=sk_c2)).
##
#cnf(i_0_14331, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c6=identity|sk_c4=sk_c2)).
#
#cnf(i_0_10985, negated_conjecture, (multiply(inverse(sk_c2),multiply(inverse(X1),sk_c2))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c6,X1)!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_14735, negated_conjecture, (sk_c4=sk_c2|sk_c6=identity)).
##
#cnf(i_0_14886, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c6=identity)).
###
#cnf(i_0_14955, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c6=identity)).
#######
#cnf(i_0_14920, negated_conjecture, (inverse(sk_c2)=inverse(sk_c3)|sk_c6=identity)).
#####
#cnf(i_0_14908, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|sk_c6=identity)).
###
#cnf(i_0_14911, negated_conjecture, (multiply(sk_c6,sk_c2)=identity|sk_c6=identity)).
#######################
#cnf(i_0_11027, negated_conjecture, (multiply(inverse(sk_c2),multiply(sk_c6,inverse(X1)))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c2)!=sk_c7|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_14913, negated_conjecture, (multiply(sk_c2,X1)=multiply(sk_c3,X1)|sk_c6=identity)).
######
#cnf(i_0_15087, negated_conjecture, (multiply(sk_c4,X1)=multiply(sk_c2,X1)|sk_c6=identity)).
########
#cnf(i_0_14914, negated_conjecture, (multiply(inverse(sk_c3),sk_c2)=identity|sk_c6=identity)).
######
#cnf(i_0_14915, negated_conjecture, (multiply(sk_c2,inverse(sk_c3))=identity|sk_c6=identity)).
#######
#cnf(i_0_14919, negated_conjecture, (multiply(sk_c3,inverse(sk_c2))=identity|sk_c6=identity)).
#####
#cnf(i_0_14921, negated_conjecture, (multiply(inverse(sk_c2),sk_c3)=identity|sk_c6=identity)).
#####
#cnf(i_0_664, negated_conjecture, (multiply(inverse(X2),multiply(inverse(X1),inverse(sk_c2)))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,X2)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_15085, negated_conjecture, (multiply(inverse(sk_c4),sk_c2)=identity|sk_c6=identity)).
#######
#cnf(i_0_15086, negated_conjecture, (multiply(sk_c2,inverse(sk_c4))=identity|sk_c6=identity)).
#
#cnf(i_0_15890, negated_conjecture, (inverse(sk_c4)=inverse(sk_c2)|sk_c6=identity)).
#######################
#cnf(i_0_15856, negated_conjecture, (multiply(inverse(sk_c2),sk_c4)=identity|sk_c6=identity)).
#########
#cnf(i_0_15891, negated_conjecture, (multiply(sk_c4,inverse(sk_c2))=identity|sk_c6=identity)).
#
#cnf(i_0_15715, negated_conjecture, (multiply(inverse(X1),multiply(X2,inverse(sk_c2)))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(inverse(X2),X1)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
##########################################################
#cnf(i_0_14909, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,X1))=X1|sk_c6=identity)).
######
#cnf(i_0_15826, negated_conjecture, (multiply(X1,multiply(inverse(X2),inverse(sk_c2)))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X2,inverse(X1))!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#############
#cnf(i_0_14912, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,X1))=X1|sk_c6=identity)).
##############################
#cnf(i_0_13346, negated_conjecture, (multiply(sk_c1,X1)=X1|sk_c7=sk_c6|sk_c6=identity)).
#############
#cnf(i_0_15635, negated_conjecture, (multiply(inverse(X1),multiply(sk_c6,inverse(sk_c2)))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c2,X1)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_13566, negated_conjecture, (multiply(sk_c7,X1)=X1|sk_c7=sk_c6|sk_c6=identity)).
#######################
#cnf(i_0_14922, negated_conjecture, (multiply(inverse(sk_c2),X1)=multiply(sk_c6,X1)|sk_c6=identity)).
################
#cnf(i_0_16257, negated_conjecture, (multiply(inverse(sk_c6),X1)=multiply(sk_c2,X1)|sk_c6=identity)).
##
#cnf(i_0_15746, negated_conjecture, (multiply(sk_c6,multiply(inverse(X1),inverse(sk_c2)))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c2)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#######################################
#cnf(i_0_14917, negated_conjecture, (multiply(inverse(sk_c2),inverse(sk_c6))=identity|sk_c6=identity)).
#####
#cnf(i_0_15784, negated_conjecture, (multiply(sk_c2,multiply(inverse(X1),inverse(sk_c2)))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c6)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_14918, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c2))=identity|sk_c6=identity)).
########################
#cnf(i_0_13369, negated_conjecture, (inverse(sk_c4)=sk_c5|inverse(sk_c1)=sk_c7|sk_c6=identity)).
##########################
#cnf(i_0_17361, negated_conjecture, (inverse(sk_c5)=sk_c4|inverse(sk_c1)=sk_c7|sk_c6=identity)).
#
#cnf(i_0_17408, negated_conjecture, (inverse(sk_c3)=sk_c5|inverse(sk_c1)=sk_c7|sk_c6=identity)).
###
#cnf(i_0_17424, negated_conjecture, (inverse(sk_c2)=sk_c5|inverse(sk_c1)=sk_c7|sk_c6=identity)).
######
#cnf(i_0_17491, negated_conjecture, (inverse(sk_c7)=sk_c1|inverse(sk_c3)=sk_c5|sk_c6=identity)).
#
#cnf(i_0_470, negated_conjecture, (multiply(inverse(X3),multiply(inverse(X2),inverse(X1)))!=sk_c6|multiply(X1,multiply(X2,multiply(X3,sk_c6)))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
##
#cnf(i_0_17540, negated_conjecture, (inverse(sk_c5)=sk_c2|inverse(sk_c1)=sk_c7|sk_c6=identity)).
#############
#cnf(i_0_2093, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c5|sk_c5=identity|sk_c7=identity)).
######
#cnf(i_0_2094, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7|sk_c5=identity|sk_c7=identity)).
#
#cnf(i_0_18016, negated_conjecture, (sk_c7=identity|sk_c5=identity)).
################
#cnf(i_0_18028, negated_conjecture, (inverse(sk_c7)=sk_c6|inverse(sk_c4)=sk_c6|sk_c7=identity)).
########
#cnf(i_0_18053, negated_conjecture, (inverse(sk_c6)=sk_c4|inverse(sk_c7)=sk_c6|sk_c7=identity)).
#
#cnf(i_0_18158, negated_conjecture, (inverse(sk_c6)=sk_c4|sk_c1=sk_c6|sk_c7=identity)).
###
#cnf(i_0_18243, negated_conjecture, (sk_c4=sk_c2|sk_c7=identity|sk_c1=sk_c6)).
##
#cnf(i_0_1312, negated_conjecture, (multiply(inverse(X3),multiply(inverse(X2),inverse(X1)))!=sk_c7|multiply(sk_c6,multiply(X1,multiply(X2,X3)))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
###
#cnf(i_0_18181, negated_conjecture, (inverse(sk_c4)=sk_c6|sk_c1=sk_c6|sk_c7=identity)).
#######
#cnf(i_0_18120, negated_conjecture, (inverse(sk_c6)=sk_c7|inverse(sk_c6)=sk_c4|sk_c7=identity)).
#
#cnf(i_0_18722, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c7=identity|sk_c4!=sk_c7)).
##
#cnf(i_0_18881, negated_conjecture, (sk_c7=sk_c3|sk_c7=identity|sk_c4!=sk_c7)).
#
#cnf(i_0_18884, negated_conjecture, (sk_c2=sk_c7|sk_c7=identity|sk_c4!=sk_c7)).
#
#cnf(i_0_18932, negated_conjecture, (sk_c2=sk_c7|sk_c7=identity|sk_c7!=sk_c3)).
##
#cnf(i_0_18821, negated_conjecture, (inverse(sk_c7)=sk_c6|sk_c7=identity|sk_c4!=sk_c7)).
#
#cnf(i_0_19120, negated_conjecture, (sk_c1=sk_c6|sk_c7=identity|sk_c4!=sk_c7)).
###
#cnf(i_0_19138, negated_conjecture, (sk_c1=sk_c6|sk_c7=identity|sk_c7!=sk_c3)).
#
#cnf(i_0_19146, negated_conjecture, (sk_c1=sk_c6|sk_c7=identity|sk_c2!=sk_c7)).
####
#cnf(i_0_18981, negated_conjecture, (inverse(sk_c7)=sk_c6|sk_c7=identity|sk_c7!=sk_c3)).
#####
#cnf(i_0_19154, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c7=identity|sk_c7!=sk_c3)).
######
#cnf(i_0_18788, negated_conjecture, (inverse(sk_c6)=sk_c7|sk_c4=sk_c2|sk_c7=identity)).
#
#cnf(i_0_19373, negated_conjecture, (sk_c7=sk_c3|sk_c7=identity|sk_c4=sk_c2)).
#
#cnf(i_0_19482, negated_conjecture, (sk_c7=sk_c3|sk_c7=identity|sk_c2!=sk_c7)).
##
#cnf(i_0_19376, negated_conjecture, (sk_c2=sk_c7|sk_c7=identity|sk_c4=sk_c2)).
####
#cnf(i_0_19313, negated_conjecture, (inverse(sk_c7)=sk_c6|sk_c4=sk_c2|sk_c7=identity)).
###
#cnf(i_0_18509, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,inverse(X1)))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c6,X1)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#########
#cnf(i_0_18818, negated_conjecture, (multiply(sk_c7,sk_c6)=identity|sk_c7=identity|sk_c4!=sk_c7)).
######
#cnf(i_0_18822, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|sk_c7=identity|sk_c4!=sk_c7)).
#
#cnf(i_0_19726, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,X1))!=sk_c7|multiply(sk_c6,inverse(X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#########
#cnf(i_0_18982, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|sk_c7=identity|sk_c7!=sk_c3)).
########
#cnf(i_0_18985, negated_conjecture, (multiply(sk_c7,sk_c6)=identity|sk_c7=identity|sk_c7!=sk_c3)).
#######
#cnf(i_0_19674, negated_conjecture, (multiply(sk_c6,multiply(sk_c2,sk_c2))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c6,sk_c6)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#######
#cnf(i_0_18726, negated_conjecture, (inverse(sk_c4)=sk_c6|inverse(sk_c6)=sk_c7|sk_c7=identity)).
#
#cnf(i_0_18521, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,inverse(X1)))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c6,X1)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
################
#cnf(i_0_20108, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,X1))!=sk_c7|multiply(sk_c6,inverse(X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
####################
#cnf(i_0_20009, negated_conjecture, (multiply(sk_c2,multiply(sk_c6,sk_c6))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c6,sk_c2)!=sk_c6|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity|sk_c4!=sk_c2)).
#######################
#cnf(i_0_2396, negated_conjecture, (multiply(sk_c4,sk_c6)=identity|sk_c2=sk_c3|sk_c5=sk_c7)).
###
#cnf(i_0_2400, negated_conjecture, (multiply(sk_c6,sk_c4)=identity|sk_c2=sk_c3|sk_c5=sk_c7)).
####
#cnf(i_0_2427, negated_conjecture, (multiply(sk_c7,sk_c1)=identity|sk_c4=sk_c3|sk_c5=sk_c7)).
##
#cnf(i_0_2075, negated_conjecture, (multiply(sk_c6,multiply(X1,multiply(X2,inverse(X3))))!=sk_c6|multiply(X3,multiply(inverse(X2),inverse(X1)))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
#
#cnf(i_0_2431, negated_conjecture, (multiply(sk_c1,sk_c7)=identity|sk_c4=sk_c3|sk_c5=sk_c7)).
####################
#cnf(i_0_4028, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c5|sk_c6=identity|sk_c7=sk_c6)).
#
#cnf(i_0_20595, negated_conjecture, (sk_c5=sk_c6|sk_c6=identity|sk_c7=sk_c6)).
###############
#cnf(i_0_4029, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7|sk_c6=identity|sk_c7=sk_c6)).
#########
#cnf(i_0_4480, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c5|sk_c5=identity|sk_c7=sk_c6)).
########
#cnf(i_0_4481, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7|sk_c5=identity|sk_c7=sk_c6)).
##
#cnf(i_0_2472, negated_conjecture, (multiply(inverse(X1),multiply(X2,multiply(X3,sk_c6)))!=sk_c7|multiply(inverse(X3),multiply(inverse(X2),X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
#################
#cnf(i_0_7180, negated_conjecture, (multiply(sk_c7,sk_c6)=sk_c5|sk_c4=sk_c6|sk_c4=sk_c3)).
####
#cnf(i_0_7181, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|sk_c4=sk_c6|sk_c4=sk_c3)).
#
#cnf(i_0_7182, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c5|sk_c4=sk_c6|sk_c4=sk_c3)).
###########################
#cnf(i_0_7540, negated_conjecture, (multiply(sk_c7,sk_c1)=identity|sk_c4=sk_c6|sk_c4=sk_c3)).
##
#cnf(i_0_7544, negated_conjecture, (multiply(sk_c1,sk_c7)=identity|sk_c4=sk_c6|sk_c4=sk_c3)).
#
#cnf(i_0_2675, negated_conjecture, (multiply(X1,multiply(X2,multiply(inverse(X3),sk_c6)))!=sk_c7|multiply(X3,multiply(inverse(X2),inverse(X1)))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
###
#cnf(i_0_7572, negated_conjecture, (multiply(sk_c6,sk_c2)=identity|sk_c4=sk_c6|sk_c4=sk_c3)).
###
#cnf(i_0_7576, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|sk_c4=sk_c6|sk_c4=sk_c3)).
################
#cnf(i_0_10123, negated_conjecture, (multiply(sk_c7,sk_c6)=identity|sk_c4=sk_c3|sk_c5=sk_c7)).
####
#cnf(i_0_10127, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|sk_c4=sk_c3|sk_c5=sk_c7)).
########
#cnf(i_0_10351, negated_conjecture, (multiply(sk_c7,sk_c6)=identity|sk_c4=sk_c3|sk_c7=identity)).
#######
#cnf(i_0_10355, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|sk_c4=sk_c3|sk_c7=identity)).
#######
#cnf(i_0_3058, negated_conjecture, (multiply(X1,multiply(X2,sk_c6))!=sk_c7|multiply(inverse(X2),inverse(X1))!=sk_c6|multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_10429, negated_conjecture, (multiply(sk_c7,sk_c6)=identity|sk_c4=sk_c2|sk_c4=sk_c3)).
####
#cnf(i_0_10433, negated_conjecture, (multiply(sk_c6,sk_c7)=identity|sk_c4=sk_c2|sk_c4=sk_c3)).
##
#cnf(i_0_21770, negated_conjecture, (multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c7!=sk_c6|sk_c6!=identity|sk_c5!=sk_c6|sk_c4!=sk_c2)).
###############
#cnf(i_0_12218, negated_conjecture, (multiply(sk_c7,sk_c1)=identity|sk_c4=sk_c7|sk_c4=sk_c3)).
###
#cnf(i_0_12222, negated_conjecture, (multiply(sk_c1,sk_c7)=identity|sk_c4=sk_c7|sk_c4=sk_c3)).
##
#cnf(i_0_21692, negated_conjecture, (multiply(sk_c6,inverse(X1))!=sk_c6|multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2|X1!=sk_c7)).
###
#cnf(i_0_12538, negated_conjecture, (multiply(sk_c6,sk_c2)=identity|sk_c4=sk_c7|sk_c4=sk_c3)).
###
#cnf(i_0_12542, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|sk_c4=sk_c7|sk_c4=sk_c3)).
##
#cnf(i_0_21920, negated_conjecture, (multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(sk_c6,X1)!=sk_c6|inverse(X1)!=sk_c7|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
###
#cnf(i_0_13300, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c5|sk_c5=sk_c6|sk_c6=identity)).
##
#cnf(i_0_22080, negated_conjecture, (sk_c5=sk_c6|sk_c6=identity)).
###
#cnf(i_0_22012, negated_conjecture, (multiply(inverse(X1),sk_c6)!=sk_c7|multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2|X1!=sk_c6)).
#
#cnf(i_0_22098, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c6=identity)).
####
#cnf(i_0_22251, negated_conjecture, (inverse(sk_c7)=sk_c1|sk_c6=identity)).
#####
#cnf(i_0_22248, negated_conjecture, (multiply(sk_c7,sk_c1)=identity|sk_c6=identity)).
#####
#cnf(i_0_22252, negated_conjecture, (multiply(sk_c1,sk_c7)=identity|sk_c6=identity)).
###########
#cnf(i_0_22394, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c5|sk_c6=identity)).
###
#cnf(i_0_22395, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7|sk_c6=identity)).
###
#cnf(i_0_21624, negated_conjecture, (multiply(inverse(sk_c2),inverse(X1))!=sk_c6|multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2|X1!=sk_c7)).
#
#cnf(i_0_22412, negated_conjecture, (multiply(sk_c3,sk_c5)=sk_c7|sk_c6=identity)).
#
#cnf(i_0_22419, negated_conjecture, (multiply(sk_c4,sk_c5)=sk_c7|sk_c6=identity)).
#
#cnf(i_0_22421, negated_conjecture, (multiply(sk_c2,sk_c5)=sk_c7|sk_c6=identity)).
##
#cnf(i_0_22448, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c6|sk_c6=identity)).
#######
#cnf(i_0_22569, negated_conjecture, (multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|inverse(sk_c2)!=sk_c7|sk_c6!=identity|sk_c5!=sk_c6|sk_c4!=sk_c2)).
#########
#cnf(i_0_22635, negated_conjecture, (multiply(sk_c4,sk_c6)=sk_c7|sk_c6=identity)).
########
#cnf(i_0_22678, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c7|sk_c6=identity)).
###
#cnf(i_0_22550, negated_conjecture, (multiply(inverse(sk_c2),X1)!=sk_c6|multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|inverse(X1)!=sk_c7|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
############################
#cnf(i_0_22457, negated_conjecture, (multiply(inverse(sk_c2),sk_c6)!=sk_c6|multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2|sk_c2!=sk_c7)).
####
#cnf(i_0_22892, negated_conjecture, (multiply(sk_c6,sk_c6)!=sk_c6|multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2|sk_c2!=sk_c7)).
############
#cnf(i_0_22015, negated_conjecture, (multiply(inverse(X1),inverse(sk_c2))!=sk_c7|multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2|X1!=sk_c6)).
##
#cnf(i_0_22400, negated_conjecture, (multiply(inverse(sk_c6),sk_c5)=sk_c7|sk_c6=identity)).
#############
#cnf(i_0_22433, negated_conjecture, (multiply(inverse(sk_c3),sk_c7)=sk_c6|sk_c6=identity)).
########
#cnf(i_0_22579, negated_conjecture, (multiply(inverse(sk_c3),sk_c7)=sk_c5|sk_c6=identity)).
###
#cnf(i_0_23174, negated_conjecture, (inverse(sk_c3)=sk_c5|sk_c6=identity)).
##
#cnf(i_0_22946, negated_conjecture, (multiply(sk_c6,inverse(sk_c2))!=sk_c7|multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2|sk_c2!=sk_c6)).
#
#cnf(i_0_23182, negated_conjecture, (inverse(sk_c5)=sk_c3|sk_c6=identity)).
#########
#cnf(i_0_23179, negated_conjecture, (multiply(sk_c5,sk_c3)=identity|sk_c6=identity)).
###
#cnf(i_0_23183, negated_conjecture, (multiply(sk_c3,sk_c5)=identity|sk_c6=identity)).
#####
#cnf(i_0_23247, negated_conjecture, (multiply(sk_c5,sk_c4)=identity|sk_c6=identity)).
#
#cnf(i_0_23385, negated_conjecture, (inverse(sk_c5)=sk_c4|sk_c6=identity)).
#
#cnf(i_0_23039, negated_conjecture, (multiply(X1,inverse(sk_c2))!=sk_c7|multiply(sk_c2,sk_c6)!=sk_c5|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|inverse(X1)!=sk_c6|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_23419, negated_conjecture, (inverse(sk_c4)=sk_c5|sk_c6=identity)).
#######
#cnf(i_0_23627, negated_conjecture, (inverse(sk_c2)=sk_c5|sk_c6=identity)).
#####
#cnf(i_0_23646, negated_conjecture, (inverse(sk_c5)=sk_c2|sk_c6=identity)).
###########
#cnf(i_0_23248, negated_conjecture, (multiply(sk_c4,sk_c5)=identity|sk_c6=identity)).
########
#cnf(i_0_23251, negated_conjecture, (multiply(sk_c5,sk_c2)=identity|sk_c6=identity)).
##
#cnf(i_0_23252, negated_conjecture, (multiply(sk_c2,sk_c5)=identity|sk_c6=identity)).
#####
#cnf(i_0_3241, negated_conjecture, (multiply(X1,multiply(X2,sk_c6))!=sk_c5|multiply(inverse(X2),inverse(X1))!=sk_c6|multiply(sk_c2,sk_c6)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
##
#cnf(i_0_23253, negated_conjecture, (multiply(sk_c5,sk_c7)=sk_c6|sk_c6=identity)).
##
#cnf(i_0_24095, negated_conjecture, (multiply(sk_c2,sk_c6)!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c6!=identity|sk_c4!=sk_c2)).
#
#cnf(i_0_23254, negated_conjecture, (multiply(sk_c5,sk_c7)=sk_c5|sk_c6=identity)).
#######
#cnf(i_0_1069, negated_conjecture, (multiply(sk_c6,multiply(inverse(X2),multiply(inverse(X1),X3)))!=sk_c6|multiply(inverse(X3),multiply(X1,X2))!=sk_c7|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c5!=identity|sk_c6!=identity)).
##############################
#cnf(i_0_23843, negated_conjecture, (multiply(sk_c7,sk_c6)=sk_c5|sk_c6=identity)).
###
#cnf(i_0_23844, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|sk_c6=identity)).
#
#cnf(i_0_24550, negated_conjecture, (sk_c6=identity)).
#
#cnf(i_0_33, negated_conjecture, (sk_c5=sk_c7)).
#
#cnf(i_0_3038, negated_conjecture, (sk_c7!=identity)).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(associativity, axiom, (multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), file('./p/GRP287-1.tptp', associativity)).
cnf(left_inverse, axiom, (multiply(inverse(X1),X1)=identity), file('./p/GRP287-1.tptp', left_inverse)).
cnf(left_identity, axiom, (multiply(identity,X1)=X1), file('./p/GRP287-1.tptp', left_identity)).
cnf(prove_this_8, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|inverse(sk_c3)=sk_c6), file('./p/GRP287-1.tptp', prove_this_8)).
cnf(prove_this_23, negated_conjecture, (inverse(sk_c2)=sk_c6|inverse(sk_c3)=sk_c6), file('./p/GRP287-1.tptp', prove_this_23)).
cnf(prove_this_13, negated_conjecture, (inverse(sk_c1)=sk_c7|inverse(sk_c3)=sk_c6), file('./p/GRP287-1.tptp', prove_this_13)).
cnf(prove_this_18, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c5|inverse(sk_c3)=sk_c6), file('./p/GRP287-1.tptp', prove_this_18)).
cnf(prove_this_3, negated_conjecture, (multiply(sk_c7,sk_c6)=sk_c5|inverse(sk_c3)=sk_c6), file('./p/GRP287-1.tptp', prove_this_3)).
cnf(prove_this_22, negated_conjecture, (inverse(sk_c2)=sk_c6|multiply(sk_c3,sk_c6)=sk_c7), file('./p/GRP287-1.tptp', prove_this_22)).
cnf(prove_this_12, negated_conjecture, (inverse(sk_c1)=sk_c7|multiply(sk_c3,sk_c6)=sk_c7), file('./p/GRP287-1.tptp', prove_this_12)).
cnf(prove_this_7, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|multiply(sk_c3,sk_c6)=sk_c7), file('./p/GRP287-1.tptp', prove_this_7)).
cnf(prove_this_15, negated_conjecture, (inverse(sk_c1)=sk_c7|multiply(sk_c4,sk_c5)=sk_c6), file('./p/GRP287-1.tptp', prove_this_15)).
cnf(prove_this_14, negated_conjecture, (inverse(sk_c1)=sk_c7|inverse(sk_c4)=sk_c6), file('./p/GRP287-1.tptp', prove_this_14)).
cnf(prove_this_1, negated_conjecture, (multiply(sk_c7,sk_c6)=sk_c5|multiply(sk_c6,sk_c7)=sk_c5), file('./p/GRP287-1.tptp', prove_this_1)).
cnf(prove_this_6, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|multiply(sk_c6,sk_c7)=sk_c5), file('./p/GRP287-1.tptp', prove_this_6)).
cnf(prove_this_9, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|inverse(sk_c4)=sk_c6), file('./p/GRP287-1.tptp', prove_this_9)).
cnf(prove_this_26, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(X1,sk_c7)!=sk_c6|inverse(X1)!=sk_c7|multiply(X2,sk_c6)!=sk_c5|inverse(X2)!=sk_c6|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X3,sk_c6)!=sk_c7|inverse(X3)!=sk_c6|inverse(X4)!=sk_c6|multiply(X4,sk_c5)!=sk_c6), file('./p/GRP287-1.tptp', prove_this_26)).
cnf(prove_this_10, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|multiply(sk_c4,sk_c5)=sk_c6), file('./p/GRP287-1.tptp', prove_this_10)).
cnf(prove_this_17, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c5|multiply(sk_c3,sk_c6)=sk_c7), file('./p/GRP287-1.tptp', prove_this_17)).
cnf(c_0_19, axiom, (multiply(multiply(X1,X2),X3)=multiply(X1,multiply(X2,X3))), associativity).
cnf(c_0_20, axiom, (multiply(inverse(X1),X1)=identity), left_inverse).
cnf(c_0_21, axiom, (multiply(identity,X1)=X1), left_identity).
cnf(c_0_22, plain, (multiply(inverse(X1),multiply(X1,X2))=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_23, plain, (multiply(inverse(inverse(X1)),identity)=X1), inference(pm,[status(thm)],[c_0_22, c_0_20])).
cnf(c_0_24, plain, (multiply(inverse(inverse(X1)),X2)=multiply(X1,X2)), inference(pm,[status(thm)],[c_0_22, c_0_22])).
cnf(c_0_25, plain, (multiply(X1,identity)=X1), inference(rw,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_26, plain, (inverse(inverse(X1))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_24, c_0_25]), c_0_25])).
cnf(c_0_27, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|inverse(sk_c3)=sk_c6), prove_this_8).
cnf(c_0_28, plain, (multiply(X1,inverse(X1))=identity), inference(pm,[status(thm)],[c_0_20, c_0_26])).
cnf(c_0_29, negated_conjecture, (inverse(sk_c2)=sk_c6|inverse(sk_c3)=sk_c6), prove_this_23).
cnf(c_0_30, negated_conjecture, (multiply(inverse(sk_c1),sk_c6)=sk_c7|inverse(sk_c3)=sk_c6), inference(pm,[status(thm)],[c_0_22, c_0_27])).
cnf(c_0_31, negated_conjecture, (inverse(sk_c1)=sk_c7|inverse(sk_c3)=sk_c6), prove_this_13).
cnf(c_0_32, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c5|inverse(sk_c3)=sk_c6), prove_this_18).
cnf(c_0_33, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|inverse(sk_c3)=sk_c6), inference(pm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, negated_conjecture, (multiply(sk_c7,sk_c6)=sk_c5|inverse(sk_c3)=sk_c6), prove_this_3).
cnf(c_0_35, negated_conjecture, (multiply(sk_c7,sk_c6)=sk_c7|inverse(sk_c3)=sk_c6), inference(pm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_36, negated_conjecture, (inverse(sk_c3)=sk_c6|sk_c5=identity), inference(pm,[status(thm)],[c_0_32, c_0_33])).
cnf(c_0_37, negated_conjecture, (inverse(sk_c3)=sk_c6|sk_c5=sk_c7), inference(pm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38, negated_conjecture, (inverse(sk_c2)=sk_c6|multiply(sk_c3,sk_c6)=sk_c7), prove_this_22).
cnf(c_0_39, negated_conjecture, (multiply(sk_c3,sk_c6)=identity|sk_c5=identity), inference(pm,[status(thm)],[c_0_28, c_0_36])).
cnf(c_0_40, negated_conjecture, (multiply(sk_c3,sk_c6)=identity|sk_c5=sk_c7), inference(pm,[status(thm)],[c_0_28, c_0_37])).
cnf(c_0_41, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c7=identity|sk_c5=identity), inference(pm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_42, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c7=identity|sk_c5=sk_c7), inference(pm,[status(thm)],[c_0_38, c_0_40])).
cnf(c_0_43, negated_conjecture, (inverse(sk_c6)=sk_c3|sk_c5=identity), inference(pm,[status(thm)],[c_0_26, c_0_36])).
cnf(c_0_44, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c7=identity|sk_c5=identity), inference(pm,[status(thm)],[c_0_26, c_0_41])).
cnf(c_0_45, negated_conjecture, (inverse(sk_c6)=sk_c3|sk_c5=sk_c7), inference(pm,[status(thm)],[c_0_26, c_0_37])).
cnf(c_0_46, negated_conjecture, (inverse(sk_c6)=sk_c2|sk_c7=identity|sk_c5=sk_c7), inference(pm,[status(thm)],[c_0_26, c_0_42])).
cnf(c_0_47, negated_conjecture, (sk_c2=sk_c3|sk_c5=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_48, negated_conjecture, (sk_c2=sk_c3|sk_c5=sk_c7|sk_c7=identity), inference(pm,[status(thm)],[c_0_45, c_0_46])).
cnf(c_0_49, negated_conjecture, (sk_c7=identity|sk_c2=sk_c3), inference(pm,[status(thm)],[c_0_47, c_0_48])).
cnf(c_0_50, negated_conjecture, (inverse(sk_c3)=sk_c6|sk_c7=identity), inference(pm,[status(thm)],[c_0_29, c_0_49])).
cnf(c_0_51, negated_conjecture, (inverse(sk_c1)=sk_c7|multiply(sk_c3,sk_c6)=sk_c7), prove_this_12).
cnf(c_0_52, negated_conjecture, (multiply(sk_c3,sk_c6)=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_28, c_0_50])).
cnf(c_0_53, negated_conjecture, (multiply(sk_c1,sk_c7)=identity|inverse(sk_c3)=sk_c6), inference(pm,[status(thm)],[c_0_28, c_0_31])).
cnf(c_0_54, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c7=identity), inference(pm,[status(thm)],[c_0_51, c_0_52])).
cnf(c_0_55, negated_conjecture, (inverse(sk_c3)=sk_c6|sk_c6=identity), inference(pm,[status(thm)],[c_0_27, c_0_53])).
cnf(c_0_56, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|multiply(sk_c3,sk_c6)=sk_c7), prove_this_7).
cnf(c_0_57, negated_conjecture, (multiply(sk_c1,sk_c7)=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_28, c_0_54])).
cnf(c_0_58, negated_conjecture, (inverse(sk_c1)=sk_c7|multiply(sk_c4,sk_c5)=sk_c6), prove_this_15).
cnf(c_0_59, negated_conjecture, (multiply(sk_c3,sk_c6)=identity|sk_c6=identity), inference(pm,[status(thm)],[c_0_28, c_0_55])).
cnf(c_0_60, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7|sk_c6=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_61, negated_conjecture, (multiply(inverse(sk_c4),sk_c6)=sk_c5|inverse(sk_c1)=sk_c7), inference(pm,[status(thm)],[c_0_22, c_0_58])).
cnf(c_0_62, negated_conjecture, (inverse(sk_c1)=sk_c7|inverse(sk_c4)=sk_c6), prove_this_14).
cnf(c_0_63, negated_conjecture, (multiply(sk_c7,sk_c6)=sk_c5|multiply(sk_c6,sk_c7)=sk_c5), prove_this_1).
cnf(c_0_64, negated_conjecture, (sk_c7=identity|sk_c6=identity), inference(pm,[status(thm)],[c_0_59, c_0_60])).
cnf(c_0_65, negated_conjecture, (multiply(sk_c6,sk_c6)=sk_c5|inverse(sk_c1)=sk_c7), inference(pm,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_66, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c5|sk_c5=sk_c6|sk_c6=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_63, c_0_64]), c_0_21])).
cnf(c_0_67, negated_conjecture, (multiply(inverse(sk_c6),sk_c5)=sk_c6|inverse(sk_c1)=sk_c7), inference(pm,[status(thm)],[c_0_22, c_0_65])).
cnf(c_0_68, negated_conjecture, (sk_c5=sk_c6|sk_c6=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_66, c_0_64]), c_0_25])])).
cnf(c_0_69, negated_conjecture, (inverse(sk_c1)=sk_c7|sk_c6=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_67, c_0_68]), c_0_20])])).
cnf(c_0_70, plain, (multiply(X1,multiply(X2,inverse(multiply(X1,X2))))=identity), inference(pm,[status(thm)],[c_0_19, c_0_28])).
cnf(c_0_71, plain, (multiply(inverse(identity),X1)=X1), inference(pm,[status(thm)],[c_0_22, c_0_21])).
cnf(c_0_72, negated_conjecture, (inverse(sk_c6)=sk_c4|inverse(sk_c1)=sk_c7), inference(pm,[status(thm)],[c_0_26, c_0_62])).
cnf(c_0_73, plain, (multiply(X1,multiply(inverse(X1),X2))=X2), inference(pm,[status(thm)],[c_0_22, c_0_26])).
cnf(c_0_74, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|multiply(sk_c6,sk_c7)=sk_c5), prove_this_6).
cnf(c_0_75, negated_conjecture, (multiply(sk_c1,sk_c7)=identity|sk_c6=identity), inference(pm,[status(thm)],[c_0_28, c_0_69])).
cnf(c_0_76, plain, (multiply(X1,inverse(multiply(X2,X1)))=inverse(X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_22, c_0_70]), c_0_25])).
cnf(c_0_77, plain, (multiply(inverse(inverse(identity)),X1)=X1), inference(pm,[status(thm)],[c_0_22, c_0_71])).
cnf(c_0_78, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|inverse(sk_c4)=sk_c6), prove_this_9).
cnf(c_0_79, negated_conjecture, (multiply(sk_c1,sk_c7)=identity|inverse(sk_c6)=sk_c4), inference(pm,[status(thm)],[c_0_28, c_0_72])).
cnf(c_0_80, negated_conjecture, (multiply(sk_c3,multiply(sk_c6,X1))=X1|sk_c6=identity), inference(pm,[status(thm)],[c_0_73, c_0_55])).
cnf(c_0_81, negated_conjecture, (multiply(sk_c6,sk_c7)=sk_c5|sk_c6=identity), inference(pm,[status(thm)],[c_0_74, c_0_75])).
cnf(c_0_82, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(X1,sk_c7)!=sk_c6|inverse(X1)!=sk_c7|multiply(X2,sk_c6)!=sk_c5|inverse(X2)!=sk_c6|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X3,sk_c6)!=sk_c7|inverse(X3)!=sk_c6|inverse(X4)!=sk_c6|multiply(X4,sk_c5)!=sk_c6), prove_this_26).
cnf(c_0_83, plain, (inverse(multiply(X1,X2))=multiply(inverse(X2),inverse(X1))), inference(pm,[status(thm)],[c_0_22, c_0_76])).
cnf(c_0_84, plain, (inverse(identity)=identity), inference(pm,[status(thm)],[c_0_20, c_0_77])).
cnf(c_0_85, negated_conjecture, (inverse(sk_c4)=sk_c6|inverse(sk_c6)=sk_c4|sk_c6=identity), inference(pm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_86, negated_conjecture, (multiply(sk_c3,sk_c5)=sk_c7|sk_c6=identity), inference(pm,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_87, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|inverse(inverse(sk_c7))!=sk_c7|multiply(X1,sk_c5)!=sk_c6|multiply(X2,sk_c6)!=sk_c7|multiply(X3,sk_c6)!=sk_c5|inverse(X1)!=sk_c6|inverse(X2)!=sk_c6|inverse(X3)!=sk_c6|sk_c6!=identity), inference(pm,[status(thm)],[c_0_82, c_0_20])).
cnf(c_0_88, negated_conjecture, (multiply(inverse(sk_c6),inverse(sk_c3))=identity|sk_c6=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_83, c_0_59]), c_0_84])).
cnf(c_0_89, negated_conjecture, (inverse(sk_c6)=sk_c4|sk_c6=identity), inference(pm,[status(thm)],[c_0_26, c_0_85])).
cnf(c_0_90, negated_conjecture, (multiply(inverse(sk_c3),sk_c7)=sk_c5|sk_c6=identity), inference(pm,[status(thm)],[c_0_22, c_0_86])).
cnf(c_0_91, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|inverse(inverse(sk_c7))!=sk_c7|inverse(inverse(sk_c5))!=sk_c6|multiply(X1,sk_c6)!=sk_c7|multiply(X2,sk_c6)!=sk_c5|inverse(X1)!=sk_c6|inverse(X2)!=sk_c6|sk_c6!=identity), inference(pm,[status(thm)],[c_0_87, c_0_20])).
cnf(c_0_92, negated_conjecture, (multiply(sk_c4,inverse(sk_c3))=identity|sk_c6=identity), inference(pm,[status(thm)],[c_0_88, c_0_89])).
cnf(c_0_93, negated_conjecture, (inverse(sk_c3)=sk_c5|sk_c6=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_90, c_0_64]), c_0_25])).
cnf(c_0_94, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c6)!=sk_c7|multiply(X2,sk_c6)!=sk_c5|inverse(X1)!=sk_c6|inverse(X2)!=sk_c6|sk_c5!=sk_c6|sk_c6!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_26]), c_0_26])])).
cnf(c_0_95, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|multiply(sk_c4,sk_c5)=sk_c6), prove_this_10).
cnf(c_0_96, negated_conjecture, (multiply(sk_c4,sk_c5)=identity|sk_c6=identity), inference(pm,[status(thm)],[c_0_92, c_0_93])).
cnf(c_0_97, negated_conjecture, (inverse(sk_c2)=sk_c6|sk_c7=identity), inference(pm,[status(thm)],[c_0_38, c_0_52])).
cnf(c_0_98, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|multiply(X1,sk_c6)!=sk_c5|inverse(X1)!=sk_c6|sk_c7!=identity|sk_c5!=sk_c6|sk_c6!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_94, c_0_20]), c_0_26])])).
cnf(c_0_99, negated_conjecture, (multiply(sk_c1,sk_c7)=sk_c6|sk_c6=identity), inference(pm,[status(thm)],[c_0_95, c_0_96])).
cnf(c_0_100, negated_conjecture, (multiply(sk_c2,sk_c6)=sk_c5|multiply(sk_c3,sk_c6)=sk_c7), prove_this_17).
cnf(c_0_101, negated_conjecture, (multiply(sk_c2,sk_c6)=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_28, c_0_97])).
cnf(c_0_102, negated_conjecture, (multiply(X1,multiply(X2,sk_c6))!=sk_c5|multiply(inverse(X2),inverse(X1))!=sk_c6|multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c7!=identity|sk_c5!=sk_c6|sk_c6!=identity), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_98, c_0_19]), c_0_83])).
cnf(c_0_103, negated_conjecture, (sk_c6=identity), inference(pm,[status(thm)],[c_0_75, c_0_99])).
cnf(c_0_104, negated_conjecture, (multiply(sk_c3,sk_c6)=sk_c7|sk_c5=identity|sk_c7=identity), inference(pm,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_105, negated_conjecture, (multiply(sk_c7,sk_c6)!=sk_c5|multiply(sk_c6,sk_c7)!=sk_c5|sk_c5!=sk_c6|sk_c6!=identity|sk_c7!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_102, c_0_22]), c_0_26]), c_0_20])])).
cnf(c_0_106, negated_conjecture, (sk_c5=sk_c7), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_103]), c_0_25]), c_0_103]), c_0_21])])).
cnf(c_0_107, negated_conjecture, (sk_c7=identity|sk_c5=identity), inference(pm,[status(thm)],[c_0_39, c_0_104])).
cnf(c_0_108, negated_conjecture, (sk_c7!=identity), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_105, c_0_103]), c_0_25]), c_0_103]), c_0_21]), c_0_103]), c_0_103])]), c_0_106]), c_0_106])])).
cnf(c_0_109, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[c_0_107, c_0_108]), c_0_106]), c_0_108]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 110
# Proof object clause steps            : 91
# Proof object formula steps           : 19
# Proof object conjectures             : 91
# Proof object clause conjectures      : 75
# Proof object formula conjectures     : 16
# Proof object initial clauses used    : 19
# Proof object initial formulas used   : 19
# Proof object generating inferences   : 67
# Proof object simplifying inferences  : 38
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 29
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 29
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 29
# Processed clauses                    : 3136
# ...of these trivial                  : 10
# ...subsumed                          : 2476
# ...remaining for further processing  : 650
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 117
# Backward-rewritten                   : 488
# Generated clauses                    : 24518
# ...of the previous two non-trivial   : 18476
# Contextual simplify-reflections      : 0
# Paramodulations                      : 24489
# Factorizations                       : 8
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 24
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 11
# Current number of unprocessed clauses: 15126
# ...number of literals in the above   : 98263
# Current number of archived formulas  : 0
# Current number of archived clauses   : 626
# Clause-clause subsumption calls (NU) : 15267
# Rec. Clause-clause subsumption calls : 4440
# Non-unit clause-clause subsumptions  : 2590
# Unit Clause-clause subsumption calls : 12
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 12
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 126829
