# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_7, plain, (divide(divide(X2,X2),X1)=inverse(X1))).
#
#cnf(i_0_9, plain, (divide(inverse(divide(X1,X1)),X2)=inverse(X2))).
#
#cnf(i_0_6, plain, (multiply(X1,X2)=divide(X1,inverse(X2)))).
#
#cnf(i_0_8, negated_conjecture, (divide(inverse(a1),inverse(a1))!=divide(inverse(b1),inverse(b1)))).
#
#cnf(i_0_10, plain, (divide(inverse(inverse(divide(X1,X1))),X2)=inverse(X2))).
##
#cnf(i_0_11, plain, (divide(inverse(inverse(inverse(divide(X1,X1)))),X2)=inverse(X2))).
###
#cnf(i_0_14, plain, (divide(inverse(inverse(inverse(inverse(divide(X1,X1))))),X2)=inverse(X2))).
####
#cnf(i_0_15, plain, (divide(inverse(inverse(inverse(inverse(inverse(divide(X1,X1)))))),X2)=inverse(X2))).
#####
#cnf(i_0_20, plain, (divide(inverse(inverse(inverse(inverse(inverse(inverse(divide(X1,X1))))))),X2)=inverse(X2))).
######
#cnf(i_0_21, plain, (divide(inverse(inverse(inverse(inverse(inverse(inverse(inverse(divide(X1,X1)))))))),X2)=inverse(X2))).
#######
#cnf(i_0_28, plain, (divide(inverse(inverse(inverse(inverse(inverse(inverse(inverse(inverse(divide(X1,X1))))))))),X2)=inverse(X2))).
########
#cnf(i_0_5, plain, (divide(inverse(divide(X1,divide(X2,divide(inverse(X1),X3)))),X3)=X2)).
#
#cnf(i_0_92, plain, (divide(inverse(inverse(divide(X2,inverse(X3)))),X3)=X2)).
################
#cnf(i_0_130, plain, (divide(inverse(inverse(inverse(inverse(X2)))),X2)=divide(X1,X1))).
#
#cnf(i_0_161, plain, (divide(X3,X3)=divide(X2,X2))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(inverse, axiom, (inverse(X1)=divide(divide(X2,X2),X1)), file('./p/GRP451-1.tptp', inverse)).
cnf(single_axiom, axiom, (divide(divide(divide(X1,X1),divide(X1,divide(X2,divide(divide(divide(X1,X1),X1),X3)))),X3)=X2), file('./p/GRP451-1.tptp', single_axiom)).
cnf(multiply, axiom, (multiply(X1,X2)=divide(X1,divide(divide(X3,X3),X2))), file('./p/GRP451-1.tptp', multiply)).
cnf(prove_these_axioms_1, negated_conjecture, (multiply(inverse(a1),a1)!=multiply(inverse(b1),b1)), file('./p/GRP451-1.tptp', prove_these_axioms_1)).
cnf(c_0_4, axiom, (inverse(X1)=divide(divide(X2,X2),X1)), inverse).
cnf(c_0_5, axiom, (divide(divide(divide(X1,X1),divide(X1,divide(X2,divide(divide(divide(X1,X1),X1),X3)))),X3)=X2), single_axiom).
cnf(c_0_6, plain, (divide(inverse(divide(X1,X1)),X2)=inverse(X2)), inference(pm,[status(thm)],[c_0_4, c_0_4])).
cnf(c_0_7, plain, (divide(inverse(divide(X1,divide(X2,divide(inverse(X1),X3)))),X3)=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_4]), c_0_4])).
cnf(c_0_8, plain, (divide(inverse(inverse(divide(X1,X1))),X2)=inverse(X2)), inference(pm,[status(thm)],[c_0_4, c_0_6])).
cnf(c_0_9, axiom, (multiply(X1,X2)=divide(X1,divide(divide(X3,X3),X2))), multiply).
cnf(c_0_10, plain, (divide(inverse(inverse(divide(X1,inverse(X2)))),X2)=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_8]), c_0_6])).
cnf(c_0_11, negated_conjecture, (multiply(inverse(a1),a1)!=multiply(inverse(b1),b1)), prove_these_axioms_1).
cnf(c_0_12, plain, (multiply(X1,X2)=divide(X1,inverse(X2))), inference(rw,[status(thm)],[c_0_9, c_0_4])).
cnf(c_0_13, plain, (divide(inverse(inverse(inverse(inverse(X1)))),X1)=divide(X2,X2)), inference(pm,[status(thm)],[c_0_10, c_0_4])).
cnf(c_0_14, negated_conjecture, (divide(inverse(a1),inverse(a1))!=divide(inverse(b1),inverse(b1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_12]), c_0_12])).
cnf(c_0_15, plain, (divide(X1,X1)=divide(X2,X2)), inference(pm,[status(thm)],[c_0_13, c_0_13])).
cnf(c_0_16, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_14, c_0_15]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 17
# Proof object clause steps            : 13
# Proof object formula steps           : 4
# Proof object conjectures             : 4
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 4
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 4
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 4
# Processed clauses                    : 58
# ...of these trivial                  : 43
# ...subsumed                          : 0
# ...remaining for further processing  : 15
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 202
# ...of the previous two non-trivial   : 153
# Contextual simplify-reflections      : 0
# Paramodulations                      : 201
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 14
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 99
# ...number of literals in the above   : 99
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 18
# Rewrite failures with RHS unbound    : 35
# BW rewrite match attempts            : 95
# BW rewrite match successes           : 11
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2342
