# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_9, plain, (divide(X1,X1)=identity)).
#
#cnf(i_0_8, plain, (divide(identity,X1)=inverse(X1))).
#
#cnf(i_0_11, plain, (inverse(identity)=identity)).
#
#cnf(i_0_10, negated_conjecture, (multiply(a,b)!=multiply(b,a))).
#
#cnf(i_0_7, plain, (multiply(X1,X2)=divide(X1,inverse(X2)))).
#
#cnf(i_0_6, plain, (divide(divide(X1,divide(divide(X1,X2),X3)),X2)=X3)).
#
#cnf(i_0_15, plain, (divide(divide(X1,inverse(X2)),X1)=X2)).
#
#cnf(i_0_10, negated_conjecture, (divide(a,inverse(b))!=divide(b,inverse(a)))).
#
#cnf(i_0_24, plain, (inverse(inverse(X1))=X1)).
#
#cnf(i_0_25, plain, (divide(divide(X1,X2),X1)=inverse(X2))).
#
#cnf(i_0_31, plain, (divide(inverse(X1),identity)=inverse(X1))).
#
#cnf(i_0_37, plain, (divide(X1,identity)=X1)).
#
#cnf(i_0_39, plain, (divide(X1,divide(X1,X2))=X2)).
##
#cnf(i_0_47, plain, (inverse(divide(X1,X2))=divide(X2,X1))).
#
#cnf(i_0_30, plain, (divide(divide(X1,inverse(X2)),X2)=X1)).
#
#cnf(i_0_50, plain, (divide(divide(X1,X2),inverse(X2))=X1)).
###
#cnf(i_0_28, plain, (divide(inverse(X2),divide(X1,X2))=inverse(X1))).
#
#cnf(i_0_64, plain, (divide(X2,divide(X1,inverse(X2)))=inverse(X1))).
######
#cnf(i_0_73, plain, (divide(inverse(X2),X1)=divide(inverse(X1),X2))).
#
#cnf(i_0_86, plain, (divide(inverse(X2),inverse(X1))=divide(X1,X2))).
##
#cnf(i_0_114, plain, (divide(X1,inverse(X2))=divide(X2,inverse(X1)))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(inverse, axiom, (inverse(X1)=divide(divide(X2,X2),X1)), file('./p/GRP536-1.tptp', inverse)).
cnf(identity, axiom, (identity=divide(X1,X1)), file('./p/GRP536-1.tptp', identity)).
cnf(single_axiom, axiom, (divide(divide(X1,divide(divide(X1,X2),X3)),X2)=X3), file('./p/GRP536-1.tptp', single_axiom)).
cnf(multiply, axiom, (multiply(X1,X2)=divide(X1,divide(divide(X3,X3),X2))), file('./p/GRP536-1.tptp', multiply)).
cnf(prove_these_axioms_4, negated_conjecture, (multiply(a,b)!=multiply(b,a)), file('./p/GRP536-1.tptp', prove_these_axioms_4)).
cnf(c_0_5, axiom, (inverse(X1)=divide(divide(X2,X2),X1)), inverse).
cnf(c_0_6, axiom, (identity=divide(X1,X1)), identity).
cnf(c_0_7, axiom, (divide(divide(X1,divide(divide(X1,X2),X3)),X2)=X3), single_axiom).
cnf(c_0_8, plain, (divide(identity,X1)=inverse(X1)), inference(rw,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_9, plain, (divide(divide(X1,inverse(X2)),X1)=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_6]), c_0_8])).
cnf(c_0_10, plain, (inverse(inverse(X1))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_9, c_0_6]), c_0_8])).
cnf(c_0_11, plain, (divide(divide(X1,X2),X1)=inverse(X2)), inference(pm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_12, plain, (divide(inverse(X1),identity)=inverse(X1)), inference(pm,[status(thm)],[c_0_11, c_0_8])).
cnf(c_0_13, plain, (divide(X1,identity)=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_12, c_0_10]), c_0_10])).
cnf(c_0_14, plain, (divide(X1,divide(X1,X2))=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_13]), c_0_13])).
cnf(c_0_15, axiom, (multiply(X1,X2)=divide(X1,divide(divide(X3,X3),X2))), multiply).
cnf(c_0_16, plain, (inverse(divide(X1,X2))=divide(X2,X1)), inference(pm,[status(thm)],[c_0_11, c_0_14])).
cnf(c_0_17, plain, (divide(divide(X1,inverse(X2)),X2)=X1), inference(pm,[status(thm)],[c_0_7, c_0_11])).
cnf(c_0_18, negated_conjecture, (multiply(a,b)!=multiply(b,a)), prove_these_axioms_4).
cnf(c_0_19, plain, (multiply(X1,X2)=divide(X1,inverse(X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_6]), c_0_8])).
cnf(c_0_20, plain, (divide(X1,divide(X2,inverse(X1)))=inverse(X2)), inference(pm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21, negated_conjecture, (divide(a,inverse(b))!=divide(b,inverse(a))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19]), c_0_19])).
cnf(c_0_22, plain, (divide(X1,inverse(X2))=divide(X2,inverse(X1))), inference(pm,[status(thm)],[c_0_14, c_0_20])).
cnf(c_0_23, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_21, c_0_22])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 24
# Proof object clause steps            : 19
# Proof object formula steps           : 5
# Proof object conjectures             : 4
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 5
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 5
# Processed clauses                    : 30
# ...of these trivial                  : 8
# ...subsumed                          : 1
# ...remaining for further processing  : 21
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 8
# Generated clauses                    : 188
# ...of the previous two non-trivial   : 72
# Contextual simplify-reflections      : 0
# Paramodulations                      : 188
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 13
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 40
# ...number of literals in the above   : 40
# Current number of archived formulas  : 0
# Current number of archived clauses   : 8
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 9
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 39
# BW rewrite match successes           : 36
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1174
