# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_9, plain, (double_divide(X1,double_divide(X1,identity))=identity)).
#
#cnf(i_0_6, plain, (double_divide(double_divide(X1,double_divide(double_divide(identity,X2),double_divide(X3,double_divide(X2,X1)))),double_divide(identity,identity))=X3)).
#
#cnf(i_0_14, plain, (double_divide(double_divide(identity,double_divide(double_divide(identity,X1),identity)),double_divide(identity,identity))=X1)).
#
#cnf(i_0_10, negated_conjecture, (double_divide(identity,identity)!=identity)).
#
# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(identity, axiom, (identity=double_divide(X1,inverse(X1))), file('./p/GRP581-1.tptp', identity)).
cnf(inverse, axiom, (inverse(X1)=double_divide(X1,identity)), file('./p/GRP581-1.tptp', inverse)).
cnf(prove_these_axioms_1, negated_conjecture, (multiply(inverse(a1),a1)!=identity), file('./p/GRP581-1.tptp', prove_these_axioms_1)).
cnf(multiply, axiom, (multiply(X1,X2)=double_divide(double_divide(X2,X1),identity)), file('./p/GRP581-1.tptp', multiply)).
cnf(single_axiom, axiom, (double_divide(double_divide(X1,double_divide(double_divide(identity,X2),double_divide(X3,double_divide(X2,X1)))),double_divide(identity,identity))=X3), file('./p/GRP581-1.tptp', single_axiom)).
cnf(c_0_5, axiom, (identity=double_divide(X1,inverse(X1))), identity).
cnf(c_0_6, axiom, (inverse(X1)=double_divide(X1,identity)), inverse).
cnf(c_0_7, negated_conjecture, (multiply(inverse(a1),a1)!=identity), prove_these_axioms_1).
cnf(c_0_8, axiom, (multiply(X1,X2)=double_divide(double_divide(X2,X1),identity)), multiply).
cnf(c_0_9, axiom, (double_divide(double_divide(X1,double_divide(double_divide(identity,X2),double_divide(X3,double_divide(X2,X1)))),double_divide(identity,identity))=X3), single_axiom).
cnf(c_0_10, plain, (identity=double_divide(X1,double_divide(X1,identity))), inference(rw,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_11, negated_conjecture, (double_divide(double_divide(a1,double_divide(a1,identity)),identity)!=identity), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_6]), c_0_8])).
cnf(c_0_12, plain, (double_divide(double_divide(identity,double_divide(double_divide(identity,X1),identity)),double_divide(identity,identity))=X1), inference(pm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_13, negated_conjecture, (double_divide(identity,identity)!=identity), inference(rw,[status(thm)],[c_0_11, c_0_10])).
cnf(c_0_14, plain, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_12, c_0_10]), c_0_10]), c_0_10]), c_0_13]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 15
# Proof object clause steps            : 10
# Proof object formula steps           : 5
# Proof object conjectures             : 4
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 5
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 3
# Processed clauses                    : 5
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 4
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 8
# ...of the previous two non-trivial   : 8
# Contextual simplify-reflections      : 0
# Paramodulations                      : 8
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 4
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 6
# ...number of literals in the above   : 6
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 237
