# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_9, plain, (divide(X1,X1)=identity)).
#
#cnf(i_0_6, plain, (divide(divide(identity,divide(X1,X2)),divide(divide(X2,X3),X1))=X3)).
#
#cnf(i_0_16, plain, (divide(identity,divide(divide(X1,X2),X1))=X2)).
#
# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(prove_these_axioms_1, negated_conjecture, (multiply(inverse(a1),a1)!=multiply(inverse(b1),b1)), file('./p/GRP545-1.tptp', prove_these_axioms_1)).
cnf(inverse, axiom, (inverse(X1)=divide(identity,X1)), file('./p/GRP545-1.tptp', inverse)).
cnf(multiply, axiom, (multiply(X1,X2)=divide(X1,divide(identity,X2))), file('./p/GRP545-1.tptp', multiply)).
cnf(identity, axiom, (identity=divide(X1,X1)), file('./p/GRP545-1.tptp', identity)).
cnf(c_0_4, negated_conjecture, (multiply(inverse(a1),a1)!=multiply(inverse(b1),b1)), prove_these_axioms_1).
cnf(c_0_5, axiom, (inverse(X1)=divide(identity,X1)), inverse).
cnf(c_0_6, axiom, (multiply(X1,X2)=divide(X1,divide(identity,X2))), multiply).
cnf(c_0_7, negated_conjecture, (divide(divide(identity,b1),divide(identity,b1))!=divide(divide(identity,a1),divide(identity,a1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_4, c_0_5]), c_0_5]), c_0_6]), c_0_6])).
cnf(c_0_8, axiom, (identity=divide(X1,X1)), identity).
cnf(c_0_9, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_8]), c_0_8])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 10
# Proof object clause steps            : 6
# Proof object formula steps           : 4
# Proof object conjectures             : 4
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 4
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 0
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 5
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 3
# Processed clauses                    : 4
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 3
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 12
# ...of the previous two non-trivial   : 12
# Contextual simplify-reflections      : 0
# Paramodulations                      : 12
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 3
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 11
# ...number of literals in the above   : 11
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 195
