# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_4, plain, (double_divide(X1,inverse(double_divide(inverse(double_divide(double_divide(X1,X2),inverse(X3))),X2)))=X3)).
#
#cnf(i_0_7, plain, (double_divide(inverse(double_divide(double_divide(double_divide(X1,X2),X3),inverse(X4))),X3)=double_divide(X1,inverse(double_divide(inverse(X4),X2))))).
#
#cnf(i_0_11, plain, (double_divide(double_divide(X1,X2),inverse(double_divide(X1,inverse(double_divide(inverse(X4),X2)))))=X4)).
#
#cnf(i_0_6, negated_conjecture, (inverse(double_divide(inverse(double_divide(c3,b3)),a3))!=inverse(double_divide(c3,inverse(double_divide(b3,a3)))))).
#
#cnf(i_0_18, plain, (double_divide(double_divide(double_divide(inverse(X1),X2),inverse(double_divide(inverse(X3),X2))),inverse(X3))=X1)).
#
#cnf(i_0_36, plain, (double_divide(X3,inverse(double_divide(double_divide(double_divide(inverse(X1),X2),X2),inverse(X3))))=X1)).
#
#cnf(i_0_44, plain, (double_divide(double_divide(double_divide(inverse(X3),X4),X4),inverse(double_divide(X1,X2)))=double_divide(X1,inverse(double_divide(inverse(X3),X2))))).
#
#cnf(i_0_58, plain, (double_divide(inverse(X1),inverse(double_divide(inverse(X3),inverse(double_divide(double_divide(inverse(X1),X2),X2)))))=X3)).
#
#cnf(i_0_76, plain, (double_divide(double_divide(inverse(X1),inverse(double_divide(double_divide(inverse(X1),X2),X2))),inverse(X3))=X3)).
#
#cnf(i_0_88, plain, (double_divide(inverse(double_divide(X3,inverse(X4))),inverse(X3))=X4)).
#
#cnf(i_0_90, plain, (double_divide(X3,inverse(double_divide(inverse(X4),inverse(X3))))=X4)).
#
#cnf(i_0_147, plain, (double_divide(double_divide(X1,inverse(X1)),inverse(X2))=X2)).
###
#cnf(i_0_160, plain, (double_divide(inverse(X2),inverse(double_divide(X1,inverse(X1))))=X2)).
#
#cnf(i_0_183, plain, (double_divide(double_divide(X1,inverse(X3)),inverse(X3))=X1)).
#
#cnf(i_0_205, plain, (double_divide(double_divide(inverse(X1),X2),X2)=inverse(X1))).
##
#cnf(i_0_211, plain, (inverse(double_divide(X1,inverse(X2)))=double_divide(X2,inverse(X1)))).
#
#cnf(i_0_6, negated_conjecture, (inverse(double_divide(inverse(double_divide(c3,b3)),a3))!=double_divide(double_divide(b3,a3),inverse(c3)))).
#
#cnf(i_0_246, plain, (double_divide(X2,double_divide(X1,inverse(X1)))=inverse(X2))).
#
#cnf(i_0_250, plain, (inverse(inverse(inverse(X1)))=inverse(X1))).
#
#cnf(i_0_256, plain, (inverse(inverse(X1))=X1)).
#
#cnf(i_0_269, plain, (double_divide(double_divide(X1,X2),X2)=X1)).
#
#cnf(i_0_270, plain, (double_divide(double_divide(X1,inverse(X1)),X2)=inverse(X2))).
#
#cnf(i_0_274, plain, (double_divide(X1,double_divide(inverse(X2),X2))=inverse(X1))).
#
#cnf(i_0_276, plain, (double_divide(X2,double_divide(X2,inverse(X1)))=inverse(X1))).
#
#cnf(i_0_305, plain, (double_divide(X1,double_divide(X1,X2))=X2)).
#
#cnf(i_0_313, plain, (double_divide(X2,double_divide(X1,X2))=X1)).
#
#cnf(i_0_317, plain, (double_divide(double_divide(X1,X2),X1)=X2)).
### double_divide is commutative
# AC handling enabled dynamically

#cnf(i_0_324, plain, (double_divide(X1,X2)=double_divide(X2,X1))).
##
#cnf(i_0_273, plain, (inverse(double_divide(X1,X2))=double_divide(inverse(X2),inverse(X1)))).
##
#cnf(i_0_6, negated_conjecture, (double_divide(inverse(a3),double_divide(c3,b3))!=double_divide(inverse(c3),double_divide(b3,a3)))).
##
#cnf(i_0_284, plain, (double_divide(X2,inverse(X2))=double_divide(X1,inverse(X1)))).
###############
#cnf(i_0_11, plain, (double_divide(double_divide(X1,X2),double_divide(double_divide(inverse(X3),X2),inverse(X1)))=X3)).
#
#cnf(i_0_388, plain, (double_divide(X1,double_divide(double_divide(inverse(X1),X2),inverse(X3)))=double_divide(X3,X2))).
####
#cnf(i_0_420, plain, (double_divide(double_divide(inverse(X1),X2),inverse(X3))=double_divide(X1,double_divide(X3,X2)))).
###
#cnf(i_0_451, plain, (double_divide(double_divide(inverse(X1),X2),X3)=double_divide(X1,double_divide(inverse(X3),X2)))).
####
#cnf(i_0_475, plain, (double_divide(double_divide(X1,X2),X3)=double_divide(inverse(X1),double_divide(inverse(X3),X2)))).
####
#cnf(i_0_494, plain, (double_divide(X1,double_divide(inverse(X2),double_divide(inverse(X1),X3)))=double_divide(X2,X3))).
#####
#cnf(i_0_520, plain, (double_divide(X1,double_divide(X2,double_divide(inverse(X1),X3)))=double_divide(inverse(X2),X3))).
#####
#cnf(i_0_527, plain, (double_divide(inverse(X2),double_divide(X3,double_divide(X2,X4)))=double_divide(inverse(X3),X4))).
####
#cnf(i_0_540, plain, (double_divide(X1,double_divide(X2,double_divide(X3,inverse(X1))))=double_divide(inverse(X2),X3))).
############
#cnf(i_0_564, plain, (double_divide(inverse(X1),double_divide(X2,double_divide(X3,X1)))=double_divide(inverse(X2),X3))).
########################
#cnf(i_0_492, plain, (double_divide(inverse(X1),double_divide(inverse(X3),X2))=double_divide(X3,double_divide(X1,X2)))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(single_axiom, axiom, (double_divide(X1,inverse(double_divide(inverse(double_divide(double_divide(X1,X2),inverse(X3))),X2)))=X3), file('./p/GRP587-1.tptp', single_axiom)).
cnf(prove_these_axioms_3, negated_conjecture, (multiply(multiply(a3,b3),c3)!=multiply(a3,multiply(b3,c3))), file('./p/GRP587-1.tptp', prove_these_axioms_3)).
cnf(multiply, axiom, (multiply(X1,X2)=inverse(double_divide(X2,X1))), file('./p/GRP587-1.tptp', multiply)).
cnf(c_0_3, axiom, (double_divide(X1,inverse(double_divide(inverse(double_divide(double_divide(X1,X2),inverse(X3))),X2)))=X3), single_axiom).
cnf(c_0_4, plain, (double_divide(inverse(double_divide(double_divide(double_divide(X1,X2),X3),inverse(X4))),X3)=double_divide(X1,inverse(double_divide(inverse(X4),X2)))), inference(pm,[status(thm)],[c_0_3, c_0_3])).
cnf(c_0_5, plain, (double_divide(double_divide(X1,X2),inverse(double_divide(X1,inverse(double_divide(inverse(X3),X2)))))=X3), inference(pm,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_6, plain, (double_divide(double_divide(double_divide(inverse(X1),X2),inverse(double_divide(inverse(X3),X2))),inverse(X3))=X1), inference(pm,[status(thm)],[c_0_5, c_0_5])).
cnf(c_0_7, plain, (double_divide(X1,inverse(double_divide(double_divide(double_divide(inverse(X2),X3),X3),inverse(X1))))=X2), inference(pm,[status(thm)],[c_0_6, c_0_3])).
cnf(c_0_8, plain, (double_divide(double_divide(double_divide(inverse(X1),X2),X2),inverse(double_divide(X3,X4)))=double_divide(X3,inverse(double_divide(inverse(X1),X4)))), inference(pm,[status(thm)],[c_0_3, c_0_7])).
cnf(c_0_9, plain, (double_divide(inverse(X1),inverse(double_divide(inverse(X2),inverse(double_divide(double_divide(inverse(X1),X3),X3)))))=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_8]), c_0_8])).
cnf(c_0_10, plain, (double_divide(double_divide(inverse(X1),inverse(double_divide(double_divide(inverse(X1),X2),X2))),inverse(X3))=X3), inference(pm,[status(thm)],[c_0_3, c_0_9])).
cnf(c_0_11, plain, (double_divide(X1,inverse(double_divide(inverse(X2),inverse(X1))))=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_5, c_0_10]), c_0_10])).
cnf(c_0_12, plain, (double_divide(double_divide(X1,inverse(X1)),inverse(X2))=X2), inference(pm,[status(thm)],[c_0_3, c_0_11])).
cnf(c_0_13, plain, (double_divide(inverse(X1),inverse(double_divide(X2,inverse(X2))))=X1), inference(pm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_14, plain, (double_divide(double_divide(X1,inverse(X2)),inverse(X2))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_6, c_0_13]), c_0_13])).
cnf(c_0_15, plain, (double_divide(inverse(double_divide(X1,inverse(X2))),inverse(X1))=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_3, c_0_10]), c_0_10])).
cnf(c_0_16, plain, (inverse(double_divide(X1,inverse(X2)))=double_divide(X2,inverse(X1))), inference(pm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_17, plain, (double_divide(double_divide(inverse(X1),X2),X2)=inverse(X1)), inference(pm,[status(thm)],[c_0_10, c_0_14])).
cnf(c_0_18, plain, (double_divide(X1,double_divide(X2,inverse(X2)))=inverse(X1)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_16, c_0_12]), c_0_16])).
cnf(c_0_19, plain, (inverse(inverse(inverse(X1)))=inverse(X1)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_17, c_0_18]), c_0_18])).
cnf(c_0_20, plain, (inverse(inverse(X1))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_6, c_0_19]), c_0_6])).
cnf(c_0_21, plain, (double_divide(double_divide(X1,X2),X2)=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_17, c_0_20]), c_0_20])).
cnf(c_0_22, plain, (double_divide(X1,double_divide(X1,inverse(X2)))=inverse(X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_16, c_0_21]), c_0_16])).
cnf(c_0_23, plain, (double_divide(X1,double_divide(X1,X2))=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_22, c_0_20]), c_0_20])).
cnf(c_0_24, plain, (double_divide(X1,double_divide(X2,X1))=X2), inference(pm,[status(thm)],[c_0_21, c_0_23])).
cnf(c_0_25, plain, (double_divide(double_divide(X1,X2),double_divide(double_divide(inverse(X3),X2),inverse(X1)))=X3), inference(rw,[status(thm)],[c_0_5, c_0_16])).
cnf(c_0_26, plain, (double_divide(X1,X2)=double_divide(X2,X1)), inference(pm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_27, plain, (double_divide(X1,double_divide(double_divide(inverse(X1),X2),inverse(X3)))=double_divide(X3,X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_24, c_0_25]), c_0_26])).
cnf(c_0_28, negated_conjecture, (multiply(multiply(a3,b3),c3)!=multiply(a3,multiply(b3,c3))), prove_these_axioms_3).
cnf(c_0_29, axiom, (multiply(X1,X2)=inverse(double_divide(X2,X1))), multiply).
cnf(c_0_30, plain, (double_divide(double_divide(inverse(X1),X2),inverse(X3))=double_divide(X1,double_divide(X3,X2))), inference(pm,[status(thm)],[c_0_23, c_0_27])).
cnf(c_0_31, negated_conjecture, (inverse(double_divide(c3,inverse(double_divide(b3,a3))))!=inverse(double_divide(inverse(double_divide(c3,b3)),a3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_28, c_0_29]), c_0_29]), c_0_29]), c_0_29])).
cnf(c_0_32, plain, (double_divide(double_divide(inverse(X1),X2),X3)=double_divide(X1,double_divide(inverse(X3),X2))), inference(pm,[status(thm)],[c_0_30, c_0_20])).
cnf(c_0_33, negated_conjecture, (inverse(double_divide(inverse(double_divide(c3,b3)),a3))!=double_divide(double_divide(b3,a3),inverse(c3))), inference(rw,[status(thm)],[c_0_31, c_0_16])).
cnf(c_0_34, plain, (double_divide(double_divide(X1,X2),X3)=double_divide(inverse(X1),double_divide(inverse(X3),X2))), inference(pm,[status(thm)],[c_0_32, c_0_20])).
cnf(c_0_35, negated_conjecture, (double_divide(inverse(a3),double_divide(c3,b3))!=double_divide(inverse(c3),double_divide(b3,a3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_26]), c_0_16]), c_0_26]), c_0_26])).
cnf(c_0_36, plain, (double_divide(inverse(X1),double_divide(inverse(X2),X3))=double_divide(X2,double_divide(X1,X3))), inference(pm,[status(thm)],[c_0_26, c_0_34])).
cnf(c_0_37, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_36]), c_0_20]), c_0_26])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 38
# Proof object clause steps            : 35
# Proof object formula steps           : 3
# Proof object conjectures             : 6
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 3
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 27
# Proof object simplifying inferences  : 25
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 3
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 2
# Processed clauses                    : 119
# ...of these trivial                  : 23
# ...subsumed                          : 54
# ...remaining for further processing  : 42
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 29
# Generated clauses                    : 697
# ...of the previous two non-trivial   : 452
# Contextual simplify-reflections      : 0
# Paramodulations                      : 697
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 13
#    Positive orientable unit clauses  : 10
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 238
# ...number of literals in the above   : 238
# Current number of archived formulas  : 0
# Current number of archived clauses   : 30
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 13
# Rewrite failures with RHS unbound    : 16
# BW rewrite match attempts            : 137
# BW rewrite match successes           : 86
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 7602
