# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_5, plain, (mult(X1,unit)=X1)).
#
#cnf(i_0_6, plain, (mult(unit,X1)=X1)).
#
#cnf(i_0_2, plain, (ld(X1,mult(X1,X2))=X2)).
#
#cnf(i_0_13, negated_conjecture, (mult(mult(mult(a,b),a),c)!=mult(a,mult(b,mult(a,c))))).
#
#cnf(i_0_15, plain, (ld(X1,X1)=unit)).
#
#cnf(i_0_14, plain, (ld(unit,X1)=X1)).
#
#cnf(i_0_1, plain, (mult(X1,ld(X1,X2))=X2)).
#
#cnf(i_0_3, plain, (mult(rd(X1,X2),X2)=X1)).
#
#cnf(i_0_20, plain, (rd(X1,unit)=X1)).
#
#cnf(i_0_4, plain, (rd(mult(X1,X2),X2)=X1)).
#
#cnf(i_0_24, plain, (rd(X1,X1)=unit)).
#
#cnf(i_0_21, plain, (ld(rd(X1,X2),X1)=X2)).
#
#cnf(i_0_26, plain, (rd(X2,ld(X1,X2))=X1)).
#
#cnf(i_0_9, plain, (mult(f(X1),f(X1))=X1)).
#
#cnf(i_0_36, plain, (ld(f(X1),X1)=f(X1))).
#
#cnf(i_0_37, plain, (rd(X1,f(X1))=f(X1))).
##
#cnf(i_0_8, plain, (mult(mult(X1,X2),X1)=mult(X1,mult(X2,X1)))).
#
#cnf(i_0_46, plain, (mult(f(X1),X1)=mult(X1,f(X1)))).
#
#cnf(i_0_13, negated_conjecture, (mult(mult(a,mult(b,a)),c)!=mult(a,mult(b,mult(a,c))))).
#
#cnf(i_0_53, plain, (rd(mult(X1,f(X1)),X1)=f(X1))).
#
#cnf(i_0_52, plain, (ld(f(X1),mult(X1,f(X1)))=X1)).
##
#cnf(i_0_44, plain, (rd(mult(X1,mult(X2,X1)),X1)=mult(X1,X2))).
#
#cnf(i_0_68, plain, (rd(mult(X1,X2),X1)=mult(X1,rd(X2,X1)))).
#
#cnf(i_0_82, plain, (mult(X1,rd(unit,X1))=unit)).
#
#cnf(i_0_85, plain, (rd(unit,X1)=ld(X1,unit))).
#
#cnf(i_0_96, plain, (mult(ld(X1,unit),X1)=unit)).
#
#cnf(i_0_95, plain, (ld(ld(X1,unit),unit)=X1)).
###
#cnf(i_0_53, plain, (mult(X1,rd(f(X1),X1))=f(X1))).
#
#cnf(i_0_116, plain, (rd(f(X1),X1)=ld(X1,f(X1)))).
#
#cnf(i_0_122, plain, (mult(ld(X1,f(X1)),X1)=f(X1))).
#
#cnf(i_0_123, plain, (ld(ld(X1,f(X1)),f(X1))=X1)).
##
#cnf(i_0_50, plain, (mult(X1,mult(ld(X1,X2),X1))=mult(X2,X1))).
#
#cnf(i_0_134, plain, (ld(X1,mult(X2,X1))=mult(ld(X1,X2),X1))).
#
#cnf(i_0_83, plain, (mult(X1,rd(ld(X1,X2),X1))=rd(X2,X1))).
#
#cnf(i_0_166, plain, (rd(ld(X1,X2),X1)=ld(X1,rd(X2,X1)))).
#
#cnf(i_0_157, plain, (mult(ld(X1,rd(X2,X1)),X1)=ld(X1,X2))).
##
#cnf(i_0_186, plain, (rd(X1,ld(X1,unit))=ld(ld(X1,unit),X1))).
#
#cnf(i_0_43, plain, (ld(mult(X1,X2),mult(X1,mult(X2,X1)))=X1)).
#
#cnf(i_0_75, plain, (ld(mult(X1,rd(X2,X1)),mult(X1,X2))=X1)).
#
#cnf(i_0_135, plain, (rd(mult(X2,X1),mult(ld(X1,X2),X1))=X1)).
##
#cnf(i_0_181, plain, (ld(ld(X1,rd(X2,X1)),ld(X1,X2))=X1)).
##
#cnf(i_0_220, plain, (mult(ld(mult(X1,X1),X1),mult(X1,X1))=X1)).
#
#cnf(i_0_331, plain, (rd(X1,mult(X1,X1))=ld(mult(X1,X1),X1))).
#
#cnf(i_0_330, plain, (ld(ld(mult(X1,X1),X1),X1)=mult(X1,X1))).
###
#cnf(i_0_160, plain, (mult(ld(X1,ld(X1,unit)),X1)=ld(X1,unit))).
##
#cnf(i_0_209, plain, (ld(ld(ld(X1,unit),X1),X1)=ld(X1,unit))).
##
#cnf(i_0_185, plain, (rd(X1,ld(X1,f(X1)))=ld(ld(X1,f(X1)),X1))).
#
#cnf(i_0_208, plain, (mult(ld(ld(X1,unit),X1),ld(X1,unit))=X1)).
##
#cnf(i_0_317, plain, (ld(ld(X1,ld(X1,unit)),ld(X1,unit))=X1)).
#######
#cnf(i_0_159, plain, (mult(ld(X1,ld(X1,f(X1))),X1)=ld(X1,f(X1)))).
##
#cnf(i_0_311, plain, (ld(ld(ld(X1,f(X1)),X1),X1)=ld(X1,f(X1)))).
##
#cnf(i_0_332, plain, (mult(ld(mult(X1,X1),X1),X1)=mult(X1,ld(mult(X1,X1),X1)))).
##
#cnf(i_0_48, plain, (mult(rd(X1,X2),mult(X2,rd(X1,X2)))=mult(X1,rd(X1,X2)))).
#
#cnf(i_0_80, plain, (mult(rd(X1,X2),rd(X2,rd(X1,X2)))=rd(X1,rd(X1,X2)))).
##
#cnf(i_0_136, plain, (mult(X1,mult(mult(ld(X1,X2),X1),X1))=mult(mult(X2,X1),X1))).
###
#cnf(i_0_163, plain, (mult(ld(ld(X1,X2),X1),ld(X1,X2))=ld(ld(X1,X2),X2))).
#
#cnf(i_0_184, plain, (ld(rd(X1,X2),rd(X1,rd(X1,X2)))=rd(X2,rd(X1,X2)))).
#
#cnf(i_0_191, plain, (mult(ld(X1,ld(X1,rd(X2,X1))),X1)=ld(X1,ld(X1,X2)))).
###
#cnf(i_0_222, plain, (mult(ld(mult(X1,f(X1)),X1),mult(X1,f(X1)))=X1)).
#
#cnf(i_0_718, plain, (rd(X1,mult(X1,f(X1)))=ld(mult(X1,f(X1)),X1))).
#
#cnf(i_0_717, plain, (ld(ld(mult(X1,f(X1)),X1),X1)=mult(X1,f(X1)))).
###
#cnf(i_0_316, plain, (ld(ld(X1,ld(X1,f(X1))),ld(X1,f(X1)))=X1)).
#
#cnf(i_0_333, plain, (mult(ld(X1,ld(mult(X1,X1),X1)),X1)=ld(mult(X1,X1),X1))).
#
#cnf(i_0_774, plain, (rd(ld(mult(X1,X1),X1),X1)=ld(X1,ld(mult(X1,X1),X1)))).
############
#cnf(i_0_54, plain, (mult(mult(X1,f(X1)),f(X1))=mult(f(X1),mult(X1,f(X1))))).
#
#cnf(i_0_127, plain, (mult(ld(X1,f(X1)),f(X1))=mult(f(X1),ld(X1,f(X1))))).
##
#cnf(i_0_390, plain, (mult(ld(ld(X1,unit),X1),X1)=mult(X1,ld(ld(X1,unit),X1)))).
#
#cnf(i_0_435, plain, (rd(X1,ld(X1,ld(X1,unit)))=ld(ld(X1,ld(X1,unit)),X1))).
#######
#cnf(i_0_800, plain, (rd(ld(X1,f(X1)),f(X1))=ld(f(X1),ld(X1,f(X1))))).
#
#cnf(i_0_231, plain, (ld(mult(ld(X1,X2),X1),mult(ld(X1,X2),X2))=ld(X1,X2))).
####
#cnf(i_0_277, plain, (rd(mult(X1,rd(X1,X2)),mult(X2,rd(X1,X2)))=rd(X1,X2))).
###
#cnf(i_0_325, plain, (ld(ld(ld(X1,X2),X1),ld(ld(X1,X2),X2))=ld(X1,X2))).
##
#cnf(i_0_534, plain, (rd(rd(X1,rd(X1,X2)),rd(X2,rd(X1,X2)))=rd(X1,X2))).
###
#cnf(i_0_366, plain, (mult(ld(X1,ld(X1,ld(X1,unit))),X1)=ld(X1,ld(X1,unit)))).
#
#cnf(i_0_388, plain, (mult(X1,rd(ld(ld(X1,unit),X1),X1))=ld(ld(X1,unit),X1))).
#
#cnf(i_0_1138, plain, (rd(ld(ld(X1,unit),X1),X1)=ld(X1,ld(ld(X1,unit),X1)))).
#
#cnf(i_0_436, plain, (ld(ld(ld(X1,ld(X1,unit)),X1),X1)=ld(X1,ld(X1,unit)))).
####
#cnf(i_0_825, plain, (mult(ld(X1,ld(ld(X1,unit),X1)),X1)=ld(ld(X1,unit),X1))).
#######
#cnf(i_0_225, plain, (ld(mult(X1,mult(X1,X2)),mult(X1,mult(X1,mult(X2,X1))))=X1)).
##
#cnf(i_0_263, plain, (ld(mult(X1,mult(X1,rd(X2,X1))),mult(X1,mult(X1,X2)))=X1)).
#
#cnf(i_0_283, plain, (rd(mult(mult(X1,X2),X2),mult(mult(ld(X2,X1),X2),X2))=X2)).
##
#cnf(i_0_321, plain, (ld(ld(X1,ld(X1,rd(X2,X1))),ld(X1,ld(X1,X2)))=X1)).
###
#cnf(i_0_773, plain, (ld(ld(X1,ld(mult(X1,X1),X1)),ld(mult(X1,X1),X1))=X1)).
##########
#cnf(i_0_1260, plain, (mult(ld(mult(X1,mult(X1,X1)),X1),mult(X1,mult(X1,X1)))=X1)).
#
#cnf(i_0_1559, plain, (rd(X1,mult(X1,mult(X1,X1)))=ld(mult(X1,mult(X1,X1)),X1))).
#
#cnf(i_0_1558, plain, (ld(ld(mult(X1,mult(X1,X1)),X1),X1)=mult(X1,mult(X1,X1)))).
##
#cnf(i_0_1599, plain, (rd(f(X1),mult(X1,f(X1)))=ld(mult(X1,f(X1)),f(X1)))).
##
#cnf(i_0_239, plain, (mult(ld(f(X1),ld(X1,f(X1))),f(X1))=ld(X1,f(X1)))).
##
#cnf(i_0_477, plain, (mult(ld(ld(X1,f(X1)),X1),X1)=mult(X1,ld(ld(X1,f(X1)),X1)))).
###
#cnf(i_0_719, plain, (mult(ld(mult(X1,f(X1)),X1),X1)=mult(X1,ld(mult(X1,f(X1)),X1)))).
####
#cnf(i_0_763, plain, (rd(X1,ld(X1,ld(X1,f(X1))))=ld(ld(X1,ld(X1,f(X1))),X1))).
######
#cnf(i_0_1631, plain, (ld(ld(mult(X1,f(X1)),f(X1)),f(X1))=mult(X1,f(X1)))).
###
#cnf(i_0_235, plain, (mult(ld(mult(X1,f(X1)),f(X1)),mult(X1,f(X1)))=f(X1))).
#
#cnf(i_0_454, plain, (mult(ld(X1,ld(X1,ld(X1,f(X1)))),X1)=ld(X1,ld(X1,f(X1))))).
#
#cnf(i_0_475, plain, (mult(X1,rd(ld(ld(X1,f(X1)),X1),X1))=ld(ld(X1,f(X1)),X1))).
#
#cnf(i_0_1788, plain, (rd(ld(ld(X1,f(X1)),X1),X1)=ld(X1,ld(ld(X1,f(X1)),X1)))).
##
#cnf(i_0_720, plain, (mult(ld(X1,ld(mult(X1,f(X1)),X1)),X1)=ld(mult(X1,f(X1)),X1))).
#
#cnf(i_0_1829, plain, (rd(ld(mult(X1,f(X1)),X1),X1)=ld(X1,ld(mult(X1,f(X1)),X1)))).
####
#cnf(i_0_766, plain, (ld(ld(ld(X1,ld(X1,f(X1))),X1),X1)=ld(X1,ld(X1,f(X1))))).
###
#cnf(i_0_870, plain, (ld(ld(f(X1),ld(X1,f(X1))),ld(X1,f(X1)))=f(X1))).
#################
#cnf(i_0_1675, plain, (mult(ld(X1,ld(ld(X1,f(X1)),X1)),X1)=ld(ld(X1,f(X1)),X1))).
#######
#cnf(i_0_7, plain, (mult(mult(X1,mult(X2,mult(X3,X2))),X3)=mult(mult(X1,X2),mult(X3,mult(X2,X3))))).
#
#cnf(i_0_1934, plain, (mult(mult(X1,X2),X2)=mult(X1,mult(X2,X2)))).
####
#cnf(i_0_2047, plain, (mult(ld(X1,unit),mult(X1,X1))=X1)).
#
#cnf(i_0_2065, plain, (ld(mult(X1,X1),X1)=ld(X1,unit))).
#
#cnf(i_0_331, plain, (rd(X1,mult(X1,X1))=ld(X1,unit))).
#
#cnf(i_0_2064, plain, (ld(ld(X1,unit),X1)=mult(X1,X1))).
####
#cnf(i_0_186, plain, (rd(X1,ld(X1,unit))=mult(X1,X1))).
#
#cnf(i_0_2106, plain, (mult(mult(X1,X1),ld(X1,unit))=X1)).
###
#cnf(i_0_2075, plain, (mult(ld(f(X1),unit),X1)=f(X1))).
#
#cnf(i_0_2236, plain, (mult(X1,ld(f(X1),unit))=f(X1))).
#
#cnf(i_0_2017, plain, (rd(mult(X1,mult(X2,X2)),X2)=mult(X1,X2))).
#
#cnf(i_0_2032, plain, (mult(rd(X1,X2),mult(X2,X2))=mult(X1,X2))).
#
#cnf(i_0_2298, plain, (mult(rd(X1,mult(X2,X2)),X2)=rd(X1,X2))).
#
#cnf(i_0_2365, plain, (rd(rd(X1,X2),X2)=rd(X1,mult(X2,X2)))).
#
#cnf(i_0_2319, plain, (ld(rd(X1,X2),mult(X1,X2))=mult(X2,X2))).
#
#cnf(i_0_2320, plain, (rd(mult(X1,X2),mult(X2,X2))=rd(X1,X2))).
###
#cnf(i_0_2117, plain, (ld(X1,f(X1))=ld(f(X1),unit))).
#
#cnf(i_0_2571, plain, (ld(ld(X1,f(X1)),unit)=f(X1))).
##
#cnf(i_0_2573, plain, (mult(f(X1),ld(X1,f(X1)))=unit)).
####
#cnf(i_0_127, plain, (mult(ld(X1,f(X1)),f(X1))=unit)).
####
#cnf(i_0_2293, plain, (rd(mult(X1,X2),f(X2))=mult(X1,f(X2)))).
#
#cnf(i_0_2755, plain, (rd(X1,f(unit))=mult(X1,f(unit)))).
#
#cnf(i_0_2777, plain, (ld(mult(X1,f(unit)),X1)=f(unit))).
#
#cnf(i_0_2331, plain, (mult(rd(X1,f(X2)),X2)=mult(X1,f(X2)))).
#
#cnf(i_0_2375, plain, (mult(rd(X1,X2),f(X2))=rd(X1,f(X2)))).
#
#cnf(i_0_2532, plain, (rd(mult(X1,f(X2)),X2)=rd(X1,f(X2)))).
####################
#cnf(i_0_2412, plain, (ld(mult(X1,X1),unit)=ld(X1,ld(X1,unit)))).
##
#cnf(i_0_54, plain, (mult(f(X1),mult(X1,f(X1)))=mult(X1,X1))).
#
#cnf(i_0_3012, plain, (ld(f(X1),mult(X1,X1))=mult(X1,f(X1)))).
#
#cnf(i_0_2016, plain, (ld(mult(X1,X2),mult(X1,mult(X2,X2)))=X2)).
##
#cnf(i_0_2364, plain, (ld(rd(X1,mult(X2,X2)),rd(X1,X2))=X2)).
##
#cnf(i_0_2707, plain, (ld(mult(X1,f(X2)),mult(X1,X2))=f(X2))).
#
#cnf(i_0_2885, plain, (ld(rd(X1,X2),rd(X1,f(X2)))=f(X2))).
##
#cnf(i_0_3013, plain, (rd(mult(X1,X1),mult(X1,f(X1)))=f(X1))).
##########
#cnf(i_0_2694, plain, (ld(f(X1),ld(X1,f(X1)))=ld(X1,unit))).
#
#cnf(i_0_239, plain, (mult(ld(X1,unit),f(X1))=ld(X1,f(X1)))).
#
#cnf(i_0_3351, plain, (mult(f(X1),ld(X1,unit))=ld(X1,f(X1)))).
##
#cnf(i_0_800, plain, (rd(ld(X1,f(X1)),f(X1))=ld(X1,unit))).
#
#cnf(i_0_2470, plain, (ld(ld(X1,ld(X1,unit)),unit)=mult(X1,X1))).
#
#cnf(i_0_2722, plain, (rd(X1,f(f(X1)))=mult(f(X1),f(f(X1))))).
#
#cnf(i_0_3515, plain, (mult(ld(X1,f(f(X1))),X1)=f(f(X1)))).
#
#cnf(i_0_3552, plain, (rd(f(f(X1)),X1)=ld(X1,f(f(X1))))).
#
#cnf(i_0_3567, plain, (mult(f(f(X1)),X1)=mult(X1,f(f(X1))))).
######
#cnf(i_0_3043, plain, (ld(f(f(X1)),X1)=mult(f(X1),f(f(X1))))).
##
#cnf(i_0_2028, plain, (mult(f(X1),mult(X1,X1))=mult(X1,mult(X1,f(X1))))).
#
#cnf(i_0_2411, plain, (rd(f(X1),mult(X1,X1))=ld(X1,ld(X1,f(X1))))).
##
#cnf(i_0_2701, plain, (ld(f(X1),mult(X2,X1))=mult(ld(f(X1),X2),X1))).
#
#cnf(i_0_2718, plain, (mult(mult(X1,f(X2)),X2)=mult(mult(X1,X2),f(X2)))).
###
#cnf(i_0_2895, plain, (rd(rd(X1,f(X2)),X2)=rd(rd(X1,X2),f(X2)))).
####
#cnf(i_0_2970, plain, (rd(ld(f(X1),X2),X1)=ld(f(X1),rd(X2,X1)))).
##
#cnf(i_0_3014, plain, (mult(mult(X1,X1),f(X1))=mult(X1,mult(X1,f(X1))))).
#####
#cnf(i_0_870, plain, (ld(ld(X1,unit),ld(X1,f(X1)))=f(X1))).
#
#cnf(i_0_3349, plain, (rd(ld(X1,f(X1)),ld(X1,unit))=f(X1))).
#####
#cnf(i_0_2050, plain, (mult(ld(X1,ld(X1,unit)),mult(X1,X1))=unit)).
#
#cnf(i_0_2969, plain, (ld(f(X1),ld(f(X1),unit))=ld(X1,unit))).
###
#cnf(i_0_2982, plain, (mult(mult(X1,X1),ld(X1,ld(X1,unit)))=unit)).
#########
#cnf(i_0_3551, plain, (ld(ld(X1,f(f(X1))),f(f(X1)))=X1)).
####
#cnf(i_0_2046, plain, (mult(ld(X1,f(X1)),mult(X1,X1))=mult(X1,f(X1)))).
##
#cnf(i_0_2536, plain, (rd(mult(X1,f(X1)),mult(X1,X1))=ld(X1,f(X1)))).
##
#cnf(i_0_3238, plain, (ld(mult(X1,f(ld(X1,X2))),X2)=f(ld(X1,X2)))).
#
#cnf(i_0_3278, plain, (ld(ld(X1,unit),ld(f(X1),unit))=f(X1))).
###
#cnf(i_0_3319, plain, (ld(X2,rd(X1,f(ld(X2,X1))))=f(ld(X2,X1)))).
###
#cnf(i_0_3739, plain, (mult(ld(f(X1),rd(X2,X1)),X1)=ld(f(X1),X2))).
###
#cnf(i_0_4016, plain, (rd(ld(f(X1),unit),ld(X1,unit))=f(X1))).
#####
#cnf(i_0_3199, plain, (ld(mult(f(X1),f(f(X1))),X1)=f(f(X1)))).
##
#cnf(i_0_3346, plain, (rd(X1,mult(f(X1),f(f(X1))))=f(f(X1)))).
############
#cnf(i_0_1991, plain, (mult(mult(X1,mult(X2,X1)),X2)=mult(X1,mult(X2,mult(X1,X2))))).
#
#cnf(i_0_4619, plain, (mult(mult(X1,f(X1)),ld(X1,f(X1)))=X1)).
#
#cnf(i_0_4643, plain, (ld(mult(X1,f(X1)),X1)=ld(X1,f(X1)))).
#
#cnf(i_0_4644, plain, (ld(ld(X1,f(X1)),X1)=mult(X1,f(X1)))).
##
#cnf(i_0_718, plain, (rd(X1,mult(X1,f(X1)))=ld(X1,f(X1)))).
##
#cnf(i_0_185, plain, (rd(X1,ld(X1,f(X1)))=mult(X1,f(X1)))).
#
#cnf(i_0_4740, plain, (ld(ld(f(X1),unit),X1)=mult(X1,f(X1)))).
##
#cnf(i_0_4795, plain, (rd(X1,ld(f(X1),unit))=mult(X1,f(X1)))).
##
#cnf(i_0_222, plain, (mult(ld(X1,f(X1)),mult(X1,f(X1)))=X1)).
####
#cnf(i_0_4615, plain, (mult(mult(X1,f(X1)),ld(f(X1),unit))=X1)).
##
#cnf(i_0_4632, plain, (mult(ld(f(X1),unit),mult(X1,f(X1)))=X1)).
#######
#cnf(i_0_2014, plain, (ld(X1,mult(X2,mult(X1,X1)))=mult(ld(X1,X2),mult(X1,X1)))).
#
#cnf(i_0_2025, plain, (mult(mult(X1,mult(X2,X2)),X2)=mult(mult(X1,X2),mult(X2,X2)))).
#
#cnf(i_0_2034, plain, (mult(mult(X1,X2),mult(X1,X1))=mult(X1,mult(X2,mult(X1,X1))))).
#
#cnf(i_0_5108, plain, (mult(mult(f(X1),X2),X1)=mult(f(X1),mult(X2,X1)))).
#
#cnf(i_0_5178, plain, (rd(mult(f(X1),mult(X2,X1)),X1)=mult(f(X1),X2))).
#
#cnf(i_0_5255, plain, (rd(mult(f(X1),X2),X1)=mult(f(X1),rd(X2,X1)))).
###
#cnf(i_0_2409, plain, (rd(rd(X1,mult(X2,X2)),X2)=rd(rd(X1,X2),mult(X2,X2)))).
#
#cnf(i_0_2418, plain, (rd(mult(X1,X2),mult(X1,X1))=mult(X1,rd(X2,mult(X1,X1))))).
#
#cnf(i_0_2422, plain, (rd(ld(X1,X2),mult(X1,X1))=ld(X1,rd(X2,mult(X1,X1))))).
##
#cnf(i_0_2448, plain, (ld(ld(X1,f(X1)),mult(X1,f(X1)))=mult(X1,X1))).
#
#cnf(i_0_2469, plain, (ld(ld(X1,ld(X1,f(X1))),f(X1))=mult(X1,X1))).
#
#cnf(i_0_2482, plain, (ld(rd(X1,ld(X1,X2)),X2)=mult(ld(X1,X2),ld(X1,X2)))).
##
#cnf(i_0_2506, plain, (ld(X2,mult(X1,ld(X2,X1)))=mult(ld(X2,X1),ld(X2,X1)))).
#####
#cnf(i_0_3353, plain, (rd(ld(X1,unit),f(X1))=ld(f(X1),ld(X1,unit)))).
#
#cnf(i_0_3661, plain, (ld(f(X1),mult(X1,mult(X1,f(X1))))=mult(X1,X1))).
#
#cnf(i_0_3666, plain, (rd(mult(X1,mult(X1,f(X1))),f(X1))=mult(X1,X1))).
#
#cnf(i_0_5867, plain, (ld(mult(X1,X1),f(X1))=ld(X1,ld(X1,f(X1))))).
###
#cnf(i_0_3713, plain, (mult(f(X1),mult(ld(f(X1),X2),X1))=mult(X2,X1))).
########
#cnf(i_0_4079, plain, (rd(f(X1),ld(X1,unit))=ld(ld(X1,unit),f(X1)))).
####
#cnf(i_0_4598, plain, (mult(mult(X1,X2),rd(X2,X1))=mult(X1,mult(X2,rd(X2,X1))))).
#
#cnf(i_0_4600, plain, (mult(ld(X1,unit),mult(X1,mult(X1,X1)))=mult(X1,X1))).
#
#cnf(i_0_6123, plain, (ld(ld(X1,unit),mult(X1,X1))=mult(X1,mult(X1,X1)))).
#
#cnf(i_0_6124, plain, (rd(mult(X1,X1),mult(X1,mult(X1,X1)))=ld(X1,unit))).
####
#cnf(i_0_4612, plain, (mult(mult(X1,mult(X1,X1)),ld(X1,unit))=mult(X1,X1))).
#
#cnf(i_0_6255, plain, (rd(mult(X1,X1),ld(X1,unit))=mult(X1,mult(X1,X1)))).
#
#cnf(i_0_6254, plain, (ld(mult(X1,mult(X1,X1)),mult(X1,X1))=ld(X1,unit))).
####
#cnf(i_0_4633, plain, (mult(mult(ld(X1,X2),X2),X1)=mult(ld(X1,X2),mult(X2,X1)))).
#
#cnf(i_0_4636, plain, (mult(mult(X1,rd(X1,X2)),X2)=mult(rd(X1,X2),mult(X2,X1)))).
#
#cnf(i_0_4637, plain, (mult(mult(X2,X1),ld(X1,X2))=mult(X1,mult(ld(X1,X2),X2)))).
####
#cnf(i_0_5816, plain, (rd(X1,f(ld(X1,unit)))=ld(f(ld(X1,unit)),X1))).
##
#cnf(i_0_5871, plain, (rd(f(X1),ld(X1,ld(X1,f(X1))))=mult(X1,X1))).
###
#cnf(i_0_5979, plain, (rd(f(ld(X1,unit)),X1)=ld(X1,f(ld(X1,unit))))).
#########
#cnf(i_0_2049, plain, (mult(ld(X1,ld(X1,f(X1))),mult(X1,X1))=f(X1))).
#
#cnf(i_0_2242, plain, (mult(ld(X1,ld(f(X1),unit)),X1)=ld(X1,f(X1)))).
#
#cnf(i_0_2253, plain, (mult(ld(f(X1),unit),mult(X1,X1))=mult(X1,f(X1)))).
##
#cnf(i_0_2266, plain, (mult(X1,ld(f(X1),ld(X1,unit)))=ld(X1,f(X1)))).
########
#cnf(i_0_3335, plain, (mult(mult(X1,X1),ld(X1,ld(X1,f(X1))))=f(X1))).
####
#cnf(i_0_3665, plain, (ld(mult(X1,X1),mult(X1,mult(X1,f(X1))))=f(X1))).
##
#cnf(i_0_3711, plain, (rd(mult(X1,X2),mult(ld(f(X2),X1),X2))=f(X2))).
######
#cnf(i_0_5801, plain, (mult(ld(f(X1),ld(X1,unit)),X1)=ld(X1,f(X1)))).
########
#cnf(i_0_2169, plain, (mult(ld(X1,unit),ld(X1,unit))=ld(X1,ld(X1,unit)))).
#####
#cnf(i_0_2677, plain, (mult(ld(X1,f(X1)),ld(X1,f(X1)))=ld(X1,unit))).
#
#cnf(i_0_7349, plain, (ld(ld(X1,f(X1)),ld(X1,unit))=ld(X1,f(X1)))).
##
#cnf(i_0_7354, plain, (rd(ld(X1,unit),ld(X1,f(X1)))=ld(X1,f(X1)))).
#########
#cnf(i_0_5339, plain, (mult(f(X1),ld(X1,ld(X1,f(X1))))=ld(X1,unit))).
#######
#cnf(i_0_136, plain, (mult(X1,mult(ld(X1,X2),mult(X1,X1)))=mult(X2,mult(X1,X1)))).
#
#cnf(i_0_7617, plain, (mult(mult(X1,f(X1)),ld(X1,unit))=f(X1))).
#
#cnf(i_0_7650, plain, (ld(ld(X1,unit),f(X1))=mult(X1,f(X1)))).
#
#cnf(i_0_4079, plain, (rd(f(X1),ld(X1,unit))=mult(X1,f(X1)))).
#
#cnf(i_0_7649, plain, (ld(mult(X1,f(X1)),f(X1))=ld(X1,unit))).
###
#cnf(i_0_1599, plain, (rd(f(X1),mult(X1,f(X1)))=ld(X1,unit))).
#
#cnf(i_0_7696, plain, (mult(ld(X1,unit),mult(X1,f(X1)))=f(X1))).
###
#cnf(i_0_2048, plain, (mult(ld(X1,rd(X2,X1)),mult(X1,X1))=mult(ld(X1,X2),X1))).
#
#cnf(i_0_2060, plain, (mult(X1,mult(ld(X1,X2),ld(X1,X2)))=mult(X2,ld(X1,X2)))).
####
#cnf(i_0_2362, plain, (mult(ld(X1,rd(X2,mult(X1,X1))),X1)=ld(X1,rd(X2,X1)))).
##
#cnf(i_0_2432, plain, (rd(X1,mult(ld(X2,X1),ld(X2,X1)))=rd(X2,ld(X2,X1)))).
#
#cnf(i_0_8260, plain, (ld(ld(X1,ld(X1,unit)),X1)=mult(X1,mult(X1,X1)))).
#
#cnf(i_0_436, plain, (ld(mult(X1,mult(X1,X1)),X1)=ld(X1,ld(X1,unit)))).
###
#cnf(i_0_1559, plain, (rd(X1,mult(X1,mult(X1,X1)))=ld(X1,ld(X1,unit)))).
##
#cnf(i_0_435, plain, (rd(X1,ld(X1,ld(X1,unit)))=mult(X1,mult(X1,X1)))).
#
#cnf(i_0_8257, plain, (rd(mult(X1,f(X1)),ld(X1,f(X1)))=mult(X1,X1))).
#
#cnf(i_0_8610, plain, (mult(mult(X1,X1),ld(X1,f(X1)))=mult(X1,f(X1)))).
#
#cnf(i_0_8611, plain, (ld(mult(X1,X1),mult(X1,f(X1)))=ld(X1,f(X1)))).
###
#cnf(i_0_8659, plain, (mult(mult(X1,X1),ld(f(X1),unit))=mult(X1,f(X1)))).
#####
#cnf(i_0_2466, plain, (ld(ld(f(X1),ld(X1,unit)),f(X1))=mult(X1,X1))).
####
#cnf(i_0_2719, plain, (ld(mult(X1,f(X2)),mult(mult(X1,X2),f(X2)))=X2)).
#
#cnf(i_0_2897, plain, (ld(rd(rd(X1,X2),f(X2)),rd(X1,f(X2)))=X2)).
#########
#cnf(i_0_3887, plain, (ld(ld(f(X1),rd(X2,X1)),ld(f(X1),X2))=X1)).
###
#cnf(i_0_4567, plain, (rd(mult(X1,mult(X2,mult(X1,X2))),X2)=mult(X1,mult(X2,X1)))).
#####
#cnf(i_0_5177, plain, (ld(mult(f(X1),X2),mult(f(X1),mult(X2,X1)))=X1)).
#
#cnf(i_0_5310, plain, (ld(mult(f(X1),rd(X2,X1)),mult(f(X1),X2))=X1)).
###
#cnf(i_0_5587, plain, (ld(ld(f(X1),unit),mult(X1,f(X1)))=mult(X1,X1))).
###
#cnf(i_0_5605, plain, (ld(ld(X1,ld(f(X1),unit)),f(X1))=mult(X1,X1))).
#####
#cnf(i_0_6395, plain, (rd(mult(ld(X1,X2),mult(X2,X1)),X1)=mult(ld(X1,X2),X2))).
#
#cnf(i_0_6494, plain, (rd(mult(rd(X1,X2),mult(X2,X1)),X2)=mult(X1,rd(X1,X2)))).
#
#cnf(i_0_6706, plain, (mult(ld(X1,f(ld(X1,unit))),X1)=f(ld(X1,unit)))).
#
#cnf(i_0_9912, plain, (mult(f(ld(X1,unit)),X1)=mult(X1,f(ld(X1,unit))))).
#
#cnf(i_0_6715, plain, (ld(ld(f(ld(X1,unit)),X1),X1)=f(ld(X1,unit)))).
#
#cnf(i_0_6787, plain, (rd(f(X1),ld(X1,ld(f(X1),unit)))=mult(X1,X1))).
###
#cnf(i_0_7967, plain, (mult(mult(X1,mult(X1,X1)),ld(X1,ld(X1,unit)))=X1)).
###
#cnf(i_0_8396, plain, (mult(ld(X1,ld(X1,unit)),mult(X1,mult(X1,X1)))=X1)).
######
#cnf(i_0_8630, plain, (rd(mult(X1,f(X1)),ld(f(X1),unit))=mult(X1,X1))).
############
#cnf(i_0_2729, plain, (mult(rd(X1,mult(X2,X2)),f(X2))=rd(rd(X1,X2),f(X2)))).
##
#cnf(i_0_10340, plain, (mult(ld(X1,ld(X1,f(X1))),f(X1))=ld(X1,unit))).
#
#cnf(i_0_2733, plain, (rd(mult(X1,mult(X2,X1)),f(X1))=mult(mult(X1,X2),f(X1)))).
##
#cnf(i_0_2734, plain, (rd(mult(X1,mult(X2,X2)),f(X2))=mult(mult(X1,X2),f(X2)))).
###
#cnf(i_0_2747, plain, (mult(ld(X1,rd(X2,X1)),f(X1))=rd(ld(X1,X2),f(X1)))).
##
#cnf(i_0_2758, plain, (rd(X2,f(ld(X1,X2)))=mult(X1,f(ld(X1,X2))))).
##
#cnf(i_0_2852, plain, (mult(ld(X1,rd(X2,f(X1))),X1)=ld(X1,mult(X2,f(X1))))).
#
#cnf(i_0_2863, plain, (mult(rd(X1,f(X2)),mult(X2,X2))=mult(mult(X1,X2),f(X2)))).
#
#cnf(i_0_2864, plain, (rd(mult(X1,f(X2)),mult(X2,X2))=rd(rd(X1,X2),f(X2)))).
####
#cnf(i_0_2911, plain, (mult(mult(X1,rd(X2,X1)),f(X1))=rd(mult(X1,X2),f(X1)))).
######
#cnf(i_0_3458, plain, (ld(ld(X1,ld(X1,f(X1))),ld(X1,unit))=f(X1))).
#######
#cnf(i_0_6880, plain, (mult(ld(X1,ld(f(X1),unit)),mult(X1,X1))=f(X1))).
####
#cnf(i_0_7036, plain, (mult(mult(X1,X1),ld(X1,ld(f(X1),unit)))=f(X1))).
#
#cnf(i_0_7201, plain, (mult(ld(f(X1),ld(X1,unit)),mult(X1,X1))=f(X1))).
#
#cnf(i_0_7462, plain, (ld(ld(f(X1),unit),ld(X1,unit))=ld(X1,f(X1)))).
##
#cnf(i_0_11413, plain, (rd(ld(f(X1),unit),ld(X1,f(X1)))=unit)).
##############
#cnf(i_0_11465, plain, (ld(ld(X1,f(X1)),ld(f(X1),unit))=unit)).
###########
#cnf(i_0_7491, plain, (rd(ld(X1,unit),ld(f(X1),unit))=ld(X1,f(X1)))).
###
#cnf(i_0_11563, plain, (rd(ld(X1,f(X1)),ld(f(X1),unit))=unit)).
############
#cnf(i_0_11597, plain, (ld(ld(f(X1),unit),ld(X1,f(X1)))=unit)).
############
#cnf(i_0_7520, plain, (rd(ld(X1,unit),ld(X1,ld(X1,f(X1))))=f(X1))).
##
#cnf(i_0_9178, plain, (rd(mult(X1,X1),ld(X1,f(X1)))=mult(X1,mult(X1,f(X1))))).
###
#cnf(i_0_9555, plain, (ld(ld(X1,ld(X1,f(X1))),X1)=mult(X1,mult(X1,f(X1))))).
###############
#cnf(i_0_763, plain, (rd(X1,ld(X1,ld(X1,f(X1))))=mult(X1,mult(X1,f(X1))))).
#
#cnf(i_0_766, plain, (ld(mult(X1,mult(X1,f(X1))),X1)=ld(X1,ld(X1,f(X1))))).
#
#cnf(i_0_11818, plain, (rd(X1,mult(X1,mult(X1,f(X1))))=ld(X1,ld(X1,f(X1))))).
#####
#cnf(i_0_2108, plain, (ld(mult(X1,X1),ld(X1,unit))=ld(X1,ld(X1,ld(X1,unit))))).
#
#cnf(i_0_12056, plain, (ld(f(X1),ld(X1,unit))=ld(X1,ld(f(X1),unit)))).
##
#cnf(i_0_3353, plain, (rd(ld(X1,unit),f(X1))=ld(X1,ld(f(X1),unit)))).
##########
#cnf(i_0_3365, plain, (mult(ld(X1,ld(f(X1),unit)),f(X1))=ld(X1,unit))).
#
#cnf(i_0_3410, plain, (mult(f(X1),ld(X1,ld(f(X1),unit)))=ld(X1,unit))).
##
#cnf(i_0_3562, plain, (rd(f(f(X1)),mult(X1,X1))=ld(X1,ld(X1,f(f(X1)))))).
#
#cnf(i_0_3589, plain, (rd(X1,ld(X1,f(f(X1))))=ld(ld(X1,f(f(X1))),X1))).
###
#cnf(i_0_3620, plain, (mult(f(f(X1)),mult(X1,X1))=mult(X1,mult(X1,f(f(X1)))))).
############
#cnf(i_0_4380, plain, (rd(ld(f(f(X1)),X2),X1)=ld(f(f(X1)),rd(X2,X1)))).
#
#cnf(i_0_4415, plain, (ld(f(f(X1)),mult(X2,X1))=mult(ld(f(f(X1)),X2),X1))).
######
#cnf(i_0_7238, plain, (ld(ld(X1,unit),ld(X1,ld(X1,unit)))=ld(X1,unit))).
#####
#cnf(i_0_7243, plain, (rd(ld(X1,ld(X1,unit)),ld(X1,unit))=ld(X1,unit))).
#######
#cnf(i_0_7306, plain, (mult(ld(f(X1),unit),ld(X1,f(X1)))=ld(X1,unit))).
######
#cnf(i_0_7328, plain, (mult(ld(X1,f(X1)),ld(f(X1),unit))=ld(X1,unit))).
#########################
#cnf(i_0_12445, plain, (ld(f(f(X1)),ld(X1,f(f(X1))))=ld(X1,unit))).
#
#cnf(i_0_12819, plain, (mult(f(f(X1)),ld(X1,unit))=ld(X1,f(f(X1))))).
#
#cnf(i_0_12452, plain, (ld(f(f(X1)),ld(X1,f(X1)))=ld(X1,f(f(X1))))).
###
#cnf(i_0_12813, plain, (rd(ld(X1,f(f(X1))),ld(X1,unit))=f(f(X1)))).
######
#cnf(i_0_2457, plain, (ld(mult(X1,rd(X2,X1)),mult(X1,mult(X2,X1)))=mult(X1,X1))).
#
#cnf(i_0_2471, plain, (ld(ld(X1,rd(X2,mult(X1,X1))),ld(X1,X2))=mult(X1,X1))).
###
#cnf(i_0_2496, plain, (ld(ld(X1,rd(X2,X1)),mult(ld(X1,X2),X1))=mult(X1,X1))).
###
#cnf(i_0_2606, plain, (ld(ld(X1,ld(X1,f(X1))),ld(f(X1),unit))=X1)).
##
#cnf(i_0_2620, plain, (ld(ld(X1,ld(f(X1),unit)),ld(X1,f(X1)))=X1)).
#######
#cnf(i_0_3926, plain, (mult(mult(X1,mult(X1,f(X1))),f(X1))=mult(X1,mult(X1,X1)))).
#
#cnf(i_0_13615, plain, (ld(f(X1),mult(X1,mult(X1,X1)))=mult(X1,mult(X1,f(X1))))).
#
#cnf(i_0_3959, plain, (ld(X1,ld(f(ld(X1,unit)),unit))=f(ld(X1,unit)))).
##
#cnf(i_0_4001, plain, (rd(ld(f(ld(X1,unit)),unit),X1)=f(ld(X1,unit)))).
########
#cnf(i_0_4634, plain, (mult(f(X1),mult(X1,mult(X1,f(X1))))=mult(X1,mult(X1,X1)))).
#
#cnf(i_0_4968, plain, (mult(ld(X1,rd(X2,mult(X1,X1))),mult(X1,X1))=ld(X1,X2))).
#
#cnf(i_0_5057, plain, (ld(mult(X1,X2),mult(X1,mult(X2,mult(X1,X1))))=mult(X1,X1))).
####
#cnf(i_0_5405, plain, (ld(mult(X1,rd(X2,mult(X1,X1))),mult(X1,X2))=mult(X1,X1))).
############
#cnf(i_0_6001, plain, (ld(mult(X1,X2),mult(X1,mult(X2,rd(X2,X1))))=rd(X2,X1))).
#
#cnf(i_0_6002, plain, (rd(mult(X1,mult(X2,rd(X2,X1))),rd(X2,X1))=mult(X1,X2))).
##
#cnf(i_0_6567, plain, (ld(mult(X1,X2),mult(X2,mult(ld(X2,X1),X1)))=ld(X2,X1))).
#
#cnf(i_0_6568, plain, (rd(mult(X2,mult(ld(X2,X1),X1)),ld(X2,X1))=mult(X1,X2))).
####
#cnf(i_0_6953, plain, (mult(ld(f(f(X1)),unit),X1)=mult(f(X1),f(f(X1))))).
#
#cnf(i_0_7710, plain, (mult(mult(X1,f(X1)),mult(X1,f(X1)))=mult(X1,mult(X1,X1)))).
#######
#cnf(i_0_8727, plain, (mult(X1,ld(f(f(X1)),unit))=mult(f(X1),f(f(X1))))).
############################
#cnf(i_0_2302, plain, (mult(ld(f(mult(X1,X1)),unit),X1)=rd(f(mult(X1,X1)),X1))).
###
#cnf(i_0_2736, plain, (mult(ld(X1,f(X1)),f(f(X1)))=ld(f(f(X1)),unit))).
##
#cnf(i_0_2740, plain, (mult(ld(X1,unit),f(mult(X1,X1)))=rd(X1,f(mult(X1,X1))))).
##
#cnf(i_0_3356, plain, (ld(ld(X1,ld(f(X1),unit)),ld(X1,unit))=f(X1))).
#####
#cnf(i_0_3613, plain, (mult(f(f(X1)),mult(X1,f(f(X1))))=mult(X1,f(X1)))).
#######
#cnf(i_0_3989, plain, (rd(ld(X1,unit),ld(X1,ld(f(X1),unit)))=f(X1))).
#########
#cnf(i_0_5328, plain, (mult(f(mult(X1,X1)),ld(X1,unit))=rd(f(mult(X1,X1)),X1))).
#
#cnf(i_0_5340, plain, (mult(f(X1),ld(X1,ld(X1,unit)))=ld(X1,ld(X1,f(X1))))).
#
#cnf(i_0_5426, plain, (mult(X1,ld(f(mult(X1,X1)),unit))=rd(X1,f(mult(X1,X1))))).
#
#cnf(i_0_15854, plain, (ld(X1,rd(X1,f(mult(X1,X1))))=ld(f(mult(X1,X1)),unit))).
##########
#cnf(i_0_8436, plain, (mult(ld(X1,f(X1)),ld(X1,unit))=ld(X1,ld(X1,f(X1))))).
######
#cnf(i_0_9177, plain, (rd(mult(X1,X1),ld(f(X1),unit))=mult(X1,mult(X1,f(X1))))).
#######################
#cnf(i_0_11741, plain, (ld(mult(X1,mult(X1,f(X1))),mult(X1,X1))=ld(X1,f(X1)))).
##
#cnf(i_0_11848, plain, (ld(ld(X1,ld(f(X1),unit)),X1)=mult(X1,mult(X1,f(X1))))).
#
#cnf(i_0_11887, plain, (rd(X1,ld(X1,ld(f(X1),unit)))=mult(X1,mult(X1,f(X1))))).
##############
#cnf(i_0_12453, plain, (mult(f(X1),ld(X1,f(f(X1))))=ld(f(f(X1)),unit))).
#
#cnf(i_0_12913, plain, (mult(f(f(X1)),ld(X1,f(f(X1))))=ld(X1,f(X1)))).
#
#cnf(i_0_16285, plain, (rd(f(X1),ld(X1,f(f(X1))))=mult(X1,f(f(X1))))).
#
#cnf(i_0_16293, plain, (rd(ld(X1,f(X1)),f(f(X1)))=ld(X1,f(f(X1))))).
#
#cnf(i_0_16361, plain, (mult(ld(X1,unit),f(f(X1)))=ld(X1,f(f(X1))))).
#
#cnf(i_0_16339, plain, (ld(mult(X1,f(f(X1))),f(X1))=ld(X1,f(f(X1))))).
#
#cnf(i_0_16362, plain, (ld(ld(X1,unit),ld(X1,f(f(X1))))=f(f(X1)))).
#
#cnf(i_0_16392, plain, (rd(ld(X1,f(f(X1))),f(f(X1)))=ld(X1,unit))).
##############
#cnf(i_0_15190, plain, (ld(mult(X1,f(X1)),mult(X1,mult(X1,X1)))=mult(X1,f(X1)))).
#
#cnf(i_0_15198, plain, (rd(mult(X1,mult(X1,X1)),mult(X1,f(X1)))=mult(X1,f(X1)))).
#########
#cnf(i_0_15308, plain, (ld(X1,mult(f(X1),f(f(X1))))=ld(f(f(X1)),unit))).
#######
#cnf(i_0_16366, plain, (mult(ld(X1,f(f(X1))),f(f(X1)))=ld(X1,f(X1)))).
##
#cnf(i_0_16390, plain, (mult(ld(X1,f(f(X1))),f(X1))=ld(f(f(X1)),unit))).
####################
#cnf(i_0_2965, plain, (ld(f(X1),ld(f(f(X1)),unit))=ld(X1,f(f(X1))))).
####
#cnf(i_0_3549, plain, (mult(ld(X1,ld(X1,f(f(X1)))),X1)=ld(X1,f(f(X1))))).
##
#cnf(i_0_3560, plain, (mult(ld(X1,f(f(X1))),mult(X1,X1))=mult(X1,f(f(X1))))).
######
#cnf(i_0_3868, plain, (rd(ld(X1,ld(X1,f(X1))),f(X1))=ld(X1,ld(X1,unit)))).
#
#cnf(i_0_17012, plain, (mult(ld(X1,ld(X1,unit)),f(X1))=ld(X1,ld(X1,f(X1))))).
#
#cnf(i_0_3910, plain, (ld(f(X1),ld(X1,ld(X1,f(X1))))=ld(X1,ld(X1,unit)))).
#######
#cnf(i_0_4136, plain, (ld(ld(ld(X1,f(f(X1))),X1),X1)=ld(X1,f(f(X1))))).
######
#cnf(i_0_7757, plain, (ld(mult(X1,f(X1)),ld(X1,f(X1)))=ld(X1,ld(X1,unit)))).
##############
#cnf(i_0_12418, plain, (mult(ld(f(f(X1)),rd(X2,X1)),X1)=ld(f(f(X1)),X2))).
####
#cnf(i_0_12907, plain, (rd(ld(X1,f(X1)),ld(X1,f(f(X1))))=f(f(X1)))).
#
#cnf(i_0_17443, plain, (mult(f(f(X1)),ld(X1,f(X1)))=ld(f(f(X1)),unit))).
#####
#cnf(i_0_12916, plain, (ld(ld(X1,f(f(X1))),ld(X1,f(X1)))=f(f(X1)))).
#######
#cnf(i_0_15592, plain, (rd(mult(X1,f(X1)),mult(X1,f(f(X1))))=f(f(X1)))).
#####################
#cnf(i_0_283, plain, (rd(mult(X1,mult(X2,X2)),mult(ld(X2,X1),mult(X2,X2)))=X2)).
####
#cnf(i_0_2013, plain, (ld(mult(X1,mult(X2,X1)),mult(X1,mult(X2,mult(X1,X1))))=X1)).
#########
#cnf(i_0_2361, plain, (ld(mult(X1,rd(X2,mult(X1,X1))),mult(X1,rd(X2,X1)))=X1)).
#####
#cnf(i_0_2392, plain, (ld(ld(X1,rd(X2,mult(X1,X1))),ld(X1,rd(X2,X1)))=X1)).
###
#cnf(i_0_2731, plain, (mult(mult(X1,X1),f(ld(X1,unit)))=ld(f(ld(X1,unit)),X1))).
#####
#cnf(i_0_2866, plain, (ld(rd(rd(X1,X2),f(X2)),mult(X1,f(X2)))=mult(X2,X2))).
#
#cnf(i_0_2868, plain, (mult(ld(f(X1),f(f(X1))),X1)=mult(f(X1),f(f(X1))))).
##
#cnf(i_0_2905, plain, (rd(mult(X1,X2),f(f(X2)))=mult(mult(X1,f(X2)),f(f(X2))))).
##
#cnf(i_0_2929, plain, (mult(rd(X1,f(f(X2))),X2)=mult(mult(X1,f(X2)),f(f(X2))))).
#
#cnf(i_0_18881, plain, (mult(mult(X1,f(f(X2))),X2)=mult(mult(X1,X2),f(f(X2))))).
##
#cnf(i_0_18963, plain, (mult(mult(X1,X1),f(f(X1)))=mult(X1,mult(X1,f(f(X1)))))).
#
#cnf(i_0_18933, plain, (rd(mult(mult(X1,X2),f(f(X2))),X2)=mult(X1,f(f(X2))))).
#
#cnf(i_0_19163, plain, (rd(mult(X1,f(f(X2))),X2)=mult(rd(X1,X2),f(f(X2))))).
#
#cnf(i_0_19330, plain, (rd(rd(X1,f(f(X2))),X2)=rd(rd(X1,X2),f(f(X2))))).
##
#cnf(i_0_2951, plain, (ld(rd(X1,f(X2)),mult(mult(X1,X2),f(X2)))=mult(X2,X2))).
#######
#cnf(i_0_3102, plain, (ld(mult(X1,mult(X2,X2)),mult(mult(X1,X2),mult(X2,X2)))=X2)).
##########
#cnf(i_0_3144, plain, (ld(rd(rd(X1,X2),mult(X2,X2)),rd(X1,mult(X2,X2)))=X2)).
############
#cnf(i_0_4566, plain, (ld(mult(X1,mult(X2,X1)),mult(X1,mult(X2,mult(X1,X2))))=X2)).
#########
#cnf(i_0_4870, plain, (mult(mult(X1,mult(X1,f(X1))),ld(X1,f(X1)))=mult(X1,X1))).
#######
#cnf(i_0_5680, plain, (ld(ld(X1,ld(f(X1),unit)),ld(f(X1),unit))=X1)).
#####
#cnf(i_0_6058, plain, (mult(ld(X1,f(X1)),mult(X1,mult(X1,f(X1))))=mult(X1,X1))).
#
#cnf(i_0_20451, plain, (ld(ld(X1,f(X1)),mult(X1,X1))=mult(X1,mult(X1,f(X1))))).
#
#cnf(i_0_20452, plain, (rd(mult(X1,X1),mult(X1,mult(X1,f(X1))))=ld(X1,f(X1)))).
##
#cnf(i_0_20527, plain, (ld(ld(f(X1),unit),mult(X1,X1))=mult(X1,mult(X1,f(X1))))).
###
#cnf(i_0_6394, plain, (ld(mult(ld(X1,X2),X2),mult(ld(X1,X2),mult(X2,X1)))=X1)).
##############
#cnf(i_0_6493, plain, (ld(mult(X1,rd(X1,X2)),mult(rd(X1,X2),mult(X2,X1)))=X2)).
##############
#cnf(i_0_6714, plain, (mult(ld(f(ld(X1,unit)),X1),f(ld(X1,unit)))=X1)).
##
#cnf(i_0_6806, plain, (ld(ld(X1,f(ld(X1,unit))),f(ld(X1,unit)))=X1)).
#####
#cnf(i_0_7716, plain, (ld(ld(X1,f(X1)),f(f(X1)))=mult(f(X1),f(f(X1))))).
####
#cnf(i_0_7765, plain, (rd(f(f(X1)),ld(X1,f(X1)))=mult(f(X1),f(f(X1))))).
########
#cnf(i_0_8660, plain, (mult(X1,ld(f(X1),f(f(X1))))=mult(f(X1),f(f(X1))))).
##########
#cnf(i_0_9197, plain, (mult(ld(X1,unit),ld(X1,f(X1)))=ld(X1,ld(f(X1),unit)))).
###########
#cnf(i_0_9972, plain, (mult(f(ld(X1,unit)),mult(X1,X1))=ld(f(ld(X1,unit)),X1))).
##################
#cnf(i_0_10934, plain, (ld(f(X1),mult(X1,f(f(X1))))=mult(f(X1),f(f(X1))))).
#####################################
#cnf(i_0_21821, plain, (rd(ld(f(ld(X1,unit)),X1),X1)=mult(X1,f(ld(X1,unit))))).
#
#cnf(i_0_21929, plain, (ld(X1,ld(f(ld(X1,unit)),X1))=mult(X1,f(ld(X1,unit))))).
#
#cnf(i_0_21949, plain, (mult(X1,mult(X1,f(ld(X1,unit))))=ld(f(ld(X1,unit)),X1))).
###
#cnf(i_0_2728, plain, (rd(mult(X1,X2),f(mult(X2,X2)))=mult(rd(X1,X2),f(mult(X2,X2))))).
####
#cnf(i_0_2759, plain, (ld(f(unit),f(f(unit)))=mult(f(unit),f(f(unit))))).
#
#cnf(i_0_2850, plain, (rd(mult(X1,f(mult(X2,X2))),X2)=mult(rd(X1,f(mult(X2,X2))),X2))).
####
#cnf(i_0_3213, plain, (ld(mult(mult(X1,X2),f(X1)),mult(X1,mult(X2,X1)))=f(X1))).
#
#cnf(i_0_3214, plain, (ld(mult(mult(X1,X2),f(X2)),mult(X1,mult(X2,X2)))=f(X2))).
###
#cnf(i_0_3296, plain, (ld(rd(X1,mult(X2,X2)),rd(rd(X1,X2),f(X2)))=f(X2))).
##
#cnf(i_0_3304, plain, (ld(mult(X1,rd(X2,X1)),rd(mult(X1,X2),f(X1)))=f(X1))).
#
#cnf(i_0_3308, plain, (ld(ld(X1,rd(X2,X1)),rd(ld(X1,X2),f(X1)))=f(X1))).
#
#cnf(i_0_3314, plain, (ld(ld(X1,unit),rd(X1,f(mult(X1,X1))))=f(mult(X1,X1)))).
#########
#cnf(i_0_3742, plain, (ld(f(X1),mult(X1,mult(X2,X1)))=mult(ld(f(X1),mult(X1,X2)),X1))).
##
#cnf(i_0_3743, plain, (ld(f(X1),mult(X2,mult(X1,X1)))=mult(ld(f(X1),X2),mult(X1,X1)))).
#
#cnf(i_0_3778, plain, (ld(X1,mult(mult(X2,X1),f(X1)))=mult(ld(X1,mult(X2,f(X1))),X1))).
###
#cnf(i_0_3790, plain, (mult(mult(X1,mult(X2,X2)),f(X2))=mult(mult(X1,f(X2)),mult(X2,X2)))).
#
#cnf(i_0_3852, plain, (rd(rd(X1,mult(X2,X2)),f(X2))=rd(rd(X1,f(X2)),mult(X2,X2)))).
###
#cnf(i_0_3900, plain, (rd(ld(f(X1),X2),mult(X1,X1))=ld(f(X1),rd(X2,mult(X1,X1))))).
##
#cnf(i_0_3927, plain, (ld(mult(X1,mult(X1,f(X1))),mult(X1,mult(X1,X1)))=f(X1))).
############################
#cnf(i_0_5113, plain, (mult(f(X1),mult(X1,mult(X1,X1)))=mult(X1,mult(X1,mult(X1,f(X1)))))).
#
#cnf(i_0_5173, plain, (ld(X1,mult(f(X1),mult(X2,X1)))=mult(ld(X1,mult(f(X1),X2)),X1))).
###
#cnf(i_0_5186, plain, (mult(mult(f(X1),X2),mult(X1,X1))=mult(f(X1),mult(X2,mult(X1,X1))))).
#
#cnf(i_0_25020, plain, (mult(mult(f(f(X1)),X2),X1)=mult(f(f(X1)),mult(X2,X1)))).
#
#cnf(i_0_25104, plain, (rd(mult(f(f(X1)),mult(X2,X1)),X1)=mult(f(f(X1)),X2))).
#
#cnf(i_0_25285, plain, (rd(mult(f(f(X1)),X2),X1)=mult(f(f(X1)),rd(X2,X1)))).
####
#cnf(i_0_5323, plain, (rd(mult(f(X1),X2),mult(X1,X1))=mult(f(X1),rd(X2,mult(X1,X1))))).
###################
#cnf(i_0_7658, plain, (mult(mult(X1,f(X1)),ld(X1,ld(X1,unit)))=ld(X1,f(X1)))).
################
#cnf(i_0_8234, plain, (rd(ld(X1,f(X1)),ld(X1,ld(X1,unit)))=mult(X1,f(X1)))).
####
#cnf(i_0_8612, plain, (mult(mult(X1,X1),mult(X1,f(X1)))=mult(X1,mult(X1,mult(X1,f(X1)))))).
##
#cnf(i_0_8616, plain, (ld(mult(X1,X1),ld(X1,f(X1)))=ld(X1,ld(X1,ld(X1,f(X1)))))).
#############################################################
#cnf(i_0_11752, plain, (mult(mult(X1,mult(X1,f(X1))),ld(X1,unit))=mult(X1,f(X1)))).
#
#cnf(i_0_25902, plain, (rd(mult(X1,f(X1)),ld(X1,unit))=mult(X1,mult(X1,f(X1))))).
#
#cnf(i_0_26000, plain, (ld(mult(X1,mult(X1,f(X1))),f(X1))=ld(X1,ld(X1,unit)))).
#
#cnf(i_0_26040, plain, (rd(f(X1),ld(X1,ld(X1,unit)))=mult(X1,mult(X1,f(X1))))).
#####################
#cnf(i_0_13631, plain, (rd(mult(X1,mult(X1,X1)),mult(X1,mult(X1,f(X1))))=f(X1))).
#
#cnf(i_0_26157, plain, (ld(ld(X1,ld(X1,unit)),f(X1))=mult(X1,mult(X1,f(X1))))).
#
#cnf(i_0_26214, plain, (rd(f(X1),mult(X1,mult(X1,f(X1))))=ld(X1,ld(X1,unit)))).
########
#cnf(i_0_15443, plain, (rd(ld(X1,unit),f(f(X1)))=ld(f(f(X1)),ld(X1,unit)))).
###
#cnf(i_0_15709, plain, (rd(rd(f(mult(X1,X1)),X1),ld(X1,unit))=f(mult(X1,X1)))).
#
#cnf(i_0_15856, plain, (mult(rd(X1,f(mult(X1,X1))),X1)=mult(X1,rd(f(mult(X1,X1)),X1)))).
####
#cnf(i_0_15936, plain, (ld(ld(X1,unit),ld(X1,ld(X1,f(X1))))=ld(X1,f(X1)))).
##
#cnf(i_0_15947, plain, (rd(ld(X1,ld(X1,f(X1))),ld(X1,unit))=ld(X1,f(X1)))).
###############
#cnf(i_0_16308, plain, (mult(ld(X1,f(f(X1))),mult(X1,f(f(X1))))=f(X1))).
#
#cnf(i_0_26747, plain, (ld(ld(X1,f(f(X1))),f(X1))=mult(X1,f(f(X1))))).
#
#cnf(i_0_26748, plain, (rd(f(X1),mult(X1,f(f(X1))))=ld(X1,f(f(X1))))).
##
#cnf(i_0_16338, plain, (mult(mult(X1,f(f(X1))),ld(X1,f(f(X1))))=f(X1))).
##
#cnf(i_0_16517, plain, (rd(f(f(X1)),ld(X1,unit))=ld(ld(X1,unit),f(f(X1))))).
####
#cnf(i_0_16716, plain, (ld(f(f(X1)),ld(X1,unit))=ld(X1,ld(f(f(X1)),unit)))).
########################################################
#cnf(i_0_26981, plain, (rd(f(f(ld(X1,unit))),X1)=ld(X1,f(f(ld(X1,unit)))))).
##
#cnf(i_0_15443, plain, (rd(ld(X1,unit),f(f(X1)))=ld(X1,ld(f(f(X1)),unit)))).
#
#cnf(i_0_2053, plain, (mult(ld(X1,ld(X1,ld(X1,unit))),mult(X1,X1))=ld(X1,unit))).
##########
#cnf(i_0_2189, plain, (mult(mult(X1,X1),ld(X1,ld(X1,ld(X1,unit))))=ld(X1,unit))).
#####
#cnf(i_0_2427, plain, (rd(X1,mult(X1,mult(X1,mult(X1,X1))))=ld(X1,ld(X1,ld(X1,unit))))).
###
#cnf(i_0_2896, plain, (rd(rd(X1,f(X2)),f(f(X2)))=mult(rd(X1,X2),f(f(X2))))).
###
#cnf(i_0_2968, plain, (ld(f(X1),ld(f(X1),f(f(X1))))=ld(X1,f(f(X1))))).
######
#cnf(i_0_3533, plain, (mult(f(X1),mult(f(X1),f(f(X1))))=mult(X1,f(f(X1))))).
#
#cnf(i_0_27727, plain, (ld(mult(X1,f(f(X1))),X1)=ld(f(X1),f(f(X1))))).
#
#cnf(i_0_27736, plain, (rd(X1,mult(X1,f(f(X1))))=ld(f(X1),f(f(X1))))).
##
#cnf(i_0_27745, plain, (ld(mult(X1,f(f(X1))),f(f(X1)))=ld(X1,unit))).
#
#cnf(i_0_27933, plain, (ld(ld(X1,unit),f(f(X1)))=mult(X1,f(f(X1))))).
#
#cnf(i_0_16517, plain, (rd(f(f(X1)),ld(X1,unit))=mult(X1,f(f(X1))))).
#
#cnf(i_0_27746, plain, (rd(f(f(X1)),mult(X1,f(f(X1))))=ld(X1,unit))).
#
#cnf(i_0_27845, plain, (rd(X1,ld(f(X1),f(f(X1))))=mult(X1,f(f(X1))))).
#
#cnf(i_0_28209, plain, (rd(X1,ld(f(f(X1)),unit))=mult(X1,f(f(X1))))).
#
#cnf(i_0_27907, plain, (ld(ld(f(X1),f(f(X1))),X1)=mult(X1,f(f(X1))))).
#
#cnf(i_0_28392, plain, (ld(ld(f(f(X1)),unit),X1)=mult(X1,f(f(X1))))).
#
#cnf(i_0_27941, plain, (mult(mult(X1,f(f(X1))),ld(X1,unit))=f(f(X1)))).
##
#cnf(i_0_27978, plain, (mult(ld(X1,unit),mult(X1,f(f(X1))))=f(f(X1)))).
#####
#cnf(i_0_28235, plain, (mult(mult(X1,f(f(X1))),ld(f(f(X1)),unit))=X1)).
##
#cnf(i_0_28493, plain, (mult(ld(f(f(X1)),unit),mult(X1,f(f(X1))))=X1)).
########
#cnf(i_0_4082, plain, (ld(f(X1),ld(X1,ld(f(X1),unit)))=ld(X1,ld(X1,unit)))).
##
#cnf(i_0_4202, plain, (mult(f(ld(X1,unit)),mult(X1,f(ld(X1,unit))))=unit)).
##
#cnf(i_0_4298, plain, (ld(ld(X1,f(X1)),ld(f(f(X1)),unit))=f(f(X1)))).
#####
#cnf(i_0_4471, plain, (rd(ld(f(f(X1)),unit),ld(X1,f(X1)))=f(f(X1)))).
############
#cnf(i_0_7925, plain, (mult(ld(X1,ld(X1,f(f(X1)))),mult(X1,X1))=f(f(X1)))).
##############
#cnf(i_0_11753, plain, (ld(mult(X1,mult(X1,f(X1))),mult(X1,f(X1)))=ld(X1,unit))).
##
#cnf(i_0_12144, plain, (rd(ld(X1,ld(f(X1),unit)),f(X1))=ld(X1,ld(X1,unit)))).
#####
#cnf(i_0_12454, plain, (rd(mult(X1,X2),mult(ld(f(f(X2)),X1),X2))=f(f(X2)))).
################
#cnf(i_0_15708, plain, (ld(f(mult(X1,X1)),rd(f(mult(X1,X1)),X1))=ld(X1,unit))).
###
#cnf(i_0_15946, plain, (ld(ld(X1,f(X1)),ld(X1,ld(X1,f(X1))))=ld(X1,unit))).
#
#cnf(i_0_15951, plain, (rd(ld(X1,ld(X1,f(X1))),ld(X1,f(X1)))=ld(X1,unit))).
#########
#cnf(i_0_16867, plain, (ld(ld(X1,f(f(X1))),ld(f(X1),unit))=f(f(X1)))).
####
#cnf(i_0_16869, plain, (rd(ld(f(X1),unit),ld(X1,f(f(X1))))=f(f(X1)))).
############
#cnf(i_0_17284, plain, (ld(mult(X1,f(X1)),ld(f(X1),unit))=ld(X1,ld(X1,unit)))).
##############
#cnf(i_0_19029, plain, (ld(mult(X1,X1),mult(X1,mult(X1,f(f(X1)))))=f(f(X1)))).
#
#cnf(i_0_29802, plain, (ld(mult(X1,X1),f(f(X1)))=ld(X1,ld(X1,f(f(X1)))))).
#
#cnf(i_0_19034, plain, (mult(mult(X1,X1),ld(X1,ld(X1,f(f(X1)))))=f(f(X1)))).
############################
#cnf(i_0_2111, plain, (ld(ld(X1,ld(X1,ld(X1,unit))),ld(X1,unit))=mult(X1,X1))).
#####
#cnf(i_0_2195, plain, (rd(ld(X1,unit),ld(X1,ld(X1,ld(X1,unit))))=mult(X1,X1))).
#
#cnf(i_0_30124, plain, (rd(ld(X1,f(X1)),ld(X1,ld(f(X1),unit)))=X1)).
########
#cnf(i_0_2271, plain, (mult(mult(X1,ld(f(X2),unit)),mult(X2,f(X2)))=mult(X1,X2))).
#
#cnf(i_0_30228, plain, (rd(mult(X1,X2),mult(X2,f(X2)))=mult(X1,ld(f(X2),unit)))).
##
#cnf(i_0_30282, plain, (mult(rd(X1,ld(f(X2),unit)),X2)=mult(X1,mult(X2,f(X2))))).
#
#cnf(i_0_30616, plain, (mult(rd(X1,ld(X2,f(X2))),X2)=mult(X1,mult(X2,f(X2))))).
#
#cnf(i_0_30707, plain, (rd(mult(X1,mult(X2,f(X2))),X2)=rd(X1,ld(X2,f(X2))))).
#
#cnf(i_0_30778, plain, (mult(ld(X1,ld(X1,f(X1))),mult(X1,f(X1)))=unit)).
#
#cnf(i_0_30936, plain, (ld(ld(X1,ld(X1,f(X1))),unit)=mult(X1,f(X1)))).
#
#cnf(i_0_31120, plain, (ld(ld(X1,ld(f(X1),unit)),unit)=mult(X1,f(X1)))).
#
#cnf(i_0_31123, plain, (ld(mult(X1,f(X1)),unit)=ld(X1,ld(f(X1),unit)))).
##
#cnf(i_0_31017, plain, (mult(mult(X1,f(X1)),ld(X1,ld(X1,f(X1))))=unit)).
##
#cnf(i_0_30741, plain, (mult(rd(X1,ld(X2,f(X2))),ld(f(X2),unit))=X1)).
###
#cnf(i_0_30892, plain, (rd(mult(X1,ld(f(X2),unit)),ld(X2,f(X2)))=X1)).
###
#cnf(i_0_31121, plain, (ld(ld(X1,f(f(X1))),unit)=mult(f(X1),f(f(X1))))).
##
#cnf(i_0_30419, plain, (mult(rd(X1,X2),ld(f(X2),unit))=rd(X1,mult(X2,f(X2))))).
##
#cnf(i_0_31888, plain, (mult(rd(X1,X2),ld(X2,f(X2)))=rd(X1,mult(X2,f(X2))))).
#
#cnf(i_0_30861, plain, (rd(X1,ld(f(mult(X1,X1)),unit))=mult(X1,f(mult(X1,X1))))).
#
#cnf(i_0_32223, plain, (ld(mult(X1,f(mult(X1,X1))),X1)=ld(f(mult(X1,X1)),unit))).
#################
#cnf(i_0_32033, plain, (ld(rd(X1,X2),rd(X1,mult(X2,f(X2))))=ld(X2,f(X2)))).
#####
#cnf(i_0_30977, plain, (mult(ld(X1,ld(f(X1),unit)),mult(X1,f(X1)))=unit)).
##
#cnf(i_0_31160, plain, (mult(mult(X1,f(X1)),ld(X1,ld(f(X1),unit)))=unit)).
####
#cnf(i_0_30270, plain, (mult(mult(X1,ld(X2,f(X2))),mult(X2,f(X2)))=mult(X1,X2))).
#
#cnf(i_0_32740, plain, (ld(mult(X1,ld(X2,f(X2))),mult(X1,X2))=mult(X2,f(X2)))).
####
#cnf(i_0_30879, plain, (rd(rd(X1,mult(X2,f(X2))),ld(X2,f(X2)))=rd(X1,X2))).
#####
#cnf(i_0_30227, plain, (ld(mult(X1,ld(f(X2),unit)),mult(X1,X2))=mult(X2,f(X2)))).
###
#cnf(i_0_2464, plain, (ld(ld(X1,ld(X1,ld(X1,unit))),X1)=mult(X1,mult(X1,mult(X1,X1))))).
##
#cnf(i_0_2501, plain, (ld(ld(X1,unit),mult(X1,mult(X1,X1)))=mult(X1,mult(X1,mult(X1,X1))))).
#
#cnf(i_0_2851, plain, (ld(mult(X1,rd(X2,f(X1))),mult(X1,mult(X2,f(X1))))=X1)).
###
#cnf(i_0_2934, plain, (ld(ld(X1,rd(X2,f(X1))),ld(X1,mult(X2,f(X1))))=X1)).
#
#cnf(i_0_3001, plain, (ld(ld(X1,ld(X1,unit)),mult(X1,X1))=mult(X1,mult(X1,mult(X1,X1))))).
#
#cnf(i_0_3004, plain, (rd(mult(X1,X1),ld(X1,ld(X1,unit)))=mult(X1,mult(X1,mult(X1,X1))))).
####
#cnf(i_0_3316, plain, (ld(mult(X1,X1),ld(f(ld(X1,unit)),X1))=f(ld(X1,unit)))).
#
#cnf(i_0_3564, plain, (ld(ld(X1,ld(X1,f(f(X1)))),f(f(X1)))=mult(X1,X1))).
#
#cnf(i_0_3596, plain, (ld(ld(X1,f(f(X1))),mult(X1,f(f(X1))))=mult(X1,X1))).
####
#cnf(i_0_4919, plain, (mult(mult(X1,mult(X1,f(X1))),ld(f(X1),unit))=mult(X1,X1))).
#############
#cnf(i_0_6057, plain, (mult(ld(f(X1),unit),mult(X1,mult(X1,f(X1))))=mult(X1,X1))).
##############
#cnf(i_0_7538, plain, (mult(ld(X1,ld(X1,f(X1))),mult(X1,mult(X1,f(X1))))=X1)).
#
#cnf(i_0_7719, plain, (mult(ld(X1,unit),f(ld(X1,unit)))=ld(X1,f(ld(X1,unit))))).
#######
#cnf(i_0_9627, plain, (mult(ld(X1,unit),ld(f(X1),unit))=ld(X1,ld(f(X1),unit)))).
########
#cnf(i_0_11845, plain, (mult(mult(X1,mult(X1,f(X1))),ld(X1,ld(X1,f(X1))))=X1)).
#######
#cnf(i_0_12354, plain, (ld(f(f(X1)),mult(X1,mult(X1,f(f(X1)))))=mult(X1,X1))).
##
#cnf(i_0_12460, plain, (mult(f(f(X1)),mult(ld(f(f(X1)),X2),X1))=mult(X2,X1))).
##############
#cnf(i_0_14576, plain, (mult(X1,ld(f(ld(X1,unit)),unit))=ld(f(ld(X1,unit)),X1))).
######
#cnf(i_0_15006, plain, (mult(ld(f(X1),unit),ld(X1,unit))=ld(X1,ld(f(X1),unit)))).
###
#cnf(i_0_15029, plain, (mult(ld(f(ld(X1,unit)),unit),X1)=ld(f(ld(X1,unit)),X1))).
##############
#cnf(i_0_19030, plain, (rd(mult(X1,mult(X1,f(f(X1)))),f(f(X1)))=mult(X1,X1))).
########
#cnf(i_0_21789, plain, (rd(ld(f(ld(X1,unit)),X1),mult(X1,X1))=f(ld(X1,unit)))).
#########################
#cnf(i_0_27704, plain, (ld(mult(X1,f(f(X1))),mult(X1,X1))=mult(f(X1),f(f(X1))))).
#
#cnf(i_0_27723, plain, (mult(mult(X1,f(f(X1))),ld(f(X1),f(f(X1))))=X1)).
##
#cnf(i_0_27739, plain, (mult(ld(f(X1),f(f(X1))),mult(X1,f(f(X1))))=X1)).
#############
#cnf(i_0_29837, plain, (rd(f(f(X1)),ld(X1,ld(X1,f(f(X1)))))=mult(X1,X1))).
######
#cnf(i_0_30706, plain, (ld(rd(X1,ld(X2,f(X2))),mult(X1,mult(X2,f(X2))))=X2)).
###
#cnf(i_0_31848, plain, (rd(rd(X1,mult(X2,f(X2))),ld(f(X2),unit))=rd(X1,X2))).
#########
#cnf(i_0_35079, plain, (rd(X1,ld(f(ld(X1,unit)),unit))=mult(X1,f(ld(X1,unit))))).
#
#cnf(i_0_36018, plain, (ld(mult(X1,f(ld(X1,unit))),X1)=ld(f(ld(X1,unit)),unit))).
###
#cnf(i_0_35108, plain, (ld(ld(f(ld(X1,unit)),X1),unit)=ld(X1,f(ld(X1,unit))))).
####
#cnf(i_0_36269, plain, (ld(ld(X1,f(ld(X1,unit))),unit)=ld(f(ld(X1,unit)),X1))).
#####
#cnf(i_0_2916, plain, (mult(mult(X1,f(unit)),f(f(unit)))=rd(X1,f(f(unit))))).
####
#cnf(i_0_3389, plain, (mult(X1,f(ld(X1,unit)))=ld(f(ld(X1,unit)),unit))).
###
#cnf(i_0_37009, plain, (rd(X1,mult(X1,f(ld(X1,unit))))=mult(X1,f(ld(X1,unit))))).
#######
#cnf(i_0_3738, plain, (mult(ld(f(X1),rd(X2,mult(X1,X1))),X1)=ld(f(X1),rd(X2,X1)))).
##
#cnf(i_0_3757, plain, (mult(ld(f(X1),ld(X1,rd(X2,X1))),X1)=ld(f(X1),ld(X1,X2)))).
##
#cnf(i_0_3768, plain, (mult(ld(f(ld(X1,X2)),X1),ld(X1,X2))=ld(f(ld(X1,X2)),X2))).
####
#cnf(i_0_3881, plain, (mult(ld(X1,ld(f(X1),rd(X2,X1))),X1)=ld(X1,ld(f(X1),X2)))).
#
#cnf(i_0_3898, plain, (mult(ld(f(X1),rd(X2,X1)),mult(X1,X1))=mult(ld(f(X1),X2),X1))).
####
#cnf(i_0_4321, plain, (ld(X1,rd(rd(X2,f(ld(X1,X2))),X1))=rd(f(ld(X1,X2)),X1))).
#####
#cnf(i_0_4424, plain, (mult(ld(f(X1),mult(X1,rd(X2,X1))),X1)=ld(f(X1),mult(X1,X2)))).
####
#cnf(i_0_5303, plain, (mult(ld(X1,mult(f(X1),rd(X2,X1))),X1)=ld(X1,mult(f(X1),X2)))).
##
#cnf(i_0_8233, plain, (rd(ld(f(X1),unit),ld(X1,ld(X1,unit)))=mult(X1,f(X1)))).
######################
#cnf(i_0_12559, plain, (mult(ld(X1,ld(f(f(X1)),unit)),X1)=ld(f(f(X1)),unit))).
#
#cnf(i_0_12927, plain, (rd(ld(X1,f(f(X1))),f(X1))=ld(X1,ld(f(f(X1)),unit)))).
####
#cnf(i_0_15782, plain, (rd(ld(X1,ld(X1,f(X1))),ld(X1,ld(X1,unit)))=f(X1))).
##
#cnf(i_0_16245, plain, (rd(ld(f(f(X1)),unit),ld(X1,f(f(X1))))=f(X1))).
#
#cnf(i_0_16391, plain, (ld(ld(X1,f(f(X1))),ld(f(f(X1)),unit))=f(X1))).
#####
#cnf(i_0_39612, plain, (ld(mult(X1,X1),mult(X1,f(f(X1))))=ld(X1,f(f(X1))))).
#
#cnf(i_0_39646, plain, (mult(mult(X1,X1),ld(X1,f(f(X1))))=mult(X1,f(f(X1))))).
#
#cnf(i_0_39638, plain, (rd(mult(X1,f(f(X1))),ld(X1,f(f(X1))))=mult(X1,X1))).
###############
#cnf(i_0_17013, plain, (ld(ld(X1,ld(X1,unit)),ld(X1,ld(X1,f(X1))))=f(X1))).
###
#cnf(i_0_17269, plain, (mult(ld(X1,f(X1)),mult(X1,mult(X1,X1)))=mult(X1,mult(X1,f(X1))))).
#########
#cnf(i_0_21714, plain, (ld(ld(X1,unit),ld(X1,ld(f(X1),unit)))=ld(X1,f(X1)))).
#
#cnf(i_0_21719, plain, (rd(ld(X1,ld(f(X1),unit)),ld(X1,unit))=ld(X1,f(X1)))).
####################
#cnf(i_0_25910, plain, (mult(mult(X1,mult(X1,f(X1))),ld(X1,ld(X1,unit)))=f(X1))).
####
#cnf(i_0_26156, plain, (mult(ld(X1,ld(X1,unit)),mult(X1,mult(X1,f(X1))))=f(X1))).
######
#cnf(i_0_27117, plain, (mult(ld(X1,f(f(ld(X1,unit)))),X1)=f(f(ld(X1,unit))))).
#
#cnf(i_0_40450, plain, (mult(f(f(ld(X1,unit))),X1)=mult(X1,f(f(ld(X1,unit)))))).
##################################
#cnf(i_0_39569, plain, (ld(f(X1),ld(X1,f(f(X1))))=ld(X1,ld(f(f(X1)),unit)))).
########
#cnf(i_0_2877, plain, (mult(ld(f(X1),rd(X2,f(X1))),X1)=mult(ld(f(X1),X2),f(X1)))).
#
#cnf(i_0_2883, plain, (mult(ld(f(X1),rd(X2,X1)),f(X1))=ld(f(X1),rd(X2,f(X1))))).
#########
#cnf(i_0_3957, plain, (ld(ld(X1,f(X1)),ld(f(X1),f(f(X1))))=f(f(X1)))).
#####
#cnf(i_0_3999, plain, (rd(ld(f(X1),f(f(X1))),ld(X1,f(X1)))=f(f(X1)))).
###########
#cnf(i_0_5079, plain, (ld(mult(X1,mult(X1,mult(X1,X1))),mult(X1,X1))=ld(X1,ld(X1,unit)))).
#
#cnf(i_0_5085, plain, (rd(mult(X1,X1),mult(X1,mult(X1,mult(X1,X1))))=ld(X1,ld(X1,unit)))).
####################
#cnf(i_0_7444, plain, (ld(ld(f(ld(X1,unit)),unit),X1)=ld(f(ld(X1,unit)),unit))).
##########
#cnf(i_0_7672, plain, (mult(mult(f(X1),f(f(X1))),ld(X1,f(X1)))=f(f(X1)))).
##
#cnf(i_0_7864, plain, (mult(ld(X1,f(X1)),mult(f(X1),f(f(X1))))=f(f(X1)))).
########
#cnf(i_0_10829, plain, (ld(f(ld(X1,f(X1))),unit)=mult(f(X1),f(ld(f(X1),unit))))).
#########################################
#cnf(i_0_21700, plain, (ld(ld(X1,f(X1)),ld(X1,ld(f(X1),unit)))=ld(X1,unit))).
##
#cnf(i_0_21715, plain, (rd(ld(X1,ld(f(X1),unit)),ld(X1,f(X1)))=ld(X1,unit))).
##############
#cnf(i_0_29599, plain, (ld(ld(f(X1),unit),ld(X1,ld(X1,f(X1))))=ld(X1,unit))).
####
#cnf(i_0_29651, plain, (rd(ld(X1,ld(X1,f(X1))),ld(f(X1),unit))=ld(X1,unit))).
#################
#cnf(i_0_30978, plain, (mult(ld(X1,f(f(X1))),mult(f(X1),f(f(X1))))=unit)).
#
#cnf(i_0_31464, plain, (mult(mult(f(X1),f(f(X1))),ld(X1,f(f(X1))))=unit)).
##############
#cnf(i_0_34994, plain, (rd(mult(X1,X1),mult(f(X1),f(f(X1))))=mult(X1,f(f(X1))))).
################
#cnf(i_0_45, plain, (mult(mult(X1,mult(X2,X1)),mult(X1,X2))=mult(mult(X1,X2),mult(X1,mult(X1,X2))))).
#####
#cnf(i_0_81, plain, (rd(mult(X1,mult(X2,X1)),mult(X1,X2))=mult(mult(X1,X2),rd(X1,mult(X1,X2))))).
#################
#cnf(i_0_194, plain, (mult(ld(X1,rd(X2,X1)),rd(X2,X1))=mult(ld(X1,X2),ld(X1,rd(X2,X1))))).
###############
#cnf(i_0_1019, plain, (ld(ld(X1,ld(X1,ld(X1,unit))),ld(X1,ld(X1,unit)))=X1)).
#
#cnf(i_0_43592, plain, (ld(mult(X1,mult(X1,mult(X1,X1))),X1)=ld(X1,ld(X1,ld(X1,unit))))).
###
#cnf(i_0_43589, plain, (rd(X1,ld(X1,ld(X1,ld(X1,unit))))=mult(X1,mult(X1,mult(X1,X1))))).
#####
#cnf(i_0_1950, plain, (mult(mult(X1,X2),mult(X1,mult(X2,X1)))=mult(X1,mult(X2,mult(X1,mult(X2,X1)))))).
#
#cnf(i_0_44239, plain, (mult(X1,mult(ld(X1,rd(X2,X1)),X2))=mult(rd(X2,X1),X2))).
###
#cnf(i_0_44499, plain, (ld(X1,mult(rd(X2,X1),X2))=mult(ld(X1,rd(X2,X1)),X2))).
#
#cnf(i_0_44838, plain, (ld(X1,mult(X2,mult(X2,X1)))=mult(ld(X1,X2),mult(X2,X1)))).
#
#cnf(i_0_44908, plain, (ld(ld(X1,X2),mult(X1,X2))=mult(ld(ld(X1,X2),X1),X2))).
##
#cnf(i_0_44270, plain, (mult(X1,mult(ld(X1,X2),mult(X2,X1)))=mult(X2,mult(X2,X1)))).
##
#cnf(i_0_45643, plain, (rd(mult(X2,mult(X2,X1)),X1)=mult(X1,mult(ld(X1,X2),X2)))).
#####
#cnf(i_0_44216, plain, (mult(ld(X1,X2),mult(ld(ld(X1,X2),X1),X2))=mult(X1,X2))).
##
#cnf(i_0_45176, plain, (rd(mult(X1,X2),mult(ld(ld(X1,X2),X1),X2))=ld(X1,X2))).
#################
#cnf(i_0_44500, plain, (rd(mult(rd(X2,X1),X2),mult(ld(X1,rd(X2,X1)),X2))=X1)).
#################
#cnf(i_0_44935, plain, (rd(mult(X1,mult(X1,X2)),mult(ld(X2,X1),mult(X1,X2)))=X2)).
###############
#cnf(i_0_45642, plain, (ld(mult(X1,mult(ld(X1,X2),X2)),mult(X2,mult(X2,X1)))=X1)).
############
#cnf(i_0_3596, plain, (mult(ld(ld(X1,f(f(X1))),X1),f(f(X1)))=mult(X1,X1))).
#
#cnf(i_0_48832, plain, (ld(ld(X1,f(f(X1))),X1)=mult(mult(X1,f(X1)),f(f(X1))))).
#
#cnf(i_0_3589, plain, (rd(X1,ld(X1,f(f(X1))))=mult(mult(X1,f(X1)),f(f(X1))))).
####
#cnf(i_0_4136, plain, (ld(mult(mult(X1,f(X1)),f(f(X1))),X1)=ld(X1,f(f(X1))))).
#
#cnf(i_0_49022, plain, (rd(X1,mult(mult(X1,f(X1)),f(f(X1))))=ld(X1,f(f(X1))))).
###
#cnf(i_0_1956, plain, (mult(mult(X1,mult(X2,X3)),rd(X3,X2))=mult(mult(X1,X2),mult(X3,rd(X3,X2))))).
#
#cnf(i_0_49471, plain, (mult(mult(X1,ld(X2,f(X2))),X2)=mult(X1,f(X2)))).
#
#cnf(i_0_49987, plain, (mult(X1,ld(X2,f(X2)))=rd(X1,f(X2)))).
#
#cnf(i_0_50144, plain, (mult(X1,ld(X2,unit))=rd(X1,X2))).
#
#cnf(i_0_50515, plain, (ld(X1,rd(X1,X2))=ld(X2,unit))).
#
#cnf(i_0_50889, plain, (ld(mult(X1,X2),X1)=ld(X2,unit))).
##
#cnf(i_0_50695, plain, (rd(X1,ld(X2,unit))=mult(X1,X2))).
##
#cnf(i_0_50960, plain, (ld(ld(X2,X1),unit)=ld(X1,X2))).
###
#cnf(i_0_51429, plain, (mult(ld(X1,X2),ld(X2,X1))=unit)).
#####
#cnf(i_0_51562, plain, (rd(X1,ld(X2,X3))=mult(X1,ld(X3,X2)))).
#####################
#cnf(i_0_52314, plain, (mult(mult(X1,ld(X2,X3)),ld(X3,X2))=X1)).
#######
#cnf(i_0_31888, plain, (rd(rd(X1,X2),f(X2))=rd(X1,mult(X2,f(X2))))).
####
#cnf(i_0_30282, plain, (mult(mult(X1,X2),f(X2))=mult(X1,mult(X2,f(X2))))).
##
#cnf(i_0_2895, plain, (rd(rd(X1,f(X2)),X2)=rd(X1,mult(X2,f(X2))))).
##
#cnf(i_0_2718, plain, (mult(mult(X1,f(X2)),X2)=mult(X1,mult(X2,f(X2))))).
#
#cnf(i_0_50759, plain, (ld(X1,rd(rd(X1,X2),X1))=rd(ld(X2,unit),X1))).
#
#cnf(i_0_50867, plain, (ld(mult(X1,X2),rd(X1,X2))=ld(X2,ld(X2,unit)))).
######
#cnf(i_0_50437, plain, (mult(X1,ld(X2,ld(X2,unit)))=rd(X1,mult(X2,X2)))).
#
#cnf(i_0_50517, plain, (mult(X1,mult(ld(X2,unit),X1))=mult(rd(X1,X2),X1))).
#
#cnf(i_0_50521, plain, (mult(X1,rd(ld(X2,unit),X1))=rd(rd(X1,X2),X1))).
#######
#cnf(i_0_50860, plain, (ld(mult(ld(X2,X1),X2),unit)=ld(mult(X1,X2),X2))).
#################
#cnf(i_0_32740, plain, (ld(rd(X1,f(X2)),mult(X1,X2))=mult(X2,f(X2)))).
#
#cnf(i_0_30228, plain, (rd(mult(X1,X2),mult(X2,f(X2)))=rd(X1,f(X2)))).
###
#cnf(i_0_30707, plain, (rd(mult(X1,mult(X2,f(X2))),X2)=mult(X1,f(X2)))).
###
#cnf(i_0_53104, plain, (mult(rd(X1,mult(X2,f(X2))),X2)=rd(X1,f(X2)))).
###
#cnf(i_0_53322, plain, (mult(rd(X1,X2),mult(X2,f(X2)))=mult(X1,f(X2)))).
####
#cnf(i_0_55221, plain, (ld(rd(X1,X2),mult(X1,f(X2)))=mult(X2,f(X2)))).
###
#cnf(i_0_50842, plain, (ld(rd(X1,X2),rd(X1,mult(X2,X2)))=ld(X2,unit))).
##
#cnf(i_0_50885, plain, (ld(mult(X1,mult(X2,X2)),mult(X1,X2))=ld(X2,unit))).
#####
#cnf(i_0_50890, plain, (ld(mult(X1,X2),mult(X1,rd(X2,X1)))=ld(X1,unit))).
##########
#cnf(i_0_50931, plain, (ld(ld(X1,X2),ld(X1,rd(X2,X1)))=ld(X1,unit))).
######
#cnf(i_0_51198, plain, (ld(mult(X1,mult(X2,X1)),mult(X1,X2))=ld(X1,unit))).
#####################
#cnf(i_0_30270, plain, (mult(rd(X1,f(X2)),mult(X2,f(X2)))=mult(X1,X2))).
##
#cnf(i_0_50847, plain, (ld(mult(X1,X2),mult(X1,f(X2)))=ld(f(X2),unit))).
###################
#cnf(i_0_51147, plain, (ld(rd(X1,f(X2)),rd(X1,X2))=ld(f(X2),unit))).
########
#cnf(i_0_31848, plain, (mult(rd(X1,mult(X2,f(X2))),f(X2))=rd(X1,X2))).
##
#cnf(i_0_51431, plain, (ld(ld(X1,X2),ld(X2,X1))=mult(ld(X2,X1),ld(X2,X1)))).
#####
#cnf(i_0_51748, plain, (ld(mult(X2,ld(X1,X2)),X1)=mult(ld(X2,X1),ld(X2,X1)))).
####
#cnf(i_0_52555, plain, (rd(X1,mult(ld(X2,X3),X2))=mult(X1,ld(mult(X3,X2),X2)))).
####
#cnf(i_0_53177, plain, (rd(mult(X1,f(X2)),mult(X2,f(X2)))=rd(X1,X2))).
###
#cnf(i_0_53243, plain, (rd(mult(X1,mult(X2,f(X2))),f(X2))=mult(X1,X2))).
################################
#cnf(i_0_51609, plain, (ld(rd(X1,mult(X2,f(X2))),rd(X1,X2))=f(X2))).
###
#cnf(i_0_53242, plain, (ld(mult(X1,X2),mult(X1,mult(X2,f(X2))))=f(X2))).
#############
#cnf(i_0_52115, plain, (mult(ld(mult(X1,X2),X2),mult(ld(X2,X1),X2))=unit)).
#######
#cnf(i_0_52290, plain, (mult(mult(ld(X1,X2),X1),ld(mult(X2,X1),X1))=unit)).
#########
#cnf(i_0_2432, plain, (rd(X1,mult(ld(X2,X1),ld(X2,X1)))=mult(X2,ld(X1,X2)))).
#######
#cnf(i_0_30706, plain, (ld(mult(X1,f(X2)),mult(X1,mult(X2,f(X2))))=X2)).
########################
#cnf(i_0_53054, plain, (mult(ld(ld(X1,X2),X2),ld(X2,X1))=ld(ld(X1,X2),X1))).
#######
#cnf(i_0_2897, plain, (ld(rd(X1,mult(X2,f(X2))),rd(X1,f(X2)))=X2)).
######
#cnf(i_0_53763, plain, (rd(mult(rd(X1,X2),X1),mult(ld(X2,unit),X1))=X1)).
####
#cnf(i_0_54722, plain, (rd(rd(rd(X1,X2),X1),rd(ld(X2,unit),X1))=X1)).
######
#cnf(i_0_50224, plain, (mult(X1,mult(ld(X2,f(X2)),X1))=mult(rd(X1,f(X2)),X1))).
####
#cnf(i_0_50228, plain, (mult(X1,rd(ld(X2,f(X2)),X1))=rd(rd(X1,f(X2)),X1))).
#######
#cnf(i_0_52536, plain, (mult(X1,ld(X2,ld(X2,f(X2))))=rd(X1,mult(X2,f(X2))))).
##
#cnf(i_0_2864, plain, (rd(mult(X1,f(X2)),mult(X2,X2))=rd(X1,mult(X2,f(X2))))).
##
#cnf(i_0_63602, plain, (rd(rd(X1,X2),mult(X2,X2))=rd(X1,mult(X2,mult(X2,X2))))).
##
#cnf(i_0_2409, plain, (rd(rd(X1,mult(X2,X2)),X2)=rd(X1,mult(X2,mult(X2,X2))))).
#
#cnf(i_0_63759, plain, (mult(mult(X1,X2),mult(X2,X2))=mult(X1,mult(X2,mult(X2,X2))))).
##
#cnf(i_0_2025, plain, (mult(mult(X1,mult(X2,X2)),X2)=mult(X1,mult(X2,mult(X2,X2))))).
#
#cnf(i_0_63739, plain, (mult(rd(X1,mult(X2,mult(X2,X2))),X2)=rd(X1,mult(X2,X2)))).
#
#cnf(i_0_63818, plain, (rd(mult(X1,X2),mult(X2,mult(X2,X2)))=rd(X1,mult(X2,X2)))).
####
#cnf(i_0_64192, plain, (rd(mult(X1,mult(X2,mult(X2,X2))),X2)=mult(X1,mult(X2,X2)))).
#
#cnf(i_0_64240, plain, (mult(rd(X1,X2),mult(X2,mult(X2,X2)))=mult(X1,mult(X2,X2)))).
######
#cnf(i_0_64772, plain, (ld(rd(X1,mult(X2,X2)),mult(X1,X2))=mult(X2,mult(X2,X2)))).
####
#cnf(i_0_65207, plain, (ld(rd(X1,X2),mult(X1,mult(X2,X2)))=mult(X2,mult(X2,X2)))).
###
#cnf(i_0_2729, plain, (mult(rd(X1,mult(X2,X2)),f(X2))=rd(X1,mult(X2,f(X2))))).
####
#cnf(i_0_2863, plain, (mult(rd(X1,f(X2)),mult(X2,X2))=mult(X1,mult(X2,f(X2))))).
#
#cnf(i_0_2734, plain, (rd(mult(X1,mult(X2,X2)),f(X2))=mult(X1,mult(X2,f(X2))))).
#
#cnf(i_0_53804, plain, (ld(rd(rd(X1,X2),X1),X1)=ld(rd(ld(X2,unit),X1),unit))).
####
#cnf(i_0_54571, plain, (ld(mult(rd(X1,X2),X1),X1)=ld(mult(ld(X2,unit),X1),unit))).
##########
#cnf(i_0_50436, plain, (rd(mult(ld(X1,unit),X2),X1)=mult(ld(X1,unit),rd(X2,X1)))).
#
#cnf(i_0_50493, plain, (rd(ld(ld(X1,unit),X2),X1)=ld(ld(X1,unit),rd(X2,X1)))).
####
#cnf(i_0_50944, plain, (ld(mult(X1,mult(X1,X1)),unit)=ld(X1,ld(X1,ld(X1,unit))))).
###########
#cnf(i_0_62070, plain, (ld(ld(X2,unit),mult(X1,X2))=mult(ld(ld(X2,unit),X1),X2))).
#####
#cnf(i_0_50935, plain, (ld(X1,mult(X2,f(ld(X2,X1))))=ld(f(ld(X2,X1)),unit))).
###########
#cnf(i_0_51277, plain, (ld(rd(X2,f(ld(X1,X2))),X1)=ld(f(ld(X1,X2)),unit))).
###########
#cnf(i_0_51425, plain, (mult(ld(X1,X2),f(ld(X2,X1)))=ld(f(ld(X2,X1)),unit))).
#######################################
#cnf(i_0_51441, plain, (mult(ld(X1,X2),mult(ld(X2,X1),ld(X2,X1)))=ld(X2,X1))).
####################
#cnf(i_0_51630, plain, (ld(mult(ld(X1,X2),X2),mult(ld(X1,X2),X1))=ld(X2,X1))).
#####
#cnf(i_0_69489, plain, (mult(mult(ld(X1,X2),X2),ld(X2,X1))=mult(ld(X1,X2),X1))).
########
#cnf(i_0_70216, plain, (mult(mult(ld(X2,unit),X1),X2)=mult(ld(X2,unit),mult(X1,X2)))).
##
#cnf(i_0_51659, plain, (ld(mult(X2,mult(ld(X2,X1),X1)),mult(X1,X2))=ld(X1,X2))).
#####
#cnf(i_0_51695, plain, (ld(ld(ld(X1,X2),X2),ld(ld(X1,X2),X1))=ld(X2,X1))).
##########
#cnf(i_0_6568, plain, (mult(mult(X1,mult(ld(X1,X2),X2)),ld(X2,X1))=mult(X2,X1))).
###
#cnf(i_0_52297, plain, (mult(f(ld(X1,X2)),ld(X2,X1))=ld(f(ld(X1,X2)),unit))).
#########################################
#cnf(i_0_53106, plain, (rd(rd(X1,mult(X2,f(X2))),f(X2))=rd(X1,mult(X2,X2)))).
##
#cnf(i_0_53254, plain, (mult(mult(X1,mult(X2,f(X2))),f(X2))=mult(X1,mult(X2,X2)))).
#
#cnf(i_0_53545, plain, (rd(rd(X1,f(X2)),mult(X2,f(X2)))=rd(X1,mult(X2,X2)))).
#
#cnf(i_0_53700, plain, (mult(mult(X1,f(X2)),mult(X2,f(X2)))=mult(X1,mult(X2,X2)))).
##
#cnf(i_0_55192, plain, (ld(mult(X1,X2),rd(X1,f(X2)))=ld(X2,ld(f(X2),unit)))).
##
#cnf(i_0_55990, plain, (ld(mult(X1,f(X2)),rd(X1,X2))=ld(X2,ld(f(X2),unit)))).
##########################
#cnf(i_0_63701, plain, (mult(rd(X1,mult(X2,mult(X2,X2))),mult(X2,X2))=rd(X1,X2))).
#
#cnf(i_0_63702, plain, (ld(rd(X1,mult(X2,mult(X2,X2))),rd(X1,X2))=mult(X2,X2))).
#
#cnf(i_0_63814, plain, (rd(mult(X1,mult(X2,X2)),mult(X2,mult(X2,X2)))=rd(X1,X2))).
####
#cnf(i_0_64039, plain, (ld(mult(X1,X2),mult(X1,mult(X2,mult(X2,X2))))=mult(X2,X2))).
#
#cnf(i_0_64040, plain, (rd(mult(X1,mult(X2,mult(X2,X2))),mult(X2,X2))=mult(X1,X2))).
#
#cnf(i_0_64238, plain, (mult(rd(X1,mult(X2,X2)),mult(X2,mult(X2,X2)))=mult(X1,X2))).
######################################################
#cnf(i_0_50411, plain, (mult(f(X1),mult(ld(X2,unit),X1))=mult(rd(f(X1),X2),X1))).
#
#cnf(i_0_50414, plain, (mult(f(X1),rd(ld(X2,unit),X1))=rd(rd(f(X1),X2),X1))).
###
#cnf(i_0_50708, plain, (mult(X1,ld(X2,ld(f(X2),unit)))=rd(X1,mult(X2,f(X2))))).
##
#cnf(i_0_50861, plain, (ld(mult(ld(f(X2),X1),X2),unit)=ld(mult(X1,X2),f(X2)))).
################
#cnf(i_0_53180, plain, (rd(mult(X1,mult(X2,X2)),mult(X2,f(X2)))=mult(X1,f(X2)))).
####
#cnf(i_0_53321, plain, (mult(rd(X1,mult(X2,X2)),mult(X2,f(X2)))=rd(X1,f(X2)))).
#
#cnf(i_0_53521, plain, (mult(rd(X1,mult(X2,f(X2))),mult(X2,X2))=mult(X1,f(X2)))).
##
#cnf(i_0_53645, plain, (rd(mult(X1,mult(X2,f(X2))),mult(X2,X2))=rd(X1,f(X2)))).
###
#cnf(i_0_55261, plain, (ld(mult(X1,f(X2)),mult(X1,mult(X2,X2)))=mult(X2,f(X2)))).
########
#cnf(i_0_56144, plain, (ld(rd(X1,mult(X2,X2)),rd(X1,f(X2)))=mult(X2,f(X2)))).
################
#cnf(i_0_62111, plain, (ld(ld(X1,unit),mult(X1,f(X1)))=mult(X1,mult(X1,f(X1))))).
#####################################################################################################################
#cnf(i_0_50734, plain, (ld(f(X1),rd(rd(f(X1),X2),X1))=rd(ld(X2,unit),X1))).
###
#cnf(i_0_53570, plain, (rd(ld(X1,f(X1)),mult(X1,f(X1)))=ld(X1,ld(X1,unit)))).
######
#cnf(i_0_61931, plain, (mult(ld(ld(X2,unit),rd(X1,X2)),X2)=ld(ld(X2,unit),X1))).
################
#cnf(i_0_50418, plain, (rd(rd(X1,f(ld(X2,unit))),X2)=mult(X1,f(ld(X2,unit))))).
#
#cnf(i_0_50531, plain, (rd(rd(X1,X2),f(ld(X2,unit)))=mult(X1,f(ld(X2,unit))))).
#
#cnf(i_0_50691, plain, (mult(X1,mult(X2,f(ld(X2,unit))))=rd(X1,f(ld(X2,unit))))).
##
#cnf(i_0_50700, plain, (rd(X1,mult(X2,f(ld(X2,unit))))=mult(X1,f(ld(X2,unit))))).
##
#cnf(i_0_51081, plain, (mult(mult(X1,X2),f(ld(X2,unit)))=rd(X1,f(ld(X2,unit))))).
############
#cnf(i_0_52806, plain, (mult(mult(X1,ld(mult(X2,X3),X3)),mult(ld(X3,X2),X3))=X1)).
#######
#cnf(i_0_52979, plain, (mult(mult(X1,mult(ld(X2,X3),X2)),ld(mult(X3,X2),X2))=X1)).
#
#cnf(i_0_2866, plain, (ld(rd(X1,mult(X2,f(X2))),mult(X1,f(X2)))=mult(X2,X2))).
#
#cnf(i_0_2951, plain, (ld(rd(X1,f(X2)),mult(X1,mult(X2,f(X2))))=mult(X2,X2))).
#
#cnf(i_0_48832, plain, (ld(ld(X1,f(f(X1))),X1)=mult(X1,mult(f(X1),f(f(X1)))))).
#
#cnf(i_0_2929, plain, (mult(rd(X1,f(f(X2))),X2)=mult(X1,mult(f(X2),f(f(X2)))))).
#
#cnf(i_0_2905, plain, (rd(mult(X1,X2),f(f(X2)))=mult(X1,mult(f(X2),f(f(X2)))))).
#
#cnf(i_0_53281, plain, (ld(mult(X1,mult(X2,f(X2))),mult(X1,X2))=ld(f(X2),unit))).
######################
#cnf(i_0_58964, plain, (mult(mult(X1,ld(X2,X1)),mult(ld(X1,X2),ld(X1,X2)))=X2)).
################################
#cnf(i_0_3144, plain, (ld(rd(X1,mult(X2,mult(X2,X2))),rd(X1,mult(X2,X2)))=X2)).
######
#cnf(i_0_3102, plain, (ld(mult(X1,mult(X2,X2)),mult(X1,mult(X2,mult(X2,X2))))=X2)).
#################
#cnf(i_0_77709, plain, (mult(mult(X1,f(ld(X2,unit))),X2)=rd(X1,f(ld(X2,unit))))).
###############
#cnf(i_0_50814, plain, (ld(mult(X1,f(f(X1))),unit)=ld(X1,ld(f(f(X1)),unit)))).
######
#cnf(i_0_51450, plain, (ld(ld(X1,X2),ld(f(ld(X2,X1)),unit))=f(ld(X2,X1)))).
##################
#cnf(i_0_51510, plain, (rd(ld(X1,X2),f(ld(X2,X1)))=ld(f(ld(X2,X1)),ld(X1,X2)))).
##
#cnf(i_0_51513, plain, (ld(ld(X1,X2),f(ld(X2,X1)))=mult(ld(X2,X1),f(ld(X2,X1))))).
##########
#cnf(i_0_52394, plain, (rd(X1,mult(ld(f(X2),X3),X2))=mult(X1,ld(mult(X3,X2),f(X2))))).
##
#cnf(i_0_3296, plain, (ld(rd(X1,mult(X2,X2)),rd(X1,mult(X2,f(X2))))=f(X2))).
###
#cnf(i_0_53147, plain, (rd(rd(X1,f(X2)),mult(X2,X2))=rd(rd(X1,X2),mult(X2,f(X2))))).
###
#cnf(i_0_53184, plain, (rd(mult(X1,rd(X2,X1)),f(X1))=rd(mult(X1,X2),mult(X1,f(X1))))).
#
#cnf(i_0_53200, plain, (rd(ld(X1,rd(X2,X1)),f(X1))=rd(ld(X1,X2),mult(X1,f(X1))))).
##
#cnf(i_0_3778, plain, (ld(X1,mult(X2,mult(X1,f(X1))))=mult(ld(X1,mult(X2,f(X1))),X1))).
###
#cnf(i_0_3214, plain, (ld(mult(X1,mult(X2,f(X2))),mult(X1,mult(X2,X2)))=f(X2))).
###
#cnf(i_0_53360, plain, (mult(mult(X1,mult(X2,X1)),f(X1))=mult(mult(X1,X2),mult(X1,f(X1))))).
#
#cnf(i_0_53361, plain, (mult(mult(X1,f(X2)),mult(X2,X2))=mult(mult(X1,X2),mult(X2,f(X2))))).
#
#cnf(i_0_53432, plain, (mult(ld(X1,ld(X1,unit)),mult(X1,f(X1)))=ld(X1,f(X1)))).
####
#cnf(i_0_53523, plain, (rd(rd(X1,mult(X2,f(X2))),X2)=rd(rd(X1,X2),mult(X2,f(X2))))).
####
#cnf(i_0_53642, plain, (mult(mult(X1,mult(X2,f(X2))),X2)=mult(mult(X1,X2),mult(X2,f(X2))))).
###
#cnf(i_0_55074, plain, (ld(ld(X1,ld(X1,unit)),ld(X1,f(X1)))=mult(X1,f(X1)))).
####
#cnf(i_0_55354, plain, (ld(rd(X1,f(ld(X1,X2))),X2)=mult(ld(X1,X2),f(ld(X1,X2))))).
########################
#cnf(i_0_63771, plain, (rd(f(X1),mult(X1,mult(X1,X1)))=ld(X1,ld(X1,ld(X1,f(X1)))))).
#########################################
#cnf(i_0_70288, plain, (mult(ld(X1,unit),mult(X1,mult(X1,f(X1))))=mult(X1,f(X1)))).
######################
#cnf(i_0_3852, plain, (rd(rd(X1,mult(X2,X2)),f(X2))=rd(rd(X1,X2),mult(X2,f(X2))))).
#####
#cnf(i_0_3790, plain, (mult(mult(X1,mult(X2,X2)),f(X2))=mult(mult(X1,X2),mult(X2,f(X2))))).
#######
#cnf(i_0_50771, plain, (ld(X1,rd(rd(X1,X2),mult(X1,X1)))=rd(ld(X2,unit),mult(X1,X1)))).
#
#cnf(i_0_50835, plain, (ld(rd(X1,rd(X1,X2)),rd(X1,X2))=ld(rd(X2,rd(X1,X2)),unit))).
###
#cnf(i_0_50853, plain, (ld(mult(X1,rd(X1,X2)),rd(X1,X2))=ld(mult(X2,rd(X1,X2)),unit))).
#######
#cnf(i_0_50878, plain, (ld(mult(f(X1),X2),mult(f(X1),rd(X2,X1)))=ld(X1,unit))).
###
#cnf(i_0_50928, plain, (ld(ld(f(X1),X2),ld(f(X1),rd(X2,X1)))=ld(X1,unit))).
###
#cnf(i_0_51087, plain, (ld(mult(X1,X2),rd(X1,mult(X1,X2)))=rd(ld(X2,unit),mult(X1,X2)))).
#########
#cnf(i_0_51189, plain, (ld(mult(f(X1),mult(X2,X1)),mult(f(X1),X2))=ld(X1,unit))).
##########
#cnf(i_0_52391, plain, (rd(X1,mult(f(X2),f(f(X2))))=mult(X1,ld(X2,f(f(X2)))))).
###
#cnf(i_0_90438, plain, (mult(mult(X1,X2),ld(X2,f(f(X2))))=mult(X1,f(f(X2))))).
#
#cnf(i_0_90695, plain, (mult(rd(X1,X2),f(f(X2)))=mult(X1,ld(X2,f(f(X2)))))).
#
#cnf(i_0_19163, plain, (rd(mult(X1,f(f(X2))),X2)=mult(X1,ld(X2,f(f(X2)))))).
##
#cnf(i_0_90552, plain, (mult(mult(X1,ld(X2,f(f(X2)))),X2)=mult(X1,f(f(X2))))).
#
#cnf(i_0_90584, plain, (ld(mult(X1,X2),mult(X1,f(f(X2))))=ld(X2,f(f(X2))))).
########
#cnf(i_0_53536, plain, (ld(rd(X1,f(X2)),rd(X1,mult(X2,f(X2))))=ld(X2,unit))).
####
#cnf(i_0_53669, plain, (ld(mult(X1,mult(X2,f(X2))),mult(X1,f(X2)))=ld(X2,unit))).
######################
#cnf(i_0_59622, plain, (rd(ld(f(X1),unit),mult(X1,f(X1)))=ld(X1,ld(X1,unit)))).
##################
#cnf(i_0_64797, plain, (ld(mult(X1,X2),rd(X1,mult(X2,X2)))=ld(X2,ld(X2,ld(X2,unit))))).
####
#cnf(i_0_65249, plain, (ld(mult(X1,mult(X2,X2)),rd(X1,X2))=ld(X2,ld(X2,ld(X2,unit))))).
##########################
#cnf(i_0_83693, plain, (ld(f(ld(X2,f(X2))),f(X2))=mult(X2,f(ld(X2,f(X2)))))).
#
#cnf(i_0_93156, plain, (rd(X1,f(ld(X1,f(X1))))=ld(f(ld(X1,f(X1))),X1))).
##
#cnf(i_0_93221, plain, (mult(ld(X1,f(ld(X1,f(X1)))),X1)=f(ld(X1,f(X1))))).
#
#cnf(i_0_93435, plain, (rd(f(ld(X1,f(X1))),X1)=ld(X1,f(ld(X1,f(X1)))))).
#
#cnf(i_0_93498, plain, (mult(f(ld(X1,f(X1))),X1)=mult(X1,f(ld(X1,f(X1)))))).
#
#cnf(i_0_93262, plain, (ld(ld(f(ld(X1,f(X1))),X1),X1)=f(ld(X1,f(X1))))).
#############
#cnf(i_0_87749, plain, (rd(mult(X1,f(X1)),mult(X1,mult(X1,f(X1))))=ld(X1,unit))).
###
#cnf(i_0_90546, plain, (ld(rd(X1,X2),mult(X1,ld(X2,f(f(X2)))))=f(f(X2)))).
######
#cnf(i_0_93511, plain, (mult(ld(X1,f(ld(f(X1),unit))),X1)=f(ld(X1,f(X1))))).
#
#cnf(i_0_93768, plain, (rd(f(ld(X1,f(X1))),f(ld(f(X1),unit)))=unit)).
############
#cnf(i_0_94037, plain, (ld(ld(f(ld(f(X1),unit)),X1),X1)=f(ld(X1,f(X1))))).
#
#cnf(i_0_94597, plain, (ld(f(ld(f(X1),unit)),f(ld(X1,f(X1))))=unit)).
########
#cnf(i_0_94917, plain, (ld(f(ld(X1,f(X1))),f(ld(f(X1),unit)))=unit)).
########################
#cnf(i_0_50529, plain, (mult(X1,mult(ld(X2,unit),mult(X1,X1)))=mult(rd(X1,X2),mult(X1,X1)))).
#
#cnf(i_0_50533, plain, (mult(X1,rd(ld(X2,unit),mult(X1,X1)))=rd(rd(X1,X2),mult(X1,X1)))).
#
#cnf(i_0_50534, plain, (ld(mult(X1,f(ld(X2,unit))),rd(X1,X2))=f(ld(X2,unit)))).
###
#cnf(i_0_50754, plain, (mult(X1,mult(ld(X2,unit),rd(X1,X2)))=rd(mult(rd(X1,X2),X1),X2))).
#
#cnf(i_0_95862, plain, (ld(X1,rd(mult(rd(X1,X2),X1),X2))=mult(ld(X2,unit),rd(X1,X2)))).
#
#cnf(i_0_96291, plain, (mult(ld(X2,unit),X1)=ld(X2,X1))).
#
#cnf(i_0_96673, plain, (rd(ld(X1,X2),X2)=ld(X1,unit))).
#
#cnf(i_0_97322, plain, (rd(X2,mult(X1,X2))=ld(X1,unit))).
#
#cnf(i_0_96909, plain, (ld(ld(X1,unit),X2)=mult(X1,X2))).
##
#cnf(i_0_97157, plain, (ld(rd(X1,X2),unit)=rd(X2,X1))).
###
#cnf(i_0_98052, plain, (mult(rd(X1,X2),rd(X2,X1))=unit)).
#####
#cnf(i_0_50517, plain, (mult(rd(X1,X2),X1)=mult(X1,ld(X2,X1)))).
###
#cnf(i_0_96683, plain, (ld(X1,mult(X2,X2))=mult(ld(X1,X2),X2))).
#
#cnf(i_0_96904, plain, (ld(f(unit),X1)=mult(f(unit),X1))).
#####
#cnf(i_0_96937, plain, (ld(ld(X1,X2),X3)=mult(ld(X2,X1),X3))).
###
#cnf(i_0_97611, plain, (ld(rd(X2,X1),X1)=mult(X1,ld(X2,X1)))).
#
#cnf(i_0_98124, plain, (rd(X1,rd(X3,X2))=mult(X1,rd(X2,X3)))).
#
#cnf(i_0_98178, plain, (ld(rd(X1,X2),X3)=mult(rd(X2,X1),X3))).
#
#cnf(i_0_99129, plain, (mult(mult(X2,X2),X1)=mult(X2,mult(X2,X1)))).
#
#cnf(i_0_100969, plain, (rd(mult(X1,X1),X2)=mult(X1,rd(X1,X2)))).
####
#cnf(i_0_97393, plain, (mult(rd(X1,X2),mult(X2,X1))=mult(X1,X1))).
#
#cnf(i_0_97429, plain, (mult(X1,ld(mult(X2,X1),X1))=ld(X2,X1))).
#
#cnf(i_0_101792, plain, (ld(mult(X2,X1),X1)=ld(X1,ld(X2,X1)))).
#
#cnf(i_0_53804, plain, (ld(rd(rd(X1,X2),X1),X1)=mult(X1,X2))).
#
#cnf(i_0_98834, plain, (rd(X1,mult(X1,ld(X2,X1)))=rd(X2,X1))).
##
#cnf(i_0_99128, plain, (mult(X1,mult(ld(X1,X2),X2))=mult(X2,X2))).
####
#cnf(i_0_4636, plain, (mult(mult(X1,rd(X1,X2)),X2)=mult(X1,X1))).
##
#cnf(i_0_45643, plain, (rd(mult(X1,mult(X1,X2)),X2)=mult(X1,X1))).
#
#cnf(i_0_4637, plain, (mult(mult(X1,X2),ld(X2,X1))=mult(X1,X1))).
#
#cnf(i_0_51659, plain, (ld(mult(X2,X2),mult(X2,X1))=ld(X2,X1))).
#
#cnf(i_0_104086, plain, (ld(mult(X1,X1),X2)=ld(X1,ld(X1,X2)))).
###
#cnf(i_0_96687, plain, (rd(ld(X1,X2),f(X2))=ld(X1,f(X2)))).
###
#cnf(i_0_96907, plain, (mult(ld(X1,f(X1)),X2)=ld(f(X1),X2))).
########################################################
#cnf(i_0_96633, plain, (rd(ld(X1,unit),X2)=ld(X1,ld(X2,unit)))).
#
#cnf(i_0_97229, plain, (ld(mult(X1,X2),unit)=ld(X2,ld(X1,unit)))).
#
#cnf(i_0_105206, plain, (ld(f(X2),rd(X2,X1))=rd(f(X2),X1))).
########
#cnf(i_0_96642, plain, (mult(rd(X2,X1),mult(X1,ld(X2,X1)))=X1)).
###
#cnf(i_0_96943, plain, (mult(ld(X1,f(X2)),f(X2))=ld(X1,X2))).
#
#cnf(i_0_105833, plain, (mult(f(X1),ld(X1,X2))=ld(f(X1),X2))).
#
#cnf(i_0_105934, plain, (ld(mult(f(X1),X2),X1)=ld(X2,f(X1)))).
####
#cnf(i_0_106139, plain, (ld(f(X1),mult(X1,X2))=mult(f(X1),X2))).
#####
#cnf(i_0_99392, plain, (mult(f(unit),mult(f(unit),X1))=X1)).
#####
#cnf(i_0_99737, plain, (mult(ld(X2,X1),mult(ld(X1,X2),X3))=X3)).
##
#cnf(i_0_54722, plain, (mult(rd(rd(X1,X2),X1),mult(X1,X2))=X1)).
#
#cnf(i_0_107174, plain, (rd(rd(X1,X2),X1)=rd(X1,mult(X1,X2)))).
###
#cnf(i_0_100233, plain, (mult(mult(X1,rd(X2,X3)),rd(X3,X2))=X1)).
###
#cnf(i_0_100544, plain, (mult(rd(X2,X1),mult(rd(X1,X2),X3))=X3)).
####
#cnf(i_0_101185, plain, (mult(f(X1),mult(f(X1),X2))=mult(X1,X2))).
#
#cnf(i_0_108701, plain, (mult(X1,ld(f(X1),X2))=mult(f(X1),X2))).
#
#cnf(i_0_108763, plain, (ld(X1,mult(f(X1),X2))=ld(f(X1),X2))).
####
#cnf(i_0_101340, plain, (mult(f(X1),rd(f(X1),X2))=rd(X1,X2))).
###
#cnf(i_0_109116, plain, (rd(X1,mult(X2,f(X1)))=rd(f(X1),X2))).
##
#cnf(i_0_102096, plain, (mult(mult(X1,X2),ld(X2,ld(X1,X2)))=X2)).
#######
#cnf(i_0_104341, plain, (ld(f(X1),ld(f(X1),X2))=ld(X1,X2))).
###################
#cnf(i_0_96444, plain, (mult(ld(f(X2),X1),ld(X1,X2))=f(X2))).
##
#cnf(i_0_2885, plain, (mult(rd(X2,X1),rd(X1,f(X2)))=f(X2))).
#######
#cnf(i_0_108529, plain, (rd(mult(X1,X2),mult(f(X1),X2))=f(X1))).
#######
#cnf(i_0_105032, plain, (mult(ld(X1,ld(X2,unit)),X2)=ld(X1,unit))).
######
#cnf(i_0_50411, plain, (mult(rd(f(X1),X2),X1)=mult(f(X1),ld(X2,X1)))).
#
#cnf(i_0_2302, plain, (rd(f(mult(X1,X1)),X1)=ld(f(mult(X1,X1)),X1))).
#
#cnf(i_0_2740, plain, (rd(X1,f(mult(X1,X1)))=ld(X1,f(mult(X1,X1))))).
#
#cnf(i_0_96450, plain, (ld(X1,mult(X2,f(X2)))=mult(ld(X1,f(X2)),X2))).
###
#cnf(i_0_96719, plain, (mult(ld(X1,f(X2)),X2)=mult(ld(X1,X2),f(X2)))).
####
#cnf(i_0_97599, plain, (mult(rd(X1,X2),f(X1))=mult(X1,ld(X2,f(X1))))).
####
#cnf(i_0_50861, plain, (ld(mult(X2,X1),f(X1))=ld(X1,ld(X2,f(X1))))).
##
#cnf(i_0_106250, plain, (ld(mult(X2,f(X1)),X1)=ld(f(X1),ld(X2,X1)))).
##
#cnf(i_0_108530, plain, (mult(mult(X1,X2),f(X1))=mult(X1,mult(X2,f(X1))))).
#
#cnf(i_0_108534, plain, (rd(mult(X1,X2),f(X1))=mult(X1,rd(X2,f(X1))))).
#
#cnf(i_0_108682, plain, (mult(f(X1),mult(X1,X2))=mult(X1,mult(f(X1),X2)))).
#
#cnf(i_0_108854, plain, (mult(f(X1),rd(X1,X2))=mult(X1,rd(f(X1),X2)))).
##
#cnf(i_0_108980, plain, (ld(f(X1),ld(X1,X2))=ld(X1,ld(f(X1),X2)))).
####
#cnf(i_0_109016, plain, (rd(rd(X1,X2),f(X1))=rd(X1,mult(f(X1),X2)))).
#####
#cnf(i_0_109650, plain, (rd(ld(X1,X2),f(X1))=ld(X1,rd(X2,f(X1))))).
###
#cnf(i_0_109734, plain, (ld(X1,mult(X2,f(X1)))=mult(ld(X1,X2),f(X1)))).
#####
#cnf(i_0_96450, plain, (ld(X1,mult(X2,f(X2)))=mult(ld(X1,X2),f(X2)))).
############
#cnf(i_0_96686, plain, (rd(ld(X1,X2),mult(X2,X2))=ld(X1,ld(X2,unit)))).
###
#cnf(i_0_97213, plain, (rd(X1,mult(X1,mult(X2,X1)))=ld(X2,ld(X1,unit)))).
#
#cnf(i_0_97214, plain, (rd(X2,mult(X1,mult(X2,X2)))=ld(X2,ld(X1,unit)))).
####
#cnf(i_0_101007, plain, (rd(X1,mult(X2,mult(X2,X1)))=ld(X2,ld(X2,unit)))).
####
#cnf(i_0_103440, plain, (rd(ld(X1,ld(X1,X2)),X2)=ld(X1,ld(X1,unit)))).
#######
#cnf(i_0_105031, plain, (mult(mult(X1,X2),ld(X2,ld(X1,unit)))=unit)).
#
#cnf(i_0_105055, plain, (mult(ld(X1,ld(X2,unit)),mult(X2,X1))=unit)).
##################
#cnf(i_0_114588, plain, (rd(ld(X1,X2),mult(X1,X2))=ld(X1,ld(X1,unit)))).
#######
#cnf(i_0_96443, plain, (rd(ld(X1,f(X2)),X2)=ld(X1,ld(f(X2),unit)))).
##
#cnf(i_0_96645, plain, (mult(ld(X1,ld(X2,unit)),X1)=ld(X1,ld(X2,X1)))).
###
#cnf(i_0_96915, plain, (mult(ld(X1,ld(X1,unit)),X2)=ld(X1,ld(X1,X2)))).
##
#cnf(i_0_97189, plain, (rd(f(X2),mult(X1,X2))=ld(f(X2),ld(X1,unit)))).
########
#cnf(i_0_50860, plain, (ld(mult(ld(X1,X2),X1),unit)=ld(X1,ld(X2,X1)))).
#####
#cnf(i_0_102297, plain, (mult(X2,ld(X1,ld(X2,unit)))=rd(X2,mult(X2,X1)))).
#######
#cnf(i_0_104736, plain, (rd(ld(X1,f(X1)),X2)=ld(f(X1),ld(X2,unit)))).
#####
#cnf(i_0_105168, plain, (mult(ld(X2,ld(X1,unit)),X3)=ld(mult(X1,X2),X3))).
#
#cnf(i_0_105171, plain, (mult(X1,ld(X3,ld(X2,unit)))=rd(X1,mult(X2,X3)))).
#####
#cnf(i_0_108562, plain, (ld(mult(X1,X2),f(X1))=ld(X2,ld(f(X1),unit)))).
############
#cnf(i_0_99037, plain, (mult(ld(f(mult(X1,X1)),X1),X1)=f(mult(X1,X1)))).
###
#cnf(i_0_55221, plain, (mult(rd(X2,X1),mult(X1,f(X2)))=mult(X2,f(X2)))).
#
#cnf(i_0_32740, plain, (mult(rd(f(X2),X1),mult(X1,X2))=mult(X2,f(X2)))).
###
#cnf(i_0_104873, plain, (mult(rd(f(X2),X1),rd(X1,X2))=ld(X2,f(X2)))).
###########
#cnf(i_0_105222, plain, (ld(X2,rd(mult(X2,f(X2)),X1))=rd(f(X2),X1))).
#
#cnf(i_0_119144, plain, (rd(mult(X1,f(X1)),X2)=mult(X1,rd(f(X1),X2)))).
######
#cnf(i_0_109936, plain, (rd(ld(f(rd(X1,X2)),X1),X2)=f(rd(X1,X2)))).
#
#cnf(i_0_109943, plain, (mult(ld(f(mult(X1,X2)),X1),X2)=f(mult(X1,X2)))).
#
#cnf(i_0_110190, plain, (mult(X1,rd(X2,f(mult(X1,X2))))=f(mult(X1,X2)))).
#
#cnf(i_0_110396, plain, (rd(X1,mult(f(rd(X1,X2)),X2))=f(rd(X1,X2)))).
#
#cnf(i_0_110788, plain, (rd(X1,mult(f(X1),ld(X2,X1)))=rd(X2,f(X1)))).
#####
#cnf(i_0_112267, plain, (mult(X1,mult(ld(X1,X2),f(X1)))=mult(X2,f(X1)))).
#####
#cnf(i_0_113737, plain, (mult(X1,mult(ld(X1,X2),f(X2)))=mult(X2,f(X2)))).
#
#cnf(i_0_120816, plain, (ld(mult(X2,f(X2)),X1)=ld(X2,ld(f(X2),X1)))).
###
#cnf(i_0_113870, plain, (mult(ld(X1,ld(X2,f(X2))),X2)=ld(X1,f(X2)))).
###
#cnf(i_0_118419, plain, (mult(mult(X1,rd(f(X1),X2)),X2)=mult(X1,f(X1)))).
##
#cnf(i_0_118458, plain, (mult(mult(X1,X2),ld(X2,f(X1)))=mult(X1,f(X1)))).
#
#cnf(i_0_118470, plain, (mult(mult(f(X1),X2),ld(X2,X1))=mult(X1,f(X1)))).
######
#cnf(i_0_119046, plain, (rd(mult(f(X1),X2),mult(X1,X2))=ld(X1,f(X1)))).
##############
#cnf(i_0_121424, plain, (rd(mult(X1,mult(f(X1),X2)),X2)=mult(X1,f(X1)))).
#
#cnf(i_0_122392, plain, (mult(mult(X1,f(X1)),X2)=mult(X1,mult(f(X1),X2)))).
#######
#cnf(i_0_96639, plain, (mult(ld(X1,ld(X2,X3)),ld(X3,X2))=ld(X1,unit))).
##
#cnf(i_0_96667, plain, (ld(mult(X1,ld(X2,X1)),rd(X1,X2))=ld(X1,unit))).
############
#cnf(i_0_97323, plain, (rd(mult(ld(X1,X2),X1),mult(X2,X1))=ld(X1,unit))).
################
#cnf(i_0_50931, plain, (mult(ld(X2,X1),ld(X1,rd(X2,X1)))=ld(X1,unit))).
###
#cnf(i_0_99728, plain, (mult(ld(X2,X1),rd(ld(X1,X2),X3))=ld(X3,unit))).
#
#cnf(i_0_100242, plain, (mult(ld(X1,rd(X2,X3)),rd(X3,X2))=ld(X1,unit))).
##
#cnf(i_0_50842, plain, (mult(rd(X2,X1),rd(X1,mult(X2,X2)))=ld(X2,unit))).
#####
#cnf(i_0_100535, plain, (mult(rd(X2,X1),rd(rd(X1,X2),X3))=ld(X3,unit))).
#######
#cnf(i_0_107616, plain, (mult(rd(X2,X1),rd(X1,mult(X1,X2)))=ld(X1,unit))).
###############
#cnf(i_0_50529, plain, (mult(rd(X1,X2),mult(X1,X1))=mult(X1,mult(ld(X2,X1),X1)))).
#
#cnf(i_0_50754, plain, (rd(mult(X1,ld(X2,X1)),X2)=mult(X1,ld(X2,rd(X1,X2))))).
#
#cnf(i_0_96447, plain, (mult(ld(X2,X1),ld(X1,f(X2)))=ld(f(X2),unit))).
#########
#cnf(i_0_96607, plain, (ld(X1,mult(X2,mult(X2,X2)))=mult(ld(X1,X2),mult(X2,X2)))).
###############
#cnf(i_0_44499, plain, (mult(ld(X1,rd(X2,X1)),X2)=mult(ld(X1,X2),ld(X1,X2)))).
####
#cnf(i_0_99130, plain, (rd(mult(ld(X1,X2),X2),X1)=ld(X1,mult(X2,rd(X2,X1))))).
#
#cnf(i_0_99187, plain, (rd(X1,mult(ld(X2,X3),X3))=mult(X1,ld(X3,ld(X3,X2))))).
#
#cnf(i_0_99544, plain, (rd(X1,mult(f(unit),X2))=mult(X1,ld(X2,f(unit))))).
#
#cnf(i_0_99836, plain, (ld(mult(ld(X2,X1),X3),X4)=mult(ld(X3,ld(X1,X2)),X4))).
#
#cnf(i_0_99838, plain, (rd(X1,mult(ld(X3,X2),X4))=mult(X1,ld(X4,ld(X2,X3))))).
#
#cnf(i_0_99868, plain, (ld(mult(f(unit),X1),X2)=mult(ld(X1,f(unit)),X2))).
####
#cnf(i_0_100316, plain, (rd(X1,mult(X2,rd(X4,X3)))=mult(X1,rd(rd(X3,X4),X2)))).
###
#cnf(i_0_100402, plain, (rd(X1,mult(X2,f(unit)))=mult(X1,rd(f(unit),X2)))).
#
#cnf(i_0_100419, plain, (rd(X1,mult(X2,ld(X4,X3)))=mult(X1,rd(ld(X3,X4),X2)))).
##
#cnf(i_0_100641, plain, (ld(mult(rd(X2,X1),X3),X4)=mult(ld(X3,rd(X1,X2)),X4))).
#
#cnf(i_0_100643, plain, (rd(X1,mult(rd(X3,X2),X4))=mult(X1,ld(X4,rd(X2,X3))))).
##
#cnf(i_0_100714, plain, (ld(mult(X1,rd(X2,X1)),X3)=mult(rd(X1,mult(X1,X2)),X3))).
#
#cnf(i_0_100734, plain, (ld(mult(X1,f(unit)),X2)=mult(rd(f(unit),X1),X2))).
#
#cnf(i_0_100740, plain, (ld(mult(X1,rd(X3,X2)),X4)=mult(rd(rd(X2,X3),X1),X4))).
#
#cnf(i_0_100752, plain, (ld(mult(X1,ld(X3,X2)),X4)=mult(rd(ld(X2,X3),X1),X4))).
#
#cnf(i_0_101019, plain, (mult(mult(X1,mult(X1,X2)),X2)=mult(X1,mult(X1,mult(X2,X2))))).
#
#cnf(i_0_101289, plain, (rd(mult(X1,rd(X1,X2)),X2)=mult(X1,rd(X1,mult(X2,X2))))).
####
#cnf(i_0_101707, plain, (mult(mult(X1,X2),mult(X1,X2))=mult(X1,mult(X2,mult(X1,X2))))).
##
#cnf(i_0_52555, plain, (rd(X1,mult(ld(X2,X3),X2))=mult(X1,ld(X2,ld(X3,X2))))).
#
#cnf(i_0_102198, plain, (ld(mult(X2,mult(X1,X1)),X1)=ld(X1,ld(X1,ld(X2,X1))))).
###
#cnf(i_0_102219, plain, (ld(mult(X1,mult(X1,X2)),X2)=ld(X2,ld(X1,ld(X1,X2))))).
###
#cnf(i_0_102926, plain, (mult(rd(X1,X2),rd(X1,X2))=mult(X1,ld(X2,rd(X1,X2))))).
#########
#cnf(i_0_103518, plain, (mult(mult(X1,X2),ld(X1,X2))=mult(X1,mult(X2,ld(X1,X2))))).
########
#cnf(i_0_104223, plain, (ld(X1,mult(ld(X1,X2),X2))=mult(ld(X1,ld(X1,X2)),X2))).
####
#cnf(i_0_105250, plain, (ld(mult(X2,X1),ld(X2,X1))=ld(X1,ld(X2,ld(X2,X1))))).
##########
#cnf(i_0_107613, plain, (rd(rd(X1,X2),mult(X1,X1))=rd(X1,mult(X1,mult(X1,X2))))).
#
#cnf(i_0_134875, plain, (rd(rd(f(X1),X2),X1)=rd(f(X1),mult(X1,X2)))).
#######
#cnf(i_0_109708, plain, (mult(ld(X1,rd(X2,f(X1))),f(X1))=ld(X1,X2))).
###########################################
#cnf(i_0_122784, plain, (mult(mult(X1,mult(X1,X1)),X2)=mult(X1,mult(X1,mult(X1,X2))))).
#
#cnf(i_0_122785, plain, (rd(mult(X1,mult(X1,X1)),X2)=mult(X1,mult(X1,rd(X1,X2))))).
#
#cnf(i_0_122786, plain, (ld(mult(X1,mult(X1,X1)),X2)=ld(X1,ld(X1,ld(X1,X2))))).
#######
#cnf(i_0_126288, plain, (rd(rd(X1,X2),mult(X1,X2))=rd(X1,mult(X1,mult(X2,X2))))).
##
#cnf(i_0_126787, plain, (rd(ld(X1,rd(X2,X1)),X2)=rd(ld(X1,X2),mult(X2,X1)))).
#########
#cnf(i_0_131845, plain, (mult(mult(X1,ld(X2,X1)),X2)=mult(rd(X1,X2),mult(X1,X2)))).
#####
#cnf(i_0_132145, plain, (rd(mult(X1,mult(X2,X1)),X2)=mult(mult(X1,X2),rd(X1,X2)))).
########
#cnf(i_0_134218, plain, (ld(X1,mult(X2,mult(X1,X2)))=mult(ld(X1,X2),mult(X1,X2)))).
#
#cnf(i_0_134404, plain, (ld(mult(X2,mult(X1,X2)),X1)=ld(mult(X1,X2),ld(X2,X1)))).
##
#cnf(i_0_134411, plain, (ld(mult(X2,X1),rd(X1,X2))=ld(X1,ld(X2,rd(X1,X2))))).
######
#cnf(i_0_136793, plain, (rd(rd(X2,X1),mult(X1,X2))=rd(X2,mult(X1,mult(X2,X1))))).
##########
#cnf(i_0_51609, plain, (mult(mult(X2,rd(f(X2),X1)),rd(X1,X2))=f(X2))).
#######
#cnf(i_0_108531, plain, (ld(mult(X1,X2),mult(X1,mult(X2,f(X1))))=f(X1))).
#######
#cnf(i_0_109013, plain, (mult(rd(X2,X1),mult(X1,ld(X2,f(X1))))=f(X1))).
#######
#cnf(i_0_110337, plain, (rd(mult(X1,mult(f(X1),X2)),mult(X1,X2))=f(X1))).
###
#cnf(i_0_111624, plain, (mult(mult(X1,X2),ld(X2,ld(X1,f(X2))))=f(X2))).
#####################
#cnf(i_0_98543, plain, (mult(mult(X1,rd(X1,X2)),rd(X2,mult(X1,X1)))=unit)).
###
#cnf(i_0_98587, plain, (mult(rd(X1,mult(X1,X2)),mult(X1,rd(X2,X1)))=unit)).
#
#cnf(i_0_98604, plain, (mult(mult(X1,ld(X2,X1)),ld(X1,rd(X2,X1)))=unit)).
##
#cnf(i_0_98624, plain, (mult(rd(ld(X1,X2),X3),mult(X3,ld(X2,X1)))=unit)).
#
#cnf(i_0_98652, plain, (mult(rd(X1,mult(X2,X2)),mult(X2,rd(X2,X1)))=unit)).
###
#cnf(i_0_98696, plain, (mult(mult(X1,rd(X2,X1)),rd(X1,mult(X1,X2)))=unit)).
#
#cnf(i_0_98713, plain, (mult(ld(X1,rd(X2,X1)),mult(X1,ld(X2,X1)))=unit)).
##
#cnf(i_0_98733, plain, (mult(mult(X1,ld(X3,X2)),rd(ld(X2,X3),X1))=unit)).
#
#cnf(i_0_99114, plain, (mult(ld(X1,ld(X1,X2)),mult(ld(X2,X1),X1))=unit)).
#
#cnf(i_0_99164, plain, (mult(mult(ld(X1,X2),X2),ld(X2,ld(X2,X1)))=unit)).
#
#cnf(i_0_99471, plain, (mult(ld(X1,f(unit)),mult(f(unit),X1))=unit)).
##
#cnf(i_0_99788, plain, (mult(ld(X1,ld(X2,X3)),mult(ld(X3,X2),X1))=unit)).
#
#cnf(i_0_99821, plain, (mult(mult(ld(X2,X1),X3),ld(X3,ld(X1,X2)))=unit)).
#
#cnf(i_0_100270, plain, (mult(rd(rd(X1,X2),X3),mult(X3,rd(X2,X1)))=unit)).
#
#cnf(i_0_100313, plain, (mult(mult(X1,rd(X3,X2)),rd(rd(X2,X3),X1))=unit)).
#
#cnf(i_0_100595, plain, (mult(ld(X1,rd(X2,X3)),mult(rd(X3,X2),X1))=unit)).
#
#cnf(i_0_100627, plain, (mult(mult(rd(X2,X1),X3),ld(X3,rd(X1,X2)))=unit)).
#################
#cnf(i_0_121276, plain, (mult(ld(X1,ld(X2,f(X2))),f(X2))=ld(X1,unit))).
###########################################################################
#cnf(i_0_96449, plain, (mult(ld(f(X2),X1),mult(ld(X1,X2),f(X2)))=X2)).
#
#cnf(i_0_96499, plain, (mult(mult(X1,X2),mult(ld(X2,X1),X2))=mult(X1,mult(X1,X2)))).
####
#cnf(i_0_96564, plain, (mult(mult(X1,ld(X2,rd(X1,X2))),X2)=mult(X1,ld(X2,X1)))).
##
#cnf(i_0_96658, plain, (mult(rd(X2,f(X1)),mult(f(X1),ld(X2,X1)))=X1)).
#######
#cnf(i_0_50835, plain, (mult(rd(X1,mult(X1,X2)),rd(X1,X2))=rd(X1,mult(X2,X2)))).
#####
#cnf(i_0_98986, plain, (mult(mult(X1,rd(X2,X1)),mult(X1,X2))=mult(X1,mult(X2,X2)))).
#
#cnf(i_0_99004, plain, (mult(ld(X1,X2),ld(X2,ld(X1,X2)))=ld(X1,ld(X1,X2)))).
####
#cnf(i_0_80, plain, (mult(rd(X1,X2),mult(X2,rd(X2,X1)))=mult(X1,rd(X2,X1)))).
###
#cnf(i_0_65207, plain, (mult(rd(X2,X1),mult(X1,mult(X2,X2)))=mult(X2,mult(X2,X2)))).
#
#cnf(i_0_64772, plain, (mult(mult(X2,rd(X2,X1)),mult(X1,X2))=mult(X2,mult(X2,X2)))).
####
#cnf(i_0_100939, plain, (rd(mult(X1,mult(X1,mult(X2,X2))),X2)=mult(X1,mult(X1,X2)))).
#
#cnf(i_0_101022, plain, (rd(mult(X1,mult(X1,X2)),mult(X2,X2))=mult(X1,rd(X1,X2)))).
##
#cnf(i_0_101241, plain, (mult(mult(X1,rd(X1,mult(X2,X2))),X2)=mult(X1,rd(X1,X2)))).
#
#cnf(i_0_101256, plain, (mult(ld(X1,mult(X2,rd(X2,X1))),X1)=mult(ld(X1,X2),X2))).
##
#cnf(i_0_101288, plain, (mult(mult(X1,rd(X1,X2)),mult(X2,X2))=mult(X1,mult(X1,X2)))).
###
#cnf(i_0_102301, plain, (mult(ld(X2,X1),ld(X1,ld(X1,X2)))=ld(X2,ld(X1,X2)))).
##################################
#cnf(i_0_107671, plain, (mult(rd(X1,mult(X2,X2)),rd(X2,X1))=rd(X1,mult(X1,X2)))).
##
#cnf(i_0_107708, plain, (rd(mult(X1,X2),mult(X1,mult(X2,X2)))=rd(X1,mult(X1,X2)))).
###############
#cnf(i_0_109698, plain, (mult(ld(X1,f(f(f(X1)))),X1)=f(f(f(X1))))).
#
#cnf(i_0_153760, plain, (rd(X1,f(f(f(X1))))=ld(f(f(f(X1))),X1))).
##
#cnf(i_0_153770, plain, (rd(f(f(f(X1))),X1)=ld(X1,f(f(f(X1)))))).
##
#cnf(i_0_153849, plain, (mult(f(f(f(X1))),X1)=mult(X1,f(f(f(X1)))))).
######
#cnf(i_0_111882, plain, (mult(mult(X1,f(X2)),ld(f(X2),ld(X1,X2)))=X2)).
####
#cnf(i_0_113770, plain, (mult(X1,mult(ld(X1,X2),mult(X2,X2)))=mult(X2,mult(X2,X2)))).
##########
#cnf(i_0_114081, plain, (rd(X1,mult(X1,mult(X1,ld(X2,X1))))=rd(X2,mult(X1,X1)))).
###################
#cnf(i_0_117635, plain, (rd(mult(X1,mult(X2,X2)),mult(X1,X2))=mult(X1,rd(X2,X1)))).
####
#cnf(i_0_122550, plain, (mult(mult(X1,mult(X1,rd(X1,X2))),X2)=mult(X1,mult(X1,X1)))).
#
#cnf(i_0_122614, plain, (rd(mult(X1,mult(X1,mult(X1,X2))),X2)=mult(X1,mult(X1,X1)))).
###############
#cnf(i_0_134281, plain, (mult(X1,mult(ld(X1,ld(X1,X2)),X2))=mult(ld(X1,X2),X2))).
#######
#cnf(i_0_136665, plain, (mult(rd(ld(X1,X2),mult(X2,X1)),X2)=ld(X1,rd(X2,X1)))).
###
#cnf(i_0_136954, plain, (rd(mult(rd(X1,X2),mult(X1,X2)),X2)=mult(X1,ld(X2,X1)))).
#
#cnf(i_0_137201, plain, (mult(mult(mult(X1,X2),rd(X1,X2)),X2)=mult(X1,mult(X2,X1)))).
#
#cnf(i_0_137488, plain, (mult(X1,mult(ld(X1,X2),mult(X1,X2)))=mult(X2,mult(X1,X2)))).
#####################
#cnf(i_0_150358, plain, (mult(mult(X1,X2),mult(ld(X2,X1),X1))=mult(X1,mult(X1,X1)))).
##
#cnf(i_0_150810, plain, (mult(mult(X1,mult(X1,X2)),ld(X2,X1))=mult(X1,mult(X1,X1)))).
##########################
#cnf(i_0_97391, plain, (ld(X2,f(mult(X2,X1)))=rd(X1,f(mult(X2,X1))))).
##########
#cnf(i_0_100801, plain, (rd(mult(X1,mult(X1,f(X2))),X2)=mult(X1,rd(X1,f(X2))))).
######
#cnf(i_0_101023, plain, (rd(mult(X1,mult(X1,X2)),f(X2))=mult(X1,mult(X1,f(X2))))).
##
#cnf(i_0_101211, plain, (mult(mult(X1,rd(X1,f(X2))),X2)=mult(X1,mult(X1,f(X2))))).
#
#cnf(i_0_101290, plain, (mult(mult(X1,rd(X1,X2)),f(X2))=mult(X1,rd(X1,f(X2))))).
#####
#cnf(i_0_102176, plain, (mult(ld(X2,rd(f(X2),X1)),X2)=ld(f(X2),ld(X1,X2)))).
#####
#cnf(i_0_104402, plain, (rd(ld(X1,X2),mult(X2,f(X2)))=ld(X1,ld(f(X2),unit)))).
######
#cnf(i_0_104520, plain, (rd(ld(X1,ld(X1,X2)),f(X2))=ld(X1,ld(X1,f(X2))))).
####
#cnf(i_0_104749, plain, (mult(ld(X1,ld(X2,f(X2))),X1)=ld(X1,ld(f(X2),X1)))).
#########
#cnf(i_0_105349, plain, (mult(rd(X2,f(X1)),rd(X1,X2))=mult(X2,rd(f(X1),X2)))).
###
#cnf(i_0_105476, plain, (ld(f(mult(X1,X2)),X1)=rd(f(mult(X1,X2)),X2))).
#####
#cnf(i_0_105840, plain, (mult(ld(X1,X2),ld(f(X2),X1))=mult(ld(X1,f(X2)),X1))).
#####
#cnf(i_0_105981, plain, (mult(f(X1),mult(ld(X1,X2),X1))=mult(ld(f(X1),X2),X1))).
####
#cnf(i_0_106141, plain, (mult(f(X1),mult(ld(X1,X2),X2))=mult(ld(f(X1),X2),X2))).
############
#cnf(i_0_106394, plain, (ld(f(rd(X1,X2)),X1)=mult(f(rd(X1,X2)),X2))).
#####
#cnf(i_0_108506, plain, (ld(mult(f(X1),X2),mult(X1,X2))=mult(ld(X2,f(X1)),X2))).
#
#cnf(i_0_108579, plain, (ld(mult(X1,X2),mult(f(X1),X2))=ld(X2,ld(f(X1),X2)))).
########
#cnf(i_0_108690, plain, (mult(X1,mult(ld(f(X1),X2),X1))=mult(f(X1),mult(X2,X1)))).
##
#cnf(i_0_108694, plain, (mult(X1,mult(ld(f(X1),X2),X2))=mult(f(X1),mult(X2,X2)))).
##########
#cnf(i_0_109061, plain, (mult(rd(X2,X1),rd(f(X1),X2))=rd(X2,mult(X2,f(X1))))).
########
#cnf(i_0_109188, plain, (rd(X1,mult(X1,mult(X2,f(X1))))=rd(f(X1),mult(X1,X2)))).
#
#cnf(i_0_109240, plain, (rd(X1,mult(X2,mult(X1,f(X1))))=ld(f(X1),ld(X2,unit)))).
#####
#cnf(i_0_109245, plain, (rd(X1,mult(X2,mult(X2,f(X1))))=rd(f(X1),mult(X2,X2)))).
############
#cnf(i_0_2733, plain, (rd(mult(X1,mult(X2,X1)),f(X1))=mult(X1,mult(X2,f(X1))))).
##
#cnf(i_0_2747, plain, (mult(ld(X1,rd(X2,X1)),f(X1))=ld(X1,rd(X2,f(X1))))).
#
#cnf(i_0_2852, plain, (mult(ld(X1,rd(X2,f(X1))),X1)=mult(ld(X1,X2),f(X1)))).
####
#cnf(i_0_114312, plain, (ld(f(X1),ld(X2,unit))=ld(X1,rd(f(X1),X2)))).
########
#cnf(i_0_168031, plain, (mult(ld(X1,rd(f(X1),X2)),X2)=ld(f(X1),unit))).
#############
#cnf(i_0_168099, plain, (mult(ld(X1,rd(f(X1),X2)),mult(X2,X1))=f(X1))).
#####
#cnf(i_0_117554, plain, (mult(ld(X1,ld(X2,f(X2))),X3)=ld(mult(f(X2),X1),X3))).
##
#cnf(i_0_117860, plain, (mult(X1,ld(X2,ld(X3,f(X3))))=rd(X1,mult(f(X3),X2)))).
#############
#cnf(i_0_129903, plain, (mult(ld(X2,rd(f(X2),X1)),X3)=ld(mult(X1,f(X2)),X3))).
#######
#cnf(i_0_130245, plain, (mult(X1,ld(X3,rd(f(X3),X2)))=rd(X1,mult(X2,f(X3))))).
#######
#cnf(i_0_149885, plain, (rd(mult(X1,X2),mult(X1,f(X2)))=mult(X1,rd(f(X2),X1)))).
###
#cnf(i_0_149965, plain, (rd(mult(X1,f(X2)),mult(X1,X2))=rd(X1,mult(X1,f(X2))))).
##############################
#cnf(i_0_152763, plain, (mult(ld(X1,f(X2)),ld(X2,X1))=ld(X1,ld(f(X2),X1)))).
###############################################################
#cnf(i_0_96389, plain, (mult(ld(X1,f(f(X2))),X2)=mult(ld(X1,X2),f(f(X2))))).
#
#cnf(i_0_96391, plain, (rd(ld(X1,f(f(X2))),X2)=ld(X1,ld(X2,f(f(X2)))))).
#
#cnf(i_0_97206, plain, (ld(mult(X2,X1),ld(X1,unit))=ld(X1,ld(X2,ld(X1,unit))))).
##
#cnf(i_0_97600, plain, (mult(rd(X1,X2),f(f(X1)))=mult(X1,ld(X2,f(f(X1)))))).
######
#cnf(i_0_99203, plain, (rd(ld(X1,ld(X2,unit)),X2)=ld(X1,ld(X2,ld(X2,unit))))).
##
#cnf(i_0_101348, plain, (rd(ld(X1,ld(X1,unit)),X2)=ld(X1,ld(X1,ld(X2,unit))))).
##########
#cnf(i_0_105169, plain, (rd(ld(X2,ld(X1,unit)),X3)=ld(mult(X1,X2),ld(X3,unit)))).
##
#cnf(i_0_105262, plain, (ld(mult(X1,X2),ld(X1,unit))=ld(X2,ld(X1,ld(X1,unit))))).
########
#cnf(i_0_110777, plain, (mult(rd(f(f(X1)),X2),X1)=mult(f(f(X1)),ld(X2,X1)))).
#
#cnf(i_0_111742, plain, (ld(mult(X2,X1),f(f(X1)))=ld(X1,ld(X2,f(f(X1)))))).
###
#cnf(i_0_112118, plain, (mult(mult(X1,X2),f(f(X1)))=mult(X1,mult(X2,f(f(X1)))))).
#
#cnf(i_0_112361, plain, (rd(mult(X1,X2),f(f(X1)))=mult(X1,rd(X2,f(f(X1)))))).
#
#cnf(i_0_112669, plain, (mult(f(f(X1)),mult(X1,X2))=mult(X1,mult(f(f(X1)),X2)))).
###
#cnf(i_0_113110, plain, (ld(f(f(X1)),ld(X1,X2))=ld(X1,ld(f(f(X1)),X2)))).
#
#cnf(i_0_113458, plain, (rd(ld(X1,X2),f(f(X1)))=ld(X1,rd(X2,f(f(X1)))))).
#
#cnf(i_0_113542, plain, (ld(X1,mult(X2,f(f(X1))))=mult(ld(X1,X2),f(f(X1))))).
########
#cnf(i_0_120536, plain, (ld(mult(X2,f(f(X1))),X1)=ld(f(f(X1)),ld(X2,X1)))).
###################################
#cnf(i_0_172954, plain, (rd(rd(X1,X2),f(f(X1)))=rd(X1,mult(f(f(X1)),X2)))).
##############
#cnf(i_0_176189, plain, (mult(f(f(X1)),rd(X1,X2))=mult(X1,rd(f(f(X1)),X2)))).
###
#cnf(i_0_96453, plain, (mult(ld(X1,f(X2)),mult(X2,f(X2)))=mult(ld(X1,X2),X2))).
####
#cnf(i_0_96626, plain, (mult(rd(X2,X1),mult(X1,mult(ld(X2,X1),X1)))=mult(X1,X1))).
########
#cnf(i_0_97182, plain, (rd(X1,mult(f(X1),mult(X2,X1)))=ld(X2,ld(f(X1),unit)))).
#########
#cnf(i_0_98308, plain, (mult(mult(X2,ld(X1,rd(X2,X1))),rd(X1,X2))=rd(X2,X1))).
##
#cnf(i_0_98341, plain, (mult(rd(X1,X2),mult(X2,ld(X1,rd(X2,X1))))=rd(X2,X1))).
###
#cnf(i_0_98350, plain, (mult(rd(X1,X2),f(rd(X2,X1)))=ld(f(rd(X2,X1)),unit))).
##########
#cnf(i_0_2471, plain, (mult(mult(X1,mult(X1,ld(X2,X1))),ld(X1,X2))=mult(X1,X1))).
##
#cnf(i_0_51695, plain, (mult(ld(X2,ld(X1,X2)),mult(ld(X2,X1),X1))=ld(X2,X1))).
###
#cnf(i_0_325, plain, (mult(ld(X1,ld(X1,X2)),mult(ld(X2,X1),X2))=ld(X1,X2))).
####
#cnf(i_0_2496, plain, (mult(mult(X1,ld(X2,X1)),mult(ld(X1,X2),X1))=mult(X1,X1))).
##
#cnf(i_0_6002, plain, (mult(mult(X1,mult(X2,rd(X2,X1))),rd(X1,X2))=mult(X1,X2))).
#
#cnf(i_0_534, plain, (mult(mult(X1,rd(X2,X1)),rd(X1,mult(X2,X2)))=rd(X1,X2))).
#####
#cnf(i_0_100209, plain, (mult(f(rd(X1,X2)),rd(X2,X1))=ld(f(rd(X1,X2)),unit))).
#######
#cnf(i_0_51277, plain, (mult(rd(f(ld(X2,X1)),X1),X2)=ld(f(ld(X2,X1)),unit))).
#######
#cnf(i_0_63702, plain, (mult(mult(X2,mult(X2,rd(X2,X1))),rd(X1,X2))=mult(X2,X2))).
###
#cnf(i_0_100975, plain, (mult(mult(X1,mult(X1,ld(X2,X3))),ld(X3,X2))=mult(X1,X1))).
##
#cnf(i_0_101522, plain, (mult(mult(X1,rd(X2,X1)),mult(X1,ld(X2,X1)))=mult(X1,X1))).
##
#cnf(i_0_184446, plain, (mult(mult(X1,X2),ld(X2,rd(X1,X2)))=mult(X1,rd(X1,X2)))).
#
#cnf(i_0_184763, plain, (ld(mult(X1,X2),mult(X1,rd(X1,X2)))=ld(X2,rd(X1,X2)))).
##
#cnf(i_0_184790, plain, (mult(rd(X2,mult(X1,X1)),mult(X1,X2))=mult(X2,ld(X1,X2)))).
##
#cnf(i_0_184809, plain, (rd(mult(X2,ld(X1,X2)),mult(X1,X2))=rd(X2,mult(X1,X1)))).
##
#cnf(i_0_184958, plain, (mult(mult(X1,mult(X2,X2)),ld(X2,X1))=mult(X1,mult(X2,X1)))).
##
#cnf(i_0_184982, plain, (rd(mult(ld(X1,X2),X2),mult(X1,X2))=ld(X1,rd(X2,X1)))).
#####
#cnf(i_0_185270, plain, (mult(rd(X1,X2),mult(X2,mult(X1,X2)))=mult(X1,mult(X1,X2)))).
#
#cnf(i_0_185792, plain, (mult(rd(X2,X1),mult(X1,mult(X1,X2)))=mult(X2,mult(X1,X2)))).
######
#cnf(i_0_186325, plain, (ld(mult(X2,X1),mult(X2,ld(X1,X2)))=ld(X1,ld(X1,X2)))).
##
#cnf(i_0_186705, plain, (mult(mult(X1,mult(X2,X1)),ld(X1,X2))=mult(X1,mult(X2,X2)))).
######
#cnf(i_0_187305, plain, (mult(ld(X1,rd(X2,X1)),mult(X1,X2))=mult(ld(X1,X2),X2))).
#######
#cnf(i_0_188715, plain, (mult(mult(X1,X2),ld(X2,ld(X2,X1)))=mult(X1,ld(X2,X1)))).
###########
#cnf(i_0_185191, plain, (rd(ld(X1,X2),mult(X1,f(X2)))=ld(X1,rd(f(X2),X1)))).
######
#cnf(i_0_185543, plain, (ld(mult(f(X1),X2),rd(X1,X2))=ld(X2,rd(f(X1),X2)))).
###
#cnf(i_0_186177, plain, (mult(mult(X1,X2),ld(f(X2),X1))=mult(X1,mult(f(X2),X1)))).
##
#cnf(i_0_186222, plain, (mult(rd(X1,X2),mult(f(X2),X1))=mult(X1,ld(f(X2),X1)))).
#########
#cnf(i_0_186659, plain, (mult(mult(X2,f(X1)),ld(X1,X2))=mult(X2,ld(f(X1),X2)))).
#########
#cnf(i_0_188753, plain, (mult(rd(X1,f(X2)),mult(X2,X1))=mult(X1,mult(f(X2),X1)))).
###########
#cnf(i_0_101560, plain, (mult(rd(X1,mult(X1,X2)),mult(X1,mult(X2,X1)))=mult(X1,X1))).
#
#cnf(i_0_101562, plain, (mult(rd(X1,mult(X2,X2)),mult(X2,mult(X2,X1)))=mult(X1,X1))).
##
#cnf(i_0_101729, plain, (mult(mult(X1,f(unit)),mult(f(unit),X1))=mult(X1,X1))).
#
#cnf(i_0_101735, plain, (mult(mult(X1,rd(X3,X2)),mult(rd(X2,X3),X1))=mult(X1,X1))).
#
#cnf(i_0_101747, plain, (mult(mult(X1,ld(X3,X2)),mult(ld(X2,X3),X1))=mult(X1,X1))).
#####
#cnf(i_0_102919, plain, (mult(mult(X1,X2),mult(ld(X2,ld(X1,X2)),X2))=mult(X2,X2))).
######
#cnf(i_0_103187, plain, (mult(mult(X1,mult(X1,rd(X3,X2))),rd(X2,X3))=mult(X1,X1))).
######
#cnf(i_0_103463, plain, (mult(mult(X1,mult(X2,X1)),ld(X1,ld(X2,X1)))=mult(X1,X1))).
##
#cnf(i_0_103464, plain, (mult(mult(X1,mult(X2,X2)),ld(X2,ld(X2,X1)))=mult(X1,X1))).
########
#cnf(i_0_104129, plain, (ld(X1,ld(X1,f(mult(X1,X1))))=ld(f(mult(X1,X1)),unit))).
############
#cnf(i_0_105691, plain, (mult(rd(X1,mult(X1,X2)),mult(X1,mult(X2,X2)))=mult(X1,X2))).
####
#cnf(i_0_105857, plain, (mult(ld(X1,f(X2)),mult(f(X2),X1))=mult(ld(X1,X2),X1))).
###########
#cnf(i_0_105935, plain, (mult(ld(X1,ld(X1,f(X2))),f(X2))=ld(X1,ld(X1,X2)))).
##########################
#cnf(i_0_109456, plain, (mult(mult(X1,rd(X1,X2)),rd(X2,mult(X2,X1)))=rd(X1,X2))).
#
#cnf(i_0_109463, plain, (mult(mult(X1,mult(X1,X2)),ld(X2,ld(X1,X2)))=mult(X1,X2))).
########
#cnf(i_0_111665, plain, (mult(ld(X2,ld(X1,f(X2))),f(X2))=ld(X2,ld(X1,X2)))).
#########
#cnf(i_0_119070, plain, (rd(rd(f(mult(X1,X2)),X2),X1)=ld(f(mult(X1,X2)),unit))).
#########
#cnf(i_0_119729, plain, (ld(X2,ld(X1,f(mult(X1,X2))))=ld(f(mult(X1,X2)),unit))).
##########
#cnf(i_0_121119, plain, (rd(ld(X1,ld(f(X1),X2)),X2)=ld(X1,ld(f(X1),unit)))).
#
#cnf(i_0_121636, plain, (rd(ld(f(X1),X2),mult(X1,X2))=ld(X1,ld(f(X1),unit)))).
#####
#cnf(i_0_121901, plain, (rd(ld(X1,X2),mult(f(X1),X2))=ld(X1,ld(f(X1),unit)))).
###
#cnf(i_0_122192, plain, (rd(mult(f(rd(X1,X2)),X2),X1)=ld(f(rd(X1,X2)),unit))).
########
#cnf(i_0_122683, plain, (rd(X1,mult(X2,mult(f(X2),X1)))=ld(X2,ld(f(X2),unit)))).
###
#cnf(i_0_125063, plain, (ld(mult(X2,X1),mult(X2,mult(X1,ld(X2,X1))))=ld(X2,X1))).
#######################
#cnf(i_0_132204, plain, (rd(mult(X1,mult(X2,mult(X1,X2))),mult(X1,X2))=mult(X1,X2))).
###
#cnf(i_0_132205, plain, (ld(mult(X1,X2),mult(X1,mult(X2,mult(X1,X2))))=mult(X1,X2))).
###################################
#cnf(i_0_133919, plain, (mult(mult(X1,mult(X2,ld(X1,X2))),ld(X2,X1))=mult(X1,X2))).
##
#cnf(i_0_133945, plain, (ld(mult(X1,mult(X2,ld(X1,X2))),mult(X1,X2))=ld(X2,X1))).
#######
#cnf(i_0_134454, plain, (mult(mult(X1,X2),ld(X2,ld(X1,ld(X1,X2))))=ld(X1,X2))).
##########
#cnf(i_0_136237, plain, (mult(rd(X1,mult(X1,mult(X2,X2))),mult(X1,X2))=rd(X1,X2))).
###
#cnf(i_0_138054, plain, (mult(mult(X1,X2),ld(X2,ld(X1,rd(X2,X1))))=rd(X2,X1))).
#####
#cnf(i_0_138427, plain, (mult(rd(X1,mult(X2,mult(X1,X2))),mult(X2,X1))=rd(X1,X2))).
###########################################################
#cnf(i_0_164106, plain, (mult(X1,ld(X2,f(rd(X2,X1))))=ld(f(rd(X2,X1)),unit))).
###############################
#cnf(i_0_181743, plain, (ld(mult(X1,mult(X2,rd(X2,X1))),mult(X1,X2))=rd(X1,X2))).
#################
#cnf(i_0_184764, plain, (mult(mult(X1,rd(X1,X2)),mult(X2,ld(X1,X2)))=mult(X1,X2))).
####
#cnf(i_0_185484, plain, (ld(mult(X1,mult(X2,X2)),mult(X1,mult(X2,X1)))=ld(X2,X1))).
#########
#cnf(i_0_186846, plain, (rd(mult(X1,mult(X1,X2)),mult(X2,mult(X1,X2)))=rd(X1,X2))).
#####
#cnf(i_0_186931, plain, (ld(mult(X1,mult(X2,X1)),mult(X1,mult(X2,X2)))=ld(X1,X2))).
####
#cnf(i_0_187260, plain, (mult(mult(X2,ld(X1,X2)),mult(ld(X2,X1),X1))=mult(X2,X1))).
###
#cnf(i_0_187773, plain, (rd(mult(X2,mult(X1,X2)),mult(X1,mult(X1,X2)))=rd(X2,X1))).
#######################################################
#cnf(i_0_96455, plain, (mult(ld(f(X2),X1),mult(ld(X1,X2),X2))=mult(X2,f(X2)))).
############
#cnf(i_0_32033, plain, (mult(rd(X2,X1),rd(X1,mult(X2,f(X2))))=ld(X2,f(X2)))).
##################
#cnf(i_0_56144, plain, (mult(mult(X2,rd(X2,X1)),rd(X1,f(X2)))=mult(X2,f(X2)))).
#############
#cnf(i_0_105869, plain, (mult(ld(X2,X1),mult(ld(X1,f(X2)),X1))=ld(f(X2),X1))).
##########
#cnf(i_0_108520, plain, (rd(mult(ld(X2,X1),X2),mult(f(X1),X2))=ld(X2,f(X1)))).
#########
#cnf(i_0_109002, plain, (mult(mult(X2,rd(X1,X2)),rd(X2,f(X1)))=mult(X2,f(X1)))).
#
#cnf(i_0_109081, plain, (ld(mult(X1,f(X2)),mult(X1,rd(X2,X1)))=rd(f(X2),X1))).
######
#cnf(i_0_109479, plain, (mult(rd(X1,X2),rd(X2,mult(X2,f(X1))))=rd(f(X1),X2))).
#
#cnf(i_0_109484, plain, (mult(mult(X1,X2),ld(X2,ld(f(X1),X2)))=mult(f(X1),X2))).
####################################
#cnf(i_0_121675, plain, (rd(mult(X1,mult(X1,X2)),mult(f(X1),X2))=mult(X1,f(X1)))).
#####
#cnf(i_0_122321, plain, (rd(mult(X1,X2),mult(X1,mult(f(X1),X2)))=ld(X1,f(X1)))).
##################
#cnf(i_0_123542, plain, (mult(mult(X2,rd(f(X1),X2)),rd(X2,X1))=rd(X2,f(X1)))).
#################################
#cnf(i_0_154964, plain, (mult(rd(X2,X1),rd(X1,mult(f(X1),X2)))=ld(X1,f(X1)))).
##############################
#cnf(i_0_163181, plain, (mult(ld(f(X1),X2),mult(ld(X2,X1),X2))=mult(f(X1),X2))).
###
#cnf(i_0_163233, plain, (mult(ld(X1,ld(f(X2),X1)),ld(X1,X2))=ld(X1,f(X2)))).
####
#cnf(i_0_163594, plain, (rd(mult(ld(f(X1),X2),X1),mult(X2,X1))=ld(X1,f(X1)))).
##############################################################
#cnf(i_0_171379, plain, (mult(rd(X1,mult(X1,f(X2))),mult(X1,X2))=mult(X1,f(X2)))).
######################
#cnf(i_0_182092, plain, (mult(rd(X1,mult(X1,X2)),mult(X1,f(X2)))=rd(X1,f(X2)))).
#######################
#cnf(i_0_186138, plain, (mult(mult(X2,ld(f(X1),X2)),ld(X2,X1))=mult(X2,f(X1)))).
##
#cnf(i_0_187436, plain, (rd(mult(X2,mult(f(X1),X2)),mult(X1,X2))=rd(X2,f(X1)))).
###########
#cnf(i_0_190827, plain, (mult(rd(X1,X2),mult(X2,ld(f(X1),X2)))=mult(f(X1),X2))).
########
#cnf(i_0_191805, plain, (ld(mult(X1,X2),mult(X1,mult(f(X2),X1)))=ld(f(X2),X1))).
#####
#cnf(i_0_191832, plain, (ld(mult(X1,mult(f(X2),X1)),mult(X1,X2))=ld(X1,f(X2)))).
##########
#cnf(i_0_192987, plain, (rd(mult(X1,X2),mult(X2,mult(f(X1),X2)))=rd(f(X1),X2))).
##########
#cnf(i_0_199596, plain, (mult(mult(f(X1),X2),ld(X2,ld(X1,X2)))=ld(f(X1),X2))).
##################################################################################
#cnf(i_0_219781, plain, (ld(mult(X2,X1),mult(X2,rd(X1,f(X2))))=ld(X2,f(X2)))).
################
#cnf(i_0_221028, plain, (mult(ld(X2,X1),ld(X1,rd(X2,f(X1))))=ld(X1,f(X1)))).
##############
#cnf(i_0_221191, plain, (mult(ld(X1,ld(f(X1),X2)),ld(X2,X1))=ld(X1,f(X1)))).
###################
#cnf(i_0_225971, plain, (ld(mult(X1,mult(X2,f(X1))),mult(X1,X2))=ld(X1,f(X1)))).
###########################################
#cnf(i_0_96390, plain, (mult(ld(X2,X1),ld(X1,f(f(X2))))=ld(X2,f(f(X2))))).
#
#cnf(i_0_96448, plain, (rd(ld(X1,f(X2)),mult(X2,f(X2)))=ld(X1,ld(X2,unit)))).
#
#cnf(i_0_96632, plain, (mult(ld(X1,ld(X2,f(f(X2)))),X2)=ld(X1,f(f(X2))))).
###
#cnf(i_0_96730, plain, (mult(ld(X1,X2),ld(X2,f(f(X2))))=ld(X1,f(f(X2))))).
#
#cnf(i_0_228971, plain, (ld(mult(X1,X2),f(f(X1)))=ld(X2,ld(X1,f(f(X1)))))).
########
#cnf(i_0_97219, plain, (rd(f(X2),mult(X1,mult(X2,f(X2))))=ld(X2,ld(X1,unit)))).
######
#cnf(i_0_99206, plain, (rd(ld(X1,ld(X2,f(X2))),f(X2))=ld(X1,ld(X2,unit)))).
########
#cnf(i_0_103203, plain, (mult(ld(X1,ld(X1,ld(X2,unit))),X2)=ld(X1,ld(X1,unit)))).
########
#cnf(i_0_104990, plain, (mult(ld(X1,ld(X2,ld(X2,unit))),X2)=ld(X1,ld(X2,unit)))).
###
#cnf(i_0_105005, plain, (mult(ld(X1,ld(X2,ld(X1,unit))),X1)=ld(X1,ld(X2,unit)))).
################
#cnf(i_0_112019, plain, (rd(f(X1),mult(X1,mult(X2,f(X1))))=ld(X2,ld(X1,unit)))).
###########
#cnf(i_0_122087, plain, (ld(mult(f(X1),X2),ld(X1,f(X1)))=ld(X2,ld(X1,unit)))).
###########################
#cnf(i_0_175378, plain, (mult(X1,mult(ld(X1,X2),f(f(X1))))=mult(X2,f(f(X1))))).
#####
#cnf(i_0_176193, plain, (mult(X1,mult(f(f(X1)),ld(X1,X2)))=mult(f(f(X1)),X2))).
##
#cnf(i_0_232194, plain, (ld(X1,mult(f(f(X1)),X2))=mult(f(f(X1)),ld(X1,X2)))).
#
#cnf(i_0_232221, plain, (rd(ld(X2,X1),f(f(X1)))=ld(mult(f(f(X1)),X2),X1))).
#
#cnf(i_0_232689, plain, (rd(f(f(X1)),mult(X2,X1))=ld(X1,rd(f(f(X1)),X2)))).
####
#cnf(i_0_176420, plain, (ld(X1,ld(f(f(X1)),mult(X1,X2)))=ld(f(f(X1)),X2))).
#
#cnf(i_0_233605, plain, (ld(f(f(X1)),mult(X1,X2))=mult(X1,ld(f(f(X1)),X2)))).
##
#cnf(i_0_234506, plain, (rd(X1,mult(X2,f(f(X1))))=ld(f(f(X1)),rd(X1,X2)))).
##########
#cnf(i_0_202929, plain, (mult(rd(f(f(X1)),X2),rd(X2,X1))=ld(X1,f(f(X1))))).
#
#cnf(i_0_202934, plain, (rd(mult(f(f(X1)),X2),mult(X1,X2))=ld(X1,f(f(X1))))).
#####
#cnf(i_0_227932, plain, (rd(mult(f(f(X1)),ld(X1,X2)),X2)=ld(X1,f(f(X1))))).
#
#cnf(i_0_235880, plain, (mult(ld(X1,f(f(X1))),X2)=mult(f(f(X1)),ld(X1,X2)))).
##
#cnf(i_0_236411, plain, (rd(ld(X1,f(f(X1))),X2)=ld(X1,rd(f(f(X1)),X2)))).
#########
#cnf(i_0_235122, plain, (mult(ld(X1,rd(f(f(X1)),X2)),X2)=ld(X1,f(f(X1))))).
##########
#cnf(i_0_96556, plain, (mult(ld(X2,ld(X2,X1)),mult(ld(X1,X2),mult(X2,X2)))=X2)).
######
#cnf(i_0_96646, plain, (mult(rd(X2,mult(X1,X1)),mult(X1,mult(X1,ld(X2,X1))))=X1)).
#####
#cnf(i_0_96721, plain, (ld(mult(f(f(X2)),X1),X2)=mult(ld(X1,f(X2)),f(f(X2))))).
###
#cnf(i_0_96906, plain, (mult(mult(X1,f(ld(X1,unit))),X2)=ld(f(ld(X1,unit)),X2))).
####
#cnf(i_0_96914, plain, (ld(mult(X1,f(ld(X1,unit))),X2)=mult(f(ld(X1,unit)),X2))).
#######
#cnf(i_0_99189, plain, (mult(mult(X1,mult(ld(X2,X3),X3)),ld(X3,ld(X3,X2)))=X1)).
###
#cnf(i_0_99191, plain, (mult(mult(X1,ld(X2,ld(X2,X3))),mult(ld(X3,X2),X2))=X1)).
##
#cnf(i_0_99546, plain, (mult(mult(X1,mult(f(unit),X2)),ld(X2,f(unit)))=X1)).
##
#cnf(i_0_99548, plain, (mult(mult(X1,ld(X2,f(unit))),mult(f(unit),X2))=X1)).
###
#cnf(i_0_2392, plain, (mult(mult(X1,mult(X1,ld(X2,X1))),ld(X1,rd(X2,X1)))=X1)).
##
#cnf(i_0_99840, plain, (mult(mult(X1,mult(ld(X3,X2),X4)),ld(X4,ld(X2,X3)))=X1)).
##
#cnf(i_0_99842, plain, (mult(mult(X1,ld(X2,ld(X3,X4))),mult(ld(X4,X3),X2))=X1)).
#####
#cnf(i_0_2866, plain, (mult(mult(X2,rd(f(X2),X1)),mult(X1,f(X2)))=mult(X2,X2))).
######
#cnf(i_0_2951, plain, (mult(rd(f(X2),X1),mult(X1,mult(X2,f(X2))))=mult(X2,X2))).
###
#cnf(i_0_3144, plain, (mult(mult(X2,mult(X2,rd(X2,X1))),rd(X1,mult(X2,X2)))=X2)).
##
#cnf(i_0_100645, plain, (mult(mult(X1,mult(rd(X3,X2),X4)),ld(X4,rd(X2,X3)))=X1)).
##
#cnf(i_0_100647, plain, (mult(mult(X1,ld(X2,rd(X3,X4))),mult(rd(X4,X3),X2))=X1)).
##
#cnf(i_0_100938, plain, (ld(mult(X1,mult(X1,X2)),mult(X1,mult(X1,mult(X2,X2))))=X2)).
#
#cnf(i_0_100978, plain, (ld(mult(X1,mult(X2,X2)),mult(X1,mult(X2,mult(X2,X1))))=X1)).
#############
#cnf(i_0_103013, plain, (ld(mult(X1,mult(X2,rd(X2,X1))),mult(X1,mult(X2,X2)))=X1)).
#########
#cnf(i_0_105625, plain, (mult(rd(rd(X1,X2),X3),mult(X3,mult(rd(X2,X1),X3)))=X3)).
#
#cnf(i_0_105677, plain, (mult(mult(X1,rd(X1,X2)),mult(X2,ld(X1,ld(X1,X2))))=X2)).
####
#cnf(i_0_105679, plain, (mult(rd(ld(X1,X2),X3),mult(X3,mult(ld(X2,X1),X3)))=X3)).
#####
#cnf(i_0_105785, plain, (mult(ld(X1,rd(X2,X1)),mult(X1,mult(ld(X2,X1),X1)))=X1)).
####
#cnf(i_0_106051, plain, (ld(f(f(X1)),rd(X1,X2))=mult(f(X1),rd(f(f(X1)),X2)))).
######
#cnf(i_0_106198, plain, (mult(mult(X1,mult(f(X1),X2)),ld(X2,f(X1)))=mult(X1,X1))).
#
#cnf(i_0_106233, plain, (mult(mult(f(X1),X2),mult(ld(X2,X1),f(X1)))=mult(X1,X1))).
#
#cnf(i_0_106330, plain, (mult(f(X1),mult(f(f(X1)),X2))=mult(X1,ld(f(f(X1)),X2)))).
##########
#cnf(i_0_106474, plain, (ld(f(ld(X1,unit)),ld(X1,X2))=mult(f(ld(X1,unit)),X2))).
##
#cnf(i_0_106903, plain, (mult(ld(X1,f(unit)),mult(mult(f(unit),X1),X2))=X2)).
#
#cnf(i_0_106907, plain, (mult(ld(X1,rd(X2,X3)),mult(mult(rd(X3,X2),X1),X4))=X4)).
#
#cnf(i_0_106959, plain, (mult(mult(ld(X1,X2),X2),mult(ld(X2,ld(X2,X1)),X3))=X3)).
#
#cnf(i_0_106960, plain, (mult(mult(ld(X1,X2),X1),mult(ld(X1,ld(X2,X1)),X3))=X3)).
#
#cnf(i_0_106961, plain, (mult(ld(X1,ld(X2,X3)),mult(mult(ld(X3,X2),X1),X4))=X4)).
#####
#cnf(i_0_107072, plain, (mult(mult(f(unit),X1),mult(ld(X1,f(unit)),X2))=X2)).
#
#cnf(i_0_107076, plain, (mult(mult(rd(X2,X1),X3),mult(ld(X3,rd(X1,X2)),X4))=X4)).
###
#cnf(i_0_107130, plain, (mult(mult(ld(X2,X1),X3),mult(ld(X3,ld(X1,X2)),X4))=X4)).
###
#cnf(i_0_107858, plain, (mult(mult(X1,mult(X2,rd(X2,X3))),rd(X3,mult(X2,X2)))=X1)).
##
#cnf(i_0_107859, plain, (mult(mult(X1,mult(X2,rd(X3,X2))),rd(X2,mult(X2,X3)))=X1)).
####
#cnf(i_0_107897, plain, (mult(mult(X1,rd(X2,mult(X2,X3))),mult(X2,rd(X3,X2)))=X1)).
##
#cnf(i_0_107898, plain, (mult(mult(X1,rd(X2,mult(X3,X3))),mult(X3,rd(X3,X2)))=X1)).
##
#cnf(i_0_107912, plain, (mult(mult(X1,mult(X2,ld(X3,X2))),ld(X2,rd(X3,X2)))=X1)).
##
#cnf(i_0_107915, plain, (mult(mult(X1,rd(f(unit),X2)),mult(X2,f(unit)))=X1)).
##
#cnf(i_0_107921, plain, (mult(mult(X1,rd(rd(X2,X3),X4)),mult(X4,rd(X3,X2)))=X1)).
##
#cnf(i_0_107933, plain, (mult(mult(X1,rd(ld(X2,X3),X4)),mult(X4,ld(X3,X2)))=X1)).
########
#cnf(i_0_108013, plain, (mult(mult(X1,ld(X2,rd(X3,X2))),mult(X2,ld(X3,X2)))=X1)).
###
#cnf(i_0_108022, plain, (mult(mult(X1,mult(X2,rd(X4,X3))),rd(rd(X3,X4),X2))=X1)).
##
#cnf(i_0_108034, plain, (mult(mult(X1,mult(X2,ld(X4,X3))),rd(ld(X3,X4),X2))=X1)).
##
#cnf(i_0_108286, plain, (mult(mult(X1,rd(X1,X2)),mult(rd(X2,mult(X1,X1)),X3))=X3)).
#
#cnf(i_0_108287, plain, (mult(mult(X1,rd(X2,X1)),mult(rd(X1,mult(X1,X2)),X3))=X3)).
###
#cnf(i_0_108325, plain, (mult(rd(X1,mult(X1,X2)),mult(mult(X1,rd(X2,X1)),X3))=X3)).
#
#cnf(i_0_108326, plain, (mult(rd(X1,mult(X2,X2)),mult(mult(X2,rd(X2,X1)),X3))=X3)).
#
#cnf(i_0_108340, plain, (mult(mult(X1,ld(X2,X1)),mult(ld(X1,rd(X2,X1)),X3))=X3)).
#
#cnf(i_0_108343, plain, (mult(rd(f(unit),X1),mult(mult(X1,f(unit)),X2))=X2)).
#
#cnf(i_0_108349, plain, (mult(rd(rd(X1,X2),X3),mult(mult(X3,rd(X2,X1)),X4))=X4)).
#
#cnf(i_0_108361, plain, (mult(rd(ld(X1,X2),X3),mult(mult(X3,ld(X2,X1)),X4))=X4)).
##########
#cnf(i_0_108480, plain, (mult(ld(X1,rd(X2,X1)),mult(mult(X1,ld(X2,X1)),X3))=X3)).
#
#cnf(i_0_108483, plain, (mult(mult(X1,f(unit)),mult(rd(f(unit),X1),X2))=X2)).
#
#cnf(i_0_108489, plain, (mult(mult(X1,rd(X3,X2)),mult(rd(rd(X2,X3),X1),X4))=X4)).
#
#cnf(i_0_108501, plain, (mult(mult(X1,ld(X3,X2)),mult(rd(ld(X2,X3),X1),X4))=X4)).
#####
#cnf(i_0_108736, plain, (ld(X1,ld(f(ld(X1,unit)),X2))=mult(f(ld(X1,unit)),X2))).
############
#cnf(i_0_109450, plain, (mult(mult(X1,mult(X1,X2)),ld(X2,ld(X1,ld(X1,X2))))=X2)).
##
#cnf(i_0_109451, plain, (mult(mult(X1,mult(X2,X2)),ld(X2,ld(X2,ld(X1,X2))))=X2)).
################################
#cnf(i_0_6394, plain, (mult(ld(X2,ld(X2,X1)),mult(ld(X1,X2),mult(X2,X1)))=X1)).
####
#cnf(i_0_127879, plain, (mult(ld(X3,ld(X2,X1)),mult(ld(X1,X2),mult(X3,X3)))=X3)).
########
#cnf(i_0_128601, plain, (mult(ld(X1,f(unit)),mult(f(unit),mult(X1,X1)))=X1)).
########
#cnf(i_0_129748, plain, (mult(ld(X3,rd(X2,X1)),mult(rd(X1,X2),mult(X3,X3)))=X3)).
#####################
#cnf(i_0_136953, plain, (mult(ld(X1,rd(X2,X1)),mult(rd(X1,X2),mult(X1,X2)))=X2)).
#
#cnf(i_0_137788, plain, (mult(mult(X1,mult(X2,X1)),ld(mult(X2,X1),ld(X1,X2)))=X2)).
#############################
#cnf(i_0_153740, plain, (mult(ld(f(f(f(X1))),X1),f(f(f(X1))))=X1)).
##################################
#cnf(i_0_159880, plain, (rd(X1,f(mult(X1,f(X1))))=ld(f(X1),f(mult(X1,f(X1)))))).
######
#cnf(i_0_162947, plain, (rd(f(mult(X1,f(X1))),X1)=ld(f(mult(X1,f(X1))),f(X1)))).
######
#cnf(i_0_164437, plain, (mult(f(ld(X2,unit)),mult(X2,X1))=ld(f(ld(X2,unit)),X1))).
##########################################################################
#cnf(i_0_232221, plain, (rd(ld(X1,X2),f(f(X2)))=mult(ld(X1,f(X2)),f(f(X2))))).
#################################################################################
#cnf(i_0_234506, plain, (rd(X1,mult(X2,f(f(X1))))=mult(f(X1),rd(f(f(X1)),X2)))).
#######################################################################################
#cnf(i_0_276484, plain, (mult(X1,mult(f(ld(X1,unit)),X2))=ld(f(ld(X1,unit)),X2))).
############
#cnf(i_0_96451, plain, (mult(ld(X1,f(X2)),mult(X2,X2))=mult(ld(X1,X2),mult(X2,f(X2))))).
#
#cnf(i_0_96535, plain, (mult(ld(X2,ld(f(X2),X1)),mult(ld(X1,X2),X2))=f(X2))).
######
#cnf(i_0_97608, plain, (rd(X1,f(f(ld(X1,unit))))=ld(f(f(ld(X1,unit))),X1))).
#####
#cnf(i_0_98853, plain, (mult(ld(f(X1),rd(X2,X1)),mult(X1,ld(X2,X1)))=f(X1))).
#
#cnf(i_0_286498, plain, (mult(ld(f(X1),X2),ld(X2,rd(X1,X2)))=rd(f(X1),X2))).
##
#cnf(i_0_286898, plain, (mult(ld(X2,f(X1)),rd(f(X1),X2))=ld(X2,rd(X1,X2)))).
##
#cnf(i_0_286925, plain, (mult(rd(X2,f(X1)),ld(f(X1),X2))=mult(X2,ld(X1,X2)))).
######
#cnf(i_0_286866, plain, (mult(ld(X1,rd(X2,X1)),rd(X1,f(X2)))=ld(X1,f(X2)))).
#
#cnf(i_0_286899, plain, (mult(rd(f(X1),X2),mult(X2,ld(X1,X2)))=ld(f(X1),X2))).
######
#cnf(i_0_286944, plain, (mult(mult(X2,ld(X1,X2)),ld(X2,f(X1)))=rd(X2,f(X1)))).
#########
#cnf(i_0_287068, plain, (ld(mult(X2,f(X1)),mult(X2,mult(X1,X2)))=mult(f(X1),X2))).
#
#cnf(i_0_289426, plain, (mult(mult(X1,f(X2)),mult(f(X2),X1))=mult(X1,mult(X2,X1)))).
##################
#cnf(i_0_288317, plain, (rd(mult(X1,mult(X2,X1)),mult(f(X2),X1))=mult(X1,f(X2)))).
########
#cnf(i_0_98882, plain, (mult(rd(X1,X2),mult(X1,f(X1)))=mult(X1,mult(ld(X2,X1),f(X1))))).
##
#cnf(i_0_99038, plain, (rd(mult(ld(f(X1),X2),X2),X1)=ld(f(X1),mult(X2,rd(X2,X1))))).
###
#cnf(i_0_3308, plain, (mult(mult(X1,ld(X2,X1)),ld(X1,rd(X2,f(X1))))=f(X1))).
#######
#cnf(i_0_55354, plain, (mult(rd(f(ld(X1,X2)),X1),X2)=mult(ld(X1,X2),f(ld(X1,X2))))).
#
#cnf(i_0_3296, plain, (mult(mult(X2,rd(X2,X1)),rd(X1,mult(X2,f(X2))))=f(X2))).
#
#cnf(i_0_100490, plain, (rd(rd(X1,X2),f(rd(X2,X1)))=ld(f(rd(X2,X1)),rd(X1,X2)))).
######
#cnf(i_0_100802, plain, (ld(mult(X1,mult(X1,f(X2))),mult(X1,mult(X1,X2)))=f(X2))).
#
#cnf(i_0_100807, plain, (mult(mult(X1,mult(X1,f(X2))),X2)=mult(X1,mult(X1,mult(X2,f(X2)))))).
#
#cnf(i_0_100986, plain, (ld(f(X1),mult(X2,mult(X2,X1)))=mult(mult(ld(f(X1),X2),X2),X1))).
###
#cnf(i_0_101051, plain, (mult(mult(X1,mult(X1,X2)),f(X2))=mult(X1,mult(X1,mult(X2,f(X2)))))).
#
#cnf(i_0_101213, plain, (rd(mult(X1,rd(X1,f(X2))),X2)=mult(X1,rd(X1,mult(X2,f(X2)))))).
#
#cnf(i_0_101295, plain, (rd(mult(X1,rd(X1,X2)),f(X2))=mult(X1,rd(X1,mult(X2,f(X2)))))).
##
#cnf(i_0_102209, plain, (ld(mult(X1,mult(X2,f(X2))),X2)=ld(X2,ld(f(X2),ld(X1,X2))))).
####################
#cnf(i_0_106249, plain, (mult(ld(X2,rd(f(X1),X2)),X1)=mult(ld(X2,f(X1)),ld(X2,X1)))).
##
#cnf(i_0_107615, plain, (rd(rd(X1,X2),mult(X1,f(X1)))=rd(X1,mult(X1,mult(f(X1),X2))))).
###
#cnf(i_0_108606, plain, (mult(mult(X1,mult(X2,f(X1))),X2)=mult(X1,mult(X2,mult(f(X1),X2))))).
#
#cnf(i_0_108633, plain, (mult(mult(X1,X2),rd(X2,f(X1)))=mult(X1,mult(X2,rd(X2,f(X1)))))).
##
#cnf(i_0_109732, plain, (ld(X1,mult(X2,ld(f(X1),X2)))=mult(ld(X1,X2),ld(f(X1),X2)))).
###
#cnf(i_0_109875, plain, (mult(ld(f(X1),f(unit)),mult(f(unit),X1))=f(X1))).
#
#cnf(i_0_109878, plain, (mult(ld(f(X1),rd(X2,X3)),mult(rd(X3,X2),X1))=f(X1))).
#
#cnf(i_0_109930, plain, (mult(mult(ld(f(X1),X2),X2),ld(X2,ld(X2,X1)))=f(X1))).
####
#cnf(i_0_109931, plain, (mult(mult(ld(f(X1),X2),X1),ld(X1,ld(X2,X1)))=f(X1))).
###
#cnf(i_0_109932, plain, (mult(ld(f(X1),ld(X2,X3)),mult(ld(X3,X2),X1))=f(X1))).
########
#cnf(i_0_110103, plain, (mult(rd(X1,mult(X2,X2)),mult(X2,rd(X2,f(X1))))=f(X1))).
####
#cnf(i_0_110215, plain, (mult(mult(X1,rd(X3,X2)),rd(rd(X2,X3),f(X1)))=f(X1))).
#
#cnf(i_0_110227, plain, (mult(mult(X1,ld(X3,X2)),rd(ld(X2,X3),f(X1)))=f(X1))).
######
#cnf(i_0_110800, plain, (mult(rd(f(X1),X2),mult(X1,X1))=mult(f(X1),mult(ld(X2,X1),X1)))).
#
#cnf(i_0_111283, plain, (mult(ld(X1,ld(X2,f(X1))),X1)=mult(ld(X1,ld(X2,X1)),f(X1)))).
####
#cnf(i_0_111385, plain, (mult(ld(X1,ld(X1,f(X2))),X2)=mult(ld(X1,ld(X1,X2)),f(X2)))).
#########
#cnf(i_0_52394, plain, (rd(X1,mult(ld(f(X2),X3),X2))=mult(X1,ld(X2,ld(X3,f(X2)))))).
##
#cnf(i_0_111753, plain, (ld(mult(X2,X1),ld(f(X1),unit))=ld(X1,ld(X2,ld(f(X1),unit))))).
##
#cnf(i_0_111754, plain, (ld(mult(X1,mult(X2,X2)),f(X2))=ld(X2,ld(X2,ld(X1,f(X2)))))).
#
#cnf(i_0_111755, plain, (ld(mult(X1,mult(X1,X2)),f(X2))=ld(X2,ld(X1,ld(X1,f(X2)))))).
######
#cnf(i_0_111958, plain, (ld(mult(X1,mult(X1,f(X2))),X2)=ld(f(X2),ld(X1,ld(X1,X2))))).
###
#cnf(i_0_3213, plain, (ld(mult(X1,mult(X2,f(X1))),mult(X1,mult(X2,X1)))=f(X1))).
#
#cnf(i_0_53360, plain, (mult(mult(X1,X2),mult(X1,f(X1)))=mult(X1,mult(X2,mult(X1,f(X1)))))).
####
#cnf(i_0_112087, plain, (rd(mult(X1,X2),mult(X1,f(X1)))=mult(X1,rd(X2,mult(X1,f(X1)))))).
###
#cnf(i_0_3304, plain, (mult(rd(X1,mult(X1,X2)),mult(X1,rd(X2,f(X1))))=f(X1))).
########
#cnf(i_0_112659, plain, (rd(mult(X1,mult(X1,X2)),mult(X1,mult(f(X1),X2)))=f(X1))).
######
#cnf(i_0_113181, plain, (ld(f(X1),mult(ld(X1,X2),X2))=ld(X1,mult(ld(f(X1),X2),X2)))).
#####
#cnf(i_0_113258, plain, (mult(mult(X1,rd(X2,X1)),rd(X1,mult(f(X1),X2)))=f(X1))).
####
#cnf(i_0_53200, plain, (rd(ld(X1,X2),mult(X1,f(X1)))=ld(X1,rd(X2,mult(X1,f(X1)))))).
########
#cnf(i_0_113530, plain, (rd(mult(ld(X1,X2),X2),f(X1))=ld(X1,mult(X2,rd(X2,f(X1)))))).
####
#cnf(i_0_3778, plain, (ld(X1,mult(X2,mult(X1,f(X1))))=mult(ld(X1,X2),mult(X1,f(X1))))).
#####
#cnf(i_0_113644, plain, (ld(X1,mult(X2,mult(X2,f(X1))))=mult(mult(ld(X1,X2),X2),f(X1)))).
#
#cnf(i_0_113705, plain, (ld(X1,mult(ld(X1,X2),f(X2)))=mult(ld(X1,ld(X1,X2)),f(X2)))).
#
#cnf(i_0_113739, plain, (rd(mult(ld(X1,X2),f(X2)),X1)=ld(X1,mult(X2,rd(f(X2),X1))))).
#######
#cnf(i_0_114305, plain, (rd(X1,mult(X2,f(mult(X1,X1))))=ld(X1,rd(f(mult(X1,X1)),X2)))).
###
#cnf(i_0_116555, plain, (rd(f(mult(X1,X1)),mult(X2,X1))=ld(f(mult(X1,X1)),rd(X1,X2)))).
####
#cnf(i_0_118099, plain, (mult(ld(X2,rd(X1,X2)),f(X1))=mult(ld(X2,X1),ld(X2,f(X1))))).
####
#cnf(i_0_118537, plain, (mult(mult(X1,X2),f(mult(X1,X2)))=mult(X1,mult(X2,f(mult(X1,X2)))))).
#
#cnf(i_0_118554, plain, (ld(X1,mult(X2,f(ld(X1,X2))))=mult(ld(X1,X2),f(ld(X1,X2))))).
########
#cnf(i_0_119331, plain, (rd(mult(X1,rd(f(X1),X2)),X2)=mult(X1,rd(f(X1),mult(X2,X2))))).
############
#cnf(i_0_120943, plain, (mult(rd(X1,X2),f(rd(X1,X2)))=mult(X1,ld(X2,f(rd(X1,X2)))))).
##
#cnf(i_0_121047, plain, (ld(X1,mult(ld(f(X1),X2),X2))=mult(ld(X1,ld(f(X1),X2)),X2))).
######
#cnf(i_0_121603, plain, (mult(mult(X1,X2),ld(f(X1),X2))=mult(X1,mult(X2,ld(f(X1),X2))))).
##
#cnf(i_0_122005, plain, (rd(mult(f(rd(X1,X2)),X1),X2)=mult(X1,ld(X2,f(rd(X1,X2)))))).
#
#cnf(i_0_122013, plain, (mult(mult(f(mult(X1,X2)),X1),X2)=mult(X1,mult(X2,f(mult(X1,X2)))))).
#
#cnf(i_0_122371, plain, (ld(X1,mult(X2,mult(f(X2),X1)))=mult(mult(ld(X1,X2),f(X2)),X1))).
#
#cnf(i_0_122397, plain, (mult(mult(X1,mult(f(X1),X2)),X2)=mult(X1,mult(f(X1),mult(X2,X2))))).
###
#cnf(i_0_122736, plain, (ld(mult(X1,mult(f(X1),X2)),X2)=ld(X2,ld(X1,ld(f(X1),X2))))).
#######
#cnf(i_0_126662, plain, (rd(ld(X1,ld(X2,f(X2))),X2)=ld(X1,ld(X2,ld(X2,f(X2)))))).
######################
#cnf(i_0_135074, plain, (rd(rd(f(X1),X2),mult(X1,X1))=rd(f(X1),mult(X1,mult(X1,X2))))).
#
#cnf(i_0_323971, plain, (rd(rd(f(f(X1)),X2),X1)=rd(f(f(X1)),mult(X1,X2)))).
################################
#cnf(i_0_156581, plain, (mult(rd(f(X1),X2),rd(X1,X2))=mult(f(X1),ld(X2,rd(X1,X2))))).
##
#cnf(i_0_156585, plain, (mult(mult(f(X1),X2),mult(X1,X2))=mult(f(X1),mult(X2,mult(X1,X2))))).
######
#cnf(i_0_159936, plain, (rd(ld(X2,X1),f(mult(X1,X1)))=ld(mult(X1,X2),f(mult(X1,X1))))).
##
#cnf(i_0_160281, plain, (rd(mult(X1,ld(f(X2),X1)),X2)=mult(rd(X1,f(X2)),rd(X1,X2)))).
####
#cnf(i_0_160551, plain, (rd(mult(X1,ld(X2,X1)),f(X2))=mult(rd(X1,X2),rd(X1,f(X2))))).
#####
#cnf(i_0_160944, plain, (mult(mult(X1,mult(f(X2),X1)),X2)=mult(mult(X1,f(X2)),mult(X1,X2)))).
######
#cnf(i_0_161232, plain, (mult(mult(X1,mult(X2,X1)),f(X2))=mult(mult(X1,X2),mult(X1,f(X2))))).
#########
#cnf(i_0_163004, plain, (ld(f(mult(X1,X1)),mult(X1,X2))=mult(f(mult(X1,X1)),ld(X1,X2)))).
#####
#cnf(i_0_163870, plain, (mult(mult(ld(f(X1),X2),X2),X1)=mult(ld(f(X1),X2),mult(X2,X1)))).
###################
#cnf(i_0_174173, plain, (ld(mult(X1,X2),rd(f(X2),X1))=ld(X2,ld(X1,rd(f(X2),X1))))).
#
#cnf(i_0_174182, plain, (ld(mult(X1,f(X2)),rd(X2,X1))=ld(f(X2),ld(X1,rd(X2,X1))))).
###########
#cnf(i_0_188104, plain, (ld(X1,mult(f(X2),ld(X1,X2)))=mult(ld(X1,f(X2)),ld(X1,X2)))).
##
#cnf(i_0_188524, plain, (mult(ld(X1,rd(X2,f(X1))),X2)=mult(ld(X1,X2),ld(f(X1),X2)))).
###
#cnf(i_0_190706, plain, (rd(mult(X1,ld(X2,f(X1))),X2)=mult(rd(X1,X2),rd(f(X1),X2)))).
##
#cnf(i_0_190710, plain, (mult(mult(X1,X2),mult(f(X1),X2))=mult(X1,mult(X2,mult(f(X1),X2))))).
#####
#cnf(i_0_191259, plain, (rd(mult(f(X2),ld(X1,X2)),X1)=mult(f(X2),ld(X1,rd(X2,X1))))).
##
#cnf(i_0_191945, plain, (mult(mult(f(X1),mult(X2,X1)),X2)=mult(f(X1),mult(X2,mult(X1,X2))))).
####
#cnf(i_0_192434, plain, (mult(ld(f(X1),rd(X2,X1)),X2)=mult(ld(f(X1),X2),ld(X1,X2)))).
##################
#cnf(i_0_202793, plain, (ld(f(X1),mult(X2,mult(X1,X2)))=mult(ld(f(X1),X2),mult(X1,X2)))).
###
#cnf(i_0_203011, plain, (rd(ld(f(rd(X1,X2)),X2),X1)=ld(f(rd(X1,X2)),rd(X2,X1)))).
#
#cnf(i_0_203282, plain, (ld(X1,mult(X2,mult(f(X1),X2)))=mult(ld(X1,X2),mult(f(X1),X2)))).
#
#cnf(i_0_204162, plain, (rd(X1,mult(f(rd(X2,X1)),X2))=ld(f(rd(X2,X1)),rd(X1,X2)))).
##
#cnf(i_0_206421, plain, (mult(mult(X1,X2),f(ld(X2,X1)))=mult(X1,mult(X2,f(ld(X2,X1)))))).
##
#cnf(i_0_206553, plain, (rd(mult(X1,X2),f(ld(X1,X2)))=mult(X1,rd(X2,f(ld(X1,X2)))))).
##
#cnf(i_0_206571, plain, (mult(rd(X1,X2),f(mult(X2,X1)))=mult(X1,ld(X2,f(mult(X2,X1)))))).
##################
#cnf(i_0_212244, plain, (rd(mult(X1,f(mult(X1,X2))),X2)=mult(X1,rd(f(mult(X1,X2)),X2)))).
#
#cnf(i_0_212260, plain, (mult(rd(X1,f(rd(X2,X1))),X2)=mult(X1,ld(f(rd(X2,X1)),X2)))).
#####
#cnf(i_0_213378, plain, (mult(mult(X1,f(X2)),mult(X1,X2))=mult(X1,mult(f(X2),mult(X1,X2))))).
#
#cnf(i_0_213415, plain, (mult(mult(X1,f(ld(X1,X2))),X2)=mult(X1,mult(f(ld(X1,X2)),X2)))).
##
#cnf(i_0_213568, plain, (mult(rd(X1,f(ld(X2,X1))),X2)=mult(X1,ld(f(ld(X2,X1)),X2)))).
##
#cnf(i_0_213586, plain, (ld(X1,mult(f(mult(X1,X2)),X2))=mult(ld(X1,f(mult(X1,X2))),X2))).
#
#cnf(i_0_213798, plain, (mult(rd(X1,f(X2)),rd(X1,X2))=mult(X1,ld(f(X2),rd(X1,X2))))).
##
#cnf(i_0_214011, plain, (rd(rd(X1,f(mult(X2,X1))),X2)=rd(X1,mult(X2,f(mult(X2,X1)))))).
##
#cnf(i_0_214027, plain, (rd(ld(X1,f(rd(X1,X2))),X2)=ld(X1,rd(f(rd(X1,X2)),X2)))).
###############
#cnf(i_0_231668, plain, (rd(ld(X2,ld(X2,f(X1))),X1)=ld(X2,ld(X2,ld(X1,f(X1)))))).
#############################################
#cnf(i_0_292842, plain, (ld(f(ld(X1,X2)),ld(X2,X1))=ld(X2,rd(X1,f(ld(X1,X2)))))).
#######
#cnf(i_0_295308, plain, (mult(mult(X1,ld(f(X2),X1)),X2)=mult(rd(X1,f(X2)),mult(X1,X2)))).
###
#cnf(i_0_296390, plain, (mult(mult(X1,ld(X2,X1)),f(X2))=mult(rd(X1,X2),mult(X1,f(X2))))).
###
#cnf(i_0_296974, plain, (rd(mult(X1,mult(f(X2),X1)),X2)=mult(mult(X1,f(X2)),rd(X1,X2)))).
##
#cnf(i_0_297015, plain, (rd(ld(X1,rd(f(X2),X1)),X2)=rd(ld(X1,f(X2)),mult(X2,X1)))).
#
#cnf(i_0_297385, plain, (rd(mult(X1,mult(X2,X1)),f(X2))=mult(mult(X1,X2),rd(X1,f(X2))))).
#
#cnf(i_0_297417, plain, (rd(ld(X1,rd(X2,X1)),f(X2))=rd(ld(X1,X2),mult(f(X2),X1)))).
######
#cnf(i_0_298925, plain, (mult(rd(X1,X2),rd(f(X1),X2))=mult(X1,ld(X2,rd(f(X1),X2))))).
###
#cnf(i_0_299414, plain, (rd(rd(f(X1),X2),mult(X1,X2))=rd(f(X1),mult(X1,mult(X2,X2))))).
####
#cnf(i_0_300256, plain, (rd(ld(X1,rd(f(X1),X2)),X2)=ld(X1,rd(f(X1),mult(X2,X2))))).
########################
#cnf(i_0_315210, plain, (ld(mult(X2,X1),f(rd(X1,X2)))=ld(X1,ld(X2,f(rd(X1,X2)))))).
#
#cnf(i_0_315226, plain, (ld(X1,mult(X2,f(mult(X1,X2))))=mult(ld(X1,X2),f(mult(X1,X2))))).
#
#cnf(i_0_315529, plain, (mult(ld(X1,f(mult(X1,X1))),X1)=mult(X1,ld(f(mult(X1,X1)),X1)))).
#########
#cnf(i_0_316093, plain, (ld(X1,mult(f(mult(X1,X1)),X2))=mult(X1,ld(f(mult(X1,X1)),X2)))).
####
#cnf(i_0_316924, plain, (ld(f(X2),mult(X1,ld(X2,X1)))=mult(ld(f(X2),X1),ld(X2,X1)))).
####
#cnf(i_0_317163, plain, (ld(mult(X1,X2),f(rd(X1,X2)))=ld(X2,ld(X1,f(rd(X1,X2)))))).
#############
#cnf(i_0_113181, plain, (ld(f(X1),mult(ld(X1,X2),X2))=mult(ld(X1,ld(f(X1),X2)),X2))).
########
#cnf(i_0_324778, plain, (mult(rd(f(mult(X1,X2)),X2),X1)=mult(f(mult(X1,X2)),ld(X2,X1)))).
#
#cnf(i_0_324812, plain, (rd(rd(f(ld(X1,X2)),X2),X1)=rd(f(ld(X1,X2)),mult(X1,X2)))).
#
#cnf(i_0_325313, plain, (mult(mult(f(rd(X1,X2)),X2),X1)=mult(f(rd(X1,X2)),mult(X2,X1)))).
#
#cnf(i_0_326808, plain, (rd(ld(X2,X1),f(mult(X2,X1)))=ld(X2,rd(X1,f(mult(X2,X1)))))).
##
#cnf(i_0_100986, plain, (ld(f(X1),mult(X2,mult(X2,X1)))=mult(ld(f(X1),X2),mult(X2,X1)))).
#
#cnf(i_0_329007, plain, (ld(mult(X2,X1),ld(X2,f(X1)))=ld(X1,ld(X2,ld(X2,f(X1)))))).
#####
#cnf(i_0_330247, plain, (ld(mult(X1,f(mult(X2,X1))),X2)=ld(f(mult(X2,X1)),ld(X1,X2)))).
########
#cnf(i_0_332918, plain, (rd(rd(f(rd(X1,X2)),X2),X1)=rd(f(rd(X1,X2)),mult(X1,X2)))).
#
#cnf(i_0_332939, plain, (rd(mult(f(mult(X1,X2)),X2),X1)=mult(f(mult(X1,X2)),rd(X2,X1)))).
#
#cnf(i_0_333461, plain, (mult(mult(f(X1),X2),rd(X2,X1))=mult(f(X1),mult(X2,rd(X2,X1))))).
#
#cnf(i_0_333745, plain, (mult(mult(f(ld(X1,X2)),X2),X1)=mult(f(ld(X1,X2)),mult(X2,X1)))).
#
#cnf(i_0_334025, plain, (ld(X2,mult(X1,ld(X2,f(X1))))=mult(ld(X2,X1),ld(X2,f(X1))))).
#
#cnf(i_0_334212, plain, (ld(mult(X1,X2),rd(f(X1),X2))=ld(X2,ld(X1,rd(f(X1),X2))))).
#
#cnf(i_0_334743, plain, (ld(f(rd(X1,X2)),mult(X2,X1))=mult(ld(f(rd(X1,X2)),X2),X1))).
###
#cnf(i_0_338151, plain, (rd(ld(f(mult(X2,X1)),X2),X1)=ld(f(mult(X2,X1)),rd(X2,X1)))).
###
#cnf(i_0_339579, plain, (rd(ld(X1,f(ld(X1,X2))),X2)=ld(X1,rd(f(ld(X1,X2)),X2)))).
#
#cnf(i_0_339623, plain, (mult(mult(X1,f(rd(X1,X2))),X2)=mult(X1,mult(f(rd(X1,X2)),X2)))).
###
#cnf(i_0_340253, plain, (ld(mult(f(mult(X2,X1)),X1),X2)=ld(X1,ld(f(mult(X2,X1)),X2)))).
##
#cnf(i_0_160944, plain, (mult(mult(X1,mult(f(X2),X1)),X2)=mult(X1,mult(f(X2),mult(X1,X2))))).
#######
#cnf(i_0_160281, plain, (rd(mult(X1,ld(f(X2),X1)),X2)=mult(X1,ld(f(X2),rd(X1,X2))))).
###
#cnf(i_0_51510, plain, (rd(ld(X1,X2),f(ld(X2,X1)))=ld(X1,rd(X2,f(ld(X2,X1)))))).
######
#cnf(i_0_348091, plain, (rd(mult(X2,X1),f(ld(X1,X2)))=mult(X2,rd(X1,f(ld(X1,X2)))))).
#
#cnf(i_0_348487, plain, (rd(rd(X2,X1),f(mult(X1,X2)))=rd(X2,mult(f(mult(X1,X2)),X1)))).
##
#cnf(i_0_190706, plain, (rd(mult(X1,ld(X2,f(X1))),X2)=mult(X1,ld(X2,rd(f(X1),X2))))).
###
#cnf(i_0_350599, plain, (ld(mult(X1,mult(X2,X1)),f(X2))=ld(mult(X2,X1),ld(X1,f(X2))))).
###
#cnf(i_0_351716, plain, (mult(mult(X1,X2),f(ld(X1,X2)))=mult(X1,mult(X2,f(ld(X1,X2)))))).
######
#cnf(i_0_357786, plain, (ld(mult(X2,f(X1)),ld(X2,X1))=ld(f(X1),ld(X2,ld(X2,X1))))).
###
#cnf(i_0_357834, plain, (ld(f(ld(X1,X2)),mult(X1,X2))=mult(ld(f(ld(X1,X2)),X1),X2))).
##
#cnf(i_0_358495, plain, (mult(mult(ld(X2,X1),X1),f(X2))=mult(ld(X2,X1),mult(X1,f(X2))))).
#####
#cnf(i_0_360680, plain, (rd(rd(X1,X2),mult(f(X1),X2))=rd(X1,mult(f(X1),mult(X2,X2))))).
########
#cnf(i_0_364358, plain, (ld(f(rd(X1,X2)),mult(X1,X2))=mult(ld(f(rd(X1,X2)),X1),X2))).
##
#cnf(i_0_370683, plain, (ld(mult(X2,f(ld(X1,X2))),X1)=ld(X2,rd(X1,f(ld(X1,X2)))))).
##
#cnf(i_0_371882, plain, (ld(mult(X1,f(mult(X1,X2))),X2)=ld(f(mult(X1,X2)),ld(X1,X2)))).
###
#cnf(i_0_113644, plain, (ld(X1,mult(X2,mult(X2,f(X1))))=mult(ld(X1,X2),mult(X2,f(X1))))).
#
#cnf(i_0_373322, plain, (ld(X2,mult(X1,f(mult(X1,X2))))=mult(ld(X2,X1),f(mult(X1,X2))))).
#
#cnf(i_0_373356, plain, (mult(mult(X2,X1),mult(X2,f(X1)))=mult(X2,mult(X1,mult(X2,f(X1)))))).
#
#cnf(i_0_377184, plain, (ld(mult(X2,X1),f(ld(X2,X1)))=ld(X1,ld(X2,f(ld(X2,X1)))))).
#
#cnf(i_0_377228, plain, (mult(mult(X1,X2),f(rd(X2,X1)))=mult(X1,mult(X2,f(rd(X2,X1)))))).
##
#cnf(i_0_161232, plain, (mult(mult(X1,mult(X2,X1)),f(X2))=mult(X1,mult(X2,mult(X1,f(X2)))))).
####
#cnf(i_0_96726, plain, (rd(ld(X1,X2),mult(X2,mult(X2,X2)))=ld(X1,ld(X2,ld(X2,unit))))).
#####
#cnf(i_0_97198, plain, (rd(X1,mult(X1,mult(X1,mult(X2,X1))))=ld(X2,ld(X1,ld(X1,unit))))).
#
#cnf(i_0_97200, plain, (rd(X1,mult(X1,mult(X2,mult(X1,X1))))=ld(X1,ld(X2,ld(X1,unit))))).
#####
#cnf(i_0_97201, plain, (rd(X2,mult(X1,mult(X2,mult(X1,X2))))=ld(X1,ld(X2,ld(X1,unit))))).
##
#cnf(i_0_97204, plain, (rd(X2,mult(X1,mult(X2,mult(X2,X2))))=ld(X2,ld(X2,ld(X1,unit))))).
##############
#cnf(i_0_98517, plain, (rd(f(ld(f(X1),unit)),f(ld(X1,f(X1))))=unit)).
##################
#cnf(i_0_50928, plain, (mult(ld(X2,f(X1)),ld(f(X1),rd(X2,X1)))=ld(X1,unit))).
####
#cnf(i_0_99867, plain, (ld(f(X1),ld(f(f(X1)),X2))=mult(f(f(X1)),ld(X1,X2)))).
####
#cnf(i_0_90546, plain, (mult(rd(X2,X1),mult(X1,ld(X2,f(f(X2)))))=f(f(X2)))).
#
#cnf(i_0_53536, plain, (mult(rd(f(X2),X1),rd(X1,mult(X2,f(X2))))=ld(X2,unit))).
##########
#cnf(i_0_102098, plain, (rd(ld(X2,ld(X1,X2)),mult(X1,X2))=ld(X2,ld(X1,ld(X1,unit))))).
#############
#cnf(i_0_51087, plain, (ld(mult(X1,X2),rd(X1,mult(X1,X2)))=ld(X2,ld(X2,ld(X1,unit))))).
#######
#cnf(i_0_105701, plain, (mult(mult(X1,X2),ld(X2,ld(X1,ld(X2,unit))))=ld(X2,unit))).
##################
#cnf(i_0_113596, plain, (rd(mult(ld(X1,X2),f(X1)),mult(X2,f(X1)))=ld(X1,unit))).
###################
#cnf(i_0_113951, plain, (rd(ld(X1,ld(X1,X2)),mult(X2,X2))=ld(X1,ld(X1,ld(X2,unit))))).
#####
#cnf(i_0_114119, plain, (rd(X1,mult(X1,mult(X2,mult(X2,X1))))=ld(X2,ld(X2,ld(X1,unit))))).
###
#cnf(i_0_114332, plain, (rd(X1,mult(X2,mult(X2,mult(X1,X1))))=ld(X1,ld(X2,ld(X2,unit))))).
#####################
#cnf(i_0_121068, plain, (rd(ld(X1,ld(X1,ld(X1,X2))),X2)=ld(X1,ld(X1,ld(X1,unit))))).
##
#cnf(i_0_122647, plain, (rd(X1,mult(X2,mult(X2,mult(X2,X1))))=ld(X2,ld(X2,ld(X2,unit))))).
#####
#cnf(i_0_123081, plain, (mult(ld(X1,ld(X3,ld(X2,unit))),mult(X2,X3))=ld(X1,unit))).
#####################
#cnf(i_0_124976, plain, (mult(mult(X1,X2),ld(X2,ld(X1,ld(X1,unit))))=ld(X1,unit))).
#####################
#cnf(i_0_12454, plain, (mult(mult(X1,X2),ld(X2,ld(X1,f(f(X2)))))=f(f(X2)))).
######
#cnf(i_0_135078, plain, (mult(rd(X2,f(X1)),rd(f(X1),mult(X1,X2)))=ld(X1,unit))).
#########
#cnf(i_0_137813, plain, (rd(ld(mult(X2,X1),ld(X1,X2)),X2)=ld(X1,ld(X2,ld(X1,unit))))).
##
#cnf(i_0_139232, plain, (mult(rd(X2,X1),mult(X1,ld(X2,f(f(X1)))))=f(f(X1)))).
#
#cnf(i_0_139238, plain, (ld(mult(X1,X2),mult(X1,mult(X2,f(f(X1)))))=f(f(X1)))).
##
#cnf(i_0_139769, plain, (mult(mult(X1,rd(f(f(X1)),X2)),rd(X2,X1))=f(f(X1)))).
#
#cnf(i_0_139775, plain, (rd(mult(X1,mult(f(f(X1)),X2)),mult(X1,X2))=f(f(X1)))).
#####################
#cnf(i_0_153759, plain, (rd(X1,mult(X1,f(f(f(X1)))))=ld(f(f(f(X1))),unit))).
##
#cnf(i_0_153813, plain, (ld(X1,ld(f(f(f(X1))),X1))=ld(f(f(f(X1))),unit))).
####################
#cnf(i_0_158792, plain, (rd(ld(X1,ld(X1,X2)),mult(X1,X2))=ld(X1,ld(X1,ld(X1,unit))))).
#
#cnf(i_0_158948, plain, (rd(ld(X1,X2),mult(X1,mult(X1,X2)))=ld(X1,ld(X1,ld(X1,unit))))).
#########################################
#cnf(i_0_185540, plain, (ld(mult(X1,mult(X2,X1)),rd(X1,X2))=ld(X1,ld(X2,ld(X2,unit))))).
#
#cnf(i_0_186891, plain, (rd(ld(X1,X2),mult(X2,mult(X1,X2)))=ld(X1,ld(X1,ld(X2,unit))))).
###
#cnf(i_0_187857, plain, (ld(mult(X1,mult(X1,X2)),rd(X1,X2))=ld(X2,ld(X1,ld(X2,unit))))).
#####
#cnf(i_0_189433, plain, (rd(ld(X1,X2),mult(X1,mult(X2,X2)))=ld(X1,ld(X2,ld(X1,unit))))).
#################
#cnf(i_0_229040, plain, (mult(mult(X1,X2),ld(X2,ld(X1,f(f(X1)))))=f(f(X1)))).
##############
#cnf(i_0_233069, plain, (mult(ld(X1,rd(f(f(X1)),X2)),mult(X2,X1))=f(f(X1)))).
##########################################################
#cnf(i_0_51106, plain, (ld(mult(X1,X2),rd(X1,f(ld(X2,unit))))=f(ld(X2,unit)))).
########
#cnf(i_0_64206, plain, (mult(X1,ld(X2,ld(X2,ld(X2,unit))))=rd(X1,mult(X2,mult(X2,X2))))).
#######
#cnf(i_0_76665, plain, (ld(mult(X1,f(f(X2))),mult(X1,X2))=mult(f(X2),f(f(X2))))).
####
#cnf(i_0_90420, plain, (mult(rd(X1,f(f(X2))),ld(X2,f(f(X2))))=rd(X1,X2))).
##
#cnf(i_0_90550, plain, (rd(mult(X1,ld(X2,f(f(X2)))),f(f(X2)))=rd(X1,X2))).
###########
#cnf(i_0_96446, plain, (rd(ld(X1,f(X2)),mult(X2,X2))=ld(X1,ld(X2,ld(f(X2),unit))))).
########
#cnf(i_0_97006, plain, (mult(ld(X1,ld(X1,ld(X2,unit))),X1)=ld(X1,ld(X1,ld(X2,X1))))).
######
#cnf(i_0_97203, plain, (rd(f(X2),mult(X1,mult(X2,X2)))=ld(X2,ld(f(X2),ld(X1,unit))))).
###
#cnf(i_0_412400, plain, (rd(f(X1),mult(X2,mult(X2,X1)))=ld(X1,rd(f(X1),mult(X2,X2))))).
#################
#cnf(i_0_103316, plain, (mult(mult(f(X1),mult(X2,X1)),ld(X1,ld(X2,f(X1))))=X1)).
#######
#cnf(i_0_50533, plain, (mult(X1,ld(X2,ld(X1,ld(X1,unit))))=rd(X1,mult(X1,mult(X1,X2))))).
############
#cnf(i_0_106229, plain, (mult(mult(f(X1),rd(X2,X1)),mult(X1,ld(X2,f(X1))))=X1)).
###############################
#cnf(i_0_110479, plain, (rd(ld(X1,X2),mult(f(ld(X1,unit)),X2))=f(ld(X1,unit)))).
######
#cnf(i_0_112018, plain, (mult(rd(f(X1),mult(X1,X2)),mult(X1,mult(X2,f(X1))))=X1)).
#####
#cnf(i_0_113946, plain, (ld(mult(f(X1),X2),ld(X1,unit))=ld(X2,ld(X1,ld(f(X1),unit))))).
#
#cnf(i_0_114025, plain, (rd(f(X1),mult(X1,mult(X2,X1)))=ld(f(X1),ld(X2,ld(X1,unit))))).
##########
#cnf(i_0_114427, plain, (ld(X1,rd(f(X1),mult(X2,X2)))=ld(f(X1),ld(X2,ld(X2,unit))))).
##################################
#cnf(i_0_114635, plain, (mult(ld(X1,ld(X2,ld(X2,unit))),X1)=ld(X1,ld(X2,ld(X2,X1))))).
#########
#cnf(i_0_117355, plain, (mult(ld(X1,ld(X2,ld(X3,unit))),X1)=ld(X1,ld(mult(X3,X2),X1)))).
###
#cnf(i_0_117562, plain, (mult(ld(X1,ld(X2,ld(X2,unit))),X3)=ld(mult(X2,mult(X2,X1)),X3))).
#####
#cnf(i_0_117565, plain, (mult(ld(X1,ld(X3,ld(X2,unit))),X4)=ld(mult(mult(X2,X3),X1),X4))).
######
#cnf(i_0_117675, plain, (rd(ld(f(ld(X1,unit)),X2),mult(X1,X2))=f(ld(X1,unit)))).
##
#cnf(i_0_117868, plain, (mult(X1,ld(X2,ld(X3,ld(X3,unit))))=rd(X1,mult(X3,mult(X3,X2))))).
###
#cnf(i_0_117871, plain, (mult(X1,ld(X2,ld(X4,ld(X3,unit))))=rd(X1,mult(mult(X3,X4),X2)))).
###
#cnf(i_0_118792, plain, (mult(ld(f(f(X1)),X2),ld(X2,X1))=mult(f(X1),f(f(X1))))).
#####
#cnf(i_0_120887, plain, (rd(mult(X1,X2),mult(f(f(X1)),X2))=mult(f(X1),f(f(X1))))).
####
#cnf(i_0_121680, plain, (mult(rd(X1,X2),rd(X2,f(f(X1))))=mult(f(X1),f(f(X1))))).
################################
#cnf(i_0_135948, plain, (rd(ld(X1,ld(X1,f(X1))),X2)=ld(X1,ld(f(X1),ld(X2,unit))))).
#############################################
#cnf(i_0_169110, plain, (rd(ld(X1,ld(X2,f(X2))),X3)=ld(mult(f(X2),X1),ld(X3,unit)))).
######
#cnf(i_0_170072, plain, (rd(ld(X1,rd(f(X1),X2)),X3)=ld(mult(X2,f(X1)),ld(X3,unit)))).
######################
#cnf(i_0_176220, plain, (mult(f(f(X1)),ld(X1,ld(f(f(X1)),X2)))=ld(X1,X2))).
###
#cnf(i_0_176462, plain, (mult(ld(X1,rd(X2,f(f(X1)))),f(f(X1)))=ld(X1,X2))).
########
#cnf(i_0_185500, plain, (rd(ld(X2,f(X1)),mult(X2,X1))=ld(X2,ld(f(X1),ld(X2,unit))))).
########################
#cnf(i_0_214555, plain, (rd(mult(X1,ld(f(f(X1)),X2)),X2)=mult(f(X1),f(f(X1))))).
##########
#cnf(i_0_228221, plain, (ld(mult(X2,X1),ld(f(X2),unit))=ld(X1,ld(X2,ld(f(X2),unit))))).
###########
#cnf(i_0_231405, plain, (ld(mult(X2,f(X1)),ld(X1,unit))=ld(f(X1),ld(X2,ld(X1,unit))))).
#
#cnf(i_0_231662, plain, (ld(mult(X2,X1),ld(X1,f(X1)))=ld(X1,ld(X2,ld(f(X1),unit))))).
#################################################################
#cnf(i_0_430585, plain, (ld(mult(X1,X2),ld(X1,f(X1)))=ld(X2,ld(X1,ld(f(X1),unit))))).
#######
#cnf(i_0_50688, plain, (mult(X1,mult(f(unit),f(f(unit))))=rd(X1,f(f(unit))))).
#####
#cnf(i_0_52308, plain, (mult(mult(X1,f(ld(X2,X3))),ld(X3,X2))=rd(X1,f(ld(X2,X3))))).
#
#cnf(i_0_52370, plain, (mult(mult(X1,ld(X3,X2)),f(ld(X2,X3)))=rd(X1,f(ld(X2,X3))))).
######
#cnf(i_0_53053, plain, (mult(ld(f(ld(X1,X2)),X2),ld(X2,X1))=ld(f(ld(X1,X2)),X1))).
#######
#cnf(i_0_53427, plain, (mult(ld(X1,rd(X2,X1)),mult(X1,f(X1)))=mult(ld(X1,X2),f(X1)))).
#
#cnf(i_0_53451, plain, (mult(X1,mult(ld(X1,X2),f(ld(X1,X2))))=mult(X2,f(ld(X1,X2))))).
####
#cnf(i_0_53477, plain, (mult(ld(X1,rd(X2,mult(X1,f(X1)))),X1)=ld(X1,rd(X2,f(X1))))).
########################################################################
#cnf(i_0_97592, plain, (mult(ld(X1,X2),ld(X2,f(ld(X1,X2))))=ld(X1,f(ld(X1,X2))))).
#########
#cnf(i_0_100225, plain, (mult(mult(X1,f(rd(X2,X3))),rd(X3,X2))=rd(X1,f(rd(X2,X3))))).
#
#cnf(i_0_100281, plain, (mult(mult(X1,rd(X3,X2)),f(rd(X2,X3)))=rd(X1,f(rd(X2,X3))))).
##########
#cnf(i_0_100879, plain, (rd(mult(X1,mult(X1,mult(X2,f(X2)))),X2)=mult(X1,mult(X1,f(X2))))).
###
#cnf(i_0_101041, plain, (rd(mult(X1,mult(X1,X2)),mult(X2,f(X2)))=mult(X1,rd(X1,f(X2))))).
######
#cnf(i_0_101230, plain, (mult(mult(X1,rd(X1,mult(X2,f(X2)))),X2)=mult(X1,rd(X1,f(X2))))).
#########
#cnf(i_0_101268, plain, (mult(ld(f(X1),mult(X2,rd(X2,X1))),X1)=mult(ld(f(X1),X2),X2))).
######
#cnf(i_0_101304, plain, (mult(mult(X1,rd(X1,X2)),mult(X2,f(X2)))=mult(X1,mult(X1,f(X2))))).
###########
#cnf(i_0_104537, plain, (rd(mult(ld(X1,X2),X1),f(mult(X2,X1)))=ld(X1,f(mult(X2,X1))))).
################
#cnf(i_0_105291, plain, (mult(ld(X2,X1),ld(X1,f(ld(X1,X2))))=ld(X2,f(ld(X1,X2))))).
#####
#cnf(i_0_105439, plain, (ld(f(rd(X1,X2)),rd(X1,mult(X2,X2)))=rd(f(rd(X1,X2)),X2))).
#######
#cnf(i_0_105477, plain, (ld(f(mult(X1,X2)),mult(X1,rd(X2,X1)))=rd(f(mult(X1,X2)),X1))).
#####
#cnf(i_0_105492, plain, (ld(f(ld(X1,X2)),ld(X1,rd(X2,X1)))=rd(f(ld(X1,X2)),X1))).
####
#cnf(i_0_105682, plain, (mult(mult(X1,rd(f(ld(X1,X2)),X2)),X2)=rd(X2,f(ld(X1,X2))))).
#
#cnf(i_0_105686, plain, (mult(rd(X1,mult(X2,f(ld(X2,X1)))),X2)=mult(X2,f(ld(X2,X1))))).
##
#cnf(i_0_106073, plain, (mult(f(rd(X1,X2)),mult(rd(X2,X1),X3))=ld(f(rd(X1,X2)),X3))).
###
#cnf(i_0_106113, plain, (ld(f(mult(X1,X2)),mult(X1,mult(X2,X1)))=mult(f(mult(X1,X2)),X1))).
##
#cnf(i_0_106114, plain, (ld(f(mult(X1,X2)),mult(X1,mult(X2,X2)))=mult(f(mult(X1,X2)),X2))).
###
#cnf(i_0_106125, plain, (mult(f(mult(X1,X1)),ld(X1,ld(X1,X2)))=ld(f(mult(X1,X1)),X2))).
#####
#cnf(i_0_106126, plain, (mult(f(mult(X1,X2)),ld(X2,ld(X1,X2)))=ld(f(mult(X1,X2)),X2))).
#####
#cnf(i_0_106127, plain, (mult(f(ld(X1,X2)),mult(ld(X2,X1),X3))=ld(f(ld(X1,X2)),X3))).
#
#cnf(i_0_106181, plain, (mult(ld(mult(f(mult(X1,X1)),X2),X1),X1)=ld(X2,f(mult(X1,X1))))).
###
#cnf(i_0_455890, plain, (rd(ld(X2,f(mult(X1,X1))),X1)=ld(mult(f(mult(X1,X1)),X2),X1))).
####
#cnf(i_0_106386, plain, (mult(f(rd(X1,X2)),mult(X2,ld(X1,X2)))=ld(f(rd(X1,X2)),X2))).
######
#cnf(i_0_106395, plain, (ld(f(rd(X1,X2)),mult(X1,ld(X2,X1)))=mult(f(rd(X1,X2)),X1))).
#########
#cnf(i_0_107632, plain, (ld(f(rd(X1,X2)),rd(X1,mult(X1,X2)))=rd(f(rd(X1,X2)),X1))).
#########
#cnf(i_0_108521, plain, (mult(mult(f(X1),X2),mult(ld(X2,X1),X2))=mult(X1,mult(f(X1),X2)))).
#
#cnf(i_0_108733, plain, (mult(rd(X1,X2),mult(f(rd(X2,X1)),X3))=ld(f(rd(X2,X1)),X3))).
#
#cnf(i_0_108734, plain, (mult(X1,mult(X1,ld(f(mult(X1,X1)),X2)))=mult(f(mult(X1,X1)),X2))).
#
#cnf(i_0_108737, plain, (mult(ld(X1,X2),mult(f(ld(X2,X1)),X3))=ld(f(ld(X2,X1)),X3))).
#####
#cnf(i_0_109003, plain, (mult(rd(f(X1),X2),mult(X2,rd(X1,X2)))=mult(X1,rd(f(X1),X2)))).
#####
#cnf(i_0_109285, plain, (mult(ld(X1,X2),rd(f(ld(X1,X2)),X2))=rd(f(ld(X1,X2)),X1))).
####
#cnf(i_0_111567, plain, (mult(mult(X1,X2),ld(X2,f(mult(X1,X2))))=mult(X1,f(mult(X1,X2))))).
######
#cnf(i_0_112159, plain, (mult(rd(X1,X2),mult(X2,f(rd(X1,X2))))=mult(X1,f(rd(X1,X2))))).
#
#cnf(i_0_112258, plain, (mult(X1,mult(ld(X1,X2),mult(X1,f(X1))))=mult(X2,mult(X1,f(X1))))).
##
#cnf(i_0_112259, plain, (mult(X1,mult(ld(X1,X2),mult(X2,f(X1))))=mult(X2,mult(X2,f(X1))))).
##
#cnf(i_0_112402, plain, (mult(rd(X1,X2),rd(X2,f(rd(X1,X2))))=rd(X1,f(rd(X1,X2))))).
###
#cnf(i_0_112706, plain, (mult(rd(X1,X2),mult(f(rd(X1,X2)),X2))=mult(f(rd(X1,X2)),X1))).
#####
#cnf(i_0_112986, plain, (mult(mult(X1,X2),rd(f(mult(X1,X2)),X2))=mult(f(mult(X1,X2)),X1))).
####
#cnf(i_0_113329, plain, (rd(mult(X1,X2),mult(f(mult(X1,X2)),X2))=rd(X1,f(mult(X1,X2))))).
####
#cnf(i_0_113738, plain, (mult(X1,mult(mult(ld(X1,X2),f(X2)),X1))=mult(X2,mult(f(X2),X1)))).
###
#cnf(i_0_114074, plain, (rd(X1,mult(X1,mult(f(X1),ld(X2,X1))))=rd(X2,mult(X1,f(X1))))).
#######
#cnf(i_0_119289, plain, (mult(mult(X1,rd(f(X1),mult(X2,X2))),X2)=mult(X1,rd(f(X1),X2)))).
#
#cnf(i_0_119304, plain, (mult(ld(X1,mult(X2,rd(f(X2),X1))),X1)=mult(ld(X1,X2),f(X2)))).
#
#cnf(i_0_119330, plain, (mult(mult(X1,rd(f(X1),X2)),mult(X2,X2))=mult(X1,mult(f(X1),X2)))).
##
#cnf(i_0_119449, plain, (rd(ld(f(rd(X1,X2)),X1),mult(X2,X2))=rd(f(rd(X1,X2)),X2))).
#####
#cnf(i_0_119660, plain, (mult(ld(X1,ld(f(mult(X2,X1)),X2)),X1)=ld(X1,f(mult(X2,X1))))).
###
#cnf(i_0_119693, plain, (mult(ld(f(mult(X1,X2)),X1),mult(X2,X2))=mult(f(mult(X1,X2)),X2))).
####
#cnf(i_0_119926, plain, (mult(ld(X2,rd(f(mult(X2,X1)),X1)),X1)=rd(X1,f(mult(X2,X1))))).
#####
#cnf(i_0_119931, plain, (mult(mult(rd(X1,f(mult(X2,X1))),X2),X1)=mult(X1,f(mult(X2,X1))))).
#
#cnf(i_0_119944, plain, (ld(X1,mult(rd(X1,f(mult(X2,X1))),X2))=rd(f(mult(X2,X1)),X1))).
###
#cnf(i_0_119953, plain, (mult(X1,mult(rd(X2,f(mult(X1,X2))),X1))=mult(f(mult(X1,X2)),X1))).
####
#cnf(i_0_119957, plain, (mult(X1,rd(X2,mult(X1,f(mult(X1,X2)))))=rd(f(mult(X1,X2)),X1))).
###
#cnf(i_0_120183, plain, (rd(ld(mult(f(rd(X1,X2)),X2),X1),X2)=ld(X2,f(rd(X1,X2))))).
####
#cnf(i_0_120224, plain, (mult(X1,ld(mult(f(rd(X1,X2)),X2),X1))=mult(f(rd(X1,X2)),X1))).
###
#cnf(i_0_120233, plain, (rd(X1,mult(X1,mult(f(rd(X1,X2)),X2)))=rd(f(rd(X1,X2)),X1))).
###############
#cnf(i_0_122357, plain, (rd(mult(X1,mult(f(X1),mult(X2,X2))),X2)=mult(X1,mult(f(X1),X2)))).
#
#cnf(i_0_122398, plain, (rd(mult(X1,mult(f(X1),X2)),mult(X2,X2))=mult(X1,rd(f(X1),X2)))).
#####################################
#cnf(i_0_147703, plain, (mult(rd(X1,X2),mult(X2,rd(f(X1),X2)))=mult(X1,rd(f(X1),X2)))).
##
#cnf(i_0_147709, plain, (mult(mult(X1,X2),mult(ld(X2,f(X1)),X2))=mult(X1,mult(f(X1),X2)))).
#
#cnf(i_0_148325, plain, (mult(X1,mult(ld(X1,X2),ld(X1,f(X2))))=mult(X2,ld(X1,f(X2))))).
#
#cnf(i_0_149214, plain, (mult(mult(ld(X1,f(X2)),X1),ld(X1,X2))=mult(ld(X1,X2),f(X2)))).
####
#cnf(i_0_149623, plain, (mult(ld(f(X1),X2),ld(X2,ld(X1,X2)))=ld(X1,ld(f(X1),X2)))).
######
#cnf(i_0_151519, plain, (rd(mult(ld(X1,f(X2)),ld(X1,X2)),X2)=ld(X1,rd(f(X2),X1)))).
#########
#cnf(i_0_152994, plain, (rd(mult(X1,X2),mult(X1,mult(X2,f(X2))))=rd(X1,mult(X1,f(X2))))).
##########################
#cnf(i_0_154016, plain, (mult(ld(X1,f(f(f(f(X1))))),X1)=f(f(f(f(X1)))))).
#
#cnf(i_0_490462, plain, (rd(X1,f(f(f(f(X1)))))=ld(f(f(f(f(X1)))),X1))).
##
#cnf(i_0_490524, plain, (rd(f(f(f(f(X1)))),X1)=ld(X1,f(f(f(f(X1))))))).
##
#cnf(i_0_490753, plain, (mult(f(f(f(f(X1)))),X1)=mult(X1,f(f(f(f(X1))))))).
########
#cnf(i_0_155373, plain, (mult(rd(X1,X2),mult(X2,rd(f(X2),X1)))=mult(X1,rd(f(X2),X1)))).
############
#cnf(i_0_160310, plain, (rd(mult(X1,mult(f(X2),mult(X1,X2))),X2)=mult(X1,mult(f(X2),X1)))).
####
#cnf(i_0_160931, plain, (mult(mult(X1,ld(f(X2),rd(X1,X2))),X2)=mult(X1,ld(f(X2),X1)))).
########
#cnf(i_0_162701, plain, (ld(mult(X1,f(mult(X2,X1))),mult(X1,X2))=rd(f(mult(X2,X1)),X1))).
############
#cnf(i_0_163224, plain, (mult(mult(ld(X1,X2),X1),ld(X1,f(X2)))=mult(ld(X1,X2),f(X2)))).
##
#cnf(i_0_163350, plain, (mult(mult(ld(X2,X1),f(X1)),ld(X1,X2))=mult(ld(X2,f(X1)),X2))).
###
#cnf(i_0_164216, plain, (mult(rd(X2,X1),ld(f(rd(X1,X2)),X1))=ld(f(rd(X1,X2)),X2))).
#############
#cnf(i_0_164913, plain, (ld(mult(X1,X2),mult(X1,mult(f(X1),X2)))=mult(ld(X2,f(X1)),X2))).
###
#cnf(i_0_165063, plain, (mult(rd(X2,mult(X1,f(X1))),rd(X1,X2))=rd(X2,mult(X2,f(X1))))).
############################
#cnf(i_0_165069, plain, (ld(mult(X1,mult(f(X1),X2)),mult(X1,X2))=ld(X2,ld(f(X1),X2)))).
#######
#cnf(i_0_165076, plain, (mult(ld(X2,X1),ld(X1,ld(f(X1),X2)))=ld(X2,ld(f(X1),X2)))).
##########
#cnf(i_0_165765, plain, (mult(X1,mult(ld(X1,X2),ld(f(X1),X2)))=mult(X2,ld(f(X1),X2)))).
##
#cnf(i_0_165818, plain, (mult(rd(X1,mult(X1,f(X2))),rd(X1,X2))=rd(X1,mult(X2,f(X2))))).
##
#cnf(i_0_165834, plain, (mult(rd(X1,mult(X1,X2)),rd(X1,f(X2)))=rd(X1,mult(X2,f(X2))))).
######################
#cnf(i_0_170946, plain, (mult(mult(X1,rd(f(X2),X1)),mult(X1,X2))=mult(X1,mult(X2,f(X2))))).
####
#cnf(i_0_170947, plain, (mult(mult(X1,rd(X2,X1)),mult(X1,f(X2)))=mult(X1,mult(X2,f(X2))))).
########
#cnf(i_0_171078, plain, (rd(mult(X1,mult(X2,f(X2))),mult(X1,X2))=mult(X1,rd(f(X2),X1)))).
#############
#cnf(i_0_172850, plain, (mult(ld(X1,X2),ld(X2,ld(f(X1),X2)))=ld(X1,ld(f(X1),X2)))).
################
#cnf(i_0_180735, plain, (mult(ld(X1,ld(f(mult(X1,X2)),X1)),X2)=ld(X1,f(mult(X1,X2))))).
###
#cnf(i_0_180774, plain, (rd(mult(ld(X1,f(X2)),X1),mult(X2,X1))=ld(X1,ld(X2,f(X2))))).
########
#cnf(i_0_180793, plain, (rd(ld(X1,mult(f(rd(X1,X2)),X2)),X2)=ld(X1,f(rd(X1,X2))))).
##
#cnf(i_0_509428, plain, (ld(X1,mult(f(rd(X1,X2)),X2))=mult(ld(X1,f(rd(X1,X2))),X2))).
#
#cnf(i_0_510537, plain, (rd(ld(X1,f(mult(X1,X2))),X2)=ld(X1,rd(f(mult(X1,X2)),X2)))).
###
#cnf(i_0_511355, plain, (rd(mult(X1,f(ld(X1,X2))),X2)=mult(X1,rd(f(ld(X1,X2)),X2)))).
#####
#cnf(i_0_181205, plain, (mult(X2,mult(ld(X1,f(rd(X1,X2))),X1))=ld(f(rd(X1,X2)),X1))).
#
#cnf(i_0_181206, plain, (ld(X2,mult(ld(X1,f(mult(X1,X2))),X1))=ld(f(mult(X1,X2)),X1))).
#######
#cnf(i_0_181961, plain, (mult(mult(X1,X2),ld(X2,f(mult(X2,X1))))=mult(X1,f(mult(X2,X1))))).
##
#cnf(i_0_181964, plain, (mult(rd(X1,X2),mult(X2,f(ld(X2,X1))))=mult(X1,f(ld(X2,X1))))).
###
#cnf(i_0_181999, plain, (mult(mult(X1,X2),rd(X1,f(mult(X2,X1))))=mult(X1,f(mult(X2,X1))))).
##
#cnf(i_0_182425, plain, (rd(mult(X1,ld(f(mult(X2,X1)),X2)),X2)=rd(X1,f(mult(X2,X1))))).
######
#cnf(i_0_517467, plain, (mult(rd(X1,f(mult(X2,X1))),X2)=mult(X1,ld(f(mult(X2,X1)),X2)))).
###
#cnf(i_0_518358, plain, (rd(rd(X1,f(ld(X2,X1))),X2)=rd(X1,mult(X2,f(ld(X2,X1)))))).
###
#cnf(i_0_519225, plain, (ld(mult(f(rd(X2,X1)),X1),X2)=ld(X1,ld(f(rd(X2,X1)),X2)))).
#########
#cnf(i_0_183148, plain, (mult(rd(f(ld(X1,X2)),X2),mult(X1,X1))=ld(f(ld(X1,X2)),X1))).
#############
#cnf(i_0_184929, plain, (rd(mult(X1,mult(X2,mult(f(X1),X2))),X2)=mult(X1,mult(X2,f(X1))))).
#
#cnf(i_0_184996, plain, (rd(mult(X1,mult(f(rd(X1,X2)),X2)),X2)=mult(X1,f(rd(X1,X2))))).
#
#cnf(i_0_185038, plain, (mult(rd(X1,X2),mult(X2,rd(X2,f(X1))))=mult(X1,rd(X2,f(X1))))).
#
#cnf(i_0_185559, plain, (mult(rd(X2,X1),mult(X1,rd(X2,f(X1))))=mult(X2,rd(X2,f(X1))))).
###
#cnf(i_0_185984, plain, (mult(rd(f(X1),mult(X1,X2)),rd(X1,X2))=rd(f(X1),mult(X2,X2)))).
#
#cnf(i_0_185989, plain, (mult(mult(f(X1),rd(X2,X1)),mult(X1,X2))=mult(f(X1),mult(X2,X2)))).
#########
#cnf(i_0_186504, plain, (mult(rd(f(X1),mult(X2,X2)),rd(X2,X1))=rd(f(X1),mult(X1,X2)))).
#
#cnf(i_0_186509, plain, (rd(mult(f(X1),mult(X2,X2)),mult(X1,X2))=mult(f(X1),rd(X2,X1)))).
###
#cnf(i_0_186756, plain, (rd(mult(ld(f(X1),X2),ld(X1,X2)),X2)=ld(f(X1),rd(X2,X1)))).
#####
#cnf(i_0_187033, plain, (mult(mult(ld(f(X1),X2),X1),ld(X1,X2))=mult(ld(f(X1),X2),X2))).
##########
#cnf(i_0_187684, plain, (rd(mult(ld(X1,X2),ld(f(X1),X2)),X2)=ld(X1,rd(X2,f(X1))))).
#
#cnf(i_0_187982, plain, (mult(mult(X1,ld(X2,rd(f(X1),X2))),X2)=mult(X1,ld(X2,f(X1))))).
####
#cnf(i_0_188167, plain, (mult(mult(X1,rd(f(mult(X1,X2)),X2)),X2)=mult(X1,f(mult(X1,X2))))).
#
#cnf(i_0_188669, plain, (mult(X2,mult(ld(X2,f(X1)),ld(X2,X1)))=mult(f(X1),ld(X2,X1)))).
##
#cnf(i_0_189128, plain, (mult(ld(X2,rd(f(ld(X2,X1)),X1)),X1)=ld(X2,f(ld(X2,X1))))).
##
#cnf(i_0_189328, plain, (mult(ld(X2,X1),ld(X1,ld(X1,f(X2))))=ld(X2,ld(X1,f(X2))))).
###
#cnf(i_0_189678, plain, (mult(mult(f(X1),ld(X2,rd(X1,X2))),X2)=mult(f(X1),ld(X2,X1)))).
##
#cnf(i_0_533966, plain, (mult(rd(X1,X2),rd(X1,f(X2)))=mult(X1,ld(X2,rd(X1,f(X2)))))).
#
#cnf(i_0_160551, plain, (rd(mult(X1,ld(X2,X1)),f(X2))=mult(X1,ld(X2,rd(X1,f(X2)))))).
####
#cnf(i_0_189683, plain, (rd(mult(f(X1),mult(X2,mult(X1,X2))),X2)=mult(f(X1),mult(X2,X1)))).
#
#cnf(i_0_190298, plain, (mult(mult(ld(f(X2),X1),X1),ld(X1,X2))=mult(ld(f(X2),X1),X2))).
####
#cnf(i_0_190986, plain, (mult(ld(X1,X2),ld(X2,ld(X1,f(X2))))=ld(X1,ld(X1,f(X2))))).
#############
#cnf(i_0_192857, plain, (mult(mult(X1,X2),mult(ld(X2,X1),f(X2)))=mult(X1,mult(X1,f(X2))))).
######
#cnf(i_0_200578, plain, (mult(X1,rd(X2,mult(X2,f(mult(X1,X2)))))=rd(f(mult(X1,X2)),X2))).
#
#cnf(i_0_201043, plain, (mult(X1,ld(X2,ld(f(rd(X1,X2)),X2)))=mult(f(rd(X1,X2)),X2))).
#################
#cnf(i_0_207348, plain, (mult(mult(f(rd(X1,X2)),X1),ld(X1,X2))=ld(f(rd(X1,X2)),X1))).
#
#cnf(i_0_207349, plain, (rd(mult(f(mult(X1,X2)),X1),mult(X2,X1))=ld(f(mult(X1,X2)),X1))).
######
#cnf(i_0_207495, plain, (mult(ld(f(rd(X1,X2)),X1),ld(X1,X2))=ld(f(rd(X1,X2)),X2))).
#####
#cnf(i_0_208198, plain, (mult(rd(X1,X2),rd(X2,f(ld(X1,X2))))=rd(X1,f(ld(X1,X2))))).
###
#cnf(i_0_209854, plain, (rd(ld(X2,ld(f(rd(X1,X2)),X1)),X1)=ld(X1,f(rd(X1,X2))))).
###
#cnf(i_0_209870, plain, (mult(X1,mult(ld(X1,f(rd(X1,X2))),X2))=ld(f(rd(X1,X2)),X1))).
###
#cnf(i_0_209899, plain, (mult(X1,mult(ld(X2,f(rd(X2,X1))),X1))=ld(f(rd(X2,X1)),X1))).
######
#cnf(i_0_210014, plain, (rd(mult(X1,ld(f(rd(X2,X1)),X2)),X2)=rd(X1,f(rd(X2,X1))))).
########
#cnf(i_0_210612, plain, (ld(mult(X1,X2),mult(X1,f(mult(X2,X1))))=rd(X1,f(mult(X2,X1))))).
####
#cnf(i_0_211098, plain, (mult(X1,mult(X2,ld(f(mult(X1,X2)),X2)))=mult(f(mult(X1,X2)),X2))).
####
#cnf(i_0_212311, plain, (rd(X1,mult(X2,mult(f(rd(X1,X2)),X2)))=rd(f(rd(X1,X2)),X2))).
###
#cnf(i_0_212735, plain, (mult(ld(X1,rd(f(X2),X1)),rd(X1,X2))=ld(X1,ld(X2,f(X2))))).
#######
#cnf(i_0_213430, plain, (mult(mult(X1,ld(f(mult(X1,X2)),X1)),X2)=mult(X1,f(mult(X1,X2))))).
######
#cnf(i_0_213780, plain, (rd(mult(X1,mult(f(rd(X2,X1)),X1)),X2)=rd(X1,f(rd(X2,X1))))).
############
#cnf(i_0_216164, plain, (mult(ld(X1,X2),ld(X2,f(rd(X2,X1))))=ld(X1,f(rd(X2,X1))))).
#
#cnf(i_0_216266, plain, (rd(ld(X2,X1),mult(f(rd(X1,X2)),X1))=ld(X1,f(rd(X1,X2))))).
#
#cnf(i_0_216278, plain, (rd(mult(X2,X1),mult(f(mult(X1,X2)),X1))=ld(X1,f(mult(X1,X2))))).
#
#cnf(i_0_216308, plain, (rd(ld(X1,X2),mult(f(rd(X2,X1)),X1))=ld(X1,f(rd(X2,X1))))).
#######
#cnf(i_0_216707, plain, (mult(rd(X2,X1),rd(X1,f(ld(X1,X2))))=mult(X1,f(ld(X1,X2))))).
#
#cnf(i_0_216708, plain, (mult(rd(X1,X2),rd(X1,f(ld(X2,X1))))=mult(X1,f(ld(X2,X1))))).
####
#cnf(i_0_217022, plain, (ld(mult(X1,f(ld(X1,X2))),rd(X2,X1))=rd(f(ld(X1,X2)),X1))).
#
#cnf(i_0_217023, plain, (ld(mult(X1,f(ld(X2,X1))),rd(X1,X2))=rd(f(ld(X2,X1)),X1))).
#####
#cnf(i_0_217167, plain, (mult(rd(f(ld(X1,X2)),X2),rd(X2,X1))=rd(f(ld(X1,X2)),X1))).
######
#cnf(i_0_217591, plain, (ld(X1,rd(X2,mult(X2,f(ld(X1,X2)))))=rd(f(ld(X1,X2)),X2))).
##
#cnf(i_0_219509, plain, (rd(mult(X1,rd(f(mult(X2,X1)),X1)),X2)=rd(X1,f(mult(X2,X1))))).
#######
#cnf(i_0_220279, plain, (mult(ld(f(rd(X1,X2)),X1),ld(X2,X1))=mult(f(rd(X1,X2)),X1))).
#
#cnf(i_0_220291, plain, (mult(ld(f(mult(X1,X2)),X1),mult(X2,X1))=mult(f(mult(X1,X2)),X1))).
##
#cnf(i_0_220321, plain, (mult(ld(f(rd(X1,X2)),X2),ld(X2,X1))=mult(f(rd(X1,X2)),X2))).
###
#cnf(i_0_220376, plain, (mult(mult(f(rd(X1,X2)),X2),ld(X2,X1))=mult(f(rd(X1,X2)),X1))).
###
#cnf(i_0_220377, plain, (mult(rd(f(mult(X1,X2)),X2),mult(X2,X1))=mult(f(mult(X1,X2)),X1))).
####
#cnf(i_0_220676, plain, (rd(ld(X1,ld(f(rd(X1,X2)),X1)),X2)=ld(X1,f(rd(X1,X2))))).
####
#cnf(i_0_221812, plain, (mult(rd(X1,mult(X1,f(ld(X1,X2)))),X2)=mult(X1,f(ld(X1,X2))))).
########
#cnf(i_0_222534, plain, (mult(rd(X1,X2),mult(X1,f(ld(X1,X2))))=rd(X1,f(ld(X1,X2))))).
######
#cnf(i_0_222856, plain, (rd(mult(X1,ld(f(rd(X1,X2)),X1)),X2)=mult(X1,f(rd(X1,X2))))).
###########
#cnf(i_0_223855, plain, (ld(X1,mult(X2,ld(f(ld(X1,X2)),X2)))=mult(f(ld(X1,X2)),X2))).
#
#cnf(i_0_224393, plain, (ld(X2,mult(X1,mult(f(ld(X1,X2)),X1)))=ld(f(ld(X1,X2)),X1))).
####
#cnf(i_0_224829, plain, (ld(mult(X1,mult(f(ld(X1,X2)),X1)),X2)=ld(X1,f(ld(X1,X2))))).
##############################################################################
#cnf(i_0_287117, plain, (rd(ld(f(rd(X1,X2)),X1),mult(X1,X2))=rd(f(rd(X1,X2)),X1))).
#
#cnf(i_0_287184, plain, (mult(ld(f(mult(X1,X2)),X2),ld(X2,X1))=rd(f(mult(X1,X2)),X2))).
#
#cnf(i_0_287185, plain, (mult(ld(f(mult(X1,X2)),X1),rd(X2,X1))=rd(f(mult(X1,X2)),X1))).
#
#cnf(i_0_287216, plain, (rd(ld(f(ld(X1,X2)),X2),mult(X1,X2))=rd(f(ld(X1,X2)),X2))).
##
#cnf(i_0_287305, plain, (mult(rd(f(X1),X2),mult(X2,rd(X2,X1)))=mult(f(X1),rd(X2,X1)))).
#
#cnf(i_0_287319, plain, (rd(mult(f(rd(X1,X2)),X2),mult(X1,X2))=rd(f(rd(X1,X2)),X1))).
#
#cnf(i_0_287320, plain, (mult(rd(f(mult(X1,X2)),X2),rd(X2,X1))=rd(f(mult(X1,X2)),X1))).
#
#cnf(i_0_288330, plain, (mult(mult(ld(X2,X1),X1),ld(X1,f(X2)))=mult(ld(X2,X1),f(X2)))).
#
#cnf(i_0_288337, plain, (mult(ld(X1,X2),ld(X2,f(mult(X2,X1))))=ld(X1,f(mult(X2,X1))))).
##
#cnf(i_0_288340, plain, (ld(mult(X2,X1),mult(X2,f(ld(X2,X1))))=ld(X1,f(ld(X2,X1))))).
##
#cnf(i_0_288377, plain, (ld(mult(X1,X2),rd(X1,f(rd(X1,X2))))=ld(X1,f(rd(X1,X2))))).
#
#cnf(i_0_288444, plain, (mult(ld(X1,X2),rd(X1,f(mult(X2,X1))))=ld(X1,f(mult(X2,X1))))).
#
#cnf(i_0_288445, plain, (mult(rd(X2,X1),rd(X1,f(mult(X1,X2))))=ld(X1,f(mult(X1,X2))))).
#
#cnf(i_0_288476, plain, (ld(mult(X2,X1),rd(X1,f(ld(X2,X1))))=ld(X1,f(ld(X2,X1))))).
#
#cnf(i_0_288860, plain, (mult(rd(f(mult(X1,X2)),X1),rd(X1,X2))=ld(f(mult(X1,X2)),X1))).
####
#cnf(i_0_288862, plain, (mult(rd(f(mult(X1,X2)),X2),ld(X1,X2))=ld(f(mult(X1,X2)),X2))).
#####
#cnf(i_0_288883, plain, (ld(mult(X2,f(X1)),mult(X2,mult(X2,X1)))=mult(ld(f(X1),X2),X1))).
#
#cnf(i_0_288901, plain, (mult(ld(f(X1),X2),ld(X2,ld(X2,X1)))=ld(f(X1),ld(X2,X1)))).
#######
#cnf(i_0_288957, plain, (mult(ld(f(mult(X1,X2)),X1),ld(X1,X2))=ld(f(mult(X1,X2)),X2))).
######
#cnf(i_0_289215, plain, (rd(mult(X1,mult(X2,X2)),mult(f(X1),X2))=mult(X1,rd(X2,f(X1))))).
####
#cnf(i_0_289218, plain, (mult(rd(X1,X2),rd(X2,f(mult(X1,X2))))=rd(X1,f(mult(X1,X2))))).
######
#cnf(i_0_289324, plain, (mult(rd(X1,X2),ld(X1,f(mult(X1,X2))))=rd(X1,f(mult(X1,X2))))).
####
#cnf(i_0_289326, plain, (mult(ld(X2,X1),ld(X1,f(mult(X2,X1))))=rd(X1,f(mult(X2,X1))))).
#####
#cnf(i_0_289333, plain, (mult(rd(X1,mult(X2,X2)),rd(X2,f(X1)))=rd(X1,mult(f(X1),X2)))).
###
#cnf(i_0_289679, plain, (ld(mult(X1,f(rd(X2,X1))),mult(X1,X2))=mult(f(rd(X2,X1)),X1))).
##
#cnf(i_0_289978, plain, (ld(mult(X1,f(ld(X1,X2))),mult(X2,X1))=mult(f(ld(X1,X2)),X1))).
####
#cnf(i_0_290084, plain, (mult(rd(f(ld(X1,X2)),X2),mult(X2,X1))=mult(f(ld(X1,X2)),X1))).
###
#cnf(i_0_290902, plain, (mult(mult(X1,X2),ld(X2,f(rd(X2,X1))))=mult(X1,f(rd(X2,X1))))).
##
#cnf(i_0_291012, plain, (rd(mult(X1,X2),mult(f(rd(X2,X1)),X1))=mult(X1,f(rd(X2,X1))))).
##
#cnf(i_0_291311, plain, (rd(mult(X2,X1),mult(f(ld(X1,X2)),X1))=mult(X1,f(ld(X1,X2))))).
######################
#cnf(i_0_306870, plain, (mult(rd(X1,X2),mult(X2,f(rd(X2,X1))))=mult(X1,f(rd(X2,X1))))).
#########
#cnf(i_0_314941, plain, (mult(X1,mult(ld(X1,ld(X1,X2)),f(X2)))=mult(ld(X1,X2),f(X2)))).
###################
#cnf(i_0_319981, plain, (mult(X1,mult(ld(X1,ld(f(X1),X2)),X2))=mult(ld(f(X1),X2),X2))).
#########
#cnf(i_0_321039, plain, (mult(mult(X1,ld(X2,f(rd(X1,X2)))),X2)=mult(f(rd(X1,X2)),X1))).
#
#cnf(i_0_321614, plain, (rd(mult(X1,mult(X2,f(mult(X1,X2)))),X2)=mult(f(mult(X1,X2)),X1))).
##########
#cnf(i_0_328962, plain, (rd(mult(ld(f(X1),X2),mult(X2,X1)),X1)=mult(ld(f(X1),X2),X2))).
####################
#cnf(i_0_335224, plain, (mult(ld(f(rd(X1,X2)),rd(X2,X1)),X1)=ld(f(rd(X1,X2)),X2))).
###
#cnf(i_0_335701, plain, (mult(X1,mult(ld(X1,X2),mult(f(X1),X2)))=mult(X2,mult(f(X1),X2)))).
#
#cnf(i_0_614319, plain, (ld(mult(X2,mult(f(X1),X2)),X1)=ld(mult(f(X1),X2),ld(X2,X1)))).
#############
#cnf(i_0_341242, plain, (ld(X1,mult(X2,mult(f(ld(X2,X1)),X1)))=ld(f(ld(X2,X1)),X1))).
#
#cnf(i_0_341300, plain, (rd(mult(X1,mult(f(ld(X1,X2)),X2)),X2)=mult(X1,f(ld(X1,X2))))).
#
#cnf(i_0_341358, plain, (ld(mult(X1,mult(f(ld(X1,X2)),X2)),X2)=ld(X2,f(ld(X1,X2))))).
######
#cnf(i_0_341805, plain, (ld(X1,mult(X2,ld(f(ld(X1,X2)),X1)))=mult(f(ld(X1,X2)),X1))).
##
#cnf(i_0_341863, plain, (rd(mult(X1,ld(f(ld(X2,X1)),X2)),X2)=rd(X1,f(ld(X2,X1))))).
##
#cnf(i_0_342325, plain, (mult(X1,mult(ld(X1,f(mult(X1,X2))),X2))=mult(f(mult(X1,X2)),X2))).
###
#cnf(i_0_343611, plain, (mult(rd(X1,mult(X2,f(mult(X2,X1)))),X2)=rd(X1,f(mult(X2,X1))))).
###
#cnf(i_0_344332, plain, (mult(ld(X1,rd(f(rd(X1,X2)),X2)),X2)=ld(X1,f(rd(X1,X2))))).
##
#cnf(i_0_344850, plain, (mult(ld(X1,ld(X1,ld(X2,f(X2)))),X2)=ld(X1,ld(X1,f(X2))))).
###
#cnf(i_0_345870, plain, (rd(mult(rd(X1,f(X2)),mult(X1,X2)),X2)=mult(X1,ld(f(X2),X1)))).
###
#cnf(i_0_346763, plain, (mult(mult(mult(X1,f(X2)),rd(X1,X2)),X2)=mult(X1,mult(f(X2),X1)))).
#
#cnf(i_0_624115, plain, (rd(rd(X1,X2),mult(X2,f(X1)))=rd(X1,mult(X2,mult(f(X1),X2))))).
##########
#cnf(i_0_347143, plain, (mult(rd(ld(X1,f(X2)),mult(X2,X1)),X2)=ld(X1,rd(f(X2),X1)))).
######
#cnf(i_0_351202, plain, (mult(X1,mult(ld(X1,X2),f(mult(X1,X2))))=mult(X2,f(mult(X1,X2))))).
##########
#cnf(i_0_354549, plain, (rd(mult(f(mult(X1,X2)),ld(X2,X1)),X1)=rd(f(mult(X1,X2)),X2))).
###
#cnf(i_0_355332, plain, (mult(rd(f(ld(X1,X2)),mult(X1,X2)),X1)=rd(f(ld(X1,X2)),X2))).
##
#cnf(i_0_355900, plain, (rd(mult(f(rd(X1,X2)),mult(X2,X1)),X1)=mult(f(rd(X1,X2)),X2))).
###
#cnf(i_0_359437, plain, (mult(rd(f(rd(X1,X2)),mult(X1,X2)),X1)=rd(f(rd(X1,X2)),X2))).
#
#cnf(i_0_632438, plain, (rd(mult(X1,X2),f(rd(X1,X2)))=mult(X1,rd(X2,f(rd(X1,X2)))))).
#
#cnf(i_0_634927, plain, (rd(rd(X1,X2),f(mult(X1,X2)))=rd(X1,mult(f(mult(X1,X2)),X2)))).
##
#cnf(i_0_359979, plain, (mult(mult(f(mult(X1,X2)),rd(X2,X1)),X1)=mult(f(mult(X1,X2)),X2))).
##
#cnf(i_0_361327, plain, (rd(mult(f(ld(X1,X2)),mult(X2,X1)),X1)=mult(f(ld(X1,X2)),X2))).
###
#cnf(i_0_363775, plain, (mult(ld(f(mult(X1,X2)),rd(X1,X2)),X2)=ld(f(mult(X1,X2)),X1))).
###############
#cnf(i_0_376649, plain, (mult(X1,mult(ld(X1,X2),f(mult(X2,X1))))=mult(X2,f(mult(X2,X1))))).
##############################
#cnf(i_0_412787, plain, (mult(ld(X2,X1),rd(f(ld(X1,X2)),X1))=rd(f(ld(X1,X2)),X2))).
#######################
#cnf(i_0_434572, plain, (mult(rd(X2,f(ld(X1,X2))),ld(X2,X1))=rd(X1,f(ld(X1,X2))))).
##########
#cnf(i_0_435492, plain, (mult(mult(f(ld(X2,X1)),X1),ld(X1,X2))=mult(f(ld(X2,X1)),X2))).
###############
#cnf(i_0_448764, plain, (ld(f(ld(X1,X2)),ld(X1,unit))=rd(f(ld(X1,X2)),X2))).
##############
#cnf(i_0_449868, plain, (mult(mult(X1,rd(f(ld(X2,X1)),X1)),X2)=mult(X2,f(ld(X2,X1))))).
###############################################
#cnf(i_0_467741, plain, (mult(rd(X2,X1),rd(X1,f(rd(X1,X2))))=rd(X2,f(rd(X1,X2))))).
###
#cnf(i_0_467789, plain, (mult(rd(f(rd(X1,X2)),X2),rd(X2,X1))=rd(f(rd(X1,X2)),X1))).
#################################
#cnf(i_0_485517, plain, (mult(X1,mult(X2,ld(f(mult(X1,X2)),X1)))=mult(f(mult(X1,X2)),X1))).
##
#cnf(i_0_485683, plain, (mult(X1,mult(X2,rd(f(mult(X1,X2)),X2)))=mult(f(mult(X1,X2)),X1))).
#####
#cnf(i_0_485742, plain, (ld(X1,mult(X2,rd(f(ld(X1,X2)),X2)))=rd(f(ld(X1,X2)),X1))).
###
#cnf(i_0_486336, plain, (mult(X1,mult(ld(X2,f(rd(X1,X2))),X2))=mult(f(rd(X1,X2)),X1))).
#####
#cnf(i_0_487168, plain, (rd(mult(X1,X2),mult(f(X1),mult(X2,X2)))=rd(X1,mult(f(X1),X2)))).
#########
#cnf(i_0_487889, plain, (rd(mult(ld(X1,f(rd(X1,X2))),X1),X2)=ld(X2,f(rd(X1,X2))))).
#
#cnf(i_0_487917, plain, (mult(mult(ld(X1,f(mult(X1,X2))),X1),X2)=mult(X2,f(mult(X1,X2))))).
#####
#cnf(i_0_488517, plain, (mult(rd(f(X1),X2),rd(X2,mult(X2,X1)))=ld(X1,rd(f(X1),X2)))).
################################
#cnf(i_0_497259, plain, (mult(ld(X1,f(rd(X2,X1))),rd(X1,X2))=ld(X2,f(rd(X2,X1))))).
##################
#cnf(i_0_504683, plain, (mult(mult(X1,rd(f(ld(X1,X2)),X1)),X2)=mult(X2,f(ld(X1,X2))))).
##############
#cnf(i_0_512387, plain, (mult(X1,ld(X2,ld(f(rd(X2,X1)),X2)))=ld(f(rd(X2,X1)),X1))).
################
#cnf(i_0_515166, plain, (mult(rd(X2,X1),mult(X1,f(ld(X2,X1))))=mult(X2,f(ld(X2,X1))))).
#############
#cnf(i_0_4321, plain, (ld(X1,rd(X2,mult(X1,f(ld(X1,X2)))))=rd(f(ld(X1,X2)),X1))).
#
#cnf(i_0_120224, plain, (mult(X1,ld(X2,ld(f(rd(X1,X2)),X1)))=mult(f(rd(X1,X2)),X1))).
##################################
#cnf(i_0_533468, plain, (ld(mult(X2,X1),mult(X2,mult(X2,f(X1))))=mult(ld(X1,X2),f(X1)))).
#########
#cnf(i_0_534553, plain, (ld(mult(X2,mult(X2,f(X1))),mult(X2,X1))=ld(f(X1),ld(X2,X1)))).
##################################
#cnf(i_0_542801, plain, (mult(mult(f(rd(X2,X1)),X1),ld(X2,X1))=ld(f(rd(X2,X1)),X1))).
#####
#cnf(i_0_543423, plain, (mult(ld(X1,X2),ld(X2,f(rd(X1,X2))))=ld(X1,f(rd(X1,X2))))).
########
#cnf(i_0_544380, plain, (mult(rd(f(ld(X1,X2)),X1),rd(X1,X2))=rd(f(ld(X1,X2)),X2))).
#########
#cnf(i_0_545797, plain, (rd(mult(ld(X1,f(rd(X2,X1))),X1),X2)=ld(X2,f(rd(X2,X1))))).
#######################################
#cnf(i_0_562100, plain, (mult(rd(X2,mult(X2,f(ld(X1,X2)))),X1)=rd(X1,f(ld(X1,X2))))).
################
#cnf(i_0_574328, plain, (mult(rd(f(rd(X1,X2)),X1),mult(X1,X2))=ld(f(rd(X1,X2)),X1))).
##
#cnf(i_0_574968, plain, (rd(mult(f(X1),X2),mult(X1,mult(X2,X2)))=rd(f(X1),mult(X1,X2)))).
########################
#cnf(i_0_577902, plain, (mult(rd(f(ld(X1,X2)),X2),mult(X1,X2))=ld(f(ld(X1,X2)),X2))).
######################
#cnf(i_0_582761, plain, (ld(mult(X2,mult(X2,X1)),mult(X2,f(X1)))=ld(X1,ld(X2,f(X1))))).
##########
#cnf(i_0_584031, plain, (mult(mult(X1,X2),ld(X2,f(ld(X1,X2))))=mult(X1,f(ld(X1,X2))))).
####
#cnf(i_0_585129, plain, (mult(mult(X1,X2),ld(X1,f(rd(X1,X2))))=rd(X1,f(rd(X1,X2))))).
################################
#cnf(i_0_593049, plain, (ld(mult(X1,f(ld(X1,X2))),mult(X1,X2))=ld(f(ld(X1,X2)),X2))).
###############
#cnf(i_0_595667, plain, (rd(mult(X1,X2),mult(f(rd(X1,X2)),X2))=rd(X1,f(rd(X1,X2))))).
###################
#cnf(i_0_602694, plain, (rd(mult(f(ld(X1,X2)),X1),mult(X2,X1))=rd(f(ld(X1,X2)),X2))).
####
#cnf(i_0_603942, plain, (ld(mult(X1,X2),mult(X1,f(rd(X2,X1))))=ld(X2,f(rd(X2,X1))))).
######################################################################
#cnf(i_0_50868, plain, (ld(mult(X1,X2),mult(X1,rd(X2,mult(X1,X1))))=ld(X1,ld(X1,unit)))).
######
#cnf(i_0_50901, plain, (mult(ld(X2,X1),ld(X1,rd(X2,mult(X1,X1))))=ld(X1,ld(X1,unit)))).
######
#cnf(i_0_51173, plain, (ld(mult(X1,mult(X2,mult(X1,X1))),mult(X1,X2))=ld(X1,ld(X1,unit)))).
####
#cnf(i_0_51233, plain, (mult(ld(X1,ld(X2,X1)),ld(X1,rd(X2,X1)))=ld(X1,ld(X1,unit)))).
##############
#cnf(i_0_49022, plain, (rd(X1,mult(X1,mult(f(X1),f(f(X1)))))=ld(X1,f(f(X1))))).
######
#cnf(i_0_53315, plain, (mult(rd(X1,X2),mult(f(X2),f(f(X2))))=rd(X1,f(f(X2))))).
####
#cnf(i_0_53585, plain, (rd(mult(X1,mult(f(X2),f(f(X2)))),X2)=rd(X1,f(f(X2))))).
#########
#cnf(i_0_54014, plain, (ld(mult(X1,mult(X2,X1)),mult(X1,rd(X2,X1)))=ld(X1,ld(X1,unit)))).
##########
#cnf(i_0_63727, plain, (mult(rd(X2,X1),rd(X1,mult(X2,mult(X2,X2))))=ld(X2,ld(X2,unit)))).
####
#cnf(i_0_64081, plain, (ld(mult(X1,mult(X2,mult(X2,X2))),mult(X1,X2))=ld(X2,ld(X2,unit)))).
###################
#cnf(i_0_90439, plain, (mult(mult(X1,f(X2)),ld(X2,f(f(X2))))=rd(X1,f(f(X2))))).
####
#cnf(i_0_90551, plain, (mult(mult(X1,ld(X2,f(f(X2)))),f(X2))=rd(X1,f(f(X2))))).
#################
#cnf(i_0_97216, plain, (rd(rd(X2,X1),mult(X1,mult(X2,rd(X2,X1))))=ld(X2,ld(X1,unit)))).
#####
#cnf(i_0_97369, plain, (ld(mult(X1,mult(X2,X2)),mult(X1,rd(X2,X1)))=ld(X2,ld(X1,unit)))).
################
#cnf(i_0_98532, plain, (mult(mult(X1,rd(f(X1),X2)),rd(X2,mult(X1,f(X1))))=unit)).
#####
#cnf(i_0_98574, plain, (mult(rd(X1,mult(f(X1),X2)),mult(f(X1),rd(X2,X1)))=unit)).
##
#cnf(i_0_98602, plain, (mult(mult(X1,ld(X2,f(X1))),ld(f(X1),rd(X2,X1)))=unit)).
#
#cnf(i_0_98610, plain, (mult(rd(f(ld(X1,X2)),X2),mult(X1,f(ld(X1,X2))))=unit)).
###########
#cnf(i_0_98641, plain, (mult(rd(X1,mult(X2,f(X2))),mult(X2,rd(f(X2),X1)))=unit)).
#####
#cnf(i_0_98683, plain, (mult(mult(f(X1),rd(X2,X1)),rd(X1,mult(f(X1),X2)))=unit)).
##
#cnf(i_0_98711, plain, (mult(ld(f(X1),rd(X2,X1)),mult(X1,ld(X2,f(X1))))=unit)).
#
#cnf(i_0_98719, plain, (mult(mult(X2,f(ld(X2,X1))),rd(f(ld(X2,X1)),X1))=unit)).
########
#cnf(i_0_98874, plain, (mult(ld(X1,rd(X2,X1)),rd(X1,mult(X1,X2)))=ld(X1,ld(X1,unit)))).
########
#cnf(i_0_101043, plain, (ld(mult(X1,mult(X1,X2)),mult(X1,rd(X1,X2)))=ld(X2,ld(X2,unit)))).
######
#cnf(i_0_104489, plain, (ld(mult(X1,f(X2)),f(mult(X1,X2)))=rd(f(X2),f(mult(X1,X2))))).
#
#cnf(i_0_104517, plain, (ld(mult(X1,f(ld(X1,X2))),f(X2))=rd(f(ld(X1,X2)),f(X2)))).
########
#cnf(i_0_105383, plain, (rd(rd(f(X1),X2),f(rd(X1,X2)))=ld(f(X1),f(rd(X1,X2))))).
####
#cnf(i_0_105446, plain, (ld(f(mult(X1,X2)),mult(X1,f(X2)))=rd(f(mult(X1,X2)),f(X2)))).
#
#cnf(i_0_105484, plain, (ld(f(ld(X1,X2)),ld(X1,f(X2)))=rd(f(ld(X1,X2)),f(X2)))).
###
#cnf(i_0_105498, plain, (ld(f(X1),mult(X2,f(ld(X2,X1))))=rd(f(X1),f(ld(X2,X1))))).
###
#cnf(i_0_105970, plain, (mult(f(X1),mult(ld(X1,X2),f(X1)))=mult(ld(f(X1),X2),f(X1)))).
###
#cnf(i_0_105983, plain, (rd(ld(f(X1),X2),f(ld(X1,X2)))=mult(f(X1),f(ld(X1,X2))))).
#
#cnf(i_0_106004, plain, (mult(ld(f(X1),X2),f(ld(X2,X1)))=rd(f(X1),f(ld(X2,X1))))).
##
#cnf(i_0_106130, plain, (ld(f(X1),rd(X2,f(ld(X1,X2))))=mult(f(X1),f(ld(X1,X2))))).
#####
#cnf(i_0_106327, plain, (rd(mult(f(X1),X2),f(mult(X1,X2)))=ld(f(X1),f(mult(X1,X2))))).
##
#cnf(i_0_106375, plain, (ld(f(rd(X1,X2)),rd(X1,f(X2)))=mult(f(rd(X1,X2)),f(X2)))).
####
#cnf(i_0_108693, plain, (mult(X1,mult(ld(f(X1),X2),f(X1)))=mult(f(X1),mult(X2,f(X1))))).
#######
#cnf(i_0_109234, plain, (rd(X1,mult(X2,mult(f(X1),f(f(X1)))))=rd(f(f(X1)),X2))).
#########
#cnf(i_0_109243, plain, (rd(X1,mult(f(X1),mult(X2,f(X1))))=rd(f(X1),mult(f(X1),X2)))).
#########
#cnf(i_0_110067, plain, (mult(f(rd(X1,X2)),rd(X2,f(X1)))=ld(f(rd(X1,X2)),f(X1)))).
##
#cnf(i_0_110283, plain, (rd(f(mult(X1,X2)),mult(f(X1),X2))=ld(f(mult(X1,X2)),f(X1)))).
##
#cnf(i_0_110349, plain, (rd(f(X1),mult(X1,f(ld(X1,f(X1)))))=f(ld(X1,f(X1))))).
#########
#cnf(i_0_112290, plain, (mult(rd(f(X1),mult(X1,X2)),mult(X1,rd(X2,f(X1))))=unit)).
#
#cnf(i_0_112318, plain, (mult(mult(X1,rd(X2,f(X1))),rd(f(X1),mult(X1,X2)))=unit)).
###
#cnf(i_0_113030, plain, (mult(mult(ld(X2,X1),f(X1)),ld(X1,ld(f(X1),X2)))=unit)).
#
#cnf(i_0_113062, plain, (mult(ld(X1,ld(f(X1),X2)),mult(ld(X2,X1),f(X1)))=unit)).
######
#cnf(i_0_113387, plain, (mult(mult(f(X1),ld(X2,X1)),ld(X1,rd(X2,f(X1))))=unit)).
#
#cnf(i_0_113415, plain, (mult(ld(X1,rd(X2,f(X1))),mult(f(X1),ld(X2,X1)))=unit)).
#######
#cnf(i_0_113696, plain, (mult(X1,mult(ld(f(X1),X2),f(X2)))=mult(f(X1),mult(X2,f(X2))))).
#####
#cnf(i_0_113773, plain, (mult(f(X1),mult(ld(X1,X2),f(X2)))=mult(ld(f(X1),X2),f(X2)))).
###############
#cnf(i_0_114579, plain, (rd(mult(ld(X1,X2),X1),mult(X1,mult(X2,X1)))=ld(X1,ld(X1,unit)))).
###
#cnf(i_0_114580, plain, (rd(mult(ld(X1,X2),X2),mult(X1,mult(X2,X2)))=ld(X1,ld(X1,unit)))).
###
#cnf(i_0_114611, plain, (mult(ld(X1,ld(X1,rd(X2,X3))),rd(X3,X2))=ld(X1,ld(X1,unit)))).
##
#cnf(i_0_114633, plain, (mult(ld(X1,ld(X1,ld(X2,X3))),ld(X3,X2))=ld(X1,ld(X1,unit)))).
##
#cnf(i_0_114784, plain, (rd(mult(ld(X1,ld(X1,X2)),X1),mult(X2,X1))=ld(X1,ld(X1,unit)))).
##############################
#cnf(i_0_119259, plain, (mult(mult(X1,rd(f(X1),f(X2))),X2)=mult(X1,mult(f(X1),f(X2))))).
###
#cnf(i_0_119332, plain, (mult(mult(X1,rd(f(X1),X2)),f(X2))=mult(X1,rd(f(X1),f(X2))))).
###
#cnf(i_0_119450, plain, (rd(ld(f(rd(X1,X2)),X1),f(X2))=mult(f(rd(X1,X2)),f(X2)))).
###
#cnf(i_0_119677, plain, (ld(f(X1),ld(X2,f(mult(X2,X1))))=rd(f(X1),f(mult(X2,X1))))).
#####
#cnf(i_0_119697, plain, (mult(ld(f(mult(X1,X2)),X1),f(X2))=rd(f(mult(X1,X2)),f(X2)))).
########
#cnf(i_0_120004, plain, (mult(f(X1),rd(X2,f(mult(X1,X2))))=ld(f(X1),f(mult(X1,X2))))).
########
#cnf(i_0_120005, plain, (rd(rd(f(mult(X1,X2)),X2),f(X1))=ld(f(mult(X1,X2)),f(X1)))).
#######
#cnf(i_0_120232, plain, (rd(f(X1),mult(f(rd(X1,X2)),X2))=ld(f(X1),f(rd(X1,X2))))).
##########
#cnf(i_0_121121, plain, (rd(ld(X1,ld(f(X1),X2)),f(X2))=ld(X1,ld(f(X1),f(X2))))).
####
#cnf(i_0_122399, plain, (rd(mult(X1,mult(f(X1),X2)),f(X2))=mult(X1,mult(f(X1),f(X2))))).
##
#cnf(i_0_122502, plain, (rd(mult(X1,mult(f(X1),f(X2))),X2)=mult(X1,rd(f(X1),f(X2))))).
##
#cnf(i_0_741136, plain, (rd(f(ld(f(X1),unit)),X1)=ld(X1,f(ld(f(X1),unit))))).
#####
#cnf(i_0_741348, plain, (rd(X1,f(ld(f(X1),unit)))=ld(f(ld(f(X1),unit)),X1))).
#####
#cnf(i_0_741942, plain, (mult(f(ld(f(X1),unit)),X1)=mult(X1,f(ld(f(X1),unit))))).
######
#cnf(i_0_122541, plain, (rd(X1,mult(X2,mult(f(X2),f(X1))))=rd(f(X1),mult(X2,f(X2))))).
#######
#cnf(i_0_123766, plain, (rd(mult(ld(X2,X1),X2),mult(X1,mult(X1,X2)))=ld(X2,ld(X1,unit)))).
##############
#cnf(i_0_125902, plain, (ld(mult(X1,mult(ld(X2,X1),X1)),rd(X1,X2))=ld(X1,ld(X1,unit)))).
####
#cnf(i_0_127856, plain, (mult(ld(X3,ld(X2,X1)),rd(ld(X1,X2),X3))=ld(X3,ld(X3,unit)))).
##
#cnf(i_0_128707, plain, (mult(ld(X1,rd(X2,X3)),rd(rd(X3,X2),X1))=ld(X1,ld(X1,unit)))).
####
#cnf(i_0_132217, plain, (ld(mult(X1,mult(X2,mult(X1,X2))),mult(X1,X2))=ld(X2,ld(X1,unit)))).
#######################
#cnf(i_0_748817, plain, (mult(rd(f(X1),X2),rd(X2,mult(X1,X1)))=ld(X1,ld(X1,f(X1))))).
##########################################
#cnf(i_0_750633, plain, (rd(mult(f(X1),X2),mult(X1,mult(X1,X2)))=ld(X1,ld(X1,f(X1))))).
#######################
#cnf(i_0_132222, plain, (rd(mult(X1,X2),mult(X1,mult(X2,mult(X1,X2))))=ld(X2,ld(X1,unit)))).
#######################
#cnf(i_0_132435, plain, (mult(rd(X1,X2),rd(X2,mult(X2,mult(X2,X1))))=ld(X2,ld(X2,unit)))).
##############
#cnf(i_0_133913, plain, (rd(ld(X1,X2),mult(X1,mult(X2,ld(X1,X2))))=ld(X2,ld(X1,unit)))).
######
#cnf(i_0_134479, plain, (mult(ld(X2,ld(X1,ld(X1,X2))),ld(X2,X1))=ld(X2,ld(X1,unit)))).
###############
#cnf(i_0_136245, plain, (mult(rd(X2,X1),rd(X1,mult(X1,mult(X2,X2))))=ld(X2,ld(X1,unit)))).
############
#cnf(i_0_138079, plain, (mult(ld(X2,ld(X1,rd(X2,X1))),rd(X1,X2))=ld(X2,ld(X1,unit)))).
######
#cnf(i_0_138435, plain, (mult(rd(X2,X1),rd(X1,mult(X2,mult(X1,X2))))=ld(X1,ld(X2,unit)))).
###############################################################################################################################################################################################
#cnf(i_0_159641, plain, (mult(rd(f(mult(X1,X2)),X2),ld(X1,f(mult(X1,X2))))=unit)).
##############
#cnf(i_0_159687, plain, (mult(ld(X2,f(mult(X2,X1))),rd(f(mult(X2,X1)),X1))=unit)).
################
#cnf(i_0_159763, plain, (mult(ld(f(mult(X1,X2)),X1),rd(X2,f(mult(X1,X2))))=unit)).
#################
#cnf(i_0_159803, plain, (mult(rd(X2,f(mult(X1,X2))),ld(f(mult(X1,X2)),X1))=unit)).
###################################
#cnf(i_0_164273, plain, (mult(ld(X1,f(rd(X1,X2))),mult(f(rd(X1,X2)),X2))=unit)).
############
#cnf(i_0_164313, plain, (mult(mult(f(rd(X1,X2)),X2),ld(X1,f(rd(X1,X2))))=unit)).
#######
#cnf(i_0_164325, plain, (rd(mult(f(rd(X1,X2)),X2),f(X1))=ld(f(rd(X1,X2)),f(X1)))).
#####################################
#cnf(i_0_184752, plain, (mult(ld(X1,rd(X2,X1)),rd(X1,mult(X2,X2)))=ld(X1,ld(X2,unit)))).
###########
#cnf(i_0_185508, plain, (mult(ld(X2,ld(X2,X1)),ld(X1,rd(X2,X1)))=ld(X2,ld(X1,unit)))).
############################################################################################################################################
#cnf(i_0_316359, plain, (rd(f(mult(X1,X1)),mult(X1,f(X1)))=ld(f(mult(X1,X1)),f(X1)))).
########
#cnf(i_0_325952, plain, (ld(X1,ld(f(X1),f(mult(X1,X1))))=rd(f(X1),f(mult(X1,X1))))).
##################################################################################################################################################################################
#cnf(i_0_522797, plain, (rd(mult(X1,f(X2)),mult(f(X1),X2))=rd(X1,mult(f(X1),f(X2))))).
#####
#cnf(i_0_522838, plain, (mult(rd(X1,f(X2)),rd(X2,f(X1)))=mult(X1,rd(f(X2),f(X1))))).
###########
#cnf(i_0_526826, plain, (rd(mult(f(X1),X2),mult(X1,f(X2)))=mult(f(X1),rd(f(X2),X1)))).
#####
#cnf(i_0_526896, plain, (mult(rd(f(X1),X2),rd(f(X2),X1))=rd(f(X1),mult(X1,f(X2))))).
############
#cnf(i_0_533469, plain, (ld(mult(f(X2),X1),mult(X2,f(X1)))=mult(ld(X1,f(X2)),f(X1)))).
###
#cnf(i_0_533538, plain, (mult(ld(X1,f(X2)),ld(X2,f(X1)))=ld(X1,ld(f(X2),f(X1))))).
####
#cnf(i_0_537541, plain, (ld(mult(X2,f(X1)),mult(f(X2),X1))=ld(f(X1),ld(f(X2),X1)))).
##
#cnf(i_0_537547, plain, (mult(ld(f(X1),X2),ld(f(X2),X1))=mult(ld(f(X1),f(X2)),X1))).
#########
#cnf(i_0_579524, plain, (rd(mult(f(X1),f(X2)),mult(X1,X2))=rd(f(X1),mult(X1,f(X2))))).
###
#cnf(i_0_579597, plain, (mult(rd(f(X1),f(X2)),rd(X2,X1))=mult(f(X1),rd(f(X2),X1)))).
##
#cnf(i_0_582760, plain, (ld(mult(X2,X1),mult(f(X2),f(X1)))=ld(X1,ld(f(X2),f(X1))))).
##
#cnf(i_0_582766, plain, (mult(ld(X1,X2),ld(f(X2),f(X1)))=mult(ld(X1,f(X2)),f(X1)))).
######
#cnf(i_0_591420, plain, (mult(ld(f(X2),f(X1)),ld(X1,X2))=ld(f(X2),ld(f(X1),X2)))).
#####
#cnf(i_0_593050, plain, (ld(mult(f(X1),f(X2)),mult(X1,X2))=mult(ld(f(X2),f(X1)),X2))).
######
#cnf(i_0_595532, plain, (mult(rd(X1,X2),rd(f(X2),f(X1)))=rd(X1,mult(f(X1),f(X2))))).
#
#cnf(i_0_595572, plain, (rd(mult(X1,X2),mult(f(X1),f(X2)))=mult(X1,rd(f(X2),f(X1))))).
###########################
#cnf(i_0_644422, plain, (rd(rd(X1,f(ld(X2,X1))),mult(X2,f(ld(X2,X1))))=unit)).
########
#cnf(i_0_644542, plain, (ld(mult(X1,f(ld(X1,X2))),rd(X2,f(ld(X1,X2))))=unit)).
############################################################################################################################################################
#cnf(i_0_713176, plain, (mult(X1,rd(f(X1),f(mult(X1,X1))))=ld(f(X1),f(mult(X1,X1))))).
#######################################################################
#cnf(i_0_765575, plain, (rd(ld(f(rd(X1,X2)),X1),mult(f(rd(X1,X2)),X2))=unit)).
############
#cnf(i_0_767360, plain, (ld(mult(f(rd(X2,X1)),X1),ld(f(rd(X2,X1)),X2))=unit)).
############################################################
#cnf(i_0_1985, plain, (mult(mult(X1,mult(ld(X2,X3),X3)),X2)=mult(mult(X1,ld(X2,X3)),mult(X3,X2)))).
#
#cnf(i_0_793737, plain, (mult(mult(ld(X1,X2),f(X2)),X1)=mult(ld(X1,f(X2)),mult(X2,X1)))).
#
#cnf(i_0_122371, plain, (ld(X1,mult(X2,mult(f(X2),X1)))=mult(ld(X1,f(X2)),mult(X2,X1)))).
#
#cnf(i_0_794082, plain, (mult(mult(X1,f(X2)),rd(X2,X1))=mult(X1,mult(X2,rd(f(X2),X1))))).
#
#cnf(i_0_794271, plain, (ld(mult(X2,X1),ld(f(X2),X1))=ld(X1,ld(X2,ld(f(X2),X1))))).
##
#cnf(i_0_794732, plain, (mult(ld(X2,f(X1)),mult(X1,X2))=mult(ld(X2,X1),mult(f(X1),X2)))).
##############
#cnf(i_0_795587, plain, (rd(rd(X1,X2),mult(X1,f(X2)))=rd(X1,mult(X1,mult(X2,f(X2)))))).
##########
#cnf(i_0_796103, plain, (mult(mult(X1,f(mult(X2,X1))),X2)=mult(X1,mult(f(mult(X2,X1)),X2)))).
#
#cnf(i_0_796173, plain, (rd(mult(X1,f(ld(X2,X1))),X2)=mult(X1,rd(f(ld(X2,X1)),X2)))).
#
#cnf(i_0_796198, plain, (mult(mult(X1,X2),rd(f(X2),X1))=mult(X1,mult(X2,rd(f(X2),X1))))).
######
#cnf(i_0_796843, plain, (rd(rd(X2,f(X1)),mult(X2,X1))=rd(X2,mult(X2,mult(X1,f(X1)))))).
####
#cnf(i_0_796845, plain, (ld(mult(f(X1),X2),ld(X1,X2))=ld(X2,ld(X1,ld(f(X1),X2))))).
#########
#cnf(i_0_122371, plain, (ld(X1,mult(X2,mult(f(X2),X1)))=mult(ld(X1,X2),mult(f(X2),X1)))).
#
#cnf(i_0_793737, plain, (mult(mult(ld(X1,X2),f(X2)),X1)=mult(ld(X1,X2),mult(f(X2),X1)))).
##
#cnf(i_0_798288, plain, (ld(X2,mult(f(rd(X1,X2)),X1))=mult(ld(X2,f(rd(X1,X2))),X1))).
#############
#cnf(i_0_804295, plain, (rd(rd(X1,f(ld(X1,X2))),X2)=rd(X1,mult(X2,f(ld(X1,X2)))))).
###
#cnf(i_0_113738, plain, (mult(X1,mult(ld(X1,X2),mult(f(X2),X1)))=mult(X2,mult(f(X2),X1)))).
#
#cnf(i_0_800175, plain, (rd(mult(X1,mult(f(mult(X2,X1)),X2)),X2)=mult(X1,f(mult(X2,X1))))).
#
#cnf(i_0_801413, plain, (mult(mult(X1,rd(f(ld(X2,X1)),X2)),X2)=mult(X1,f(ld(X2,X1))))).
###
#cnf(i_0_805445, plain, (rd(mult(ld(X1,X2),mult(f(X2),X1)),X1)=mult(ld(X1,X2),f(X2)))).
##
#cnf(i_0_806974, plain, (mult(X1,mult(ld(X1,f(rd(X2,X1))),X2))=mult(f(rd(X2,X1)),X2))).
####
#cnf(i_0_812701, plain, (ld(mult(f(rd(X2,X1)),X2),X1)=ld(X2,ld(f(rd(X2,X1)),X1)))).
#
#cnf(i_0_807851, plain, (mult(rd(X1,mult(X2,f(ld(X1,X2)))),X2)=rd(X1,f(ld(X1,X2))))).
######
#cnf(i_0_1987, plain, (mult(mult(X1,mult(X2,rd(X2,X3))),X3)=mult(mult(X1,rd(X2,X3)),mult(X3,X2)))).
#
#cnf(i_0_816408, plain, (mult(rd(f(rd(X1,X2)),X2),mult(X2,X1))=mult(f(rd(X1,X2)),X1))).
##
#cnf(i_0_816610, plain, (rd(mult(X1,X2),mult(f(rd(X2,X1)),X2))=rd(X1,f(rd(X2,X1))))).
###
#cnf(i_0_816708, plain, (rd(mult(f(rd(X1,X2)),X1),mult(X2,X1))=rd(f(rd(X1,X2)),X2))).
###
#cnf(i_0_817720, plain, (mult(mult(f(mult(X1,X2)),X2),ld(X2,X1))=mult(f(mult(X1,X2)),X1))).
#
#cnf(i_0_818026, plain, (mult(rd(f(X1),X2),mult(X2,mult(X1,X2)))=mult(X1,mult(f(X1),X2)))).
##
#cnf(i_0_820645, plain, (mult(mult(f(mult(X1,X2)),X1),ld(X1,X2))=mult(f(mult(X1,X2)),X2))).
##
#cnf(i_0_821551, plain, (rd(mult(f(ld(X1,X2)),X2),mult(X1,X2))=rd(f(ld(X1,X2)),X1))).
####
#cnf(i_0_822032, plain, (mult(mult(f(X1),X2),ld(X2,rd(X1,X2)))=mult(X1,rd(f(X1),X2)))).
########
#cnf(i_0_825133, plain, (mult(rd(f(ld(X1,X2)),X1),mult(X1,X2))=mult(f(ld(X1,X2)),X2))).
###
#cnf(i_0_825218, plain, (mult(mult(f(X2),X1),ld(X1,ld(X1,X2)))=mult(f(X2),ld(X1,X2)))).
#######
#cnf(i_0_826800, plain, (rd(mult(X2,ld(X1,X2)),mult(f(X1),X2))=rd(X2,mult(X1,f(X1))))).
########
#cnf(i_0_827601, plain, (rd(mult(X1,X2),mult(f(ld(X1,X2)),X2))=rd(X1,f(ld(X1,X2))))).
###
#cnf(i_0_829057, plain, (mult(rd(f(X2),X1),mult(X1,mult(X1,X2)))=mult(f(X2),mult(X1,X2)))).
##
#cnf(i_0_829254, plain, (rd(mult(ld(X1,X2),X2),mult(f(X1),X2))=ld(X1,rd(X2,f(X1))))).
#############
#cnf(i_0_819379, plain, (rd(ld(X1,X2),mult(f(X1),f(X2)))=ld(X1,rd(f(X2),f(X1))))).
############
#cnf(i_0_820509, plain, (mult(mult(f(X2),f(X1)),ld(X1,X2))=mult(f(X2),ld(f(X1),X2)))).
##
#cnf(i_0_828149, plain, (mult(rd(f(X2),f(X1)),mult(X1,X2))=mult(f(X2),mult(f(X1),X2)))).
###########
#cnf(i_0_1989, plain, (mult(mult(X1,mult(X3,X2)),ld(X2,X3))=mult(mult(X1,X2),mult(ld(X2,X3),X3)))).
#
#cnf(i_0_837762, plain, (rd(mult(X1,mult(X2,X3)),X3)=mult(mult(X1,X3),ld(X3,X2)))).
##
#cnf(i_0_839925, plain, (rd(mult(X1,ld(X2,X3)),X3)=rd(mult(X1,X3),mult(X2,X3)))).
#####
#cnf(i_0_840002, plain, (mult(mult(X1,rd(X2,X3)),X3)=mult(rd(X1,X3),mult(X3,X2)))).
#
#cnf(i_0_840277, plain, (rd(ld(X1,mult(X2,X3)),X3)=mult(ld(X1,X3),ld(X3,X2)))).
##
#cnf(i_0_841372, plain, (rd(ld(X1,ld(X2,X3)),X3)=rd(ld(X1,X3),mult(X2,X3)))).
#
#cnf(i_0_841448, plain, (mult(rd(X1,mult(X3,X2)),X3)=mult(rd(X1,X3),rd(X3,X2)))).
#####
#cnf(i_0_841836, plain, (ld(mult(X3,X2),rd(X3,X1))=ld(X3,rd(rd(X3,X2),X1)))).
##
#cnf(i_0_842663, plain, (ld(mult(X3,X1),mult(X3,X2))=mult(ld(X1,rd(X2,X3)),X3))).
###
#cnf(i_0_845320, plain, (ld(X2,rd(mult(X2,X3),X1))=mult(rd(X3,X2),rd(X2,X1)))).
######
#cnf(i_0_846260, plain, (mult(rd(X1,mult(X3,X2)),X2)=mult(rd(X1,X2),ld(X3,X2)))).
##
#cnf(i_0_846475, plain, (ld(mult(mult(X2,X3),X1),X2)=ld(X2,rd(rd(X2,X1),X3)))).
######
#cnf(i_0_847031, plain, (rd(ld(X1,X2),mult(X2,X1))=ld(X1,rd(X2,mult(X2,X1))))).
####
#cnf(i_0_847308, plain, (rd(ld(X2,X1),mult(X3,X1))=rd(X1,mult(X3,mult(X2,X1))))).
##
#cnf(i_0_848007, plain, (ld(mult(X2,X1),rd(X1,X3))=ld(X1,rd(ld(X2,X1),X3)))).
##
#cnf(i_0_848440, plain, (mult(mult(ld(X2,X1),X2),X1)=mult(ld(X2,X1),mult(X2,X1)))).
#
#cnf(i_0_849115, plain, (ld(mult(X1,X2),mult(X3,X1))=mult(ld(X2,ld(X1,X3)),X1))).
##
#cnf(i_0_850915, plain, (rd(mult(X2,X1),mult(X3,X1))=mult(X1,rd(ld(X1,X2),X3)))).
###
#cnf(i_0_850972, plain, (ld(X1,rd(mult(X2,X1),X3))=mult(ld(X1,X2),rd(X1,X3)))).
#
#cnf(i_0_851342, plain, (rd(mult(X1,rd(X3,X2)),X3)=rd(mult(X1,X3),mult(X3,X2)))).
###
#cnf(i_0_852343, plain, (mult(mult(X1,ld(X3,X2)),X3)=mult(rd(X1,X3),mult(X2,X3)))).
###
#cnf(i_0_852570, plain, (mult(mult(X1,rd(X2,X1)),X2)=mult(X1,mult(X2,ld(X1,X2))))).
#
#cnf(i_0_852579, plain, (mult(rd(X1,X2),ld(X2,X1))=mult(X1,ld(X2,ld(X2,X1))))).
#
#cnf(i_0_852615, plain, (ld(mult(mult(X2,X3),X1),X3)=ld(mult(X3,X1),ld(X2,X3)))).
##
#cnf(i_0_852676, plain, (ld(X2,mult(ld(X1,X2),X1))=mult(ld(X2,ld(X1,X2)),X1))).
###
#cnf(i_0_853901, plain, (ld(X1,mult(rd(X1,X3),X2))=mult(ld(X3,rd(X2,X1)),X1))).
##
#cnf(i_0_126787, plain, (rd(ld(X1,rd(X2,X1)),X2)=ld(X1,rd(X2,mult(X2,X1))))).
###
#cnf(i_0_854628, plain, (ld(mult(X2,X1),ld(X1,X2))=ld(X1,ld(X2,ld(X1,X2))))).
####
#cnf(i_0_854964, plain, (ld(X2,mult(X1,rd(X1,X2)))=mult(ld(X2,X1),rd(X1,X2)))).
#####
#cnf(i_0_855293, plain, (mult(mult(X1,X2),rd(X1,X2))=mult(X1,mult(X2,rd(X1,X2))))).
##
#cnf(i_0_841372, plain, (rd(ld(X1,ld(X2,X3)),X3)=rd(X3,mult(X2,mult(X1,X3))))).
#
#cnf(i_0_856537, plain, (mult(ld(X1,X2),ld(X2,X3))=mult(X2,ld(mult(X1,X2),X3)))).
####
#cnf(i_0_857524, plain, (rd(rd(X2,X1),mult(X3,X2))=rd(X2,mult(X3,mult(X2,X1))))).
###
#cnf(i_0_858397, plain, (rd(ld(X1,ld(X2,X1)),X2)=ld(X1,ld(X2,rd(X1,X2))))).
####
#cnf(i_0_859008, plain, (mult(mult(X2,X1),mult(X1,X2))=mult(X2,mult(X1,mult(X1,X2))))).
#####
#cnf(i_0_860161, plain, (rd(ld(X2,mult(X1,X3)),X1)=mult(ld(X2,X1),rd(X3,X1)))).
###
#cnf(i_0_839925, plain, (rd(mult(X1,ld(X2,X3)),X3)=mult(X3,rd(ld(X3,X1),X2)))).
#
#cnf(i_0_861873, plain, (mult(mult(X1,X2),ld(X2,X3))=mult(X2,mult(ld(X2,X1),X3)))).
#
#cnf(i_0_863787, plain, (rd(mult(X1,X2),mult(X3,X1))=mult(X1,rd(rd(X2,X1),X3)))).
#
#cnf(i_0_863959, plain, (ld(X1,mult(X2,rd(X1,X2)))=mult(ld(X1,X2),rd(X1,X2)))).
#
#cnf(i_0_864801, plain, (rd(rd(X1,mult(X1,X2)),X2)=rd(X1,mult(X2,mult(X1,X2))))).
###
#cnf(i_0_864937, plain, (rd(ld(X1,rd(X2,X3)),X2)=rd(ld(X1,X2),mult(X2,X3)))).
#
#cnf(i_0_865131, plain, (rd(mult(X1,mult(X2,X3)),X2)=mult(mult(X1,X2),rd(X3,X2)))).
###############
#cnf(i_0_869679, plain, (mult(rd(X1,X2),ld(X1,X2))=mult(X1,ld(X2,ld(X1,X2))))).
#
#cnf(i_0_869772, plain, (rd(mult(ld(X1,X2),X1),X2)=mult(ld(X1,X2),rd(X1,X2)))).
###
#cnf(i_0_134404, plain, (ld(mult(X1,mult(X2,X1)),X2)=ld(X1,ld(X2,ld(X1,X2))))).
##
#cnf(i_0_872865, plain, (mult(rd(X2,X1),mult(X2,X1))=mult(X2,mult(ld(X1,X2),X1)))).
##
#cnf(i_0_99130, plain, (rd(mult(ld(X1,X2),X2),X1)=mult(ld(X1,X2),rd(X2,X1)))).
###
#cnf(i_0_132145, plain, (rd(mult(X1,mult(X2,X1)),X2)=mult(X1,mult(X2,rd(X1,X2))))).
#####
#cnf(i_0_875062, plain, (ld(mult(X3,mult(X2,X1)),X1)=ld(X1,ld(X2,ld(X3,X1))))).
####
#cnf(i_0_840277, plain, (rd(ld(X1,mult(X2,X3)),X3)=mult(X3,ld(mult(X1,X3),X2)))).
#
#cnf(i_0_877687, plain, (mult(rd(X1,X2),ld(X1,X3))=mult(X1,ld(mult(X1,X2),X3)))).
#
#cnf(i_0_877688, plain, (mult(rd(X1,X2),ld(X3,X1))=mult(X1,ld(mult(X3,X2),X1)))).
#####
#cnf(i_0_881860, plain, (ld(X1,mult(X2,ld(X3,X1)))=mult(rd(ld(X1,X2),X3),X1))).
#
#cnf(i_0_837762, plain, (rd(mult(X1,mult(X2,X3)),X3)=mult(X3,mult(ld(X3,X1),X2)))).
###
#cnf(i_0_883857, plain, (ld(X2,mult(mult(X2,X1),X3))=mult(rd(X1,X2),mult(X2,X3)))).
###
#cnf(i_0_884871, plain, (mult(mult(X1,X2),ld(X1,X3))=mult(X1,mult(rd(X2,X1),X3)))).
#######
#cnf(i_0_896396, plain, (rd(rd(rd(X3,X2),X1),X3)=rd(X3,mult(mult(X3,X1),X2)))).
#
#cnf(i_0_896653, plain, (ld(X3,ld(X2,mult(X1,X3)))=mult(ld(mult(X2,X3),X1),X3))).
##
#cnf(i_0_898765, plain, (ld(mult(X2,X1),mult(X1,X3))=ld(X1,mult(ld(X2,X1),X3)))).
##
#cnf(i_0_898945, plain, (ld(X1,rd(mult(X3,X2),X1))=mult(ld(X1,X3),rd(X2,X1)))).
#
#cnf(i_0_900453, plain, (rd(rd(mult(X1,X2),X3),X1)=mult(X1,rd(X2,mult(X1,X3))))).
#
#cnf(i_0_900559, plain, (ld(X1,mult(X2,mult(X3,X1)))=mult(mult(ld(X1,X2),X3),X1))).
###
#cnf(i_0_902855, plain, (ld(X1,mult(mult(X2,X1),X3))=mult(ld(X1,X2),mult(X1,X3)))).
##
#cnf(i_0_906201, plain, (rd(mult(rd(X1,X2),X3),X1)=mult(X1,ld(X2,rd(X3,X1))))).
#####
#cnf(i_0_908761, plain, (ld(mult(X2,X1),rd(X3,X1))=ld(X1,rd(ld(X2,X3),X1)))).
#
#cnf(i_0_908837, plain, (rd(ld(X1,X2),mult(X1,X3))=ld(X1,rd(rd(X2,X3),X1)))).
##
#cnf(i_0_911115, plain, (ld(X1,mult(X2,ld(X1,X3)))=mult(ld(X1,rd(X2,X1)),X3))).
##
#cnf(i_0_911190, plain, (rd(mult(mult(X1,X2),X3),X1)=mult(X1,mult(X2,rd(X3,X1))))).
######
#cnf(i_0_915024, plain, (mult(rd(X2,X1),mult(X3,X2))=mult(X2,mult(ld(X1,X3),X2)))).
#
#cnf(i_0_915993, plain, (mult(mult(X1,X2),ld(X3,X1))=mult(X1,mult(rd(X2,X3),X1)))).
#
#cnf(i_0_916954, plain, (ld(mult(X1,mult(X2,X1)),X3)=ld(X1,ld(X2,ld(X1,X3))))).
#
#cnf(i_0_916999, plain, (rd(ld(X1,rd(X2,X1)),X3)=ld(X1,rd(X2,mult(X3,X1))))).
#
#cnf(i_0_922077, plain, (mult(mult(X1,ld(X2,X1)),X3)=mult(X1,ld(X2,mult(X1,X3))))).
##
#cnf(i_0_923492, plain, (rd(mult(X1,ld(X2,X1)),X3)=mult(X1,ld(X2,rd(X1,X3))))).
##
#cnf(i_0_923648, plain, (rd(ld(mult(X1,X3),X2),X1)=ld(X3,ld(X1,rd(X2,X1))))).
#
#cnf(i_0_925991, plain, (ld(X2,mult(X1,mult(X3,X1)))=mult(mult(ld(X2,X1),X3),X1))).
#
#cnf(i_0_926067, plain, (rd(mult(rd(X2,X1),X3),X1)=mult(X2,ld(X1,rd(X3,X1))))).
###
#cnf(i_0_927106, plain, (ld(X1,mult(X2,ld(X3,X2)))=mult(rd(ld(X1,X2),X3),X2))).
#
#cnf(i_0_928082, plain, (rd(rd(rd(X1,X2),X3),X2)=rd(X1,mult(X2,mult(X3,X2))))).
###
#cnf(i_0_842340, plain, (mult(rd(X1,X3),mult(X3,mult(X2,X3)))=mult(mult(X1,X2),X3))).
#
#cnf(i_0_930031, plain, (mult(mult(mult(X1,X3),X2),X3)=mult(X1,mult(X3,mult(X2,X3))))).
###
#cnf(i_0_932164, plain, (mult(rd(mult(X1,X2),X3),X2)=mult(X1,mult(X2,ld(X3,X2))))).
##
#cnf(i_0_842411, plain, (mult(rd(X1,X3),mult(X3,ld(X2,X3)))=mult(rd(X1,X2),X3))).
#
#cnf(i_0_845178, plain, (mult(X1,mult(rd(X3,X1),rd(X1,X2)))=rd(mult(X1,X3),X2))).
##
#cnf(i_0_935364, plain, (rd(mult(X1,mult(X2,X1)),X3)=mult(X1,mult(X2,rd(X1,X3))))).
#####
#cnf(i_0_184763, plain, (mult(ld(X2,rd(X1,mult(X1,X2))),X1)=ld(X2,rd(X1,X2)))).
####
#cnf(i_0_848482, plain, (mult(ld(X2,ld(X1,rd(X2,X1))),X1)=ld(X2,ld(X1,X2)))).
###
#cnf(i_0_849143, plain, (mult(ld(X2,ld(X1,rd(X3,X1))),X1)=ld(mult(X1,X2),X3))).
##
#cnf(i_0_851000, plain, (mult(ld(X1,rd(X2,X1)),rd(X1,X3))=ld(X1,rd(X2,X3)))).
##
#cnf(i_0_851379, plain, (mult(X1,mult(ld(X1,X3),rd(X1,X2)))=rd(mult(X3,X1),X2))).
###########
#cnf(i_0_860829, plain, (mult(X1,mult(ld(mult(X2,X1),X3),X1))=ld(X2,mult(X3,X1)))).
###
#cnf(i_0_861276, plain, (mult(X2,mult(rd(ld(X2,X1),X3),X2))=mult(X1,ld(X3,X2)))).
#########
#cnf(i_0_864744, plain, (rd(mult(X1,X2),mult(X2,mult(X3,X2)))=rd(rd(X1,X3),X2))).
###
#cnf(i_0_866143, plain, (rd(X1,mult(X2,mult(X1,ld(X2,X1))))=rd(X2,mult(X2,X1)))).
###
#cnf(i_0_869317, plain, (mult(X1,mult(ld(X1,ld(X2,X1)),X2))=mult(ld(X2,X1),X2))).
####
#cnf(i_0_870011, plain, (mult(X1,mult(ld(X2,rd(X3,X1)),X1))=mult(rd(X1,X2),X3))).
#####
#cnf(i_0_873053, plain, (mult(X1,mult(ld(X1,X2),rd(X2,X1)))=mult(X2,rd(X2,X1)))).
######
#cnf(i_0_873995, plain, (mult(mult(X1,mult(X2,rd(X1,X2))),X2)=mult(X1,mult(X2,X1)))).
############
#cnf(i_0_883037, plain, (mult(X2,mult(mult(ld(X2,X1),X3),X2))=mult(X1,mult(X3,X2)))).
#
#cnf(i_0_949625, plain, (mult(mult(X1,X2),mult(X3,X1))=mult(X1,mult(mult(X2,X3),X1)))).
##
#cnf(i_0_883269, plain, (mult(X2,mult(ld(X2,X1),mult(X2,X3)))=mult(mult(X1,X2),X3))).
#
#cnf(i_0_952842, plain, (mult(mult(X1,mult(X2,X1)),X3)=mult(X1,mult(X2,mult(X1,X3))))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(f04, axiom, ![X1, X2]:rd(mult(X2,X1),X1)=X2, file('./p/GRP667+1.tptp', f04)).
fof(f08, axiom, ![X1, X2]:mult(mult(X2,X1),X2)=mult(X2,mult(X1,X2)), file('./p/GRP667+1.tptp', f08)).
fof(f03, axiom, ![X1, X2]:mult(rd(X2,X1),X1)=X2, file('./p/GRP667+1.tptp', f03)).
fof(f06, axiom, ![X2]:mult(unit,X2)=X2, file('./p/GRP667+1.tptp', f06)).
fof(f05, axiom, ![X2]:mult(X2,unit)=X2, file('./p/GRP667+1.tptp', f05)).
fof(f01, axiom, ![X1, X2]:mult(X2,ld(X2,X1))=X1, file('./p/GRP667+1.tptp', f01)).
fof(f07, axiom, ![X3, X1, X2]:mult(mult(X2,X1),mult(mult(X3,X1),X3))=mult(mult(X2,mult(mult(X1,X3),X1)),X3), file('./p/GRP667+1.tptp', f07)).
fof(f02, axiom, ![X1, X2]:ld(X2,mult(X2,X1))=X1, file('./p/GRP667+1.tptp', f02)).
fof(f09, axiom, ![X2]:mult(f(X2),f(X2))=X2, file('./p/GRP667+1.tptp', f09)).
fof(goals, conjecture, mult(a,mult(b,mult(a,c)))=mult(mult(mult(a,b),a),c), file('./p/GRP667+1.tptp', goals)).
fof(c_0_10, plain, ![X19, X20]:rd(mult(X20,X19),X19)=X20, inference(variable_rename,[status(thm)],[f04])).
fof(c_0_11, plain, ![X26, X27]:mult(mult(X27,X26),X27)=mult(X27,mult(X26,X27)), inference(variable_rename,[status(thm)],[f08])).
cnf(c_0_12, plain, (rd(mult(X1,X2),X2)=X1), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13, plain, (mult(mult(X1,X2),X1)=mult(X1,mult(X2,X1))), inference(split_conjunct,[status(thm)],[c_0_11])).
fof(c_0_14, plain, ![X17, X18]:mult(rd(X18,X17),X17)=X18, inference(variable_rename,[status(thm)],[f03])).
fof(c_0_15, plain, ![X22]:mult(unit,X22)=X22, inference(variable_rename,[status(thm)],[f06])).
cnf(c_0_16, plain, (rd(mult(X1,mult(X2,X1)),X1)=mult(X1,X2)), inference(pm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17, plain, (mult(rd(X1,X2),X2)=X1), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_18, plain, ![X21]:mult(X21,unit)=X21, inference(variable_rename,[status(thm)],[f05])).
cnf(c_0_19, plain, (mult(unit,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_20, plain, ![X13, X14]:mult(X14,ld(X14,X13))=X13, inference(variable_rename,[status(thm)],[f01])).
fof(c_0_21, plain, ![X23, X24, X25]:mult(mult(X25,X24),mult(mult(X23,X24),X23))=mult(mult(X25,mult(mult(X24,X23),X24)),X23), inference(variable_rename,[status(thm)],[f07])).
fof(c_0_22, plain, ![X15, X16]:ld(X16,mult(X16,X15))=X15, inference(variable_rename,[status(thm)],[f02])).
cnf(c_0_23, plain, (rd(mult(X1,X2),X1)=mult(X1,rd(X2,X1))), inference(pm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_24, plain, (mult(X1,unit)=X1), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (rd(X1,X1)=unit), inference(pm,[status(thm)],[c_0_12, c_0_19])).
cnf(c_0_26, plain, (mult(X1,ld(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (mult(mult(X1,X2),mult(mult(X3,X2),X3))=mult(mult(X1,mult(mult(X2,X3),X2)),X3)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (ld(X1,mult(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_29, plain, (mult(X1,rd(unit,X1))=unit), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_23, c_0_24]), c_0_25])).
cnf(c_0_30, plain, (mult(X1,mult(ld(X1,X2),X1))=mult(X2,X1)), inference(pm,[status(thm)],[c_0_13, c_0_26])).
fof(c_0_31, plain, ![X28]:mult(f(X28),f(X28))=X28, inference(variable_rename,[status(thm)],[f09])).
cnf(c_0_32, plain, (mult(mult(X1,mult(X2,mult(X3,X2))),X3)=mult(mult(X1,X2),mult(X3,mult(X2,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_13]), c_0_13])).
cnf(c_0_33, plain, (rd(unit,X1)=ld(X1,unit)), inference(pm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_34, plain, (ld(X1,mult(X2,X1))=mult(ld(X1,X2),X1)), inference(pm,[status(thm)],[c_0_28, c_0_30])).
cnf(c_0_35, plain, (ld(mult(X1,X2),mult(X1,mult(X2,X1)))=X1), inference(pm,[status(thm)],[c_0_28, c_0_13])).
cnf(c_0_36, plain, (mult(f(X1),f(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_31])).
cnf(c_0_37, plain, (mult(mult(X1,X2),X2)=mult(X1,mult(X2,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_24, c_0_32]), c_0_24]), c_0_19]), c_0_19])).
cnf(c_0_38, plain, (mult(ld(X1,unit),X1)=unit), inference(pm,[status(thm)],[c_0_17, c_0_33])).
cnf(c_0_39, plain, (mult(ld(mult(X1,X1),X1),mult(X1,X1))=X1), inference(pm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_40, plain, (mult(f(X1),X1)=mult(X1,f(X1))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_13, c_0_36]), c_0_36])).
cnf(c_0_41, plain, (mult(ld(X1,unit),mult(X1,X1))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_37, c_0_38]), c_0_19])).
cnf(c_0_42, plain, (rd(X1,mult(X1,X1))=ld(mult(X1,X1),X1)), inference(pm,[status(thm)],[c_0_12, c_0_39])).
cnf(c_0_43, plain, (mult(X1,rd(ld(X1,X2),X1))=rd(X2,X1)), inference(pm,[status(thm)],[c_0_23, c_0_26])).
cnf(c_0_44, plain, (rd(mult(X1,f(X1)),X1)=f(X1)), inference(pm,[status(thm)],[c_0_12, c_0_40])).
cnf(c_0_45, plain, (ld(mult(X1,X1),X1)=ld(X1,unit)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_12, c_0_41]), c_0_42])).
cnf(c_0_46, plain, (ld(rd(X1,X2),X1)=X2), inference(pm,[status(thm)],[c_0_28, c_0_17])).
cnf(c_0_47, plain, (rd(ld(X1,X2),X1)=ld(X1,rd(X2,X1))), inference(pm,[status(thm)],[c_0_28, c_0_43])).
cnf(c_0_48, plain, (mult(rd(X1,X2),mult(X2,X2))=mult(X1,X2)), inference(pm,[status(thm)],[c_0_37, c_0_17])).
cnf(c_0_49, plain, (mult(X1,rd(f(X1),X1))=f(X1)), inference(rw,[status(thm)],[c_0_44, c_0_23])).
cnf(c_0_50, plain, (ld(X1,f(X1))=ld(f(X1),unit)), inference(pm,[status(thm)],[c_0_45, c_0_36])).
cnf(c_0_51, plain, (ld(ld(X1,rd(X2,X1)),ld(X1,X2))=X1), inference(pm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_52, plain, (rd(mult(X1,X2),mult(X2,X2))=rd(X1,X2)), inference(pm,[status(thm)],[c_0_12, c_0_48])).
cnf(c_0_53, plain, (mult(X1,mult(mult(ld(X1,X2),X1),X1))=mult(mult(X2,X1),X1)), inference(pm,[status(thm)],[c_0_13, c_0_30])).
cnf(c_0_54, plain, (ld(rd(X1,X2),mult(X1,X2))=mult(X2,X2)), inference(pm,[status(thm)],[c_0_28, c_0_48])).
cnf(c_0_55, plain, (rd(f(X1),X1)=ld(X1,f(X1))), inference(pm,[status(thm)],[c_0_28, c_0_49])).
cnf(c_0_56, plain, (mult(f(X1),ld(X1,f(X1)))=unit), inference(pm,[status(thm)],[c_0_26, c_0_50])).
cnf(c_0_57, plain, (rd(X1,ld(X2,X1))=X2), inference(pm,[status(thm)],[c_0_12, c_0_26])).
cnf(c_0_58, plain, (ld(ld(X1,ld(X1,unit)),ld(X1,unit))=X1), inference(pm,[status(thm)],[c_0_51, c_0_33])).
cnf(c_0_59, plain, (rd(mult(X1,f(X2)),X2)=rd(X1,f(X2))), inference(pm,[status(thm)],[c_0_52, c_0_36])).
cnf(c_0_60, plain, (mult(X1,mult(ld(X1,X2),mult(X1,X1)))=mult(X2,mult(X1,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_37]), c_0_37])).
cnf(c_0_61, plain, (ld(ld(X1,f(X1)),mult(X1,f(X1)))=mult(X1,X1)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_54, c_0_40]), c_0_55])).
cnf(c_0_62, plain, (mult(ld(X1,f(X1)),ld(X1,f(X1)))=ld(X1,unit)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_54, c_0_56]), c_0_57])).
cnf(c_0_63, plain, (mult(mult(X1,X1),ld(X1,unit))=X1), inference(pm,[status(thm)],[c_0_26, c_0_45])).
cnf(c_0_64, plain, (mult(ld(X1,f(X1)),X1)=f(X1)), inference(pm,[status(thm)],[c_0_17, c_0_55])).
cnf(c_0_65, plain, (rd(X1,ld(X1,ld(X1,unit)))=ld(ld(X1,ld(X1,unit)),X1)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_47, c_0_58]), c_0_57])).
cnf(c_0_66, plain, (ld(f(X1),ld(f(X1),unit))=ld(X1,unit)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_59, c_0_38]), c_0_33]), c_0_47]), c_0_33])).
cnf(c_0_67, plain, (mult(mult(X1,f(X1)),ld(X1,unit))=f(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_63]), c_0_64]), c_0_62])).
cnf(c_0_68, plain, (rd(f(X1),ld(X1,unit))=ld(ld(X1,unit),f(X1))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_65, c_0_66]), c_0_66])).
cnf(c_0_69, plain, (ld(ld(X1,unit),f(X1))=mult(X1,f(X1))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_12, c_0_67]), c_0_68])).
cnf(c_0_70, plain, (rd(mult(X1,mult(X2,X2)),X2)=mult(X1,X2)), inference(pm,[status(thm)],[c_0_12, c_0_37])).
cnf(c_0_71, plain, (mult(ld(X1,f(X1)),f(X1))=mult(f(X1),ld(X1,f(X1)))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_13, c_0_64]), c_0_26])).
cnf(c_0_72, plain, (mult(rd(X1,X2),mult(X2,rd(X1,X2)))=mult(X1,rd(X1,X2))), inference(pm,[status(thm)],[c_0_13, c_0_17])).
cnf(c_0_73, plain, (rd(f(X1),ld(X1,unit))=mult(X1,f(X1))), inference(rw,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_74, plain, (rd(mult(X1,X2),f(X2))=mult(X1,f(X2))), inference(pm,[status(thm)],[c_0_70, c_0_36])).
cnf(c_0_75, plain, (mult(ld(X1,f(X1)),f(X1))=unit), inference(rw,[status(thm)],[c_0_71, c_0_56])).
cnf(c_0_76, plain, (mult(X1,ld(f(X1),unit))=f(X1)), inference(pm,[status(thm)],[c_0_63, c_0_36])).
cnf(c_0_77, plain, (mult(ld(f(X1),unit),X1)=f(X1)), inference(pm,[status(thm)],[c_0_41, c_0_36])).
cnf(c_0_78, plain, (mult(mult(X1,f(X1)),f(X1))=mult(f(X1),mult(X1,f(X1)))), inference(pm,[status(thm)],[c_0_13, c_0_40])).
cnf(c_0_79, plain, (mult(mult(X1,mult(X2,X3)),rd(X3,X2))=mult(mult(X1,X2),mult(X3,rd(X3,X2)))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_32, c_0_17]), c_0_72])).
cnf(c_0_80, plain, (rd(f(f(X1)),ld(X1,f(X1)))=mult(f(X1),f(f(X1)))), inference(pm,[status(thm)],[c_0_73, c_0_50])).
cnf(c_0_81, plain, (mult(ld(X1,f(X1)),f(f(X1)))=ld(f(f(X1)),unit)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_74, c_0_75]), c_0_33])).
cnf(c_0_82, plain, (mult(mult(X1,ld(f(X2),unit)),mult(X2,f(X2)))=mult(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_32, c_0_76]), c_0_38]), c_0_24]), c_0_77])).
cnf(c_0_83, plain, (mult(f(X1),mult(X1,f(X1)))=mult(X1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_37]), c_0_36])).
cnf(c_0_84, plain, (mult(mult(X1,ld(X2,f(X2))),X2)=mult(X1,f(X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_79, c_0_80]), c_0_81]), c_0_82]), c_0_80]), c_0_83]), c_0_36])).
cnf(c_0_85, plain, (mult(rd(X1,mult(X2,X2)),X2)=rd(X1,X2)), inference(pm,[status(thm)],[c_0_70, c_0_17])).
cnf(c_0_86, plain, (mult(X1,ld(X2,f(X2)))=rd(X1,f(X2))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_12, c_0_84]), c_0_59])).
cnf(c_0_87, plain, (rd(rd(X1,X2),X2)=rd(X1,mult(X2,X2))), inference(pm,[status(thm)],[c_0_12, c_0_85])).
cnf(c_0_88, plain, (mult(mult(X1,mult(X2,X1)),X2)=mult(X1,mult(X2,mult(X1,X2)))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_32, c_0_19]), c_0_19])).
cnf(c_0_89, plain, (mult(X1,ld(X2,unit))=rd(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_37, c_0_86]), c_0_86]), c_0_87]), c_0_36]), c_0_62])).
cnf(c_0_90, plain, (mult(mult(X1,X2),ld(X2,X1))=mult(X2,mult(ld(X2,X1),X1))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_88, c_0_30]), c_0_26])).
cnf(c_0_91, plain, (ld(X1,rd(X1,X2))=ld(X2,unit)), inference(pm,[status(thm)],[c_0_28, c_0_89])).
cnf(c_0_92, plain, (mult(mult(X1,X2),mult(X1,mult(X2,X1)))=mult(X1,mult(X2,mult(X1,mult(X2,X1))))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_13, c_0_32]), c_0_88])).
cnf(c_0_93, plain, (mult(mult(ld(X1,X2),X2),X1)=mult(ld(X1,X2),mult(X2,X1))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_88, c_0_26]), c_0_30])).
cnf(c_0_94, plain, (mult(X1,mult(ld(X2,unit),rd(X1,X2)))=rd(mult(rd(X1,X2),X1),X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_90, c_0_91]), c_0_89]), c_0_91])).
cnf(c_0_95, plain, (mult(X1,mult(ld(X1,X2),mult(X2,X1)))=mult(X2,mult(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_92, c_0_30]), c_0_26]), c_0_30])).
cnf(c_0_96, plain, (rd(mult(ld(X1,X2),mult(X2,X1)),X1)=mult(ld(X1,X2),X2)), inference(pm,[status(thm)],[c_0_12, c_0_93])).
cnf(c_0_97, plain, (ld(X1,rd(mult(rd(X1,X2),X1),X2))=mult(ld(X2,unit),rd(X1,X2))), inference(pm,[status(thm)],[c_0_28, c_0_94])).
cnf(c_0_98, plain, (rd(mult(X1,mult(X1,X2)),X2)=mult(X2,mult(ld(X2,X1),X1))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_23, c_0_95]), c_0_96])).
cnf(c_0_99, plain, (ld(mult(X1,X2),mult(X2,mult(ld(X2,X1),X1)))=ld(X2,X1)), inference(pm,[status(thm)],[c_0_28, c_0_90])).
cnf(c_0_100, plain, (mult(ld(X1,unit),X2)=ld(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_97, c_0_12]), c_0_98]), c_0_99]), c_0_12])).
cnf(c_0_101, plain, (ld(X1,mult(X2,X2))=mult(ld(X1,X2),X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_37, c_0_100]), c_0_100])).
cnf(c_0_102, plain, (mult(mult(X1,X1),X2)=mult(X1,mult(X1,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_30, c_0_101]), c_0_93]), c_0_95])).
cnf(c_0_103, plain, (mult(rd(X1,f(X2)),X2)=mult(X1,f(X2))), inference(pm,[status(thm)],[c_0_48, c_0_36])).
cnf(c_0_104, plain, (rd(ld(X1,X2),X2)=ld(X1,unit)), inference(pm,[status(thm)],[c_0_12, c_0_100])).
cnf(c_0_105, plain, (ld(mult(X1,X2),X1)=ld(X2,unit)), inference(pm,[status(thm)],[c_0_91, c_0_12])).
cnf(c_0_106, plain, (rd(ld(X1,unit),X2)=ld(X1,ld(X2,unit))), inference(pm,[status(thm)],[c_0_89, c_0_100])).
cnf(c_0_107, plain, (mult(mult(X1,mult(X2,X3)),ld(X3,X2))=mult(mult(X1,X3),mult(ld(X3,X2),X2))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_32, c_0_30]), c_0_26])).
cnf(c_0_108, plain, (mult(f(X1),mult(f(X1),X2))=mult(X1,X2)), inference(pm,[status(thm)],[c_0_102, c_0_36])).
cnf(c_0_109, plain, (mult(ld(X1,f(X2)),f(X2))=ld(X1,X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_103, c_0_104]), c_0_100])).
cnf(c_0_110, plain, (ld(mult(X1,X2),unit)=ld(X2,ld(X1,unit))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_104, c_0_105]), c_0_106])).
cnf(c_0_111, plain, (ld(ld(X1,X2),unit)=ld(X2,X1)), inference(pm,[status(thm)],[c_0_91, c_0_57])).
cnf(c_0_112, plain, (rd(mult(X1,mult(X2,X3)),X3)=mult(mult(X1,X3),ld(X3,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_107, c_0_108]), c_0_105]), c_0_89]), c_0_105]), c_0_100]), c_0_107]), c_0_109])).
cnf(c_0_113, plain, (mult(X1,ld(X2,ld(X3,unit)))=rd(X1,mult(X3,X2))), inference(pm,[status(thm)],[c_0_89, c_0_110])).
cnf(c_0_114, plain, (rd(X1,ld(X2,X3))=mult(X1,ld(X3,X2))), inference(pm,[status(thm)],[c_0_89, c_0_111])).
cnf(c_0_115, plain, (ld(rd(X1,X2),unit)=rd(X2,X1)), inference(pm,[status(thm)],[c_0_104, c_0_46])).
cnf(c_0_116, plain, (ld(unit,X1)=X1), inference(pm,[status(thm)],[c_0_28, c_0_19])).
cnf(c_0_117, plain, (rd(mult(X1,ld(X2,X3)),X3)=rd(mult(X1,X3),mult(X2,X3))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_112, c_0_100]), c_0_113])).
cnf(c_0_118, plain, (rd(X1,rd(X2,X3))=mult(X1,rd(X3,X2))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_114, c_0_115]), c_0_116])).
cnf(c_0_119, plain, (ld(ld(X1,X2),X3)=mult(ld(X2,X1),X3)), inference(pm,[status(thm)],[c_0_100, c_0_111])).
cnf(c_0_120, plain, (ld(rd(X1,X2),X3)=mult(rd(X2,X1),X3)), inference(pm,[status(thm)],[c_0_100, c_0_115])).
cnf(c_0_121, plain, (mult(rd(X1,mult(X2,X3)),X2)=mult(rd(X1,X2),rd(X2,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_117, c_0_113]), c_0_114]), c_0_116]), c_0_89]), c_0_89]), c_0_118])).
cnf(c_0_122, plain, (ld(mult(rd(X1,X2),X3),X4)=mult(ld(X3,rd(X2,X1)),X4)), inference(pm,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_123, plain, (ld(X1,rd(mult(X1,X2),X3))=mult(rd(X2,X1),rd(X1,X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_110, c_0_121]), c_0_122]), c_0_120]), c_0_24]), c_0_120]), c_0_24])).
cnf(c_0_124, plain, (rd(X1,ld(X2,unit))=mult(X1,X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_89, c_0_46]), c_0_33])).
fof(c_0_125, negated_conjecture, ~(mult(a,mult(b,mult(a,c)))=mult(mult(mult(a,b),a),c)), inference(assume_negation,[status(cth)],[goals])).
cnf(c_0_126, plain, (rd(mult(X1,X2),mult(X3,X2))=mult(X2,rd(ld(X2,X1),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_123, c_0_100]), c_0_119]), c_0_116]), c_0_124]), c_0_106]), c_0_113])).
fof(c_0_127, negated_conjecture, mult(a,mult(b,mult(a,c)))!=mult(mult(mult(a,b),a),c), inference(fof_simplification,[status(thm)],[c_0_125])).
cnf(c_0_128, plain, (mult(mult(X1,X2),ld(X2,X3))=mult(X2,mult(ld(X2,X1),X3))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_126, c_0_100]), c_0_114]), c_0_114]), c_0_116])).
cnf(c_0_129, negated_conjecture, (mult(a,mult(b,mult(a,c)))!=mult(mult(mult(a,b),a),c)), inference(split_conjunct,[status(thm)],[c_0_127])).
cnf(c_0_130, plain, (mult(X1,mult(ld(X1,X2),mult(X1,X3)))=mult(mult(X2,X1),X3)), inference(pm,[status(thm)],[c_0_128, c_0_28])).
cnf(c_0_131, negated_conjecture, (mult(mult(a,mult(b,a)),c)!=mult(a,mult(b,mult(a,c)))), inference(rw,[status(thm)],[c_0_129, c_0_13])).
cnf(c_0_132, plain, (mult(mult(X1,mult(X2,X1)),X3)=mult(X1,mult(X2,mult(X1,X3)))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_130, c_0_28]), c_0_13])).
cnf(c_0_133, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_131, c_0_132])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 134
# Proof object clause steps            : 113
# Proof object formula steps           : 21
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 10
# Proof object generating inferences   : 95
# Proof object simplifying inferences  : 93
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 13
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 13
# Processed clauses                    : 12995
# ...of these trivial                  : 7147
# ...subsumed                          : 3790
# ...remaining for further processing  : 2058
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1067
# Generated clauses                    : 954480
# ...of the previous two non-trivial   : 822693
# Contextual simplify-reflections      : 0
# Paramodulations                      : 954480
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 991
#    Positive orientable unit clauses  : 982
#    Positive unorientable unit clauses: 9
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 807848
# ...number of literals in the above   : 807851
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1067
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 79
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13502
# BW rewrite match successes           : 1009
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 19392140
