# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_2, plain, (ld(X1,mult(X1,X2))=X2)).
#
#cnf(i_0_1, plain, (mult(X1,ld(X1,X2))=X2)).
#
#cnf(i_0_3, plain, (mult(rd(X1,X2),X2)=X1)).
#
#cnf(i_0_6, negated_conjecture, (mult(X1,esk2_1(X1))!=esk2_1(X1)|mult(esk1_1(X1),X1)!=esk1_1(X1))).
#
#cnf(i_0_4, plain, (rd(mult(X1,X2),X2)=X1)).
#
#cnf(i_0_7, plain, (ld(rd(X1,X2),X1)=X2)).
#
#cnf(i_0_8, plain, (rd(X2,ld(X1,X2))=X1)).
#
#cnf(i_0_5, plain, (mult(mult(X1,mult(X2,X2)),X3)=mult(mult(X1,X2),mult(X2,X3)))).
#
#cnf(i_0_17, plain, (rd(mult(mult(X1,X2),mult(X2,X3)),X3)=mult(X1,mult(X2,X2)))).
#
#cnf(i_0_24, plain, (rd(mult(mult(X1,X2),X3),ld(X2,X3))=mult(X1,mult(X2,X2)))).
#
#cnf(i_0_36, plain, (rd(X3,ld(X2,ld(mult(X1,X2),X3)))=mult(X1,mult(X2,X2)))).
#
#cnf(i_0_13, negated_conjecture, (mult(mult(X1,X2),mult(X2,esk2_1(mult(X1,mult(X2,X2)))))!=esk2_1(mult(X1,mult(X2,X2)))|mult(esk1_1(mult(X1,mult(X2,X2))),mult(X1,mult(X2,X2)))!=esk1_1(mult(X1,mult(X2,X2))))).
#
#cnf(i_0_39, plain, (ld(mult(X3,mult(X2,X2)),X1)=ld(X2,ld(mult(X3,X2),X1)))).
#
#cnf(i_0_19, plain, (mult(mult(rd(X1,mult(X2,X2)),X2),mult(X2,X3))=mult(X1,X3))).
#
#cnf(i_0_78, plain, (mult(mult(rd(X1,mult(X2,X2)),X2),X3)=mult(X1,ld(X2,X3)))).
###
#cnf(i_0_59, plain, (ld(X2,ld(mult(rd(X1,mult(X2,X2)),X2),X3))=ld(X1,X3))).
#
#cnf(i_0_110, plain, (ld(mult(rd(X2,mult(X1,X1)),X1),X3)=mult(X1,ld(X2,X3)))).
#####
#cnf(i_0_48, negated_conjecture, (mult(esk1_1(mult(rd(X1,X2),mult(X2,X2))),mult(rd(X1,X2),mult(X2,X2)))!=esk1_1(mult(rd(X1,X2),mult(X2,X2)))|mult(X1,mult(X2,esk2_1(mult(rd(X1,X2),mult(X2,X2)))))!=esk2_1(mult(rd(X1,X2),mult(X2,X2))))).
###
#cnf(i_0_25, plain, (rd(mult(X1,mult(X2,X3)),X3)=mult(rd(X1,X2),mult(X2,X2)))).
##
#cnf(i_0_152, plain, (mult(rd(mult(X1,mult(X2,X4)),X4),X3)=mult(X1,mult(X2,X3)))).
#
#cnf(i_0_211, plain, (mult(rd(mult(X1,X2),X3),X4)=mult(X1,mult(rd(X2,X3),X4)))).
#
#cnf(i_0_261, plain, (mult(X1,mult(rd(X2,X2),X3))=mult(X1,X3))).
##
#cnf(i_0_276, plain, (mult(rd(X2,X2),X3)=X3)).
#
#cnf(i_0_306, plain, (ld(rd(X1,X1),X2)=X2)).
##
#cnf(i_0_301, negated_conjecture, (mult(esk1_1(rd(X1,X1)),rd(X1,X1))!=esk1_1(rd(X1,X1)))).
#
#cnf(i_0_308, plain, (rd(X2,X2)=rd(X1,X1))).
###
#cnf(i_0_346, negated_conjecture, (mult(esk1_1(rd(X2,X2)),rd(X1,X1))!=esk1_1(rd(X1,X1)))).
####
#cnf(i_0_347, negated_conjecture, (mult(esk1_1(rd(X1,X1)),rd(X2,X2))!=esk1_1(rd(X1,X1)))).
#
#cnf(i_0_302, plain, (mult(mult(X2,X2),X3)=mult(X2,mult(X2,X3)))).
#
#cnf(i_0_303, plain, (ld(mult(X2,X2),X3)=ld(X2,ld(X2,X3)))).
#
#cnf(i_0_310, plain, (rd(mult(X2,mult(X2,X3)),X3)=mult(X2,X2))).
####
#cnf(i_0_311, plain, (rd(mult(X2,X3),ld(X2,X3))=mult(X2,X2))).
######
#cnf(i_0_323, plain, (rd(X1,ld(X2,ld(X2,X1)))=mult(X2,X2))).
#####################
#cnf(i_0_305, plain, (mult(rd(rd(X1,X1),X2),mult(X2,X2))=X2)).
###
#cnf(i_0_474, plain, (ld(rd(rd(X1,X1),X2),X2)=mult(X2,X2))).
#
#cnf(i_0_327, plain, (mult(X1,mult(rd(X2,mult(X1,X2)),X3))=X3)).
#
#cnf(i_0_526, plain, (mult(rd(X2,mult(X1,X2)),X3)=ld(X1,X3))).
##
#cnf(i_0_592, plain, (mult(ld(X1,X1),X2)=X2)).
#
#cnf(i_0_593, plain, (ld(ld(X1,X1),X2)=X2)).
##
#cnf(i_0_604, negated_conjecture, (mult(esk1_1(ld(X1,X1)),ld(X1,X1))!=esk1_1(ld(X1,X1)))).
###
#cnf(i_0_609, plain, (rd(X2,X2)=ld(X1,X1))).
##
#cnf(i_0_668, plain, (rd(X1,rd(X2,X2))=X1)).
#
#cnf(i_0_670, plain, (mult(X1,rd(X2,X2))=X1)).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(f04, axiom, ![X1, X2]:rd(mult(X2,X1),X1)=X2, file('./p/GRP658+1.tptp', f04)).
fof(f05, axiom, ![X3, X1, X2]:mult(mult(X2,mult(X1,X1)),X3)=mult(mult(X2,X1),mult(X1,X3)), file('./p/GRP658+1.tptp', f05)).
fof(f03, axiom, ![X1, X2]:mult(rd(X2,X1),X1)=X2, file('./p/GRP658+1.tptp', f03)).
fof(f02, axiom, ![X1, X2]:ld(X2,mult(X2,X1))=X1, file('./p/GRP658+1.tptp', f02)).
fof(f01, axiom, ![X1, X2]:mult(X2,ld(X2,X1))=X1, file('./p/GRP658+1.tptp', f01)).
fof(goals, conjecture, ?[X4]:![X5]:(mult(X5,X4)=X5&mult(X4,X5)=X5), file('./p/GRP658+1.tptp', goals)).
fof(c_0_6, plain, ![X12, X13]:rd(mult(X13,X12),X12)=X13, inference(variable_rename,[status(thm)],[f04])).
fof(c_0_7, plain, ![X14, X15, X16]:mult(mult(X16,mult(X15,X15)),X14)=mult(mult(X16,X15),mult(X15,X14)), inference(variable_rename,[status(thm)],[f05])).
cnf(c_0_8, plain, (rd(mult(X1,X2),X2)=X1), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, plain, (mult(mult(X1,mult(X2,X2)),X3)=mult(mult(X1,X2),mult(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_10, plain, ![X10, X11]:mult(rd(X11,X10),X10)=X11, inference(variable_rename,[status(thm)],[f03])).
cnf(c_0_11, plain, (rd(mult(mult(X1,X2),mult(X2,X3)),X3)=mult(X1,mult(X2,X2))), inference(pm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12, plain, (mult(rd(X1,X2),X2)=X1), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_13, plain, (rd(mult(X1,mult(X2,X3)),X3)=mult(rd(X1,X2),mult(X2,X2))), inference(pm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_14, plain, (mult(rd(mult(X1,mult(X2,X3)),X3),X4)=mult(X1,mult(X2,X4))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_9, c_0_13]), c_0_12])).
fof(c_0_15, plain, ![X8, X9]:ld(X9,mult(X9,X8))=X8, inference(variable_rename,[status(thm)],[f02])).
cnf(c_0_16, plain, (mult(rd(mult(X1,X2),X3),X4)=mult(X1,mult(rd(X2,X3),X4))), inference(pm,[status(thm)],[c_0_14, c_0_12])).
cnf(c_0_17, plain, (ld(X1,mult(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_18, plain, (mult(X1,mult(rd(X2,X2),X3))=mult(X1,X3)), inference(pm,[status(thm)],[c_0_16, c_0_8])).
fof(c_0_19, plain, ![X6, X7]:mult(X7,ld(X7,X6))=X6, inference(variable_rename,[status(thm)],[f01])).
cnf(c_0_20, plain, (mult(rd(X1,X1),X2)=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_17, c_0_18]), c_0_17])).
fof(c_0_21, negated_conjecture, ~(?[X4]:![X5]:(mult(X5,X4)=X5&mult(X4,X5)=X5)), inference(assume_negation,[status(cth)],[goals])).
cnf(c_0_22, plain, (mult(mult(rd(X1,mult(X2,X2)),X2),mult(X2,X3))=mult(X1,X3)), inference(pm,[status(thm)],[c_0_9, c_0_12])).
cnf(c_0_23, plain, (mult(X1,ld(X1,X2))=X2), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_24, plain, (mult(X1,mult(rd(X2,mult(X1,X2)),X3))=X3), inference(pm,[status(thm)],[c_0_16, c_0_20])).
fof(c_0_25, negated_conjecture, ![X17]:(mult(esk1_1(X17),X17)!=esk1_1(X17)|mult(X17,esk2_1(X17))!=esk2_1(X17)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_21])])])])).
cnf(c_0_26, plain, (mult(mult(rd(X1,mult(X2,X2)),X2),X3)=mult(X1,ld(X2,X3))), inference(pm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_27, plain, (mult(rd(X1,mult(X2,X1)),X3)=ld(X2,X3)), inference(pm,[status(thm)],[c_0_17, c_0_24])).
cnf(c_0_28, negated_conjecture, (mult(esk1_1(X1),X1)!=esk1_1(X1)|mult(X1,esk2_1(X1))!=esk2_1(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_29, plain, (rd(mult(X1,mult(X1,X2)),X2)=mult(X1,X1)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_11, c_0_20]), c_0_20])).
cnf(c_0_30, plain, (mult(ld(X1,X1),X2)=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_26, c_0_27]), c_0_23])).
cnf(c_0_31, negated_conjecture, (mult(esk1_1(rd(X1,X1)),rd(X1,X1))!=esk1_1(rd(X1,X1))), inference(pm,[status(thm)],[c_0_28, c_0_20])).
cnf(c_0_32, plain, (rd(X1,X1)=rd(X2,X2)), inference(pm,[status(thm)],[c_0_8, c_0_20])).
cnf(c_0_33, plain, (rd(X1,X1)=ld(X2,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_29, c_0_30]), c_0_30]), c_0_30])).
cnf(c_0_34, negated_conjecture, (mult(esk1_1(rd(X1,X1)),rd(X2,X2))!=esk1_1(rd(X1,X1))), inference(pm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_35, plain, (mult(X1,rd(X2,X2))=X1), inference(pm,[status(thm)],[c_0_23, c_0_33])).
cnf(c_0_36, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_34, c_0_35])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 37
# Proof object clause steps            : 24
# Proof object formula steps           : 13
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 6
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 6
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 6
# Processed clauses                    : 94
# ...of these trivial                  : 24
# ...subsumed                          : 27
# ...remaining for further processing  : 43
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 8
# Generated clauses                    : 761
# ...of the previous two non-trivial   : 577
# Contextual simplify-reflections      : 0
# Paramodulations                      : 761
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 35
#    Positive orientable unit clauses  : 28
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 3
# Current number of unprocessed clauses: 481
# ...number of literals in the above   : 584
# Current number of archived formulas  : 0
# Current number of archived clauses   : 8
# Clause-clause subsumption calls (NU) : 13
# Rec. Clause-clause subsumption calls : 11
# Non-unit clause-clause subsumptions  : 8
# Unit Clause-clause subsumption calls : 31
# Rewrite failures with RHS unbound    : 74
# BW rewrite match attempts            : 93
# BW rewrite match successes           : 14
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 11410
