# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_15, plain, (less_equal(X1,identity))).
#
#cnf(i_0_13, plain, (less_equal(zero,X1))).
#
#cnf(i_0_11, plain, (less_equal(divide(X1,X2),X1))).
#
#cnf(i_0_16, negated_conjecture, (~less_equal(a,a))).
#
#cnf(i_0_10, plain, (less_equal(X1,X2)|divide(X1,X2)!=zero)).
#
#cnf(i_0_9, plain, (divide(X1,X2)=zero|~less_equal(X1,X2))).
#
#cnf(i_0_14, plain, (X1=X2|~less_equal(X2,X1)|~less_equal(X1,X2))).
#
#cnf(i_0_19, plain, (zero=X1|~less_equal(X1,zero))).
#
#cnf(i_0_25, plain, (divide(zero,X1)=zero)).
#
#cnf(i_0_21, plain, (X1=identity|~less_equal(identity,X1))).
###
#cnf(i_0_20, plain, (divide(X1,X2)=X1|~less_equal(X1,divide(X1,X2)))).
##
#cnf(i_0_29, plain, (divide(X1,X2)=X1|~less_equal(X1,zero)|~less_equal(X1,X2))).
##
#cnf(i_0_33, plain, (less_equal(X1,X1)|~less_equal(X1,zero)|~less_equal(X1,X2))).
###
#cnf(i_0_35, negated_conjecture, (~less_equal(a,zero)|~less_equal(a,X1))).
#
#cnf(i_0_12, plain, (less_equal(divide(divide(X1,X2),divide(X3,X2)),divide(divide(X1,X3),X2)))).
#
#cnf(i_0_43, plain, (less_equal(divide(divide(X1,X2),divide(X3,X2)),zero)|~less_equal(X1,X3))).
#
#cnf(i_0_55, plain, (divide(divide(X1,X2),divide(X3,X2))=zero|~less_equal(X1,X3))).
##
#cnf(i_0_66, plain, (less_equal(divide(X1,X2),divide(X3,X2))|~less_equal(X1,X3))).
#
#cnf(i_0_90, plain, (less_equal(divide(X1,X2),zero)|~less_equal(X1,zero))).
#
#cnf(i_0_98, plain, (divide(X1,X2)=zero|~less_equal(X1,zero))).
#
#cnf(i_0_108, plain, (less_equal(X1,X2)|~less_equal(X1,zero))).
#
#cnf(i_0_124, plain, (X1=identity|~less_equal(identity,zero))).
#
#cnf(i_0_134, plain, (X2=X1|~less_equal(identity,zero))).
#
#cnf(i_0_161, plain, (less_equal(X1,X2)|~less_equal(identity,zero))).
#
#cnf(i_0_122, negated_conjecture, (~less_equal(a,zero))).
########
#cnf(i_0_229, negated_conjecture, (~less_equal(identity,zero))).
######
#cnf(i_0_111, plain, (divide(X1,X2)=X1|~less_equal(X1,zero))).
#########
#cnf(i_0_131, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X1,zero))).
######
#cnf(i_0_126, plain, (X1=X2|~less_equal(X1,zero)|~less_equal(X2,X1))).
##################
#cnf(i_0_289, plain, (X1=X2|~less_equal(X1,zero)|~less_equal(X2,zero))).
##################################################
#cnf(i_0_268, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X3,zero)|~less_equal(X1,X3))).
###
#cnf(i_0_80, plain, (divide(divide(X1,X2),zero)=zero|~less_equal(X1,zero))).
###
#cnf(i_0_362, plain, (X1=divide(X2,X3)|~less_equal(X1,zero)|~less_equal(X2,X1))).
######
#cnf(i_0_453, plain, (divide(divide(X1,X2),X3)=X1|~less_equal(X1,zero))).
########
#cnf(i_0_410, plain, (divide(X1,X2)=X1|~less_equal(X3,zero)|~less_equal(X1,X3))).
##
#cnf(i_0_510, plain, (less_equal(X1,X1)|~less_equal(X3,zero)|~less_equal(X1,X3))).
#
#cnf(i_0_542, negated_conjecture, (~less_equal(X1,zero)|~less_equal(a,X1))).
###############
#cnf(i_0_563, plain, (less_equal(X1,X1)|~less_equal(X1,divide(X2,X3))|~less_equal(X2,zero))).
#########
#cnf(i_0_576, negated_conjecture, (~less_equal(a,divide(X1,X2))|~less_equal(X1,zero))).
###########
#cnf(i_0_287, plain, (X1=divide(X2,X3)|~less_equal(X1,zero)|~less_equal(X2,zero))).
###################
#cnf(i_0_272, plain, (divide(divide(X1,X2),X3)=divide(X1,X2)|~less_equal(X1,zero))).
#####
#cnf(i_0_733, plain, (divide(X1,X3)=divide(X1,X2)|~less_equal(X1,zero))).
######################################
#cnf(i_0_92, plain, (less_equal(divide(X1,X2),zero)|~less_equal(X1,X3)|~less_equal(X3,X2))).
#
#cnf(i_0_827, plain, (less_equal(divide(X1,identity),zero)|~less_equal(X1,X2))).
#
#cnf(i_0_863, plain, (divide(X1,identity)=zero|~less_equal(X1,X2))).
#
#cnf(i_0_873, negated_conjecture, (~less_equal(a,divide(X1,identity))|~less_equal(X1,X2))).
#
#cnf(i_0_941, plain, (divide(divide(X1,X2),identity)=zero)).
########
#cnf(i_0_868, plain, (less_equal(divide(X1,identity),X2)|~less_equal(X1,X3))).
##################
#cnf(i_0_1002, plain, (less_equal(divide(divide(X1,identity),zero),zero))).
####
#cnf(i_0_1092, plain, (divide(divide(X1,identity),zero)=zero)).
###
#cnf(i_0_1148, plain, (less_equal(divide(X1,identity),zero))).
#
#cnf(i_0_1201, plain, (divide(X1,identity)=zero)).
#########
#cnf(i_0_817, negated_conjecture, (~less_equal(a,divide(X1,X2))|~less_equal(X1,X3)|~less_equal(X3,X2))).
#####
#cnf(i_0_1253, plain, (zero=X1|~less_equal(X2,zero)|~less_equal(X1,X2))).
##############
#cnf(i_0_807, plain, (divide(X1,X2)=zero|~less_equal(X1,X3)|~less_equal(X3,X2))).
##
#cnf(i_0_1342, plain, (less_equal(X1,X2)|~less_equal(X1,X3)|~less_equal(X3,X2))).
##
#cnf(i_0_1451, plain, (less_equal(X1,X2)|~less_equal(identity,X2))).
#
#cnf(i_0_1456, negated_conjecture, (~less_equal(identity,a))).
######
#cnf(i_0_1435, plain, (less_equal(X1,X2)|~less_equal(X1,divide(X2,X3)))).
#####
#cnf(i_0_1527, plain, (less_equal(divide(divide(X1,X2),X3),X1))).
#########
#cnf(i_0_1391, negated_conjecture, (~less_equal(a,X1)|~less_equal(X1,a))).
#####
#cnf(i_0_1448, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X1,X3))).
###
#cnf(i_0_1608, negated_conjecture, (~less_equal(a,divide(a,X1)))).
######
#cnf(i_0_1476, plain, (X1=X2|~less_equal(identity,X2)|~less_equal(X2,X1))).
#######
#cnf(i_0_1535, plain, (less_equal(X1,X2)|~less_equal(identity,divide(X2,X3)))).
#######
#cnf(i_0_1698, negated_conjecture, (~less_equal(identity,divide(a,X1)))).
#####
#cnf(i_0_1401, plain, (X1=identity|~less_equal(identity,X2)|~less_equal(X2,X1))).
####
#cnf(i_0_1806, plain, (X1=identity|~less_equal(identity,divide(X1,X2)))).
###########
#cnf(i_0_1400, negated_conjecture, (~less_equal(identity,X1)|~less_equal(X1,zero))).
##
#cnf(i_0_1403, plain, (less_equal(X1,X2)|~less_equal(X3,zero)|~less_equal(X1,X3))).
############
#cnf(i_0_1486, plain, (less_equal(X1,X2)|~less_equal(identity,X3)|~less_equal(X3,X2))).
######
#cnf(i_0_1504, negated_conjecture, (~less_equal(identity,X1)|~less_equal(X1,a))).
#########
#cnf(i_0_1373, plain, (divide(X1,X2)=zero|~less_equal(X1,divide(X2,X3)))).
##########
#cnf(i_0_1712, plain, (X1=X2|~less_equal(identity,X2)|~less_equal(identity,X1))).
#######
#cnf(i_0_2075, plain, (divide(divide(divide(X1,X2),X3),X1)=zero)).
#####
#cnf(i_0_1474, plain, (divide(X1,X2)=X1|~less_equal(identity,divide(X1,X2)))).
######
#cnf(i_0_1563, plain, (less_equal(X1,X2)|~less_equal(X1,divide(divide(X2,X3),X4)))).
##############
#cnf(i_0_1384, plain, (divide(divide(X1,X2),X3)=zero|~less_equal(X1,X3))).
#######
#cnf(i_0_1549, plain, (less_equal(divide(divide(divide(X1,X2),X3),X4),X1))).
####################################################################
#cnf(i_0_1559, plain, (less_equal(divide(divide(X1,X2),X3),X4)|~less_equal(X1,X4))).
#################
#cnf(i_0_1638, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X1,divide(X3,X4)))).
##################
#cnf(i_0_1551, plain, (divide(divide(X1,X2),X3)=zero|~less_equal(X1,zero))).
#####################
#cnf(i_0_1415, plain, (X1=X2|~less_equal(X2,X1)|~less_equal(X1,X3)|~less_equal(X3,X2))).
#########
#cnf(i_0_2816, plain, (X1=X2|~less_equal(X1,divide(X2,X3))|~less_equal(X2,X1))).
#####
#cnf(i_0_1607, negated_conjecture, (~less_equal(a,divide(divide(a,X1),X2)))).
#####
#cnf(i_0_1735, plain, (X1=X2|~less_equal(identity,divide(X2,X3))|~less_equal(X2,X1))).
#########
#cnf(i_0_1805, plain, (X1=identity|~less_equal(identity,divide(divide(X1,X2),X3)))).
#################
#cnf(i_0_2084, plain, (divide(X1,X2)=zero|~less_equal(identity,divide(X2,X3)))).
##################
#cnf(i_0_1534, plain, (less_equal(X1,X2)|~less_equal(X4,divide(X2,X3))|~less_equal(X1,X4))).
#######
#cnf(i_0_2034, negated_conjecture, (~less_equal(identity,divide(divide(a,X1),X2)))).
#######
#cnf(i_0_1653, plain, (less_equal(X1,X2)|~less_equal(X1,divide(X3,X4))|~less_equal(X3,X2))).
###########
#cnf(i_0_1734, plain, (less_equal(X1,X2)|~less_equal(identity,divide(divide(X2,X3),X4)))).
##################################################
#cnf(i_0_1459, negated_conjecture, (~less_equal(identity,divide(X1,X2))|~less_equal(X1,zero))).
#
#cnf(i_0_1927, plain, (less_equal(divide(divide(X1,X2),X3),X4)|~less_equal(X1,zero))).
#############################
#cnf(i_0_812, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X1,X4)|~less_equal(X4,X2))).
#
#cnf(i_0_3754, plain, (less_equal(divide(divide(X1,X2),X3),X4)|~less_equal(X1,X3))).
#
#cnf(i_0_1624, negated_conjecture, (~less_equal(a,divide(X1,X2))|~less_equal(X1,a))).
##########
#cnf(i_0_3731, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X1,divide(X2,X4)))).
##############
#cnf(i_0_1697, negated_conjecture, (~less_equal(X2,divide(a,X1))|~less_equal(a,X2))).
#
#cnf(i_0_1649, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X4,X3)|~less_equal(X1,X4))).
##########
#cnf(i_0_4018, plain, (less_equal(divide(X1,X2),X3)|~less_equal(identity,divide(X2,X4)))).
#############
#cnf(i_0_1799, negated_conjecture, (~less_equal(X2,divide(a,X1))|~less_equal(identity,X2))).
######
#cnf(i_0_1331, plain, (zero=X1|~less_equal(X1,divide(X2,X3))|~less_equal(X2,zero))).
############
#cnf(i_0_1488, plain, (divide(X1,X2)=zero|~less_equal(identity,X2)|~less_equal(X1,X3))).
###
#cnf(i_0_4519, plain, (divide(divide(X1,X2),X3)=zero|~less_equal(identity,X3))).
#########
#cnf(i_0_1848, plain, (X1=identity|~less_equal(X3,divide(X1,X2))|~less_equal(identity,X3))).
######################
#cnf(i_0_2037, negated_conjecture, (~less_equal(identity,divide(X1,X2))|~less_equal(X1,a))).
######
#cnf(i_0_3993, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X2,zero)|~less_equal(X1,X2))).
#######
#cnf(i_0_1436, plain, (less_equal(X1,X2)|~less_equal(X1,divide(X3,X4))|~less_equal(X3,zero))).
##########
#cnf(i_0_1776, plain, (less_equal(X1,X2)|~less_equal(X4,divide(X2,X3))|~less_equal(identity,X4))).
######################################################
#cnf(i_0_3740, plain, (less_equal(divide(X1,X2),X3)|~less_equal(identity,X2)|~less_equal(X1,X4))).
#
#cnf(i_0_5116, plain, (less_equal(divide(divide(X1,X2),X3),X4)|~less_equal(identity,X3))).
#######################
#cnf(i_0_1377, plain, (divide(X1,X2)=zero|~less_equal(X3,zero)|~less_equal(X1,X3))).
########
#cnf(i_0_1404, plain, (X1=X2|~less_equal(X3,zero)|~less_equal(X2,X1)|~less_equal(X1,X3))).
#############
#cnf(i_0_1484, plain, (divide(X1,X2)=zero|~less_equal(identity,X3)|~less_equal(X3,X2))).
#####
#cnf(i_0_1718, plain, (X1=X2|~less_equal(identity,X3)|~less_equal(X2,X1)|~less_equal(X3,X2))).
##########
#cnf(i_0_1808, plain, (X1=identity|~less_equal(identity,divide(X2,X3))|~less_equal(X2,X1))).
###
#cnf(i_0_1397, negated_conjecture, (~less_equal(X1,zero)|~less_equal(a,X2)|~less_equal(X2,X1))).
###############################
#cnf(i_0_1425, plain, (less_equal(X1,X2)|~less_equal(X3,X2)|~less_equal(X1,X4)|~less_equal(X4,X3))).
#####
#cnf(i_0_1613, negated_conjecture, (~less_equal(a,X1)|~less_equal(X2,a)|~less_equal(X1,X2))).
#######
#cnf(i_0_1746, plain, (less_equal(X1,X2)|~less_equal(identity,divide(X3,X4))|~less_equal(X3,X2))).
#########################
#cnf(i_0_1859, negated_conjecture, (~less_equal(identity,X1)|~less_equal(X2,zero)|~less_equal(X1,X2))).
#################
#cnf(i_0_1262, plain, (less_equal(divide(divide(X1,X2),divide(identity,X2)),zero))).
######
#cnf(i_0_6410, plain, (divide(divide(X1,X2),divide(identity,X2))=zero)).
########
#cnf(i_0_6538, plain, (less_equal(divide(X1,X2),divide(identity,X2)))).
##################
#cnf(i_0_6739, plain, (less_equal(divide(X1,X2),zero)|~less_equal(identity,X2))).
#
#cnf(i_0_6782, plain, (divide(X1,X2)=zero|~less_equal(identity,X2))).
########
#cnf(i_0_6794, plain, (less_equal(divide(X1,X2),X3)|~less_equal(identity,X2))).
##################
#cnf(i_0_6954, negated_conjecture, (~less_equal(a,divide(X1,X2))|~less_equal(identity,X2))).
##############
#cnf(i_0_6978, plain, (divide(divide(X1,X2),X3)=zero|~less_equal(identity,X2))).
################
#cnf(i_0_6982, plain, (less_equal(divide(divide(X1,X2),X3),X4)|~less_equal(identity,X2))).
######
#cnf(i_0_6955, negated_conjecture, (~less_equal(identity,divide(X1,X2))|~less_equal(identity,X2))).
######
#cnf(i_0_6683, plain, (divide(identity,X1)=identity|~less_equal(identity,divide(X2,X1)))).
#########
#cnf(i_0_6695, plain, (less_equal(divide(X1,X2),X3)|~less_equal(divide(identity,X2),X3))).
############
#cnf(i_0_7547, plain, (less_equal(divide(X1,X2),divide(X3,X2))|~less_equal(identity,X3))).
######################
#cnf(i_0_6703, plain, (less_equal(X1,divide(identity,X2))|~less_equal(X1,divide(X3,X2)))).
######
#cnf(i_0_7818, plain, (less_equal(divide(divide(X1,X2),X3),divide(identity,X2)))).
###############
#cnf(i_0_6971, plain, (zero=X1|~less_equal(X1,divide(X2,X3))|~less_equal(identity,X3))).
#################
#cnf(i_0_6967, plain, (less_equal(X1,X2)|~less_equal(X1,divide(X3,X4))|~less_equal(identity,X4))).
###################################
#cnf(i_0_6677, plain, (divide(X1,X2)=zero|~less_equal(divide(identity,X2),zero))).
######
#cnf(i_0_1959, negated_conjecture, (~less_equal(X1,a)|~less_equal(identity,X2)|~less_equal(X2,X1))).
###########
#cnf(i_0_8259, plain, (less_equal(X1,X2)|~less_equal(divide(identity,X2),zero))).
########################
#cnf(i_0_8469, plain, (less_equal(X1,X2)|~less_equal(divide(identity,zero),zero))).
#
#cnf(i_0_8595, negated_conjecture, (~less_equal(divide(identity,zero),zero))).
#######################################################################
#cnf(i_0_8509, plain, (X1=X2|~less_equal(divide(identity,X2),zero)|~less_equal(X2,X1))).
################################
#cnf(i_0_8498, plain, (less_equal(X1,X2)|~less_equal(divide(identity,divide(X2,X3)),zero))).
#################
#cnf(i_0_8789, negated_conjecture, (~less_equal(divide(identity,zero),X1)|~less_equal(X1,zero))).
##############################################
#cnf(i_0_6961, plain, (divide(X1,X2)=X3|~less_equal(X3,zero)|~less_equal(identity,X2))).
##################################
#cnf(i_0_9098, negated_conjecture, (~less_equal(divide(identity,X1),zero)|~less_equal(X1,zero))).
###
#cnf(i_0_1413, plain, (divide(X1,X2)=X1|~less_equal(X3,divide(X1,X2))|~less_equal(X1,X3))).
#######
#cnf(i_0_9504, plain, (divide(X1,X2)=X1|~less_equal(X1,divide(divide(X1,X2),X3)))).
#############################
#cnf(i_0_9518, plain, (divide(X1,X2)=X1|~less_equal(X1,divide(X3,X2))|~less_equal(X3,X1))).
###########
#cnf(i_0_6976, plain, (divide(divide(X1,X2),X3)=divide(X1,X2)|~less_equal(identity,X2))).
##################
#cnf(i_0_9101, negated_conjecture, (~less_equal(divide(X1,zero),zero)|~less_equal(identity,X1))).
##################
#cnf(i_0_1419, plain, (X1=X2|~less_equal(X1,zero)|~less_equal(X2,X3)|~less_equal(X3,X1))).
#######################
#cnf(i_0_10229, plain, (X1=X2|~less_equal(X2,divide(X1,X3))|~less_equal(X1,zero))).
########################
#cnf(i_0_1561, plain, (divide(X1,X2)=zero|~less_equal(X1,divide(divide(X2,X3),X4)))).
########################
#cnf(i_0_1711, plain, (X1=X2|~less_equal(identity,X2)|~less_equal(X2,X3)|~less_equal(X3,X1))).
#############
#cnf(i_0_10831, plain, (X1=X2|~less_equal(X2,divide(X1,X3))|~less_equal(identity,X2))).
###########
#cnf(i_0_1741, plain, (X1=X2|~less_equal(identity,divide(X1,X3))|~less_equal(identity,X2))).
#################
#cnf(i_0_1930, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X4,zero)|~less_equal(X1,X4))).
###############
#cnf(i_0_3762, plain, (less_equal(divide(X1,X2),X3)|~less_equal(identity,X4)|~less_equal(X4,X2))).
##################
#cnf(i_0_6694, plain, (less_equal(divide(X1,X2),X3)|~less_equal(divide(identity,X2),zero))).
############################################################################################################
#cnf(i_0_6707, plain, (less_equal(X1,divide(identity,X2))|~less_equal(identity,divide(X3,X2)))).
############################################################################################################
#cnf(i_0_8576, plain, (less_equal(X1,X2)|~less_equal(divide(identity,X2),X3)|~less_equal(X3,zero))).
###########################
#cnf(i_0_12311, plain, (less_equal(X1,X2)|~less_equal(divide(X3,X2),zero)|~less_equal(identity,X3))).
##############################################
#cnf(i_0_1284, negated_conjecture, (~less_equal(a,divide(divide(X1,X2),X3))|~less_equal(X1,X3))).
#########
#cnf(i_0_1639, plain, (divide(X1,X2)=X3|~less_equal(X3,divide(X1,X2))|~less_equal(X1,X3))).
##########
#cnf(i_0_2813, plain, (X1=X2|~less_equal(X1,divide(divide(X2,X3),X4))|~less_equal(X2,X1))).
####################
#cnf(i_0_1732, plain, (divide(X1,X2)=X1|~less_equal(identity,divide(divide(X1,X2),X3)))).
########################
#cnf(i_0_2475, plain, (less_equal(X1,X2)|~less_equal(X1,divide(divide(divide(X2,X3),X4),X5)))).
#####
#cnf(i_0_3796, negated_conjecture, (~less_equal(identity,divide(divide(X1,X2),X3))|~less_equal(X1,X3))).
############################################################################
#cnf(i_0_1273, negated_conjecture, (~less_equal(a,divide(X1,X2))|~less_equal(X1,divide(X2,X3)))).
###############################################
#cnf(i_0_2074, plain, (divide(divide(divide(divide(X1,X2),X3),X4),X1)=zero)).
######################################################################################
#cnf(i_0_1412, plain, (zero=X1|~less_equal(X3,zero)|~less_equal(X1,X2)|~less_equal(X2,X3))).
#################
#cnf(i_0_1811, plain, (X1=identity|~less_equal(identity,X2)|~less_equal(X2,X3)|~less_equal(X3,X1))).
##################################
#cnf(i_0_2442, negated_conjecture, (~less_equal(a,divide(divide(divide(a,X1),X2),X3)))).
##########################################################################################################################
#cnf(i_0_1899, plain, (less_equal(X1,X2)|~less_equal(X4,zero)|~less_equal(X1,X3)|~less_equal(X3,X4))).
####################
#cnf(i_0_1969, plain, (less_equal(X1,X2)|~less_equal(identity,X4)|~less_equal(X3,X2)|~less_equal(X4,X3))).
##############################################################
#cnf(i_0_2443, negated_conjecture, (~less_equal(identity,divide(divide(divide(a,X1),X2),X3)))).
##
#cnf(i_0_8529, plain, (less_equal(X1,X2)|~less_equal(divide(identity,X3),zero)|~less_equal(X3,X2))).
########################
#cnf(i_0_14698, plain, (X1=identity|~less_equal(divide(identity,X2),zero)|~less_equal(X2,X1))).
####################
#cnf(i_0_14962, plain, (X1=identity|~less_equal(divide(identity,divide(X1,X2)),zero))).
#################################################################################################
#cnf(i_0_1986, plain, (divide(X1,X2)=X1|~less_equal(X3,divide(X1,X2))|~less_equal(identity,X3))).
###############
#cnf(i_0_9519, plain, (divide(X1,X2)=X1|~less_equal(X1,divide(X3,X2))|~less_equal(identity,X1))).
#############################
#cnf(i_0_1557, plain, (divide(divide(divide(X1,X2),X3),X4)=zero|~less_equal(X1,X4))).
#########
#cnf(i_0_3929, negated_conjecture, (~less_equal(identity,divide(X1,X2))|~less_equal(X1,divide(X2,X3)))).
######
#cnf(i_0_3812, plain, (divide(divide(divide(X1,X2),X3),X4)=zero|~less_equal(X1,X3))).
################
#cnf(i_0_1530, plain, (less_equal(divide(divide(X1,X2),divide(X3,X2)),divide(X1,X3)))).
########
#cnf(i_0_16080, plain, (less_equal(divide(divide(X1,X2),zero),divide(X1,zero)))).
########################
#cnf(i_0_2362, plain, (less_equal(divide(divide(divide(divide(X1,X2),X3),X4),X5),X1))).
#######################################
#cnf(i_0_2077, plain, (divide(divide(X1,X2),X3)=zero|~less_equal(X1,divide(X3,X4)))).
##############
#cnf(i_0_3945, plain, (divide(divide(X1,X2),X3)=zero|~less_equal(X1,divide(X2,X5)))).
###################
#cnf(i_0_1406, plain, (X1=X2|~less_equal(X2,zero)|~less_equal(X3,zero)|~less_equal(X1,X3))).
########################################
#cnf(i_0_1995, plain, (X1=X2|~less_equal(identity,X2)|~less_equal(identity,X3)|~less_equal(X3,X1))).
########################
#cnf(i_0_8521, plain, (X1=X2|~less_equal(divide(identity,X1),zero)|~less_equal(identity,X2))).
########################################################
#cnf(i_0_2365, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X1,divide(divide(X3,X4),X5)))).
########################################################
#cnf(i_0_3727, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X1,divide(divide(X2,X4),X5)))).
###############################################################################
#cnf(i_0_1554, plain, (divide(divide(divide(X1,X2),X3),X4)=X1|~less_equal(X1,zero))).
###############
#cnf(i_0_8472, negated_conjecture, (~less_equal(divide(identity,a),zero)|~less_equal(a,X1))).
########################################################
#cnf(i_0_8473, negated_conjecture, (~less_equal(divide(identity,a),zero)|~less_equal(identity,X1))).
#####################################################################################
#cnf(i_0_273, plain, (divide(X1,X2)=X3|~less_equal(X3,divide(X1,X2))|~less_equal(X1,zero))).
#####
#cnf(i_0_19985, plain, (divide(X1,X2)=X3|~less_equal(X1,zero)|~less_equal(X3,X1))).
########################################
#cnf(i_0_20443, plain, (X4=X3|~less_equal(X1,zero)|~less_equal(X3,X1)|~less_equal(X4,X1))).
###########################################################################
#cnf(i_0_6959, plain, (divide(X1,X2)=X3|~less_equal(X3,divide(X1,X2))|~less_equal(identity,X2))).
########
#cnf(i_0_21239, plain, (divide(identity,X1)=divide(X2,X1)|~less_equal(identity,X1))).
#################
#cnf(i_0_21643, plain, (divide(X3,X1)=divide(X2,X1)|~less_equal(identity,X1))).
#######################################################################################################################################
#cnf(i_0_21238, plain, (divide(divide(X2,X1),X3)=divide(identity,X1)|~less_equal(identity,X1))).
################################################
#cnf(i_0_21740, plain, (divide(X1,X2)=X1|~less_equal(X1,divide(X3,X2))|~less_equal(identity,X2))).
############################
#cnf(i_0_10223, plain, (X1=X2|~less_equal(X2,divide(divide(X1,X3),X4))|~less_equal(X1,zero))).
############################
#cnf(i_0_10825, plain, (X1=X2|~less_equal(X2,divide(divide(X1,X3),X4))|~less_equal(identity,X2))).
############################################
#cnf(i_0_6692, plain, (divide(divide(X1,X2),X3)=zero|~less_equal(divide(identity,X2),X3))).
#############
#cnf(i_0_130, plain, (less_equal(divide(divide(X1,X2),divide(X3,X2)),X4)|~less_equal(X1,X3))).
#############
#cnf(i_0_2471, plain, (less_equal(divide(divide(divide(X1,X2),X3),X4),X5)|~less_equal(X1,X5))).
##############
#cnf(i_0_3750, plain, (less_equal(divide(divide(divide(X1,X2),X3),X4),X5)|~less_equal(X1,X4))).
##############
#cnf(i_0_3815, plain, (less_equal(divide(divide(divide(X1,X2),X3),X4),X5)|~less_equal(X1,X3))).
################################################################
#cnf(i_0_3808, plain, (zero=X1|~less_equal(X1,divide(divide(X2,X3),X4))|~less_equal(X2,X4))).
#######################
#cnf(i_0_24447, plain, (divide(divide(X1,X2),divide(X3,X2))=zero|~less_equal(X1,X2))).
#########
#cnf(i_0_24473, plain, (less_equal(divide(X1,X2),divide(X3,X2))|~less_equal(X1,X2))).
##
#cnf(i_0_24852, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X1,X2))).
#########################
#cnf(i_0_25047, plain, (divide(divide(X1,X2),X3)=zero|~less_equal(X1,X2))).
################
#cnf(i_0_25051, plain, (less_equal(divide(divide(X1,X2),X3),X4)|~less_equal(X1,X2))).
######
#cnf(i_0_25017, negated_conjecture, (~less_equal(a,divide(X1,X2))|~less_equal(X1,X2))).
###########
#cnf(i_0_25039, plain, (zero=X1|~less_equal(X1,divide(X2,X3))|~less_equal(X2,X3))).
##################
#cnf(i_0_25031, plain, (less_equal(X1,X2)|~less_equal(X1,divide(X3,X4))|~less_equal(X3,X4))).
#######
#cnf(i_0_25018, negated_conjecture, (~less_equal(identity,divide(X1,X2))|~less_equal(X1,X2))).
#########################
#cnf(i_0_25024, plain, (divide(X1,X2)=X3|~less_equal(X3,zero)|~less_equal(X1,X2))).
##################################
#cnf(i_0_25046, plain, (divide(divide(X1,X2),X3)=divide(X1,X2)|~less_equal(X1,X2))).
#################################################################################################################
#cnf(i_0_25066, negated_conjecture, (~less_equal(a,divide(divide(X1,X2),X3))|~less_equal(X1,X2))).
######################
#cnf(i_0_25022, plain, (divide(X1,X2)=X3|~less_equal(X3,divide(X1,X2))|~less_equal(X1,X2))).
############################
#cnf(i_0_25030, plain, (divide(X1,X2)=X1|~less_equal(X1,divide(X3,X4))|~less_equal(X3,X4))).
##################################
#cnf(i_0_25071, negated_conjecture, (~less_equal(identity,divide(divide(X1,X2),X3))|~less_equal(X1,X2))).
##################
#cnf(i_0_26854, plain, (divide(X1,X2)=divide(X3,X4)|~less_equal(X1,X2)|~less_equal(X3,X4))).
########################
#cnf(i_0_25076, plain, (divide(divide(divide(X1,X2),X3),X4)=zero|~less_equal(X1,X2))).
############################################
#cnf(i_0_27403, plain, (divide(X1,X2)=divide(X3,X2)|~less_equal(X1,X2)|~less_equal(X3,X1))).
###########
#cnf(i_0_29389, plain, (divide(divide(X1,X3),X2)=divide(X1,X2)|~less_equal(X1,X2))).
#
#cnf(i_0_29617, plain, (divide(divide(X1,X2),X1)=zero)).
########################################################################################################
#cnf(i_0_29910, plain, (less_equal(divide(divide(X1,X1),divide(X2,X1)),zero))).
######
#cnf(i_0_30026, negated_conjecture, (~less_equal(a,divide(divide(X1,X1),divide(X2,X1))))).
#
#cnf(i_0_30125, plain, (less_equal(divide(divide(X1,X1),zero),zero))).
####
#cnf(i_0_30540, plain, (divide(divide(X1,X1),zero)=zero)).
#
#cnf(i_0_30757, plain, (less_equal(divide(X1,X1),zero))).
#
#cnf(i_0_31039, plain, (divide(X1,X1)=zero)).
##
#cnf(i_0_31281, plain, (less_equal(X1,X1))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(less_equal_and_equal, axiom, (X1=X2|~less_equal(X1,X2)|~less_equal(X2,X1)), file('./p/HEN012-3.tptp', less_equal_and_equal)).
cnf(zero_is_smallest, axiom, (less_equal(zero,X1)), file('./p/HEN012-3.tptp', zero_is_smallest)).
cnf(quotient_smaller_than_numerator, axiom, (less_equal(divide(X1,X2),X1)), file('./p/HEN012-3.tptp', quotient_smaller_than_numerator)).
cnf(quotient_property, axiom, (less_equal(divide(divide(X1,X2),divide(X3,X2)),divide(divide(X1,X3),X2))), file('./p/HEN012-3.tptp', quotient_property)).
cnf(quotient_less_equal1, axiom, (divide(X1,X2)=zero|~less_equal(X1,X2)), file('./p/HEN012-3.tptp', quotient_less_equal1)).
cnf(quotient_less_equal2, axiom, (less_equal(X1,X2)|divide(X1,X2)!=zero), file('./p/HEN012-3.tptp', quotient_less_equal2)).
cnf(identity_is_largest, axiom, (less_equal(X1,identity)), file('./p/HEN012-3.tptp', identity_is_largest)).
cnf(prove_a_LE_a, negated_conjecture, (~less_equal(a,a)), file('./p/HEN012-3.tptp', prove_a_LE_a)).
cnf(c_0_8, axiom, (X1=X2|~less_equal(X1,X2)|~less_equal(X2,X1)), less_equal_and_equal).
cnf(c_0_9, axiom, (less_equal(zero,X1)), zero_is_smallest).
cnf(c_0_10, plain, (zero=X1|~less_equal(X1,zero)), inference(pm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_11, axiom, (less_equal(divide(X1,X2),X1)), quotient_smaller_than_numerator).
cnf(c_0_12, axiom, (less_equal(divide(divide(X1,X2),divide(X3,X2)),divide(divide(X1,X3),X2))), quotient_property).
cnf(c_0_13, axiom, (divide(X1,X2)=zero|~less_equal(X1,X2)), quotient_less_equal1).
cnf(c_0_14, plain, (divide(zero,X1)=zero), inference(pm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, plain, (less_equal(divide(divide(X1,X2),divide(X3,X2)),zero)|~less_equal(X1,X3)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_12, c_0_13]), c_0_14])).
cnf(c_0_16, axiom, (less_equal(X1,X2)|divide(X1,X2)!=zero), quotient_less_equal2).
cnf(c_0_17, plain, (divide(divide(X1,X2),divide(X3,X2))=zero|~less_equal(X1,X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_8, c_0_15]), c_0_9])])).
cnf(c_0_18, plain, (less_equal(divide(X1,X2),divide(X3,X2))|~less_equal(X1,X3)), inference(pm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_19, plain, (less_equal(divide(X1,X2),zero)|~less_equal(X1,X3)|~less_equal(X3,X2)), inference(pm,[status(thm)],[c_0_18, c_0_13])).
cnf(c_0_20, axiom, (less_equal(X1,identity)), identity_is_largest).
cnf(c_0_21, plain, (less_equal(divide(X1,identity),zero)|~less_equal(X1,X2)), inference(pm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_22, plain, (less_equal(divide(X1,X2),zero)|~less_equal(X1,zero)), inference(pm,[status(thm)],[c_0_18, c_0_14])).
cnf(c_0_23, plain, (divide(X1,identity)=zero|~less_equal(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_8, c_0_21]), c_0_9])])).
cnf(c_0_24, plain, (divide(X1,X2)=zero|~less_equal(X1,zero)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_8, c_0_22]), c_0_9])])).
cnf(c_0_25, plain, (divide(X1,X2)=X1|~less_equal(X1,divide(X1,X2))), inference(pm,[status(thm)],[c_0_8, c_0_11])).
cnf(c_0_26, plain, (divide(divide(X1,X2),identity)=zero), inference(pm,[status(thm)],[c_0_23, c_0_11])).
cnf(c_0_27, plain, (less_equal(X1,X2)|~less_equal(X1,zero)), inference(pm,[status(thm)],[c_0_16, c_0_24])).
cnf(c_0_28, plain, (divide(X1,X2)=X1|~less_equal(X1,zero)), inference(pm,[status(thm)],[c_0_25, c_0_24])).
cnf(c_0_29, plain, (less_equal(divide(divide(X1,identity),zero),zero)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_12, c_0_26]), c_0_26])).
cnf(c_0_30, plain, (X1=X2|~less_equal(X1,zero)|~less_equal(X2,X1)), inference(pm,[status(thm)],[c_0_8, c_0_27])).
cnf(c_0_31, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X3,zero)|~less_equal(X1,X3)), inference(pm,[status(thm)],[c_0_18, c_0_28])).
cnf(c_0_32, plain, (divide(divide(X1,identity),zero)=zero), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_8, c_0_29]), c_0_9])])).
cnf(c_0_33, plain, (X1=divide(X2,X3)|~less_equal(X1,zero)|~less_equal(X2,X1)), inference(pm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_34, plain, (less_equal(divide(X1,identity),zero)), inference(pm,[status(thm)],[c_0_16, c_0_32])).
cnf(c_0_35, plain, (divide(X1,X2)=zero|~less_equal(X1,X3)|~less_equal(X3,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_8, c_0_19]), c_0_9])])).
cnf(c_0_36, plain, (divide(X1,X2)=X1|~less_equal(X3,zero)|~less_equal(X1,X3)), inference(pm,[status(thm)],[c_0_25, c_0_33])).
cnf(c_0_37, plain, (divide(X1,identity)=zero), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_8, c_0_34]), c_0_9])])).
cnf(c_0_38, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X1,X4)|~less_equal(X4,X2)), inference(pm,[status(thm)],[c_0_27, c_0_19])).
cnf(c_0_39, plain, (less_equal(X1,X2)|~less_equal(X1,X3)|~less_equal(X3,X2)), inference(pm,[status(thm)],[c_0_16, c_0_35])).
cnf(c_0_40, plain, (zero=X1|~less_equal(X2,zero)|~less_equal(X1,X2)), inference(pm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_41, plain, (less_equal(divide(divide(X1,X2),X3),X4)|~less_equal(X1,X3)), inference(pm,[status(thm)],[c_0_38, c_0_11])).
cnf(c_0_42, plain, (less_equal(X1,X2)|~less_equal(X1,divide(X2,X3))), inference(pm,[status(thm)],[c_0_39, c_0_11])).
cnf(c_0_43, plain, (zero=X1|~less_equal(X1,divide(divide(X2,X3),X4))|~less_equal(X2,X4)), inference(pm,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_44, plain, (less_equal(divide(divide(X1,X2),X3),X1)), inference(pm,[status(thm)],[c_0_42, c_0_11])).
cnf(c_0_45, plain, (divide(divide(X1,X2),divide(X3,X2))=zero|~less_equal(X1,X2)), inference(pm,[status(thm)],[c_0_43, c_0_12])).
cnf(c_0_46, plain, (less_equal(X1,X2)|~less_equal(X1,divide(divide(X2,X3),X4))), inference(pm,[status(thm)],[c_0_39, c_0_44])).
cnf(c_0_47, plain, (less_equal(divide(X1,X2),divide(X3,X2))|~less_equal(X1,X2)), inference(pm,[status(thm)],[c_0_16, c_0_45])).
cnf(c_0_48, plain, (less_equal(divide(X1,X2),X3)|~less_equal(X1,X2)), inference(pm,[status(thm)],[c_0_46, c_0_47])).
cnf(c_0_49, plain, (divide(X1,X2)=X3|~less_equal(X3,divide(X1,X2))|~less_equal(X1,X2)), inference(pm,[status(thm)],[c_0_30, c_0_48])).
cnf(c_0_50, plain, (divide(X1,X2)=zero|~less_equal(X1,divide(X2,X3))), inference(pm,[status(thm)],[c_0_35, c_0_11])).
cnf(c_0_51, plain, (divide(X1,X2)=divide(X3,X2)|~less_equal(X1,X2)|~less_equal(X3,X1)), inference(pm,[status(thm)],[c_0_49, c_0_18])).
cnf(c_0_52, plain, (divide(divide(divide(X1,X2),X3),X1)=zero), inference(pm,[status(thm)],[c_0_50, c_0_11])).
cnf(c_0_53, plain, (divide(divide(X1,X2),X3)=divide(X1,X3)|~less_equal(X1,X3)), inference(pm,[status(thm)],[c_0_51, c_0_11])).
cnf(c_0_54, plain, (divide(divide(X1,X2),X1)=zero), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_52, c_0_53]), c_0_11])])).
cnf(c_0_55, plain, (less_equal(divide(divide(X1,X1),divide(X2,X1)),zero)), inference(pm,[status(thm)],[c_0_12, c_0_54])).
cnf(c_0_56, plain, (less_equal(divide(divide(X1,X1),zero),zero)), inference(pm,[status(thm)],[c_0_55, c_0_14])).
cnf(c_0_57, plain, (divide(divide(X1,X1),zero)=zero), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_8, c_0_56]), c_0_9])])).
cnf(c_0_58, plain, (less_equal(divide(X1,X1),zero)), inference(pm,[status(thm)],[c_0_16, c_0_57])).
cnf(c_0_59, plain, (divide(X1,X1)=zero), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_8, c_0_58]), c_0_9])])).
cnf(c_0_60, negated_conjecture, (~less_equal(a,a)), prove_a_LE_a).
cnf(c_0_61, plain, (less_equal(X1,X1)), inference(pm,[status(thm)],[c_0_16, c_0_59])).
cnf(c_0_62, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_61])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 63
# Proof object clause steps            : 55
# Proof object formula steps           : 8
# Proof object conjectures             : 3
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 46
# Proof object simplifying inferences  : 22
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 8
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 8
# Processed clauses                    : 4465
# ...of these trivial                  : 7
# ...subsumed                          : 4214
# ...remaining for further processing  : 244
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 21
# Backward-rewritten                   : 13
# Generated clauses                    : 31555
# ...of the previous two non-trivial   : 28168
# Contextual simplify-reflections      : 0
# Paramodulations                      : 31555
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 210
#    Positive orientable unit clauses  : 19
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 181
# Current number of unprocessed clauses: 23327
# ...number of literals in the above   : 85975
# Current number of archived formulas  : 0
# Current number of archived clauses   : 34
# Clause-clause subsumption calls (NU) : 33047
# Rec. Clause-clause subsumption calls : 28259
# Non-unit clause-clause subsumptions  : 2673
# Unit Clause-clause subsumption calls : 136
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 177
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 189797
