# Initializing proof state
# Scanning for AC axioms
# addition is AC
# multiplication is associative
# AC handling enabled
#
#cnf(i_0_24, negated_conjecture, (test(esk2_0))).
#
#cnf(i_0_4, plain, (addition(X1,X1)=X1)).
#
#cnf(i_0_14, plain, (test(X2)|~complement(X1,X2))).
#
#cnf(i_0_23, negated_conjecture, (addition(esk2_0,c(esk2_0))!=one)).
#
#cnf(i_0_3, plain, (addition(X1,zero)=X1)).
#
#cnf(i_0_6, plain, (multiplication(X1,one)=X1)).
#
#cnf(i_0_7, plain, (multiplication(one,X1)=X1)).
#
#cnf(i_0_10, plain, (multiplication(X1,zero)=zero)).
#
#cnf(i_0_11, plain, (multiplication(zero,X1)=zero)).
#
#cnf(i_0_22, plain, (c(X1)=zero|test(X1))).
#
#cnf(i_0_15, plain, (complement(esk1_1(X1),X1)|~test(X1))).
#
#cnf(i_0_1, plain, (addition(X1,X2)=addition(X2,X1))).
#
#cnf(i_0_27, plain, (addition(zero,X1)=X1)).
##
#cnf(i_0_12, plain, (leq(X1,X2)|addition(X1,X2)!=X2)).
#
#cnf(i_0_13, plain, (addition(X1,X2)=X2|~leq(X1,X2))).
#
#cnf(i_0_17, plain, (addition(X1,X2)=one|~complement(X2,X1))).
#
#cnf(i_0_19, plain, (multiplication(X1,X2)=zero|~complement(X2,X1))).
#
#cnf(i_0_18, plain, (multiplication(X1,X2)=zero|~complement(X1,X2))).
#
#cnf(i_0_34, plain, (addition(X1,esk1_1(X1))=one|~test(X1))).
#
#cnf(i_0_37, plain, (esk1_1(zero)=one|~test(zero))).
#
#cnf(i_0_21, plain, (complement(X1,X2)|c(X1)!=X2|~test(X1))).
#
#cnf(i_0_39, plain, (test(X1)|c(X2)!=X1|~test(X2))).
#
#cnf(i_0_42, plain, (test(c(X1))|~test(X1))).
#
#cnf(i_0_20, plain, (c(X1)=X2|~complement(X1,X2)|~test(X1))).
#
#cnf(i_0_35, plain, (multiplication(X1,esk1_1(X1))=zero|~test(X1))).
#
#cnf(i_0_47, plain, (esk1_1(one)=zero|~test(one))).
#
#cnf(i_0_36, plain, (multiplication(esk1_1(X1),X1)=zero|~test(X1))).
##
#cnf(i_0_2, plain, (addition(addition(X1,X2),X3)=addition(X1,addition(X2,X3)))).
#
#cnf(i_0_60, plain, (addition(X1,addition(X1,X2))=addition(X1,X2))).
#
#cnf(i_0_66, plain, (addition(X1,addition(X2,X1))=addition(X1,X2))).
##
#cnf(i_0_5, plain, (multiplication(multiplication(X1,X2),X3)=multiplication(X1,multiplication(X2,X3)))).
#
#cnf(i_0_64, plain, (addition(X1,esk1_1(X1))=addition(X1,one)|~test(X1))).
##
#cnf(i_0_94, plain, (addition(X1,one)=one|~test(X1))).
#
#cnf(i_0_96, plain, (addition(one,X1)=one|~test(X1))).
#######
#cnf(i_0_101, plain, (addition(one,X1)=one|~test(addition(one,X1)))).
#
#cnf(i_0_102, plain, (addition(one,X1)=one|~test(addition(X1,one)))).
#########
#cnf(i_0_38, plain, (multiplication(X1,X2)=zero|c(X1)!=X2|~test(X1))).
#
#cnf(i_0_128, plain, (zero=X1|c(X1)!=one|~test(X1))).
#
#cnf(i_0_126, plain, (zero=X1|c(one)!=X1|~test(one))).
#
#cnf(i_0_131, plain, (c(one)=zero|~test(one))).
#
#cnf(i_0_133, plain, (test(zero)|~test(one))).
#
#cnf(i_0_135, plain, (test(X1)|zero!=X1|~test(one))).
#
#cnf(i_0_40, plain, (addition(X1,X2)=one|c(X2)!=X1|~test(X2))).
#
#cnf(i_0_153, plain, (one=X1|c(X1)!=X1|~test(X1))).
#
#cnf(i_0_141, plain, (one=X1|c(X1)!=zero|~test(X1))).
#
#cnf(i_0_139, negated_conjecture, (c(c(esk2_0))!=esk2_0|~test(c(esk2_0)))).
#
#cnf(i_0_145, plain, (one=X1|c(zero)!=X1|~test(zero))).
#
#cnf(i_0_158, plain, (c(zero)=one|~test(zero))).
#
#cnf(i_0_161, plain, (test(one)|~test(zero))).
#
#cnf(i_0_166, plain, (esk1_1(one)=zero|~test(zero))).
#
#cnf(i_0_164, plain, (test(X1)|one!=X1|~test(zero))).
####
#cnf(i_0_41, plain, (multiplication(X1,X2)=zero|c(X2)!=X1|~test(X2))).
####
#cnf(i_0_148, plain, (addition(X2,X1)=one|c(X2)!=X1|~test(X2))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(test_2, axiom, ![X4, X5]:(complement(X5,X4)<=>((multiplication(X4,X5)=zero&multiplication(X5,X4)=zero)&addition(X4,X5)=one)), file('./p/KLE006+1.tptp', test_2)).
fof(test_3, axiom, ![X4, X5]:(test(X4)=>(c(X4)=X5<=>complement(X4,X5))), file('./p/KLE006+1.tptp', test_3)).
fof(goals, conjecture, ![X4]:(test(X4)=>one=addition(X4,c(X4))), file('./p/KLE006+1.tptp', goals)).
fof(additive_commutativity, axiom, ![X1, X2]:addition(X1,X2)=addition(X2,X1), file('./p/KLE006+1.tptp', additive_commutativity)).
fof(c_0_4, plain, ![X32, X33]:((((multiplication(X32,X33)=zero|~complement(X33,X32))&(multiplication(X33,X32)=zero|~complement(X33,X32)))&(addition(X32,X33)=one|~complement(X33,X32)))&(multiplication(X32,X33)!=zero|multiplication(X33,X32)!=zero|addition(X32,X33)!=one|complement(X33,X32))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[test_2])])])).
fof(c_0_5, plain, ![X34, X35]:((c(X34)!=X35|complement(X34,X35)|~test(X34))&(~complement(X34,X35)|c(X34)=X35|~test(X34))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[test_3])])])).
fof(c_0_6, negated_conjecture, ~(![X4]:(test(X4)=>one=addition(X4,c(X4)))), inference(assume_negation,[status(cth)],[goals])).
fof(c_0_7, plain, ![X6, X7]:addition(X6,X7)=addition(X7,X6), inference(variable_rename,[status(thm)],[additive_commutativity])).
cnf(c_0_8, plain, (addition(X1,X2)=one|~complement(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9, plain, (complement(X1,X2)|c(X1)!=X2|~test(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_10, negated_conjecture, (test(esk2_0)&one!=addition(esk2_0,c(esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_6])])])).
cnf(c_0_11, plain, (addition(X1,X2)=addition(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (addition(X1,X2)=one|c(X2)!=X1|~test(X2)), inference(pm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_13, negated_conjecture, (one!=addition(esk2_0,c(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (addition(X1,X2)=one|c(X1)!=X2|~test(X1)), inference(pm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_15, negated_conjecture, (test(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_13, c_0_14]), c_0_15])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 17
# Proof object clause steps            : 8
# Proof object formula steps           : 9
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 2
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 17
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 24
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 24
# Processed clauses                    : 77
# ...of these trivial                  : 5
# ...subsumed                          : 19
# ...remaining for further processing  : 53
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 168
# ...of the previous two non-trivial   : 100
# Contextual simplify-reflections      : 0
# Paramodulations                      : 165
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 53
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 47
# ...number of literals in the above   : 107
# Current number of archived formulas  : 0
# Current number of archived clauses   : 0
# Clause-clause subsumption calls (NU) : 118
# Rec. Clause-clause subsumption calls : 111
# Non-unit clause-clause subsumptions  : 19
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 6
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2033
