# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_56, plain, (cUnsatisfiable(i2003_11_14_17_19_21224))).
#
#cnf(i_0_16, plain, (xsd_string(X1)|xsd_integer(X1))).
#
#cnf(i_0_44, plain, (cp(esk1_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_14, plain, (~cowlNothing(X1))).
#
#cnf(i_0_32, plain, (cp(esk6_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_28, plain, (cp(esk7_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_43, plain, (cp1(esk1_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_49, plain, (~cp2(X1)|~cp1(X1))).
#
#cnf(i_0_41, plain, (cp1(esk2_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_36, plain, (cp2(esk4_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_29, plain, (cp2(esk7_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_46, plain, (~cp3(X1)|~cp1(X1))).
#
#cnf(i_0_34, plain, (cp3(esk5_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_31, plain, (cp3(esk6_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_39, plain, (cp4(esk3_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_48, plain, (~cp4(X1)|~cp1(X1))).
#
#cnf(i_0_45, plain, (rr(X1,esk1_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_42, plain, (rr(X1,esk2_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_40, plain, (rr(X1,esk3_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_47, plain, (~cp5(X1)|~cp1(X1))).
#
#cnf(i_0_37, plain, (rr(X1,esk4_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_35, plain, (rr(X1,esk5_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_33, plain, (rr(X1,esk6_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_50, plain, (~cp3(X1)|~cp2(X1))).
#
#cnf(i_0_30, plain, (rr(X1,esk7_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_38, plain, (X4=X5|X3=X5|X3=X4|X2=X5|X2=X4|X2=X3|~rr(X1,X5)|~rr(X1,X4)|~rr(X1,X3)|~rr(X1,X2)|~cUnsatisfiable(X1))).
#
#cnf(i_0_366, plain, (esk7_1(X4)=X2|esk7_1(X4)=X1|esk7_1(X4)=X3|X1=X2|X3=X2|X3=X1|~rr(X4,X2)|~rr(X4,X1)|~rr(X4,X3)|~cUnsatisfiable(X4))).
#
#cnf(i_0_52, plain, (~cp4(X1)|~cp2(X1))).
#
#cnf(i_0_367, plain, (esk6_1(X4)=X2|esk6_1(X4)=X1|esk6_1(X4)=X3|X1=X2|X3=X2|X3=X1|~rr(X4,X2)|~rr(X4,X1)|~rr(X4,X3)|~cUnsatisfiable(X4))).
#
#cnf(i_0_368, plain, (esk5_1(X4)=X2|esk5_1(X4)=X1|esk5_1(X4)=X3|X1=X2|X3=X2|X3=X1|~rr(X4,X2)|~rr(X4,X1)|~rr(X4,X3)|~cUnsatisfiable(X4))).
#
#cnf(i_0_369, plain, (esk4_1(X4)=X2|esk4_1(X4)=X1|esk4_1(X4)=X3|X1=X2|X3=X2|X3=X1|~rr(X4,X2)|~rr(X4,X1)|~rr(X4,X3)|~cUnsatisfiable(X4))).
#
#cnf(i_0_51, plain, (~cp5(X1)|~cp2(X1))).
#
#cnf(i_0_370, plain, (esk3_1(X4)=X2|esk3_1(X4)=X1|esk3_1(X4)=X3|X1=X2|X3=X2|X3=X1|~rr(X4,X2)|~rr(X4,X1)|~rr(X4,X3)|~cUnsatisfiable(X4))).
#
#cnf(i_0_371, plain, (esk2_1(X4)=X2|esk2_1(X4)=X1|esk2_1(X4)=X3|X1=X2|X3=X2|X3=X1|~rr(X4,X2)|~rr(X4,X1)|~rr(X4,X3)|~cUnsatisfiable(X4))).
#
#cnf(i_0_372, plain, (esk1_1(X4)=X2|esk1_1(X4)=X1|esk1_1(X4)=X3|X1=X2|X3=X2|X3=X1|~rr(X4,X2)|~rr(X4,X1)|~rr(X4,X3)|~cUnsatisfiable(X4))).
#
#cnf(i_0_54, plain, (~cp4(X1)|~cp3(X1))).
####
#cnf(i_0_53, plain, (~cp5(X1)|~cp3(X1))).
####
#cnf(i_0_55, plain, (~cp5(X1)|~cp4(X1))).
####
#cnf(i_0_17, plain, (~xsd_string(X1)|~xsd_integer(X1))).
####
#cnf(i_0_57, plain, (~cp1(esk4_1(X1))|~cUnsatisfiable(X1))).
####
#cnf(i_0_58, plain, (~cp1(esk7_1(X1))|~cUnsatisfiable(X1))).
####
#cnf(i_0_59, plain, (~cp1(esk5_1(X1))|~cUnsatisfiable(X1))).
####
#cnf(i_0_60, plain, (~cp1(esk6_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_526, plain, (esk7_1(X1)=esk6_1(X1)|esk7_1(X1)=X2|esk7_1(X1)=X3|esk6_1(X1)=X2|esk6_1(X1)=X3|X3=X2|~rr(X1,X2)|~rr(X1,X3)|~cUnsatisfiable(X1))).
#
#cnf(i_0_1714, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|esk7_1(X1)=X2|esk6_1(X1)=X2|esk5_1(X1)=X2|~rr(X1,X2)|~cUnsatisfiable(X1))).
#
#cnf(i_0_1724, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|esk5_1(X1)=esk4_1(X1)|~cUnsatisfiable(X1))).
#
#cnf(i_0_61, plain, (~cp1(esk3_1(X1))|~cUnsatisfiable(X1))).
##
#cnf(i_0_1730, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|cp3(esk4_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_1789, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|~cp2(esk4_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_62, plain, (~cp5(esk2_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_1904, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
##
#cnf(i_0_1962, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|cp3(esk7_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_63, plain, (~cp5(esk1_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_1965, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|~cp2(esk7_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_2136, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
##
#cnf(i_0_64, plain, (~cp2(esk6_1(X1))|~cUnsatisfiable(X1))).
###
#cnf(i_0_1725, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|esk3_1(X1)=esk5_1(X1)|~cUnsatisfiable(X1))).
#
#cnf(i_0_65, plain, (~cp2(esk5_1(X1))|~cUnsatisfiable(X1))).
###
#cnf(i_0_2198, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|cp4(esk5_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_544, plain, (~cp2(esk3_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_2258, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|~cp3(esk5_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_2433, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
##
#cnf(i_0_995, plain, (~cp5(esk7_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_2438, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|esk3_1(X1)=esk6_1(X1)|~cp2(esk7_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_2551, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|esk3_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
##
#cnf(i_0_996, plain, (~cp5(esk4_1(X1))|~cUnsatisfiable(X1))).
##
#cnf(i_0_2552, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|cp4(esk6_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_2612, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|~cp3(esk6_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_1447, plain, (~cp3(esk3_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_2728, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk7_1(X1)|esk5_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
##
#cnf(i_0_2790, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk6_1(X1)|~cp2(esk7_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_1448, plain, (~cp5(esk6_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_2848, plain, (esk7_1(X1)=esk6_1(X1)|esk5_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
###
#cnf(i_0_1449, plain, (~cp5(esk5_1(X1))|~cUnsatisfiable(X1))).
####
#cnf(i_0_1450, plain, (~cp5(esk3_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_1715, plain, (esk7_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|esk7_1(X1)=X2|esk6_1(X1)=X2|esk4_1(X1)=X2|~rr(X1,X2)|~cUnsatisfiable(X1))).
##
#cnf(i_0_3023, plain, (esk7_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|esk3_1(X1)=esk4_1(X1)|~cUnsatisfiable(X1))).
#
#cnf(i_0_3028, plain, (esk7_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|cp4(esk4_1(X1))|~cUnsatisfiable(X1))).
##
#cnf(i_0_3031, plain, (esk7_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|~cp2(esk4_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_3150, plain, (esk7_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
#
#cnf(i_0_3152, plain, (esk7_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|esk3_1(X1)=esk6_1(X1)|cp4(esk7_1(X1))|~cUnsatisfiable(X1))).
##
#cnf(i_0_3155, plain, (esk7_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|esk3_1(X1)=esk6_1(X1)|~cp2(esk7_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_3274, plain, (esk7_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|esk3_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
#
#cnf(i_0_3275, plain, (esk7_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|cp4(esk6_1(X1))|~cUnsatisfiable(X1))).
####
#cnf(i_0_3280, plain, (esk7_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|~cp3(esk6_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_3397, plain, (esk7_1(X1)=esk6_1(X1)|esk4_1(X1)=esk7_1(X1)|esk4_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
#####
#cnf(i_0_1716, plain, (esk7_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|esk7_1(X1)=X2|esk6_1(X1)=X2|esk3_1(X1)=X2|~rr(X1,X2)|~cUnsatisfiable(X1))).
###
#cnf(i_0_3405, plain, (esk7_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|esk3_1(X1)=esk2_1(X1)|~cUnsatisfiable(X1))).
#
#cnf(i_0_3411, plain, (esk7_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|~cp1(esk2_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_3471, plain, (esk7_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
#
#cnf(i_0_3473, plain, (esk7_1(X1)=esk6_1(X1)|esk3_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|cp4(esk7_1(X1))|~cUnsatisfiable(X1))).
##
#cnf(i_0_3476, plain, (esk7_1(X1)=esk6_1(X1)|esk3_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|~cp2(esk7_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_3595, plain, (esk7_1(X1)=esk6_1(X1)|esk3_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
#
#cnf(i_0_3596, plain, (esk7_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|cp4(esk6_1(X1))|~cUnsatisfiable(X1))).
####
#cnf(i_0_3601, plain, (esk7_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|~cp3(esk6_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_3718, plain, (esk7_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
##
#cnf(i_0_3720, plain, (esk7_1(X1)=esk6_1(X1)|esk2_1(X1)=esk6_1(X1)|cp1(esk7_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_3724, plain, (esk7_1(X1)=esk6_1(X1)|esk2_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
##
#cnf(i_0_3781, plain, (esk7_1(X1)=esk6_1(X1)|cp1(esk6_1(X1))|~cUnsatisfiable(X1))).
#
#cnf(i_0_3785, plain, (esk7_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1))).
###
#cnf(i_0_3843, plain, (cp2(esk6_1(X1))|~cUnsatisfiable(X1))).
##
#cnf(i_0_3849, plain, (~cUnsatisfiable(X1))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
fof(axiom_2, axiom, ![X4]:(cUnsatisfiable(X4)<=>(((((((?[X5]:((rr(X4,X5)&cp(X5))&cp1(X5))&?[X5]:(rr(X4,X5)&cp1(X5)))&?[X5]:(rr(X4,X5)&cp4(X5)))&![X6, X7, X8, X9]:((((rr(X4,X6)&rr(X4,X7))&rr(X4,X8))&rr(X4,X9))=>(X6=X7|X6=X8|X6=X9|X7=X8|X7=X9|X8=X9)))&?[X5]:(rr(X4,X5)&cp2(X5)))&?[X5]:(rr(X4,X5)&cp3(X5)))&?[X5]:((rr(X4,X5)&cp(X5))&cp3(X5)))&?[X5]:((rr(X4,X5)&cp2(X5))&cp(X5)))), file('./p/KRS080+1.tptp', axiom_2)).
fof(axiom_3, axiom, ![X4]:(cp1(X4)=>~((cp2(X4)|cp4(X4)|cp5(X4)|cp3(X4)))), file('./p/KRS080+1.tptp', axiom_3)).
fof(axiom_4, axiom, ![X4]:(cp2(X4)=>~((cp4(X4)|cp5(X4)|cp3(X4)))), file('./p/KRS080+1.tptp', axiom_4)).
fof(axiom_5, axiom, ![X4]:(cp3(X4)=>~((cp4(X4)|cp5(X4)))), file('./p/KRS080+1.tptp', axiom_5)).
fof(axiom_7, axiom, cUnsatisfiable(i2003_11_14_17_19_21224), file('./p/KRS080+1.tptp', axiom_7)).
fof(c_0_5, plain, ![X40, X44, X45, X46, X47, X52, X53, X54, X55, X60, X61, X62, X63]:(((((((((((rr(X40,esk1_1(X40))|~cUnsatisfiable(X40))&(cp(esk1_1(X40))|~cUnsatisfiable(X40)))&(cp1(esk1_1(X40))|~cUnsatisfiable(X40)))&((rr(X40,esk2_1(X40))|~cUnsatisfiable(X40))&(cp1(esk2_1(X40))|~cUnsatisfiable(X40))))&((rr(X40,esk3_1(X40))|~cUnsatisfiable(X40))&(cp4(esk3_1(X40))|~cUnsatisfiable(X40))))&(~rr(X40,X44)|~rr(X40,X45)|~rr(X40,X46)|~rr(X40,X47)|(X44=X45|X44=X46|X44=X47|X45=X46|X45=X47|X46=X47)|~cUnsatisfiable(X40)))&((rr(X40,esk4_1(X40))|~cUnsatisfiable(X40))&(cp2(esk4_1(X40))|~cUnsatisfiable(X40))))&((rr(X40,esk5_1(X40))|~cUnsatisfiable(X40))&(cp3(esk5_1(X40))|~cUnsatisfiable(X40))))&(((rr(X40,esk6_1(X40))|~cUnsatisfiable(X40))&(cp(esk6_1(X40))|~cUnsatisfiable(X40)))&(cp3(esk6_1(X40))|~cUnsatisfiable(X40))))&(((rr(X40,esk7_1(X40))|~cUnsatisfiable(X40))&(cp2(esk7_1(X40))|~cUnsatisfiable(X40)))&(cp(esk7_1(X40))|~cUnsatisfiable(X40))))&(((((rr(X52,esk8_1(X52))|(~rr(X52,X53)|~cp(X53)|~cp1(X53)|(~rr(X52,X54)|~cp1(X54))|(~rr(X52,X55)|~cp4(X55)))|(~rr(X52,X60)|~cp2(X60))|(~rr(X52,X61)|~cp3(X61))|(~rr(X52,X62)|~cp(X62)|~cp3(X62))|(~rr(X52,X63)|~cp2(X63)|~cp(X63))|cUnsatisfiable(X52))&(rr(X52,esk9_1(X52))|(~rr(X52,X53)|~cp(X53)|~cp1(X53)|(~rr(X52,X54)|~cp1(X54))|(~rr(X52,X55)|~cp4(X55)))|(~rr(X52,X60)|~cp2(X60))|(~rr(X52,X61)|~cp3(X61))|(~rr(X52,X62)|~cp(X62)|~cp3(X62))|(~rr(X52,X63)|~cp2(X63)|~cp(X63))|cUnsatisfiable(X52)))&(rr(X52,esk10_1(X52))|(~rr(X52,X53)|~cp(X53)|~cp1(X53)|(~rr(X52,X54)|~cp1(X54))|(~rr(X52,X55)|~cp4(X55)))|(~rr(X52,X60)|~cp2(X60))|(~rr(X52,X61)|~cp3(X61))|(~rr(X52,X62)|~cp(X62)|~cp3(X62))|(~rr(X52,X63)|~cp2(X63)|~cp(X63))|cUnsatisfiable(X52)))&(rr(X52,esk11_1(X52))|(~rr(X52,X53)|~cp(X53)|~cp1(X53)|(~rr(X52,X54)|~cp1(X54))|(~rr(X52,X55)|~cp4(X55)))|(~rr(X52,X60)|~cp2(X60))|(~rr(X52,X61)|~cp3(X61))|(~rr(X52,X62)|~cp(X62)|~cp3(X62))|(~rr(X52,X63)|~cp2(X63)|~cp(X63))|cUnsatisfiable(X52)))&((((((esk8_1(X52)!=esk9_1(X52)|(~rr(X52,X53)|~cp(X53)|~cp1(X53)|(~rr(X52,X54)|~cp1(X54))|(~rr(X52,X55)|~cp4(X55)))|(~rr(X52,X60)|~cp2(X60))|(~rr(X52,X61)|~cp3(X61))|(~rr(X52,X62)|~cp(X62)|~cp3(X62))|(~rr(X52,X63)|~cp2(X63)|~cp(X63))|cUnsatisfiable(X52))&(esk8_1(X52)!=esk10_1(X52)|(~rr(X52,X53)|~cp(X53)|~cp1(X53)|(~rr(X52,X54)|~cp1(X54))|(~rr(X52,X55)|~cp4(X55)))|(~rr(X52,X60)|~cp2(X60))|(~rr(X52,X61)|~cp3(X61))|(~rr(X52,X62)|~cp(X62)|~cp3(X62))|(~rr(X52,X63)|~cp2(X63)|~cp(X63))|cUnsatisfiable(X52)))&(esk8_1(X52)!=esk11_1(X52)|(~rr(X52,X53)|~cp(X53)|~cp1(X53)|(~rr(X52,X54)|~cp1(X54))|(~rr(X52,X55)|~cp4(X55)))|(~rr(X52,X60)|~cp2(X60))|(~rr(X52,X61)|~cp3(X61))|(~rr(X52,X62)|~cp(X62)|~cp3(X62))|(~rr(X52,X63)|~cp2(X63)|~cp(X63))|cUnsatisfiable(X52)))&(esk9_1(X52)!=esk10_1(X52)|(~rr(X52,X53)|~cp(X53)|~cp1(X53)|(~rr(X52,X54)|~cp1(X54))|(~rr(X52,X55)|~cp4(X55)))|(~rr(X52,X60)|~cp2(X60))|(~rr(X52,X61)|~cp3(X61))|(~rr(X52,X62)|~cp(X62)|~cp3(X62))|(~rr(X52,X63)|~cp2(X63)|~cp(X63))|cUnsatisfiable(X52)))&(esk9_1(X52)!=esk11_1(X52)|(~rr(X52,X53)|~cp(X53)|~cp1(X53)|(~rr(X52,X54)|~cp1(X54))|(~rr(X52,X55)|~cp4(X55)))|(~rr(X52,X60)|~cp2(X60))|(~rr(X52,X61)|~cp3(X61))|(~rr(X52,X62)|~cp(X62)|~cp3(X62))|(~rr(X52,X63)|~cp2(X63)|~cp(X63))|cUnsatisfiable(X52)))&(esk10_1(X52)!=esk11_1(X52)|(~rr(X52,X53)|~cp(X53)|~cp1(X53)|(~rr(X52,X54)|~cp1(X54))|(~rr(X52,X55)|~cp4(X55)))|(~rr(X52,X60)|~cp2(X60))|(~rr(X52,X61)|~cp3(X61))|(~rr(X52,X62)|~cp(X62)|~cp3(X62))|(~rr(X52,X63)|~cp2(X63)|~cp(X63))|cUnsatisfiable(X52))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[axiom_2])])])])])])).
cnf(c_0_6, plain, (X2=X3|X2=X4|X2=X5|X3=X4|X3=X5|X4=X5|~rr(X1,X2)|~rr(X1,X3)|~rr(X1,X4)|~rr(X1,X5)|~cUnsatisfiable(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_7, plain, (rr(X1,esk7_1(X1))|~cUnsatisfiable(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, plain, (esk7_1(X1)=X2|esk7_1(X1)=X3|esk7_1(X1)=X4|X3=X2|X4=X2|X4=X3|~rr(X1,X2)|~rr(X1,X3)|~rr(X1,X4)|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_9, plain, (rr(X1,esk6_1(X1))|~cUnsatisfiable(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_10, plain, ![X64]:((((~cp2(X64)|~cp1(X64))&(~cp4(X64)|~cp1(X64)))&(~cp5(X64)|~cp1(X64)))&(~cp3(X64)|~cp1(X64))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[axiom_3])])])).
cnf(c_0_11, plain, (esk7_1(X1)=esk6_1(X1)|esk7_1(X1)=X2|esk7_1(X1)=X3|esk6_1(X1)=X2|esk6_1(X1)=X3|X3=X2|~rr(X1,X2)|~rr(X1,X3)|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12, plain, (rr(X1,esk3_1(X1))|~cUnsatisfiable(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_13, plain, (~cp4(X1)|~cp1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_14, plain, (cp4(esk3_1(X1))|~cUnsatisfiable(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15, plain, (esk7_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|esk7_1(X1)=X2|esk6_1(X1)=X2|esk3_1(X1)=X2|~rr(X1,X2)|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_16, plain, (rr(X1,esk2_1(X1))|~cUnsatisfiable(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_17, plain, ![X65]:(((~cp4(X65)|~cp2(X65))&(~cp5(X65)|~cp2(X65)))&(~cp3(X65)|~cp2(X65))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[axiom_4])])])).
cnf(c_0_18, plain, (~cp1(esk3_1(X1))|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_19, plain, (esk7_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|esk3_1(X1)=esk2_1(X1)|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_20, plain, (~cp4(X1)|~cp2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_21, plain, (esk7_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|~cp1(esk2_1(X1))|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22, plain, (cp1(esk2_1(X1))|~cUnsatisfiable(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_23, plain, ![X66]:((~cp4(X66)|~cp3(X66))&(~cp5(X66)|~cp3(X66))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[axiom_5])])])).
cnf(c_0_24, plain, (~cp2(esk3_1(X1))|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_20, c_0_14])).
cnf(c_0_25, plain, (esk7_1(X1)=esk6_1(X1)|esk3_1(X1)=esk7_1(X1)|esk3_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_26, plain, (~cp4(X1)|~cp3(X1)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_27, plain, (esk7_1(X1)=esk6_1(X1)|esk3_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|~cp2(esk7_1(X1))|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_28, plain, (cp2(esk7_1(X1))|~cUnsatisfiable(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_29, plain, (~cp3(esk3_1(X1))|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_26, c_0_14])).
cnf(c_0_30, plain, (esk7_1(X1)=esk6_1(X1)|esk3_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_31, plain, (esk7_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|~cp3(esk6_1(X1))|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_32, plain, (cp3(esk6_1(X1))|~cUnsatisfiable(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_33, plain, (~cp2(X1)|~cp1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_34, plain, (esk7_1(X1)=esk6_1(X1)|esk2_1(X1)=esk7_1(X1)|esk2_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_35, plain, (~cp1(esk7_1(X1))|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_33, c_0_28])).
cnf(c_0_36, plain, (esk7_1(X1)=esk6_1(X1)|esk2_1(X1)=esk6_1(X1)|cp1(esk7_1(X1))|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_22, c_0_34])).
cnf(c_0_37, plain, (~cp3(X1)|~cp1(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_38, plain, (esk7_1(X1)=esk6_1(X1)|esk2_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_35, c_0_36])).
cnf(c_0_39, plain, (~cp1(esk6_1(X1))|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_37, c_0_32])).
cnf(c_0_40, plain, (esk7_1(X1)=esk6_1(X1)|cp1(esk6_1(X1))|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_22, c_0_38])).
cnf(c_0_41, plain, (~cp3(X1)|~cp2(X1)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_42, plain, (esk7_1(X1)=esk6_1(X1)|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_39, c_0_40])).
cnf(c_0_43, plain, (~cp2(esk6_1(X1))|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_41, c_0_32])).
cnf(c_0_44, plain, (cp2(esk6_1(X1))|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_28, c_0_42])).
cnf(c_0_45, plain, (cUnsatisfiable(i2003_11_14_17_19_21224)), inference(split_conjunct,[status(thm)],[axiom_7])).
cnf(c_0_46, plain, (~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_43, c_0_44])).
cnf(c_0_47, plain, ($false), inference(sr,[status(thm)],[c_0_45, c_0_46]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 48
# Proof object clause steps            : 39
# Proof object formula steps           : 9
# Proof object conjectures             : 0
# Proof object clause conjectures      : 0
# Proof object formula conjectures     : 0
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 22
# Proof object simplifying inferences  : 1
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 21
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 56
# Removed in clause preprocessing      : 14
# Initial clauses in saturation        : 42
# Processed clauses                    : 160
# ...of these trivial                  : 0
# ...subsumed                          : 58
# ...remaining for further processing  : 102
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 44
# Backward-rewritten                   : 0
# Generated clauses                    : 3848
# ...of the previous two non-trivial   : 3521
# Contextual simplify-reflections      : 0
# Paramodulations                      : 3522
# Factorizations                       : 325
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 57
#    Positive orientable unit clauses  : 0
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 55
# Current number of unprocessed clauses: 3359
# ...number of literals in the above   : 41996
# Current number of archived formulas  : 0
# Current number of archived clauses   : 46
# Clause-clause subsumption calls (NU) : 880
# Rec. Clause-clause subsumption calls : 340
# Non-unit clause-clause subsumptions  : 102
# Unit Clause-clause subsumption calls : 43
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 57930
