# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_27, negated_conjecture, (e(exist))).
#
#cnf(i_0_28, plain, (t2least(X1)|~c(X1))).
#
#cnf(i_0_34, plain, (t1most(X1)|~d(X1))).
#
#cnf(i_0_30, plain, (~equalish(u1r2(X1),u1r1(X1))|~t2least(X1))).
#
#cnf(i_0_29, plain, (c(X1)|~t2least(X1))).
#
#cnf(i_0_35, plain, (d(X1)|~t1most(X1))).
#
#cnf(i_0_32, plain, (t(X1,u1r2(X1))|~t2least(X1))).
#
#cnf(i_0_39, plain, (t1most(X1)|t(X1,u3r2(X1)))).
#
#cnf(i_0_31, plain, (t(X1,u1r1(X1))|~t2least(X1))).
#
#cnf(i_0_40, plain, (r(X1,u4r4(X1))|~e(X1))).
#
#cnf(i_0_38, plain, (t1most(X1)|t(X1,u3r1(X1)))).
#
#cnf(i_0_44, plain, (c(X1)|~r(X2,X1)|~e(X2))).
#
#cnf(i_0_55, plain, (c(u4r4(X1))|~e(X1))).
#
#cnf(i_0_56, plain, (t2least(u4r4(X1))|~e(X1))).
#
#cnf(i_0_43, plain, (d(X1)|~s(X2,X1)|~e(X2))).
#
#cnf(i_0_42, plain, (r(X1,X2)|~s(X1,X2)|~e(X1))).
#
#cnf(i_0_41, plain, (s(X1,X2)|~r(X1,X2)|~e(X1))).
#
#cnf(i_0_58, plain, (d(X1)|~r(X2,X1)|~e(X2))).
#
#cnf(i_0_59, plain, (d(u4r4(X1))|~e(X1))).
#
#cnf(i_0_60, plain, (t1most(u4r4(X1))|~e(X1))).
#
#cnf(i_0_37, plain, (t1most(X1)|~equalish(u3r2(X1),u3r1(X1)))).
#
#cnf(i_0_36, plain, (equalish(X1,X2)|~t1most(X3)|~t(X3,X2)|~t(X3,X1))).
#
#cnf(i_0_33, plain, (equalish(X2,X3)|t2least(X1)|~t(X1,X3)|~t(X1,X2))).
#
#cnf(i_0_62, plain, (~t1most(X2)|~t(X2,u1r1(X1))|~t(X2,u1r2(X1))|~t2least(X1))).
#
#cnf(i_0_65, plain, (equalish(u1r1(X1),X2)|~t1most(X1)|~t(X1,X2)|~t2least(X1))).
#
#cnf(i_0_66, plain, (equalish(u1r2(X1),X2)|~t1most(X1)|~t(X1,X2)|~t2least(X1))).
#
#cnf(i_0_69, plain, (equalish(X1,u1r1(X2))|~t1most(X2)|~t(X2,X1)|~t2least(X2))).
#
#cnf(i_0_81, plain, (~t1most(X1)|~t(X1,u1r2(X1))|~t2least(X1))).
#
#cnf(i_0_70, plain, (equalish(X1,u1r2(X2))|~t1most(X2)|~t(X2,X1)|~t2least(X2))).
#
#cnf(i_0_73, plain, (t1most(X1)|equalish(u3r1(X1),X2)|t2least(X1)|~t(X1,X2))).
#
#cnf(i_0_74, plain, (t1most(X1)|equalish(u3r2(X1),X2)|t2least(X1)|~t(X1,X2))).
#
#cnf(i_0_84, plain, (~t1most(X1)|~t2least(X1))).
#
#cnf(i_0_85, plain, (t1most(X1)|t2least(X1)|~t(X1,u3r1(X1)))).
#
#cnf(i_0_87, plain, (t1most(X1)|t2least(X1))).
##
#cnf(i_0_86, plain, (~t1most(u4r4(X1))|~e(X1))).
##
#cnf(i_0_61, plain, (t1most(X1)|~t1most(X2)|~t(X2,u3r1(X1))|~t(X2,u3r2(X1)))).
#
#cnf(i_0_72, plain, (t2least(X2)|~t(X2,u1r1(X1))|~t(X2,u1r2(X1))|~t2least(X1))).
#
#cnf(i_0_89, plain, (~e(X1))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(clause_4, axiom, (~t2least(X1)|~equalish(u1r2(X1),u1r1(X1))), file('./p/KRS015-1.tptp', clause_4)).
cnf(clause_10, axiom, (equalish(X1,X2)|~t1most(X3)|~t(X3,X1)|~t(X3,X2)), file('./p/KRS015-1.tptp', clause_10)).
cnf(clause_5, axiom, (t(X1,u1r1(X1))|~t2least(X1)), file('./p/KRS015-1.tptp', clause_5)).
cnf(clause_18, axiom, (c(X1)|~e(X2)|~r(X2,X1)), file('./p/KRS015-1.tptp', clause_18)).
cnf(clause_14, axiom, (r(X1,u4r4(X1))|~e(X1)), file('./p/KRS015-1.tptp', clause_14)).
cnf(clause_17, axiom, (d(X1)|~e(X2)|~s(X2,X1)), file('./p/KRS015-1.tptp', clause_17)).
cnf(clause_15, axiom, (s(X1,X2)|~e(X1)|~r(X1,X2)), file('./p/KRS015-1.tptp', clause_15)).
cnf(clause_6, axiom, (t(X1,u1r2(X1))|~t2least(X1)), file('./p/KRS015-1.tptp', clause_6)).
cnf(clause_2, axiom, (t2least(X1)|~c(X1)), file('./p/KRS015-1.tptp', clause_2)).
cnf(clause_8, axiom, (t1most(X1)|~d(X1)), file('./p/KRS015-1.tptp', clause_8)).
cnf(clause_1, negated_conjecture, (e(exist)), file('./p/KRS015-1.tptp', clause_1)).
cnf(c_0_11, axiom, (~t2least(X1)|~equalish(u1r2(X1),u1r1(X1))), clause_4).
cnf(c_0_12, axiom, (equalish(X1,X2)|~t1most(X3)|~t(X3,X1)|~t(X3,X2)), clause_10).
cnf(c_0_13, plain, (~t1most(X1)|~t(X1,u1r1(X2))|~t(X1,u1r2(X2))|~t2least(X2)), inference(pm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_14, axiom, (t(X1,u1r1(X1))|~t2least(X1)), clause_5).
cnf(c_0_15, axiom, (c(X1)|~e(X2)|~r(X2,X1)), clause_18).
cnf(c_0_16, axiom, (r(X1,u4r4(X1))|~e(X1)), clause_14).
cnf(c_0_17, axiom, (d(X1)|~e(X2)|~s(X2,X1)), clause_17).
cnf(c_0_18, axiom, (s(X1,X2)|~e(X1)|~r(X1,X2)), clause_15).
cnf(c_0_19, plain, (~t1most(X1)|~t(X1,u1r2(X1))|~t2least(X1)), inference(pm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20, axiom, (t(X1,u1r2(X1))|~t2least(X1)), clause_6).
cnf(c_0_21, axiom, (t2least(X1)|~c(X1)), clause_2).
cnf(c_0_22, plain, (c(u4r4(X1))|~e(X1)), inference(pm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_23, plain, (d(X1)|~r(X2,X1)|~e(X2)), inference(pm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, plain, (~t1most(X1)|~t2least(X1)), inference(pm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_25, plain, (t2least(u4r4(X1))|~e(X1)), inference(pm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_26, axiom, (t1most(X1)|~d(X1)), clause_8).
cnf(c_0_27, plain, (d(u4r4(X1))|~e(X1)), inference(pm,[status(thm)],[c_0_23, c_0_16])).
cnf(c_0_28, plain, (~t1most(u4r4(X1))|~e(X1)), inference(pm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, plain, (t1most(u4r4(X1))|~e(X1)), inference(pm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_30, negated_conjecture, (e(exist)), clause_1).
cnf(c_0_31, plain, (~e(X1)), inference(pm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_32, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_30, c_0_31]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 22
# Proof object formula steps           : 11
# Proof object conjectures             : 3
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 11
# Proof object initial formulas used   : 11
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 1
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 26
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 26
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 26
# Processed clauses                    : 40
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 38
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 8
# Backward-rewritten                   : 0
# Generated clauses                    : 40
# ...of the previous two non-trivial   : 25
# Contextual simplify-reflections      : 0
# Paramodulations                      : 39
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 0
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 28
# Current number of unprocessed clauses: 11
# ...number of literals in the above   : 58
# Current number of archived formulas  : 0
# Current number of archived clauses   : 9
# Clause-clause subsumption calls (NU) : 208
# Rec. Clause-clause subsumption calls : 109
# Non-unit clause-clause subsumptions  : 10
# Unit Clause-clause subsumption calls : 18
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1534
