# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_16, plain, (cUnsatisfiable(i2003_11_14_17_20_53634))).
#
#cnf(i_0_15, plain, (cd(X1)|~ca_Ax2(X1))).
#
#cnf(i_0_5, plain, (cdxcomp(X1)|~cUnsatisfiable(X1))).
#
#cnf(i_0_1, plain, (~cowlNothing(X1))).
#
#cnf(i_0_6, plain, (cc(X1)|~cUnsatisfiable(X1))).
#
#cnf(i_0_3, plain, (xsd_integer(X1)|xsd_string(X1))).
#
#cnf(i_0_10, plain, (cdxcomp(X1)|~ra_Px1(X1,X2))).
#
#cnf(i_0_4, plain, (~xsd_integer(X1)|~xsd_string(X1))).
#
#cnf(i_0_12, plain, (ca_Ax2(X1)|~cc(esk3_1(X1)))).
#
#cnf(i_0_18, plain, (ca_Ax2(X1)|~cUnsatisfiable(esk3_1(X1)))).
#
#cnf(i_0_11, plain, (ra_Px1(X1,esk2_1(X1))|~cdxcomp(X1))).
#
#cnf(i_0_9, plain, (~ra_Px1(X1,X2)|~cd(X1))).
#
#cnf(i_0_8, plain, (ra_Px1(X1,esk1_1(X1))|cd(X1))).
#
#cnf(i_0_21, plain, (cd(X1)|cdxcomp(X1))).
#
#cnf(i_0_13, plain, (ca_Ax2(X1)|rr(X1,esk3_1(X1)))).
#
#cnf(i_0_20, plain, (~cd(X1)|~cdxcomp(X1))).
#
#cnf(i_0_7, plain, (cc(X2)|~rr(X1,X2)|~cc(X1))).
#
#cnf(i_0_14, plain, (cc(X2)|~ca_Ax2(X1)|~rr(X1,X2))).
#
#cnf(i_0_25, plain, (ca_Ax2(X1)|cc(esk3_1(X1))|~cc(X1))).
#
#cnf(i_0_23, plain, (~cd(X1)|~cUnsatisfiable(X1))).
#
#cnf(i_0_27, plain, (ca_Ax2(X1)|~cc(X1))).
#
#cnf(i_0_28, plain, (ca_Ax2(X1)|~cUnsatisfiable(X1))).
#
#cnf(i_0_29, plain, (cd(X1)|~cUnsatisfiable(X1))).
#
#cnf(i_0_30, plain, (~cUnsatisfiable(X1))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
fof(axiom_4, axiom, ![X1]:(cc(X1)=>![X2]:(rr(X1,X2)=>cc(X2))), file('./p/KRS105+1.tptp', axiom_4)).
fof(axiom_7, axiom, ![X1]:(ca_Ax2(X1)<=>![X2]:(rr(X1,X2)=>cc(X2))), file('./p/KRS105+1.tptp', axiom_7)).
fof(axiom_5, axiom, ![X1]:(cd(X1)<=>~(?[X2]:ra_Px1(X1,X2))), file('./p/KRS105+1.tptp', axiom_5)).
fof(axiom_6, axiom, ![X1]:(cdxcomp(X1)<=>?[X3]:ra_Px1(X1,X3)), file('./p/KRS105+1.tptp', axiom_6)).
fof(axiom_3, axiom, ![X1]:(cUnsatisfiable(X1)=>cc(X1)), file('./p/KRS105+1.tptp', axiom_3)).
fof(axiom_2, axiom, ![X1]:(cUnsatisfiable(X1)=>cdxcomp(X1)), file('./p/KRS105+1.tptp', axiom_2)).
fof(axiom_8, axiom, ![X1]:(ca_Ax2(X1)=>cd(X1)), file('./p/KRS105+1.tptp', axiom_8)).
fof(axiom_9, axiom, cUnsatisfiable(i2003_11_14_17_20_53634), file('./p/KRS105+1.tptp', axiom_9)).
fof(c_0_8, plain, ![X8, X9]:(~cc(X8)|(~rr(X8,X9)|cc(X9))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[axiom_4])])])).
fof(c_0_9, plain, ![X18, X19, X20]:((~ca_Ax2(X18)|(~rr(X18,X19)|cc(X19)))&((rr(X20,esk3_1(X20))|ca_Ax2(X20))&(~cc(esk3_1(X20))|ca_Ax2(X20)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[axiom_7])])])])])])).
cnf(c_0_10, plain, (cc(X2)|~cc(X1)|~rr(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, plain, (rr(X1,esk3_1(X1))|ca_Ax2(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_12, plain, ![X10, X11, X12]:((~cd(X10)|~ra_Px1(X10,X11))&(ra_Px1(X12,esk1_1(X12))|cd(X12))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[axiom_5])])])])])).
fof(c_0_13, plain, ![X14, X16, X17]:((~cdxcomp(X14)|ra_Px1(X14,esk2_1(X14)))&(~ra_Px1(X16,X17)|cdxcomp(X16))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[axiom_6])])])])])).
cnf(c_0_14, plain, (ca_Ax2(X1)|~cc(esk3_1(X1))), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (ca_Ax2(X1)|cc(esk3_1(X1))|~cc(X1)), inference(pm,[status(thm)],[c_0_10, c_0_11])).
fof(c_0_16, plain, ![X7]:(~cUnsatisfiable(X7)|cc(X7)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[axiom_3])])).
cnf(c_0_17, plain, (~cd(X1)|~ra_Px1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (ra_Px1(X1,esk2_1(X1))|~cdxcomp(X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
fof(c_0_19, plain, ![X6]:(~cUnsatisfiable(X6)|cdxcomp(X6)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[axiom_2])])).
fof(c_0_20, plain, ![X22]:(~ca_Ax2(X22)|cd(X22)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[axiom_8])])).
cnf(c_0_21, plain, (ca_Ax2(X1)|~cc(X1)), inference(pm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_22, plain, (cc(X1)|~cUnsatisfiable(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_23, plain, (~cd(X1)|~cdxcomp(X1)), inference(pm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_24, plain, (cdxcomp(X1)|~cUnsatisfiable(X1)), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (cd(X1)|~ca_Ax2(X1)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, plain, (ca_Ax2(X1)|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_27, plain, (~cd(X1)|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_28, plain, (cd(X1)|~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_29, plain, (cUnsatisfiable(i2003_11_14_17_20_53634)), inference(split_conjunct,[status(thm)],[axiom_9])).
cnf(c_0_30, plain, (~cUnsatisfiable(X1)), inference(pm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_31, plain, ($false), inference(sr,[status(thm)],[c_0_29, c_0_30]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 32
# Proof object clause steps            : 17
# Proof object formula steps           : 15
# Proof object conjectures             : 0
# Proof object clause conjectures      : 0
# Proof object formula conjectures     : 0
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 1
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 10
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 15
# Processed clauses                    : 24
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 24
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 15
# ...of the previous two non-trivial   : 9
# Contextual simplify-reflections      : 0
# Paramodulations                      : 14
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 22
#    Positive orientable unit clauses  : 0
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 20
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3
# Clause-clause subsumption calls (NU) : 17
# Rec. Clause-clause subsumption calls : 15
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 916
