# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_30, plain, (cEuropeanCountry(iBE))).
#
#cnf(i_0_32, plain, (cEuropeanCountry(iFR))).
#
#cnf(i_0_31, plain, (cEuropeanCountry(iES))).
#
#cnf(i_0_13, plain, (~cowlNothing(X1))).
#
#cnf(i_0_36, plain, (cEuropeanCountry(iUK))).
#
#cnf(i_0_34, plain, (cEuropeanCountry(iNL))).
#
#cnf(i_0_35, plain, (cEuropeanCountry(iPT))).
#
#cnf(i_0_16, plain, (~xsd_string(X1)|~xsd_integer(X1))).
#
#cnf(i_0_33, plain, (cPerson(iKinnock))).
#
#cnf(i_0_15, plain, (xsd_string(X1)|xsd_integer(X1))).
#
#cnf(i_0_37, plain, (rhasEuroMP(iUK,iKinnock))).
#
#cnf(i_0_27, plain, (cEUCountry(X1)|~rhasEuroMP(X1,X2))).
#
#cnf(i_0_41, plain, (cEUCountry(iUK))).
#
#cnf(i_0_24, plain, (cEuroMP(X1)|~risEuroMPFrom(X1,X2))).
#
#cnf(i_0_22, plain, (cEUCountry(X1)|X1!=iBE)).
#
#cnf(i_0_21, plain, (cEUCountry(X1)|X1!=iFR)).
#
#cnf(i_0_20, plain, (cEUCountry(X1)|X1!=iES)).
#
#cnf(i_0_19, plain, (cEUCountry(X1)|X1!=iUK)).
#
#cnf(i_0_18, plain, (cEUCountry(X1)|X1!=iNL)).
#
#cnf(i_0_17, plain, (cEUCountry(X1)|X1!=iPT)).
#
#cnf(i_0_29, plain, (rhasEuroMP(X2,X1)|~risEuroMPFrom(X1,X2))).
#
#cnf(i_0_28, plain, (risEuroMPFrom(X2,X1)|~rhasEuroMP(X1,X2))).
#
#cnf(i_0_42, plain, (cEuroMP(X1)|~rhasEuroMP(X2,X1))).
#
#cnf(i_0_44, plain, (cEuroMP(iKinnock))).
#
#cnf(i_0_26, plain, (risEuroMPFrom(X1,esk1_1(X1))|~cEuroMP(X1))).
#
#cnf(i_0_46, plain, (rhasEuroMP(esk1_1(X1),X1)|~cEuroMP(X1))).
#
#cnf(i_0_47, plain, (cEUCountry(esk1_1(X1))|~cEuroMP(X1))).
#
#cnf(i_0_39, negated_conjecture, (xsd_integer(esk3_0)|~xsd_string(esk3_0))).
#
#cnf(i_0_49, negated_conjecture, (xsd_integer(esk3_0))).
#
#cnf(i_0_38, negated_conjecture, (xsd_string(esk3_0))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(axiom_3, axiom, ![X4]:(cEuroMP(X4)<=>?[X5]:(risEuroMPFrom(X4,X5)&cowlThing(X5))), file('./p/KRS129+1.tptp', axiom_3)).
fof(axiom_0, axiom, ![X4]:(cowlThing(X4)&~(cowlNothing(X4))), file('./p/KRS129+1.tptp', axiom_0)).
fof(the_axiom, conjecture, ((![X4]:(cowlThing(X4)&~(cowlNothing(X4)))&![X4]:(xsd_string(X4)<=>~(xsd_integer(X4))))&cEuroMP(iKinnock)), file('./p/KRS129+1.tptp', the_axiom)).
fof(axiom_5, axiom, ![X4, X5]:(risEuroMPFrom(X4,X5)<=>rhasEuroMP(X5,X4)), file('./p/KRS129+1.tptp', axiom_5)).
fof(axiom_13, axiom, rhasEuroMP(iUK,iKinnock), file('./p/KRS129+1.tptp', axiom_13)).
fof(axiom_1, axiom, ![X4]:(xsd_string(X4)<=>~(xsd_integer(X4))), file('./p/KRS129+1.tptp', axiom_1)).
fof(c_0_6, plain, ![X37, X39, X40]:(((risEuroMPFrom(X37,esk1_1(X37))|~cEuroMP(X37))&(cowlThing(esk1_1(X37))|~cEuroMP(X37)))&(~risEuroMPFrom(X39,X40)|~cowlThing(X40)|cEuroMP(X39))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[axiom_3])])])])])])).
fof(c_0_7, plain, ![X34]:(cowlThing(X34)&~cowlNothing(X34)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[axiom_0])])).
fof(c_0_8, negated_conjecture, ~(((![X4]:(cowlThing(X4)&~(cowlNothing(X4)))&![X4]:(xsd_string(X4)<=>~(xsd_integer(X4))))&cEuroMP(iKinnock))), inference(assume_negation,[status(cth)],[the_axiom])).
cnf(c_0_9, plain, (cEuroMP(X1)|~risEuroMPFrom(X1,X2)|~cowlThing(X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (cowlThing(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_11, plain, ![X43, X44]:((~risEuroMPFrom(X43,X44)|rhasEuroMP(X44,X43))&(~rhasEuroMP(X44,X43)|risEuroMPFrom(X43,X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[axiom_5])])).
fof(c_0_12, negated_conjecture, ((~xsd_string(esk3_0)|xsd_integer(esk3_0)|(~cowlThing(esk2_0)|cowlNothing(esk2_0))|~cEuroMP(iKinnock))&(xsd_string(esk3_0)|~xsd_integer(esk3_0)|(~cowlThing(esk2_0)|cowlNothing(esk2_0))|~cEuroMP(iKinnock))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_8])])])])])).
cnf(c_0_13, plain, (cEuroMP(X1)|~risEuroMPFrom(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_10])])).
cnf(c_0_14, plain, (risEuroMPFrom(X2,X1)|~rhasEuroMP(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_15, negated_conjecture, (xsd_integer(esk3_0)|cowlNothing(esk2_0)|~xsd_string(esk3_0)|~cowlThing(esk2_0)|~cEuroMP(iKinnock)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_16, plain, (cEuroMP(X1)|~rhasEuroMP(X2,X1)), inference(pm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_17, plain, (rhasEuroMP(iUK,iKinnock)), inference(split_conjunct,[status(thm)],[axiom_13])).
cnf(c_0_18, negated_conjecture, (cowlNothing(esk2_0)|xsd_integer(esk3_0)|~cEuroMP(iKinnock)|~xsd_string(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_10])])).
cnf(c_0_19, plain, (cEuroMP(iKinnock)), inference(pm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_20, plain, (~cowlNothing(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_21, plain, ![X35]:((~xsd_string(X35)|~xsd_integer(X35))&(xsd_integer(X35)|xsd_string(X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[axiom_1])])])).
cnf(c_0_22, negated_conjecture, (xsd_string(esk3_0)|cowlNothing(esk2_0)|~xsd_integer(esk3_0)|~cowlThing(esk2_0)|~cEuroMP(iKinnock)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_23, negated_conjecture, (xsd_integer(esk3_0)|~xsd_string(esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_19])]), c_0_20])).
cnf(c_0_24, plain, (xsd_integer(X1)|xsd_string(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_25, negated_conjecture, (cowlNothing(esk2_0)|xsd_string(esk3_0)|~cEuroMP(iKinnock)|~xsd_integer(esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_22, c_0_10])])).
cnf(c_0_26, negated_conjecture, (xsd_integer(esk3_0)), inference(pm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_27, plain, (~xsd_string(X1)|~xsd_integer(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, negated_conjecture, (xsd_string(esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_25, c_0_19]), c_0_26])]), c_0_20])).
cnf(c_0_29, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_27, c_0_28]), c_0_26])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 30
# Proof object clause steps            : 18
# Proof object formula steps           : 12
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 27
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 39
# Removed in clause preprocessing      : 14
# Initial clauses in saturation        : 25
# Processed clauses                    : 30
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 30
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 11
# ...of the previous two non-trivial   : 6
# Contextual simplify-reflections      : 0
# Paramodulations                      : 11
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 12
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 7
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 61
# Rec. Clause-clause subsumption calls : 61
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1898
