# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_3, plain, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(join(meet(X1,join(join(meet(X4,X2),meet(X2,X5)),X2)),meet(join(meet(X2,meet(meet(join(X4,join(X2,X5)),join(X6,X2)),X2)),meet(X7,join(X2,meet(meet(join(X4,join(X2,X5)),join(X6,X2)),X2)))),join(X1,join(join(meet(X4,X2),meet(X2,X5)),X2)))),join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2)).
#
#cnf(i_0_11, plain, (join(meet(join(meet(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),join(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))))),X4),meet(join(meet(X1,join(join(meet(X5,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),X6)),join(meet(X2,X3),meet(X3,join(X2,X3))))),meet(X3,join(X1,join(join(meet(X5,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),X6)),join(meet(X2,X3),meet(X3,join(X2,X3))))))),join(join(meet(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),join(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))))),X4)))=join(meet(X2,X3),meet(X3,join(X2,X3))))).
#
#cnf(i_0_53, plain, (join(meet(join(meet(X1,X3),meet(X3,join(X1,X3))),X8),meet(join(meet(X1,join(join(meet(X9,X3),meet(X3,X10)),X3)),meet(join(meet(X2,join(join(meet(X4,X3),meet(X3,X5)),X3)),meet(join(meet(X3,meet(meet(join(X4,join(X3,X5)),join(X6,X3)),X3)),meet(X7,join(X3,meet(meet(join(X4,join(X3,X5)),join(X6,X3)),X3)))),join(X2,join(join(meet(X4,X3),meet(X3,X5)),X3)))),join(X1,join(join(meet(X9,X3),meet(X3,X10)),X3)))),join(join(meet(X1,X3),meet(X3,join(X1,X3))),X8)))=X3)).
#
#cnf(i_0_4, negated_conjecture, (meet(meet(a,b),c)!=meet(a,meet(b,c))|join(join(a,b),c)!=join(a,join(b,c))|meet(b,a)!=meet(a,b)|join(b,a)!=join(a,b)|meet(a,join(a,b))!=a|join(a,meet(a,b))!=a|meet(a,a)!=a|join(a,a)!=a)).
#
#cnf(i_0_148, plain, (join(meet(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X2),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X2))),X3),meet(X2,join(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X2),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X2))),X3)))=X2)).
#
#cnf(i_0_217, plain, (join(meet(X2,X3),meet(X2,join(X2,X3)))=X2)).
##
#cnf(i_0_251, plain, (join(meet(meet(X1,X2),meet(X1,join(X1,X2))),meet(meet(X1,X2),X1))=meet(X1,X2))).
#####
#cnf(i_0_320, plain, (join(meet(meet(meet(X1,X2),meet(X1,join(X1,X2))),meet(meet(X1,X2),X1)),meet(meet(meet(X1,X2),meet(X1,join(X1,X2))),meet(X1,X2)))=meet(meet(X1,X2),meet(X1,join(X1,X2))))).
##
#cnf(i_0_205, plain, (join(meet(join(meet(X1,X3),meet(X3,join(X1,X3))),X4),meet(join(meet(X1,join(join(meet(X5,X3),meet(X3,X6)),X3)),meet(X3,join(X1,join(join(meet(X5,X3),meet(X3,X6)),X3)))),join(join(meet(X1,X3),meet(X3,join(X1,X3))),X4)))=X3)).
#
#cnf(i_0_426, plain, (join(meet(X1,join(join(meet(X3,X2),meet(X2,X4)),X2)),meet(X2,join(X1,join(join(meet(X3,X2),meet(X2,X4)),X2))))=X2)).
#
#cnf(i_0_536, plain, (join(meet(X1,join(X3,X3)),meet(X3,join(X1,join(X3,X3))))=X3)).
####
#cnf(i_0_551, plain, (join(meet(meet(X1,join(X2,X2)),meet(X2,join(X1,join(X2,X2)))),meet(meet(X1,join(X2,X2)),X2))=meet(X1,join(X2,X2)))).
#
#cnf(i_0_205, plain, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2)).
#
#cnf(i_0_506, plain, (join(meet(meet(X1,join(join(meet(X2,X3),meet(X3,X4)),X3)),meet(X3,join(X1,join(join(meet(X2,X3),meet(X3,X4)),X3)))),meet(meet(X1,join(join(meet(X2,X3),meet(X3,X4)),X3)),X3))=meet(X1,join(join(meet(X2,X3),meet(X3,X4)),X3)))).
#
#cnf(i_0_614, plain, (join(meet(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3))),meet(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),X2))=meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3))).
#
#cnf(i_0_684, plain, (join(meet(join(meet(meet(X1,X2),meet(X1,join(X1,X2))),meet(meet(X1,join(X1,X2)),X1)),X3),meet(meet(X1,join(X1,X2)),join(join(meet(meet(X1,X2),meet(X1,join(X1,X2))),meet(meet(X1,join(X1,X2)),X1)),X3)))=meet(X1,join(X1,X2)))).
###############
#cnf(i_0_552, plain, (join(meet(meet(meet(X1,join(X2,X2)),meet(X2,join(X1,join(X2,X2)))),meet(meet(X1,join(X2,X2)),X2)),meet(meet(meet(X1,join(X2,X2)),meet(X2,join(X1,join(X2,X2)))),meet(X1,join(X2,X2))))=meet(meet(X1,join(X2,X2)),meet(X2,join(X1,join(X2,X2)))))).
##
#cnf(i_0_530, plain, (join(meet(X1,join(join(meet(X4,X2),meet(X2,X5)),X2)),meet(join(meet(X2,meet(meet(join(X4,join(X2,X5)),join(X6,X2)),X2)),meet(X7,join(X2,meet(meet(join(X4,join(X2,X5)),join(X6,X2)),X2)))),join(X1,join(join(meet(X4,X2),meet(X2,X5)),X2))))=X2)).
#
#cnf(i_0_1029, plain, (join(meet(X2,meet(meet(join(X3,join(X2,X4)),join(X5,X2)),X2)),meet(X6,join(X2,meet(meet(join(X3,join(X2,X4)),join(X5,X2)),X2))))=X2)).
#
#cnf(i_0_1123, plain, (join(meet(X1,X2),meet(X2,join(X1,X2)))=X2)).
#
#cnf(i_0_1208, plain, (meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)=X1)).
#
#cnf(i_0_1029, plain, (join(meet(X1,X1),meet(X5,join(X1,X1)))=X1)).
#
#cnf(i_0_1290, plain, (join(X1,X1)=X1)).
#
#cnf(i_0_4, negated_conjecture, (meet(meet(a,b),c)!=meet(a,meet(b,c))|join(join(a,b),c)!=join(a,join(b,c))|meet(b,a)!=meet(a,b)|join(b,a)!=join(a,b)|meet(a,join(a,b))!=a|join(a,meet(a,b))!=a|meet(a,a)!=a)).
#
#cnf(i_0_1307, plain, (meet(X1,X1)=X1)).
##
#cnf(i_0_1029, plain, (join(X1,meet(X2,X1))=X1)).
#
#cnf(i_0_4, negated_conjecture, (meet(meet(a,b),c)!=meet(a,meet(b,c))|join(join(a,b),c)!=join(a,join(b,c))|meet(b,a)!=meet(a,b)|join(b,a)!=join(a,b)|meet(a,join(a,b))!=a|join(a,meet(a,b))!=a)).
#
#cnf(i_0_1349, plain, (join(meet(X1,meet(X2,X1)),X1)=X1)).
#
#cnf(i_0_1306, plain, (meet(meet(join(X1,X2),join(X3,X1)),X1)=X1)).
#
#cnf(i_0_1413, plain, (meet(meet(join(X1,X2),X1),X1)=X1)).
##
#cnf(i_0_1424, plain, (meet(meet(X1,join(X3,X1)),X1)=X1)).
##
#cnf(i_0_1317, plain, (meet(meet(join(X1,X2),join(X3,X2)),X2)=X2)).
#
#cnf(i_0_1548, plain, (meet(join(X1,X2),X2)=X2)).
#
#cnf(i_0_1586, plain, (meet(X1,meet(X2,X1))=meet(X2,X1))).
#
#cnf(i_0_1349, plain, (join(meet(X2,X1),X1)=X1)).
##
#cnf(i_0_1590, plain, (meet(meet(X2,X1),X1)=meet(X2,X1))).
#
#cnf(i_0_1413, plain, (meet(join(X1,X2),X1)=X1)).
###
#cnf(i_0_1717, plain, (meet(X1,meet(X1,X2))=meet(X1,X2))).
##
#cnf(i_0_1318, plain, (meet(join(X1,join(X2,X3)),X2)=X2)).
#
#cnf(i_0_1780, plain, (meet(join(X1,X3),meet(X2,X3))=meet(X2,X3))).
#
#cnf(i_0_1781, plain, (meet(join(X1,X2),meet(X2,X3))=meet(X2,X3))).
#
#cnf(i_0_1864, plain, (meet(join(X1,join(X2,X3)),X3)=X3)).
###
#cnf(i_0_1561, plain, (join(X2,meet(X2,join(join(X1,X2),X2)))=X2)).
#
#cnf(i_0_1928, plain, (meet(X1,join(join(X2,X1),X1))=X1)).
##
#cnf(i_0_1693, plain, (join(X1,meet(X1,join(join(X1,X2),X1)))=X1)).
#
#cnf(i_0_2002, plain, (meet(X1,join(join(X1,X2),X1))=X1)).
#
#cnf(i_0_1862, plain, (meet(join(X1,join(X2,join(X3,X4))),X3)=X3)).
#
#cnf(i_0_1868, plain, (meet(join(X1,meet(X2,join(X3,X2))),X2)=X2)).
#
#cnf(i_0_1903, plain, (meet(join(X1,join(X2,join(X3,X4))),X4)=X4)).
#####
#cnf(i_0_1832, plain, (meet(X1,meet(X3,meet(X2,X1)))=meet(X3,meet(X2,X1)))).
#
#cnf(i_0_1881, plain, (meet(X1,meet(meet(X2,X1),X3))=meet(meet(X2,X1),X3))).
#
#cnf(i_0_1972, plain, (join(X1,join(join(X2,X1),X1))=join(join(X2,X1),X1))).
#
#cnf(i_0_1986, plain, (join(join(join(X1,X2),X2),X2)=join(join(X1,X2),X2))).
#
#cnf(i_0_2048, plain, (join(X1,join(join(X1,X2),X1))=join(join(X1,X2),X1))).
#
#cnf(i_0_2062, plain, (join(join(join(X1,X2),X1),X1)=join(join(X1,X2),X1))).
###
#cnf(i_0_1851, plain, (meet(join(X1,join(X2,X3)),meet(X3,X4))=meet(X3,X4))).
#
#cnf(i_0_1861, plain, (meet(join(X1,join(X2,X3)),meet(X4,X3))=meet(X4,X3))).
#
#cnf(i_0_2072, plain, (meet(meet(X1,X2),join(X1,meet(X1,X2)))=meet(X1,X2))).
#
#cnf(i_0_2727, plain, (meet(X2,join(join(X1,join(X2,X3)),X2))=X2)).
#
#cnf(i_0_2729, plain, (meet(X3,join(join(X1,join(X2,X3)),X3))=X3)).
#
#cnf(i_0_2737, plain, (meet(X1,join(meet(X1,join(X2,X1)),X1))=X1)).
##########
#cnf(i_0_1722, plain, (join(meet(X1,X2),meet(X1,join(X1,meet(X1,X2))))=X1)).
##
#cnf(i_0_1866, plain, (meet(join(X1,meet(join(X2,X3),join(X4,X2))),X2)=X2)).
#
#cnf(i_0_1867, plain, (meet(join(X1,meet(join(X2,X3),join(X4,X3))),X3)=X3)).
########
#cnf(i_0_2096, plain, (meet(join(X1,join(X2,join(X3,join(X4,X5)))),X4)=X4)).
#
#cnf(i_0_2141, plain, (meet(join(X1,join(X2,meet(X3,join(X4,X3)))),X3)=X3)).
#
#cnf(i_0_2189, plain, (meet(join(X1,join(X2,join(X3,join(X4,X5)))),X5)=X5)).
####
#cnf(i_0_2728, plain, (meet(X3,join(join(X1,join(X2,join(X3,X4))),X3))=X3)).
#
#cnf(i_0_2730, plain, (meet(X4,join(join(X1,join(X2,join(X3,X4))),X4))=X4)).
#
#cnf(i_0_2731, plain, (meet(X2,join(join(X1,meet(X2,join(X3,X2))),X2))=X2)).
#
#cnf(i_0_2735, plain, (meet(X1,join(meet(join(X1,X2),join(X3,X1)),X1))=X1)).
#
#cnf(i_0_2736, plain, (meet(X2,join(meet(join(X1,X2),join(X3,X2)),X2))=X2)).
#########
#cnf(i_0_2232, plain, (meet(join(X1,X2),meet(X3,meet(X4,X2)))=meet(X3,meet(X4,X2)))).
#
#cnf(i_0_2297, plain, (meet(join(X1,X2),meet(meet(X3,X2),X4))=meet(meet(X3,X2),X4))).
#
#cnf(i_0_2483, plain, (join(meet(X1,X2),join(X1,meet(X1,X2)))=join(X1,meet(X1,X2)))).
#
#cnf(i_0_2538, plain, (join(join(X1,meet(X1,X2)),meet(X1,X2))=join(X1,meet(X1,X2)))).
########
#cnf(i_0_1414, plain, (meet(meet(join(meet(X1,X2),X3),X2),meet(X1,X2))=meet(X1,X2))).
#
#cnf(i_0_1421, plain, (meet(meet(X2,join(X3,meet(X1,X2))),meet(X1,X2))=meet(X1,X2))).
#
#cnf(i_0_1422, plain, (meet(meet(X1,join(X3,meet(X1,X2))),meet(X1,X2))=meet(X1,X2))).
#
#cnf(i_0_4662, plain, (meet(meet(join(X1,join(X2,X3)),join(X4,X3)),X3)=X3)).
#
#cnf(i_0_4678, plain, (meet(meet(meet(X1,join(X2,X1)),join(X3,X1)),X1)=X1)).
###
#cnf(i_0_1536, plain, (meet(meet(join(X1,meet(X2,X3)),X3),meet(X2,X3))=meet(X2,X3))).
##
#cnf(i_0_1552, plain, (join(X2,meet(join(X1,X2),join(join(X1,X2),X2)))=join(X1,X2))).
##
#cnf(i_0_1684, plain, (join(X1,meet(join(X1,X2),join(join(X1,X2),X1)))=join(X1,X2))).
#
#cnf(i_0_2560, plain, (meet(join(X1,join(X2,join(X3,X4))),meet(X4,X5))=meet(X4,X5))).
##
#cnf(i_0_2574, plain, (meet(join(X1,join(X2,join(X3,X4))),meet(X5,X4))=meet(X5,X4))).
##
#cnf(i_0_2725, plain, (meet(meet(X2,X3),join(join(X1,X2),meet(X2,X3)))=meet(X2,X3))).
#
#cnf(i_0_2726, plain, (meet(meet(X3,X2),join(join(X1,X2),meet(X3,X2)))=meet(X3,X2))).
########################
#cnf(i_0_1829, plain, (meet(X2,meet(X3,meet(X2,join(X1,X2))))=meet(X3,meet(X2,join(X1,X2))))).
#
#cnf(i_0_5991, plain, (meet(meet(X1,join(X2,X1)),meet(X3,X1))=meet(X3,X1))).
#
#cnf(i_0_6081, plain, (meet(join(X1,meet(X2,join(X3,X2))),meet(X4,X2))=meet(X4,X2))).
#
#cnf(i_0_1830, plain, (meet(X1,meet(X3,meet(X1,join(X1,X2))))=meet(X3,meet(X1,join(X1,X2))))).
#
#cnf(i_0_1865, plain, (meet(join(X1,meet(join(X2,join(X3,X4)),join(X5,X3))),X3)=X3)).
#
#cnf(i_0_1878, plain, (meet(X2,meet(meet(X2,join(X1,X2)),X3))=meet(meet(X2,join(X1,X2)),X3))).
#
#cnf(i_0_6657, plain, (meet(meet(X1,join(X2,X1)),meet(X1,X3))=meet(X1,X3))).
#
#cnf(i_0_6793, plain, (meet(meet(join(X1,X2),join(X3,join(X1,X2))),X1)=X1)).
#
#cnf(i_0_6794, plain, (meet(meet(join(X1,X2),join(X3,join(X1,X2))),X2)=X2)).
###
#cnf(i_0_6767, plain, (join(meet(X1,join(X1,join(X2,X1))),X1)=meet(X1,join(X2,X1)))).
#
#cnf(i_0_7164, plain, (meet(X1,join(X2,X1))=X1)).
#
#cnf(i_0_7262, plain, (join(X1,join(X2,X1))=join(X2,X1))).
#
#cnf(i_0_7291, plain, (join(join(X1,X2),X2)=join(X1,X2))).
#
#cnf(i_0_7410, plain, (join(meet(X1,join(X1,X2)),X1)=X1)).
#
#cnf(i_0_7496, plain, (join(X1,meet(X1,join(X1,X2)))=X1)).
#####
#cnf(i_0_7320, plain, (meet(meet(X1,join(X1,X2)),X1)=meet(X1,join(X1,X2)))).
####
#cnf(i_0_7407, plain, (join(meet(meet(X1,X2),X1),meet(X1,X2))=meet(X1,X2))).
#
#cnf(i_0_7790, plain, (join(meet(X2,join(X1,join(X2,X3))),X2)=X2)).
#
#cnf(i_0_7792, plain, (join(meet(X3,join(X1,join(X2,X3))),X3)=X3)).
####
#cnf(i_0_7493, plain, (join(meet(X1,X2),meet(meet(X1,X2),X1))=meet(X1,X2))).
#
#cnf(i_0_8107, plain, (join(X2,meet(X2,join(X1,join(X2,X3))))=X2)).
#
#cnf(i_0_8109, plain, (join(X3,meet(X3,join(X1,join(X2,X3))))=X3)).
########
#cnf(i_0_7791, plain, (join(meet(X3,join(X1,join(X2,join(X3,X4)))),X3)=X3)).
#
#cnf(i_0_7793, plain, (join(meet(X4,join(X1,join(X2,join(X3,X4)))),X4)=X4)).
#
#cnf(i_0_7806, plain, (join(meet(X1,meet(join(X1,X2),join(X3,X1))),X1)=X1)).
#
#cnf(i_0_7807, plain, (join(meet(X2,meet(join(X1,X2),join(X3,X2))),X2)=X2)).
#####
#cnf(i_0_8108, plain, (join(X3,meet(X3,join(X1,join(X2,join(X3,X4)))))=X3)).
#
#cnf(i_0_8110, plain, (join(X4,meet(X4,join(X1,join(X2,join(X3,X4)))))=X4)).
#
#cnf(i_0_8123, plain, (join(X1,meet(X1,meet(join(X1,X2),join(X3,X1))))=X1)).
#
#cnf(i_0_8124, plain, (join(X2,meet(X2,meet(join(X1,X2),join(X3,X2))))=X2)).
#########
#cnf(i_0_7317, plain, (meet(meet(meet(X1,X2),X1),meet(X1,X2))=meet(meet(X1,X2),X1))).
######
#cnf(i_0_7788, plain, (join(meet(meet(X2,X3),join(X1,X2)),meet(X2,X3))=meet(X2,X3))).
#
#cnf(i_0_7789, plain, (join(meet(meet(X3,X2),join(X1,X2)),meet(X3,X2))=meet(X3,X2))).
#######
#cnf(i_0_8105, plain, (join(meet(X2,X3),meet(meet(X2,X3),join(X1,X2)))=meet(X2,X3))).
#
#cnf(i_0_8106, plain, (join(meet(X3,X2),meet(meet(X3,X2),join(X1,X2)))=meet(X3,X2))).
###########
#cnf(i_0_1879, plain, (meet(X1,meet(meet(X1,join(X1,X2)),X3))=meet(meet(X1,join(X1,X2)),X3))).
#
#cnf(i_0_2279, plain, (meet(X1,meet(meet(X3,meet(X2,X1)),X4))=meet(meet(X3,meet(X2,X1)),X4))).
#
#cnf(i_0_2331, plain, (meet(X1,meet(X3,meet(X4,meet(X2,X1))))=meet(X3,meet(X4,meet(X2,X1))))).
#
#cnf(i_0_2576, plain, (meet(join(X1,join(X2,join(X3,join(X4,join(X5,X6))))),X5)=X5)).
#
#cnf(i_0_10564, plain, (join(join(meet(X5,X6),meet(X6,X7)),X6)=X6)).
#
#cnf(i_0_10905, plain, (join(meet(X1,X2),X1)=X1)).
#
#cnf(i_0_10945, plain, (meet(meet(X1,X2),X1)=meet(X1,X2))).
#
#cnf(i_0_11074, plain, (join(X1,meet(X1,X2))=X1)).
#
#cnf(i_0_4, negated_conjecture, (meet(meet(a,b),c)!=meet(a,meet(b,c))|join(join(a,b),c)!=join(a,join(b,c))|meet(b,a)!=meet(a,b)|join(b,a)!=join(a,b)|meet(a,join(a,b))!=a)).
#
#cnf(i_0_11106, plain, (meet(X1,join(X1,X2))=X1)).
#
#cnf(i_0_11019, plain, (join(X1,join(X1,X2))=join(X1,X2))).
#
#cnf(i_0_11189, plain, (join(join(X1,X2),X1)=join(X1,X2))).
#
#cnf(i_0_4, negated_conjecture, (meet(meet(a,b),c)!=meet(a,meet(b,c))|join(join(a,b),c)!=join(a,join(b,c))|meet(b,a)!=meet(a,b)|join(b,a)!=join(a,b))).
#####################
#cnf(i_0_11034, plain, (join(meet(meet(X2,X1),X3),X1)=X1)).
#
#cnf(i_0_11036, plain, (join(meet(X2,meet(X3,X1)),X1)=X1)).
#
#cnf(i_0_11096, plain, (meet(X2,join(X1,join(X2,X3)))=X2)).
#
#cnf(i_0_11098, plain, (meet(X3,join(X1,join(X2,X3)))=X3)).
#
#cnf(i_0_11203, plain, (join(X1,meet(meet(X2,X1),X3))=X1)).
#
#cnf(i_0_11205, plain, (join(X1,meet(X2,meet(X3,X1)))=X1)).
#
#cnf(i_0_11344, plain, (meet(join(join(X1,X2),X3),X1)=X1)).
#
#cnf(i_0_11347, plain, (meet(join(join(X1,X2),X3),X2)=X2)).
##
#cnf(i_0_11496, plain, (join(meet(meet(X1,X2),X3),X1)=X1)).
###
#cnf(i_0_11622, plain, (join(meet(X1,meet(X2,X3)),X2)=X2)).
##
#cnf(i_0_11719, plain, (meet(X1,join(join(X1,X2),X3))=X1)).
#
#cnf(i_0_11776, plain, (meet(X1,join(join(X2,X1),X3))=X1)).
#
#cnf(i_0_11854, plain, (join(X1,meet(meet(X1,X2),X3))=X1)).
##
#cnf(i_0_11959, plain, (join(X1,meet(X2,meet(X1,X3)))=X1)).
#########
#cnf(i_0_11007, plain, (join(meet(X2,X3),join(X1,X2))=join(X1,X2))).
#
#cnf(i_0_11008, plain, (join(meet(X3,X2),join(X1,X2))=join(X1,X2))).
#
#cnf(i_0_11094, plain, (meet(meet(X2,X3),join(X1,X2))=meet(X2,X3))).
#
#cnf(i_0_11095, plain, (meet(meet(X3,X2),join(X1,X2))=meet(X3,X2))).
#
#cnf(i_0_11169, plain, (meet(meet(X1,X2),meet(X2,X1))=meet(X2,X1))).
## meet is commutative
# AC handling enabled dynamically

#cnf(i_0_13349, plain, (meet(X1,X2)=meet(X2,X1))).
###
#cnf(i_0_4, negated_conjecture, (meet(c,meet(a,b))!=meet(a,meet(b,c))|join(join(a,b),c)!=join(a,join(b,c))|join(b,a)!=join(a,b))).
#######################
#cnf(i_0_11177, plain, (join(join(X1,X2),meet(X2,X3))=join(X1,X2))).
#
#cnf(i_0_11178, plain, (join(join(X1,X2),meet(X3,X2))=join(X1,X2))).
#
#cnf(i_0_11243, plain, (meet(join(X1,X2),meet(X1,X3))=meet(X1,X3))).
#
#cnf(i_0_11246, plain, (meet(join(X1,X2),meet(X3,X1))=meet(X3,X1))).
#####
#cnf(i_0_11498, plain, (join(meet(X1,X3),join(X1,X2))=join(X1,X2))).
##
#cnf(i_0_11624, plain, (join(meet(X1,X2),join(X2,X3))=join(X2,X3))).
######
#cnf(i_0_11858, plain, (join(join(X1,X2),meet(X1,X3))=join(X1,X2))).
##
#cnf(i_0_11963, plain, (join(join(X1,X2),meet(X3,X1))=join(X1,X2))).
##
#cnf(i_0_12501, plain, (meet(meet(X1,X2),join(X2,X3))=meet(X1,X2))).
#
#cnf(i_0_12502, plain, (meet(meet(X1,X2),join(X1,X3))=meet(X1,X2))).
#################
#cnf(i_0_13350, plain, (join(meet(X1,X2),meet(X2,X1))=meet(X1,X2))).
######################################
#cnf(i_0_10858, plain, (join(join(meet(X1,X2),meet(X3,X2)),X2)=X2)).
#
#cnf(i_0_11035, plain, (join(meet(meet(X2,meet(X3,X1)),X4),X1)=X1)).
#
#cnf(i_0_11038, plain, (join(meet(X2,meet(X3,meet(X4,X1))),X1)=X1)).
#
#cnf(i_0_11204, plain, (join(X1,meet(meet(X2,meet(X3,X1)),X4))=X1)).
#
#cnf(i_0_11207, plain, (join(X1,meet(X2,meet(X3,meet(X4,X1))))=X1)).
#
#cnf(i_0_11391, plain, (join(X2,join(meet(X1,X2),meet(X2,X3)))=X2)).
#####
#cnf(i_0_12260, plain, (join(meet(meet(meet(X2,X1),X3),X4),X1)=X1)).
##
#cnf(i_0_12271, plain, (join(meet(meet(X3,meet(X1,X2)),X4),X1)=X1)).
#
#cnf(i_0_12273, plain, (join(meet(X3,meet(X4,meet(X1,X2))),X1)=X1)).
#
#cnf(i_0_12398, plain, (join(meet(X1,meet(meet(X3,X2),X4)),X2)=X2)).
##
#cnf(i_0_12476, plain, (meet(X2,join(X1,join(join(X2,X3),X4)))=X2)).
#
#cnf(i_0_12552, plain, (meet(X3,join(X1,join(join(X2,X3),X4)))=X3)).
#
#cnf(i_0_12684, plain, (join(X1,meet(meet(meet(X2,X1),X3),X4))=X1)).
##
#cnf(i_0_12697, plain, (join(X1,meet(meet(X3,meet(X1,X2)),X4))=X1)).
#
#cnf(i_0_12699, plain, (join(X1,meet(X3,meet(X4,meet(X1,X2))))=X1)).
#
#cnf(i_0_12803, plain, (join(X1,meet(X2,meet(meet(X3,X1),X4)))=X1)).
###
#cnf(i_0_12932, plain, (join(meet(meet(meet(X1,X2),X3),X4),X1)=X1)).
####
#cnf(i_0_13068, plain, (join(meet(X1,meet(meet(X2,X3),X4)),X2)=X2)).
#####
#cnf(i_0_13334, plain, (meet(X1,join(X4,join(X2,join(X1,X3))))=X1)).
#
#cnf(i_0_13335, plain, (meet(X1,join(X4,join(X2,join(X3,X1))))=X1)).
#
#cnf(i_0_1317, plain, (meet(X2,meet(join(X1,X2),join(X3,X2)))=X2)).
#
#cnf(i_0_1306, plain, (meet(X1,meet(join(X1,X2),join(X3,X1)))=X1)).
###
#cnf(i_0_13530, plain, (join(join(meet(X2,X1),meet(X2,X3)),X2)=X2)).
#####
#cnf(i_0_13686, plain, (join(X1,meet(meet(meet(X1,X2),X3),X4))=X1)).
####
#cnf(i_0_13758, plain, (join(X1,meet(X4,meet(meet(X1,X2),X3)))=X1)).
###
#cnf(i_0_13919, plain, (meet(X1,join(join(join(X1,X2),X3),X4))=X1)).
#
#cnf(i_0_13920, plain, (meet(X2,join(join(join(X1,X2),X3),X4))=X2)).
#
#cnf(i_0_13921, plain, (meet(X2,join(join(X1,join(X2,X3)),X4))=X2)).
#
#cnf(i_0_13922, plain, (meet(X3,join(join(X1,join(X2,X3)),X4))=X3)).
#####################
#cnf(i_0_14616, plain, (join(X2,join(meet(X1,X2),meet(X3,X2)))=X2)).
#
#cnf(i_0_14696, plain, (join(join(meet(X2,X1),meet(X3,X2)),X2)=X2)).
#############################
#cnf(i_0_15231, plain, (join(X1,join(meet(X1,X2),meet(X1,X3)))=X1)).
##################################################
#cnf(i_0_17140, plain, (meet(X1,meet(join(X2,X1),join(X1,X3)))=X1)).
##
#cnf(i_0_17332, plain, (meet(X1,meet(join(X1,X2),join(X1,X3)))=X1)).
##############
#cnf(i_0_18721, plain, (join(X1,join(meet(X1,X2),meet(X3,X1)))=X1)).
###############
#cnf(i_0_11134, plain, (meet(X1,meet(X3,meet(X1,X2)))=meet(X3,meet(X1,X2)))).
###
#cnf(i_0_11135, plain, (meet(X1,meet(meet(X1,X2),X3))=meet(meet(X1,X2),X3))).
#######
#cnf(i_0_11676, plain, (join(X1,join(X2,join(X1,X3)))=join(X2,join(X1,X3)))).
#
#cnf(i_0_11696, plain, (join(join(X1,join(X2,X3)),X2)=join(X1,join(X2,X3)))).
#
#cnf(i_0_11733, plain, (join(X1,join(X2,join(X3,X1)))=join(X2,join(X3,X1)))).
#
#cnf(i_0_11753, plain, (join(join(X1,join(X2,X3)),X3)=join(X1,join(X2,X3)))).
#####
#cnf(i_0_12459, plain, (join(X1,join(join(X1,X2),X3))=join(join(X1,X2),X3))).
#
#cnf(i_0_12481, plain, (join(join(join(X1,X2),X3),X1)=join(join(X1,X2),X3))).
#
#cnf(i_0_12535, plain, (join(X1,join(join(X2,X1),X3))=join(join(X2,X1),X3))).
#
#cnf(i_0_12557, plain, (join(join(join(X1,X2),X3),X2)=join(join(X1,X2),X3))).
###########################################
#cnf(i_0_11005, plain, (join(meet(meet(X3,X2),X4),join(X1,X2))=join(X1,X2))).
#
#cnf(i_0_11006, plain, (join(meet(X3,meet(X4,X2)),join(X1,X2))=join(X1,X2))).
#
#cnf(i_0_11175, plain, (join(join(X1,X2),meet(meet(X3,X2),X4))=join(X1,X2))).
#
#cnf(i_0_11176, plain, (join(join(X1,X2),meet(X3,meet(X4,X2)))=join(X1,X2))).
#
#cnf(i_0_11229, plain, (meet(join(X1,join(X2,X3)),meet(X4,X2))=meet(X4,X2))).
#
#cnf(i_0_11237, plain, (meet(join(X1,join(X2,X3)),meet(X2,X4))=meet(X2,X4))).
###
#cnf(i_0_11345, plain, (meet(join(join(X1,X2),X3),meet(X2,X4))=meet(X2,X4))).
#
#cnf(i_0_11346, plain, (meet(join(join(X1,X2),X3),meet(X4,X2))=meet(X4,X2))).
#######
#cnf(i_0_12230, plain, (join(meet(meet(X2,X3),X4),join(X1,X2))=join(X1,X2))).
##
#cnf(i_0_12368, plain, (join(meet(X1,meet(X3,X4)),join(X2,X3))=join(X2,X3))).
##
#cnf(i_0_12483, plain, (meet(join(join(X1,X2),X3),meet(X1,X4))=meet(X1,X4))).
#
#cnf(i_0_12487, plain, (meet(join(join(X1,X2),X3),meet(X4,X1))=meet(X4,X1))).
#####
#cnf(i_0_12652, plain, (join(join(X1,X2),meet(meet(X2,X3),X4))=join(X1,X2))).
##
#cnf(i_0_12771, plain, (join(join(X1,X2),meet(X3,meet(X2,X4)))=join(X1,X2))).
##
#cnf(i_0_12862, plain, (meet(meet(X1,X2),join(join(X3,X1),X4))=meet(X1,X2))).
#
#cnf(i_0_12878, plain, (meet(meet(X1,X2),join(X3,join(X4,X1)))=meet(X1,X2))).
###
#cnf(i_0_12996, plain, (meet(meet(X1,X2),join(join(X3,X2),X4))=meet(X1,X2))).
###
#cnf(i_0_13012, plain, (meet(meet(X1,X2),join(X3,join(X4,X2)))=meet(X1,X2))).
###############
#cnf(i_0_13351, plain, (meet(meet(X2,X1),join(X3,meet(X1,X2)))=meet(X2,X1))).
######################
#cnf(i_0_13359, plain, (join(meet(meet(X2,X1),X3),meet(X1,X2))=meet(X1,X2))).
################################
#cnf(i_0_13360, plain, (join(meet(X1,X2),meet(meet(X2,X1),X3))=meet(X1,X2))).
####
#cnf(i_0_13366, plain, (join(meet(X1,meet(X3,X2)),meet(X2,X3))=meet(X2,X3))).
############################
#cnf(i_0_13375, plain, (join(meet(X1,X2),meet(X3,meet(X2,X1)))=meet(X1,X2))).
#######################
#cnf(i_0_13766, plain, (meet(meet(X1,X3),join(X4,join(X1,X2)))=meet(X1,X3))).
##
#cnf(i_0_13774, plain, (join(meet(meet(X1,X3),X4),join(X1,X2))=join(X1,X2))).
#
#cnf(i_0_13775, plain, (join(join(X1,X2),meet(meet(X1,X3),X4))=join(X1,X2))).
##
#cnf(i_0_13780, plain, (join(meet(X1,meet(X2,X4)),join(X2,X3))=join(X2,X3))).
##
#cnf(i_0_13790, plain, (join(join(X1,X2),meet(X3,meet(X1,X4)))=join(X1,X2))).
###
#cnf(i_0_13821, plain, (meet(meet(X2,X1),join(meet(X1,X2),X3))=meet(X2,X1))).
##########################
#cnf(i_0_13861, plain, (meet(meet(X3,X1),join(X4,join(X1,X2)))=meet(X3,X1))).
####
#cnf(i_0_13869, plain, (join(meet(meet(X3,X1),X4),join(X1,X2))=join(X1,X2))).
#
#cnf(i_0_13870, plain, (join(join(X1,X2),meet(meet(X3,X1),X4))=join(X1,X2))).
##
#cnf(i_0_13875, plain, (join(meet(X1,meet(X4,X2)),join(X2,X3))=join(X2,X3))).
##
#cnf(i_0_13885, plain, (join(join(X1,X2),meet(X3,meet(X4,X1)))=join(X1,X2))).
###
#cnf(i_0_13967, plain, (meet(meet(X1,X2),join(join(X1,X3),X4))=meet(X1,X2))).
########
#cnf(i_0_14051, plain, (meet(meet(X1,X2),join(join(X2,X3),X4))=meet(X1,X2))).
################################################################################################################################################################################################################################################################################################
#cnf(i_0_10853, plain, (join(join(meet(X1,X2),meet(meet(X3,X2),X4)),X2)=X2)).
##
#cnf(i_0_10855, plain, (join(join(meet(X1,X2),meet(X3,meet(X4,X2))),X2)=X2)).
##
#cnf(i_0_11510, plain, (join(meet(meet(X3,meet(X4,meet(X1,X2))),X5),X2)=X2)).
##
#cnf(i_0_11512, plain, (join(meet(X3,meet(X4,meet(X5,meet(X1,X2)))),X2)=X2)).
##
#cnf(i_0_11872, plain, (join(X1,meet(meet(X3,meet(X4,meet(X2,X1))),X5))=X1)).
##
#cnf(i_0_11874, plain, (join(X1,meet(X3,meet(X4,meet(X5,meet(X2,X1)))))=X1)).
##
#cnf(i_0_12261, plain, (join(meet(meet(meet(X2,meet(X3,X1)),X4),X5),X1)=X1)).
###
#cnf(i_0_12272, plain, (join(meet(meet(X3,meet(X4,meet(X1,X2))),X5),X1)=X1)).
##
#cnf(i_0_12274, plain, (join(meet(X3,meet(X4,meet(X5,meet(X1,X2)))),X1)=X1)).
##
#cnf(i_0_12399, plain, (join(meet(X1,meet(meet(X3,meet(X4,X2)),X5)),X2)=X2)).
###
#cnf(i_0_12478, plain, (meet(X3,join(X1,join(X2,join(join(X3,X4),X5))))=X3)).
##
#cnf(i_0_12554, plain, (meet(X4,join(X1,join(X2,join(join(X3,X4),X5))))=X4)).
##
#cnf(i_0_12685, plain, (join(X1,meet(meet(meet(X2,meet(X3,X1)),X4),X5))=X1)).
###
#cnf(i_0_12698, plain, (join(X1,meet(meet(X3,meet(X4,meet(X1,X2))),X5))=X1)).
##
#cnf(i_0_12700, plain, (join(X1,meet(X3,meet(X4,meet(X5,meet(X1,X2)))))=X1)).
##
#cnf(i_0_12804, plain, (join(X1,meet(X2,meet(meet(X3,meet(X4,X1)),X5)))=X1)).
###
#cnf(i_0_4662, plain, (meet(X3,meet(join(X1,join(X2,X3)),join(X4,X3)))=X3)).
##
#cnf(i_0_1208, plain, (meet(X2,meet(join(X1,join(X2,X3)),join(X4,X2)))=X2)).
##
#cnf(i_0_1867, plain, (meet(X3,join(X1,meet(join(X2,X3),join(X4,X3))))=X3)).
##
#cnf(i_0_1866, plain, (meet(X2,join(X1,meet(join(X2,X3),join(X4,X2))))=X2)).
##
#cnf(i_0_2189, plain, (meet(X5,join(X1,join(X2,join(X3,join(X4,X5)))))=X5)).
##
#cnf(i_0_2096, plain, (meet(X4,join(X1,join(X2,join(X3,join(X4,X5)))))=X4)).
##
#cnf(i_0_13992, plain, (join(meet(join(meet(X1,X2),meet(X2,X3)),X4),X2)=X2)).
##
#cnf(i_0_14077, plain, (join(meet(X1,join(meet(X2,X3),meet(X3,X4))),X3)=X3)).
##
#cnf(i_0_14186, plain, (join(X2,meet(join(meet(X1,X2),meet(X2,X3)),X4))=X2)).
##
#cnf(i_0_14270, plain, (join(X2,meet(X4,join(meet(X1,X2),meet(X2,X3))))=X2)).
##
#cnf(i_0_14618, plain, (join(X2,meet(join(meet(X1,X2),meet(X3,X2)),X4))=X2)).
##
#cnf(i_0_14619, plain, (join(X2,meet(X4,join(meet(X1,X2),meet(X3,X2))))=X2)).
##
#cnf(i_0_14622, plain, (join(meet(X1,join(meet(X2,X3),meet(X4,X3))),X3)=X3)).
##
#cnf(i_0_14627, plain, (join(meet(join(meet(X1,X2),meet(X3,X2)),X4),X2)=X2)).
##########################
#cnf(i_0_15195, plain, (join(X1,join(meet(X2,X1),meet(meet(X3,X1),X4)))=X1)).
##
#cnf(i_0_15197, plain, (join(X1,join(meet(X2,X1),meet(X3,meet(X4,X1))))=X1)).
##
#cnf(i_0_15264, plain, (join(X2,meet(meet(meet(meet(X1,X2),X3),X4),X5))=X2)).
##
#cnf(i_0_15265, plain, (join(X2,meet(X5,meet(meet(meet(X1,X2),X3),X4)))=X2)).
##
#cnf(i_0_15268, plain, (join(meet(X1,meet(meet(meet(X2,X3),X4),X5)),X3)=X3)).
##
#cnf(i_0_15273, plain, (join(meet(meet(meet(meet(X1,X2),X3),X4),X5),X2)=X2)).
####
#cnf(i_0_15342, plain, (join(meet(meet(X4,meet(meet(X1,X2),X3)),X5),X2)=X2)).
##
#cnf(i_0_15344, plain, (join(meet(X4,meet(X5,meet(meet(X1,X2),X3))),X2)=X2)).
##
#cnf(i_0_15376, plain, (join(X2,meet(meet(meet(X1,meet(X2,X3)),X4),X5))=X2)).
##
#cnf(i_0_15377, plain, (join(X2,meet(X5,meet(meet(X1,meet(X2,X3)),X4)))=X2)).
##
#cnf(i_0_15380, plain, (join(meet(X1,meet(meet(X2,meet(X3,X4)),X5)),X3)=X3)).
##
#cnf(i_0_15385, plain, (join(meet(meet(meet(X1,meet(X2,X3)),X4),X5),X2)=X2)).
############
#cnf(i_0_15597, plain, (join(X3,meet(meet(X1,meet(meet(X2,X3),X4)),X5))=X3)).
##
#cnf(i_0_15598, plain, (join(X3,meet(X5,meet(X1,meet(meet(X2,X3),X4))))=X3)).
#######
#cnf(i_0_15716, plain, (meet(X1,join(join(X2,join(join(X1,X3),X4)),X5))=X1)).
#####
#cnf(i_0_15867, plain, (meet(X1,join(join(X2,join(join(X3,X1),X4)),X5))=X1)).
################################
#cnf(i_0_16460, plain, (join(X1,meet(meet(meet(meet(X1,X2),X3),X4),X5))=X1)).
##
#cnf(i_0_16461, plain, (join(X1,meet(X5,meet(meet(meet(X1,X2),X3),X4)))=X1)).
##
#cnf(i_0_16464, plain, (join(meet(X1,meet(meet(meet(X2,X3),X4),X5)),X2)=X2)).
##
#cnf(i_0_16469, plain, (join(meet(meet(meet(meet(X1,X2),X3),X4),X5),X1)=X1)).
######
#cnf(i_0_16546, plain, (join(meet(meet(X4,meet(meet(X1,X2),X3)),X5),X1)=X1)).
##
#cnf(i_0_16548, plain, (join(meet(X4,meet(X5,meet(meet(X1,X2),X3))),X1)=X1)).
##
#cnf(i_0_16582, plain, (join(X2,meet(meet(X1,meet(meet(X2,X3),X4)),X5))=X2)).
##
#cnf(i_0_16583, plain, (join(X2,meet(X5,meet(X1,meet(meet(X2,X3),X4))))=X2)).
#########
#cnf(i_0_16714, plain, (meet(X1,join(join(X2,join(X3,join(X1,X4))),X5))=X1)).
#####
#cnf(i_0_16892, plain, (meet(X1,join(join(X2,join(X3,join(X4,X1))),X5))=X1)).
#####
#cnf(i_0_17070, plain, (meet(X1,join(meet(join(X2,X1),join(X3,X1)),X4))=X1)).
#####
#cnf(i_0_17262, plain, (meet(X1,join(meet(join(X1,X2),join(X3,X1)),X4))=X1)).
####
#cnf(i_0_17443, plain, (join(X1,meet(join(meet(X1,X2),meet(X1,X3)),X4))=X1)).
##
#cnf(i_0_17444, plain, (join(X1,meet(X4,join(meet(X1,X2),meet(X1,X3))))=X1)).
##
#cnf(i_0_17448, plain, (join(meet(X1,join(meet(X2,X3),meet(X2,X4))),X2)=X2)).
##
#cnf(i_0_17453, plain, (join(meet(join(meet(X1,X2),meet(X1,X3)),X4),X1)=X1)).
##
#cnf(i_0_17501, plain, (join(join(meet(X1,X2),meet(meet(X3,X1),X4)),X1)=X1)).
##
#cnf(i_0_17503, plain, (join(join(meet(X1,X2),meet(X3,meet(X4,X1))),X1)=X1)).
##
#cnf(i_0_17537, plain, (join(join(meet(meet(X2,X1),X3),meet(X1,X4)),X1)=X1)).
##
#cnf(i_0_17539, plain, (join(join(meet(X2,meet(X3,X1)),meet(X1,X4)),X1)=X1)).
####################
#cnf(i_0_17850, plain, (meet(X1,join(X5,join(join(join(X1,X2),X3),X4)))=X1)).
##
#cnf(i_0_17863, plain, (meet(X1,join(join(join(join(X1,X2),X3),X4),X5))=X1)).
####
#cnf(i_0_18040, plain, (meet(X1,join(X5,join(join(join(X2,X1),X3),X4)))=X1)).
##
#cnf(i_0_18053, plain, (meet(X1,join(join(join(join(X2,X1),X3),X4),X5))=X1)).
####
#cnf(i_0_18230, plain, (meet(X1,join(X5,join(join(X2,join(X1,X3)),X4)))=X1)).
##
#cnf(i_0_18243, plain, (meet(X1,join(join(join(X2,join(X1,X3)),X4),X5))=X1)).
####
#cnf(i_0_18420, plain, (meet(X1,join(X5,join(join(X2,join(X3,X1)),X4)))=X1)).
##
#cnf(i_0_18433, plain, (meet(X1,join(join(join(X2,join(X3,X1)),X4),X5))=X1)).
########
#cnf(i_0_18768, plain, (join(X1,meet(join(meet(X1,X2),meet(X3,X1)),X4))=X1)).
##
#cnf(i_0_18769, plain, (join(X1,meet(X4,join(meet(X1,X2),meet(X3,X1))))=X1)).
##
#cnf(i_0_18774, plain, (join(meet(X1,join(meet(X2,X3),meet(X4,X2))),X2)=X2)).
##
#cnf(i_0_18779, plain, (join(meet(join(meet(X1,X2),meet(X3,X1)),X4),X1)=X1)).
##
#cnf(i_0_18876, plain, (join(join(meet(meet(X2,X1),X3),meet(X4,X1)),X1)=X1)).
##
#cnf(i_0_18878, plain, (join(join(meet(X2,meet(X3,X1)),meet(X4,X1)),X1)=X1)).
######
#cnf(i_0_18995, plain, (join(X1,join(meet(X1,X2),meet(meet(X3,X1),X4)))=X1)).
##
#cnf(i_0_18997, plain, (join(X1,join(meet(X1,X2),meet(X3,meet(X4,X1))))=X1)).
##
#cnf(i_0_19035, plain, (join(X1,join(meet(meet(X2,X1),X3),meet(X1,X4)))=X1)).
##
#cnf(i_0_19037, plain, (join(X1,join(meet(X2,meet(X3,X1)),meet(X1,X4)))=X1)).
##
#cnf(i_0_19087, plain, (meet(X1,join(X4,meet(join(X2,X1),join(X1,X3))))=X1)).
##
#cnf(i_0_19100, plain, (meet(X1,join(meet(join(X2,X1),join(X1,X3)),X4))=X1)).
####
#cnf(i_0_19308, plain, (meet(X1,join(X4,meet(join(X1,X2),join(X1,X3))))=X1)).
##
#cnf(i_0_19321, plain, (meet(X1,join(meet(join(X1,X2),join(X1,X3)),X4))=X1)).
########
#cnf(i_0_19644, plain, (join(X1,join(meet(meet(X2,X1),X3),meet(X4,X1)))=X1)).
##
#cnf(i_0_19646, plain, (join(X1,join(meet(X2,meet(X3,X1)),meet(X4,X1)))=X1)).
##################
#cnf(i_0_19710, plain, (join(join(meet(X2,meet(X1,X3)),meet(X4,X1)),X1)=X1)).
###
#cnf(i_0_19713, plain, (join(join(meet(X2,meet(X1,X3)),meet(X1,X4)),X1)=X1)).
###
#cnf(i_0_19715, plain, (join(X1,join(meet(X2,meet(X1,X3)),meet(X1,X4)))=X1)).
##
#cnf(i_0_19716, plain, (join(X1,join(meet(X2,meet(X1,X3)),meet(X4,X1)))=X1)).
##
#cnf(i_0_19719, plain, (join(join(meet(X1,X2),meet(X3,meet(X2,X4))),X2)=X2)).
#####
#cnf(i_0_19738, plain, (join(X1,join(meet(X2,X1),meet(X3,meet(X1,X4))))=X1)).
###
#cnf(i_0_19742, plain, (join(join(meet(X1,X2),meet(X3,meet(X1,X4))),X1)=X1)).
###
#cnf(i_0_19744, plain, (join(X1,join(meet(X1,X2),meet(X3,meet(X1,X4))))=X1)).
###################
#cnf(i_0_19902, plain, (join(join(meet(meet(X1,X2),X3),meet(X4,X1)),X1)=X1)).
###
#cnf(i_0_19905, plain, (join(join(meet(meet(X1,X2),X3),meet(X1,X4)),X1)=X1)).
###
#cnf(i_0_19907, plain, (join(X1,join(meet(meet(X1,X2),X3),meet(X1,X4)))=X1)).
##
#cnf(i_0_19908, plain, (join(X1,join(meet(meet(X1,X2),X3),meet(X4,X1)))=X1)).
##
#cnf(i_0_19911, plain, (join(join(meet(X1,X2),meet(meet(X2,X3),X4)),X2)=X2)).
#####
#cnf(i_0_19930, plain, (join(X1,join(meet(X2,X1),meet(meet(X1,X3),X4)))=X1)).
###
#cnf(i_0_19934, plain, (join(join(meet(X1,X2),meet(meet(X1,X3),X4)),X1)=X1)).
###
#cnf(i_0_19936, plain, (join(X1,join(meet(X1,X2),meet(meet(X1,X3),X4)))=X1)).
#############
#cnf(i_0_20088, plain, (meet(X1,meet(join(X2,join(X1,X3)),join(X1,X4)))=X1)).
####
#cnf(i_0_20101, plain, (meet(X1,meet(join(X2,X1),join(X3,join(X1,X4))))=X1)).
##
#cnf(i_0_20102, plain, (meet(X1,meet(join(X1,X2),join(X3,join(X1,X4))))=X1)).
#############################
#cnf(i_0_20403, plain, (meet(X1,meet(join(X2,join(X3,X1)),join(X1,X4)))=X1)).
####
#cnf(i_0_20417, plain, (meet(X1,meet(join(X2,X1),join(X3,join(X4,X1))))=X1)).
##
#cnf(i_0_20418, plain, (meet(X1,meet(join(X1,X2),join(X3,join(X4,X1))))=X1)).
###########################
#cnf(i_0_20734, plain, (meet(X1,meet(join(join(X1,X2),X3),join(X4,X1)))=X1)).
###
#cnf(i_0_20736, plain, (meet(X1,meet(join(join(X1,X2),X3),join(X1,X4)))=X1)).
####
#cnf(i_0_20751, plain, (meet(X1,meet(join(X2,X1),join(join(X1,X3),X4)))=X1)).
##
#cnf(i_0_20752, plain, (meet(X1,meet(join(X1,X2),join(join(X1,X3),X4)))=X1)).
###########################
#cnf(i_0_21084, plain, (meet(X1,meet(join(join(X2,X1),X3),join(X4,X1)))=X1)).
###
#cnf(i_0_21086, plain, (meet(X1,meet(join(join(X2,X1),X3),join(X1,X4)))=X1)).
####
#cnf(i_0_21101, plain, (meet(X1,meet(join(X2,X1),join(join(X3,X1),X4)))=X1)).
##
#cnf(i_0_21102, plain, (meet(X1,meet(join(X1,X2),join(join(X3,X1),X4)))=X1)).
###############################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################
#cnf(i_0_11001, plain, (join(meet(X3,X4),join(X1,join(X2,X3)))=join(X1,join(X2,X3)))).
#
#cnf(i_0_11002, plain, (join(meet(X4,X3),join(X1,join(X2,X3)))=join(X1,join(X2,X3)))).
#
#cnf(i_0_11141, plain, (meet(join(X1,X2),meet(meet(X2,X3),X4))=meet(meet(X2,X3),X4))).
####
#cnf(i_0_11142, plain, (meet(join(X1,X2),meet(X4,meet(X2,X3)))=meet(X4,meet(X2,X3)))).
######
#cnf(i_0_11171, plain, (join(join(X1,join(X2,X3)),meet(X3,X4))=join(X1,join(X2,X3)))).
#
#cnf(i_0_11172, plain, (join(join(X1,join(X2,X3)),meet(X4,X3))=join(X1,join(X2,X3)))).
#
#cnf(i_0_11247, plain, (meet(join(X1,X2),meet(meet(X3,X1),X4))=meet(meet(X3,X1),X4))).
##
#cnf(i_0_11249, plain, (meet(join(X1,X2),meet(X3,meet(X4,X1)))=meet(X3,meet(X4,X1)))).
######
#cnf(i_0_11682, plain, (join(meet(X1,X4),join(X2,join(X1,X3)))=join(X2,join(X1,X3)))).
#
#cnf(i_0_11688, plain, (join(meet(X1,X2),join(X3,join(X2,X4)))=join(X3,join(X2,X4)))).
#
#cnf(i_0_11708, plain, (meet(meet(meet(X1,X2),X3),join(X4,X2))=meet(meet(X1,X2),X3))).
##
#cnf(i_0_11709, plain, (meet(meet(X1,meet(X2,X3)),join(X4,X3))=meet(X1,meet(X2,X3)))).
#######
#cnf(i_0_11856, plain, (join(join(X1,join(X2,X3)),meet(X2,X4))=join(X1,join(X2,X3)))).
###
#cnf(i_0_11961, plain, (join(join(X1,join(X2,X3)),meet(X4,X2))=join(X1,join(X2,X3)))).
##
#cnf(i_0_12210, plain, (meet(meet(meet(X1,X2),X3),join(X4,X1))=meet(meet(X1,X2),X3))).
####
#cnf(i_0_12349, plain, (meet(meet(X1,meet(X2,X3)),join(X4,X2))=meet(X1,meet(X2,X3)))).
######
#cnf(i_0_12465, plain, (join(meet(X1,X4),join(join(X1,X2),X3))=join(join(X1,X2),X3))).
#
#cnf(i_0_12473, plain, (join(meet(X1,X2),join(join(X2,X3),X4))=join(join(X2,X3),X4))).
#
#cnf(i_0_12484, plain, (join(join(join(X1,X2),X3),meet(X1,X4))=join(join(X1,X2),X3))).
#
#cnf(i_0_12489, plain, (join(join(join(X1,X2),X3),meet(X4,X1))=join(join(X1,X2),X3))).
#
#cnf(i_0_12497, plain, (meet(meet(meet(X1,X2),X3),join(X2,X4))=meet(meet(X1,X2),X3))).
##
#cnf(i_0_12498, plain, (meet(meet(meet(X1,X2),X3),join(X1,X4))=meet(meet(X1,X2),X3))).
####
#cnf(i_0_12499, plain, (meet(meet(X1,meet(X2,X3)),join(X3,X4))=meet(X1,meet(X2,X3)))).
####
#cnf(i_0_12500, plain, (meet(meet(X1,meet(X2,X3)),join(X2,X4))=meet(X1,meet(X2,X3)))).
######
#cnf(i_0_12541, plain, (join(meet(X1,X4),join(join(X2,X1),X3))=join(join(X2,X1),X3))).
#
#cnf(i_0_12549, plain, (join(meet(X1,X2),join(join(X3,X2),X4))=join(join(X3,X2),X4))).
#
#cnf(i_0_12560, plain, (join(join(join(X1,X2),X3),meet(X2,X4))=join(join(X1,X2),X3))).
#
#cnf(i_0_12565, plain, (join(join(join(X1,X2),X3),meet(X4,X2))=join(join(X1,X2),X3))).
###############
#cnf(i_0_13355, plain, (join(meet(X2,X1),join(X3,meet(X1,X2)))=join(X3,meet(X1,X2)))).
################################################################################
#cnf(i_0_13656, plain, (join(join(X1,meet(X2,X3)),meet(X3,X2))=join(X1,meet(X2,X3)))).
###############
#cnf(i_0_13961, plain, (meet(join(X1,X3),meet(X4,meet(X1,X2)))=meet(X4,meet(X1,X2)))).
####
#cnf(i_0_13962, plain, (meet(join(X1,X3),meet(meet(X1,X2),X4))=meet(meet(X1,X2),X4))).
##########
#cnf(i_0_14024, plain, (join(meet(X2,X1),join(meet(X1,X2),X3))=join(meet(X1,X2),X3))).
#############################################################################
#cnf(i_0_14168, plain, (join(join(meet(X1,X2),X3),meet(X2,X1))=join(meet(X1,X2),X3))).
###############################
#cnf(i_0_14506, plain, (meet(meet(X1,X2),meet(X3,meet(X2,X1)))=meet(X3,meet(X2,X1)))).
####
#cnf(i_0_14507, plain, (meet(meet(X1,X2),meet(meet(X2,X1),X3))=meet(meet(X2,X1),X3))).
##################################################################################################################################################################################################################################################################################################################################################################################################################################################################################
#cnf(i_0_10892, plain, (join(join(X1,meet(join(X2,X1),X3)),join(X2,X1))=join(X2,X1))).
###
#cnf(i_0_112606, plain, (join(join(X1,meet(X2,X3)),join(X2,X1))=join(X2,X1))).
###
#cnf(i_0_113164, plain, (join(join(X1,X2),join(X2,X1))=join(X2,X1))).
#
#cnf(i_0_113380, plain, (join(X1,X2)=join(X2,X1))).
#
#cnf(i_0_4, negated_conjecture, (meet(c,meet(a,b))!=meet(a,meet(b,c))|join(c,join(a,b))!=join(a,join(b,c)))).
#############################################################
#cnf(i_0_113370, plain, (meet(join(X1,X2),join(X2,X1))=join(X1,X2))).
#################################################################################################
#cnf(i_0_113383, plain, (meet(join(X1,X2),join(join(X2,X1),X3))=join(X1,X2))).
###
#cnf(i_0_113384, plain, (join(join(X2,X1),meet(join(X1,X2),X3))=join(X2,X1))).
################################################################
#cnf(i_0_113385, plain, (join(join(X2,X1),meet(X3,join(X1,X2)))=join(X2,X1))).
#####
#cnf(i_0_113439, plain, (meet(join(X1,X2),join(X3,join(X2,X1)))=join(X1,X2))).
#####
#cnf(i_0_112606, plain, (join(join(X2,X1),join(X1,meet(X2,X3)))=join(X2,X1))).
##################################################################################################
#cnf(i_0_117814, plain, (join(join(X1,X2),join(X2,meet(X3,X1)))=join(X1,X2))).
#####################
#cnf(i_0_117844, plain, (join(join(X1,X2),join(meet(X1,X3),X2))=join(X1,X2))).
##
#cnf(i_0_117997, plain, (join(join(X2,X1),join(X2,meet(X1,X3)))=join(X1,X2))).
###########################################################################################################
#cnf(i_0_118446, plain, (join(join(X1,X2),join(meet(X3,X1),X2))=join(X1,X2))).
##
#cnf(i_0_118600, plain, (join(join(X2,X1),join(X2,meet(X3,X1)))=join(X1,X2))).
####################################################
#cnf(i_0_119233, plain, (join(join(X2,X1),join(meet(X1,X3),X2))=join(X1,X2))).
##############
#cnf(i_0_120477, plain, (join(join(X2,X1),join(meet(X3,X1),X2))=join(X1,X2))).
##############
#cnf(i_0_113368, plain, (meet(join(X2,X1),meet(X3,join(X1,X2)))=meet(X3,join(X1,X2)))).
####
#cnf(i_0_113369, plain, (meet(join(X2,X1),meet(join(X1,X2),X3))=meet(join(X1,X2),X3))).
###################################################################
#cnf(i_0_113391, plain, (join(join(X1,X2),join(join(X2,X1),X3))=join(join(X2,X1),X3))).
####
#cnf(i_0_113441, plain, (join(join(X1,X2),join(X3,join(X2,X1)))=join(X3,join(X2,X1)))).
############################
#cnf(i_0_117523, plain, (meet(join(X2,X1),join(X1,meet(X2,X3)))=join(X1,meet(X2,X3)))).
################################################################################
#cnf(i_0_117638, plain, (meet(join(X1,X3),join(meet(X1,X2),X3))=join(meet(X1,X2),X3))).
######
#cnf(i_0_118125, plain, (meet(join(X3,X1),join(X1,meet(X2,X3)))=join(X1,meet(X2,X3)))).
########################
#cnf(i_0_126298, plain, (meet(join(X3,X1),join(join(X1,X2),X3))=join(X3,X1))).
######
#cnf(i_0_126299, plain, (meet(join(X3,X2),join(join(X1,X2),X3))=join(X3,X2))).
######
#cnf(i_0_126729, plain, (meet(join(X1,X2),join(X1,join(X2,X3)))=join(X1,X2))).
###########
#cnf(i_0_127103, plain, (meet(join(X2,X1),join(join(X2,X3),X1))=join(X1,X2))).
############################################################################
#cnf(i_0_128676, plain, (join(X1,join(X2,meet(X1,X3)))=join(X1,X2))).
##############################################################
#cnf(i_0_128677, plain, (join(X1,join(X2,meet(X3,X1)))=join(X1,X2))).
#######
#cnf(i_0_129327, plain, (join(X1,join(meet(X1,X3),X2))=join(X1,X2))).
####
#cnf(i_0_129822, plain, (join(X1,join(meet(X3,X1),X2))=join(X1,X2))).
###############################
#cnf(i_0_129202, plain, (join(X3,join(X2,X1))=join(join(X1,X2),X3))).
##########################################################################################################################
#cnf(i_0_129257, plain, (join(X3,meet(X2,X1))=join(meet(X1,X2),X3))).
########################################################################################################################################################################################################################################################################################################################################################################################################################################
#cnf(i_0_129789, plain, (join(join(X1,X2),join(X3,X1))=join(join(X1,X2),X3))).
####################################################################
#cnf(i_0_129790, plain, (join(join(X1,X2),join(X3,X2))=join(join(X1,X2),X3))).
###
#cnf(i_0_133402, plain, (join(join(X1,X2),X3)=join(join(X3,X2),X1))).
###################
#cnf(i_0_133995, plain, (join(join(X3,X2),X1)=join(X3,join(X1,X2)))).
####
#cnf(i_0_129202, plain, (join(X1,join(X2,X3))=join(X3,join(X1,X2)))).
###
#cnf(i_0_4, negated_conjecture, (meet(c,meet(a,b))!=meet(a,meet(b,c)))).
###########
#cnf(i_0_133402, plain, (join(X1,join(X3,X2))=join(X3,join(X1,X2)))).
###################
#cnf(i_0_135130, plain, (join(X2,join(X3,X1))=join(X1,join(X3,X2)))).
########################################################################################################################################################
#cnf(i_0_128006, plain, (meet(join(X1,X2),join(X1,join(X3,X2)))=join(X1,X2))).
########
#cnf(i_0_128107, plain, (meet(join(X1,X2),join(X2,join(X1,X3)))=join(X1,X2))).
#########
#cnf(i_0_128108, plain, (meet(join(X1,X2),join(X2,join(X3,X1)))=join(X1,X2))).
###########
#cnf(i_0_129316, plain, (join(X1,join(X2,meet(meet(X3,X1),X4)))=join(X1,X2))).
#########
#cnf(i_0_129318, plain, (join(X1,join(X2,meet(meet(X1,X3),X4)))=join(X1,X2))).
########
#cnf(i_0_129319, plain, (join(X1,join(X2,meet(X3,meet(X4,X1))))=join(X1,X2))).
#############
#cnf(i_0_129321, plain, (join(X1,join(X2,meet(X3,meet(X1,X4))))=join(X1,X2))).
############
#cnf(i_0_129896, plain, (join(X1,join(meet(meet(X1,X2),X3),X4))=join(X1,X4))).
##############
#cnf(i_0_130311, plain, (join(X1,join(meet(meet(X2,X1),X3),X4))=join(X1,X4))).
#################
#cnf(i_0_130314, plain, (join(X1,join(meet(X2,meet(X3,X1)),X4))=join(X1,X4))).
#######################
#cnf(i_0_130316, plain, (join(X1,join(meet(X2,meet(X1,X3)),X4))=join(X1,X4))).
#################################
#cnf(i_0_113385, plain, (join(X1,join(X2,meet(join(X2,X1),X3)))=join(X1,X2))).
####################################################
#cnf(i_0_113384, plain, (join(X1,join(X2,meet(X3,join(X2,X1))))=join(X1,X2))).
#######################
#cnf(i_0_134934, plain, (join(X1,join(X2,meet(join(X1,X2),X3)))=join(X1,X2))).
#####################################################################################################################
#cnf(i_0_140601, plain, (join(X1,join(X2,meet(X3,join(X1,X2))))=join(X1,X2))).
########################################
#cnf(i_0_125075, plain, (meet(join(X1,X2),join(X1,meet(X2,X3)))=join(X1,meet(X2,X3)))).
##############
#cnf(i_0_125619, plain, (meet(join(X1,X2),join(meet(X2,X3),X1))=join(meet(X2,X3),X1))).
#####################################################
#cnf(i_0_125753, plain, (meet(join(X1,X2),join(meet(X3,X1),X2))=join(meet(X3,X1),X2))).
##############
#cnf(i_0_126191, plain, (meet(join(X1,X2),join(X1,meet(X3,X2)))=join(X1,meet(X3,X2)))).
################
#cnf(i_0_129231, plain, (join(X1,join(X3,join(X2,meet(X1,X4))))=join(X1,join(X3,X2)))).
##########################################
#cnf(i_0_129232, plain, (join(X1,join(X3,join(X2,meet(X4,X1))))=join(X1,join(X3,X2)))).
#############################
#cnf(i_0_134939, plain, (join(X1,join(X2,join(meet(X4,X1),X3)))=join(X1,join(X2,X3)))).
############
#cnf(i_0_134941, plain, (join(X1,join(X2,join(meet(X1,X4),X3)))=join(X1,join(X2,X3)))).
########################################################################
#cnf(i_0_137369, plain, (meet(join(X3,X2),join(meet(X1,X2),X3))=join(meet(X1,X2),X3))).
################################################################
#cnf(i_0_11683, plain, (meet(join(X1,join(X2,join(X3,X4))),meet(X5,X3))=meet(X5,X3))).
#####
#cnf(i_0_11692, plain, (meet(join(X1,join(X2,join(X3,X4))),meet(X3,X5))=meet(X3,X5))).
###########
#cnf(i_0_13186, plain, (meet(meet(X3,X4),join(X5,join(X1,join(X2,X3))))=meet(X3,X4))).
######
#cnf(i_0_13187, plain, (meet(meet(X4,X3),join(X5,join(X1,join(X2,X3))))=meet(X4,X3))).
########
#cnf(i_0_13368, plain, (meet(meet(X4,X3),join(X1,join(X2,meet(X3,X4))))=meet(X4,X3))).
##########
#cnf(i_0_1422, plain, (meet(meet(X1,X3),meet(X1,join(X2,meet(X1,X3))))=meet(X1,X3))).
#########
#cnf(i_0_147566, plain, (meet(meet(X1,X3),meet(X1,meet(X2,X3)))=meet(X1,meet(X2,X3)))).
#
#cnf(i_0_147696, plain, (join(meet(X1,X2),meet(X1,meet(X3,X2)))=meet(X1,X2))).
###
#cnf(i_0_147933, plain, (join(meet(X1,X2),meet(X2,meet(X3,X1)))=meet(X1,X2))).
#####
#cnf(i_0_148023, plain, (meet(meet(X1,X2),meet(X1,join(X2,X3)))=meet(X1,X2))).
#########
#cnf(i_0_148024, plain, (meet(meet(X1,X3),meet(X1,join(X2,X3)))=meet(X1,X3))).
##########
#cnf(i_0_148066, plain, (meet(meet(X2,X3),meet(join(X1,X2),X3))=meet(X2,X3))).
########
#cnf(i_0_148068, plain, (meet(meet(X1,X3),meet(join(X1,X2),X3))=meet(X1,X3))).
#######
#cnf(i_0_148286, plain, (join(meet(X1,X3),meet(meet(X2,X3),X1))=meet(X1,X3))).
####
#cnf(i_0_148456, plain, (join(meet(X1,X2),meet(X1,meet(X2,X3)))=meet(X1,X2))).
######
#cnf(i_0_148501, plain, (join(meet(X1,X2),meet(meet(X3,X1),X2))=meet(X1,X2))).
#####
#cnf(i_0_152042, plain, (meet(meet(X1,X2),meet(X2,X3))=meet(meet(X1,X2),X3))).
###
#cnf(i_0_152779, plain, (meet(X1,meet(join(X1,X2),X3))=meet(X1,X3))).
###
#cnf(i_0_152780, plain, (meet(X1,meet(join(X2,X1),X3))=meet(X1,X3))).
######
#cnf(i_0_153213, plain, (meet(X1,meet(X3,join(X1,X2)))=meet(X1,X3))).
###########################################
#cnf(i_0_153620, plain, (meet(X1,meet(X3,join(X2,X1)))=meet(X1,X3))).
################################
#cnf(i_0_154325, plain, (meet(X3,meet(X2,X1))=meet(meet(X1,X2),X3))).
######################################################################################################################################
#cnf(i_0_155036, plain, (meet(join(X3,X2),X1)=meet(X1,join(X2,X3)))).
#######################################################################################################################################################################################################################################################
#cnf(i_0_152380, plain, (meet(meet(X2,X3),meet(X1,X2))=meet(meet(X1,X2),X3))).
#######
#cnf(i_0_152698, plain, (meet(meet(X1,X2),meet(X3,X2))=meet(meet(X1,X2),X3))).
####
#cnf(i_0_156455, plain, (meet(meet(X1,X2),X3)=meet(meet(X3,X2),X1))).
#############
#cnf(i_0_156936, plain, (meet(meet(X1,X2),X3)=meet(X1,meet(X3,X2)))).
###
#cnf(i_0_154325, plain, (meet(X1,meet(X2,X3))=meet(X3,meet(X1,X2)))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(single_axiom, axiom, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(join(meet(X1,join(join(meet(X4,X2),meet(X2,X5)),X2)),meet(join(meet(X2,meet(meet(join(X4,join(X2,X5)),join(X6,X2)),X2)),meet(X7,join(X2,meet(meet(join(X4,join(X2,X5)),join(X6,X2)),X2)))),join(X1,join(join(meet(X4,X2),meet(X2,X5)),X2)))),join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), file('./p/LAT015-1.tptp', single_axiom)).
cnf(prove_normal_axioms, negated_conjecture, (meet(a,a)!=a|meet(a,b)!=meet(b,a)|meet(meet(a,b),c)!=meet(a,meet(b,c))|join(a,a)!=a|join(a,b)!=join(b,a)|join(join(a,b),c)!=join(a,join(b,c))|meet(a,join(a,b))!=a|join(a,meet(a,b))!=a), file('./p/LAT015-1.tptp', prove_normal_axioms)).
cnf(c_0_2, axiom, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(join(meet(X1,join(join(meet(X4,X2),meet(X2,X5)),X2)),meet(join(meet(X2,meet(meet(join(X4,join(X2,X5)),join(X6,X2)),X2)),meet(X7,join(X2,meet(meet(join(X4,join(X2,X5)),join(X6,X2)),X2)))),join(X1,join(join(meet(X4,X2),meet(X2,X5)),X2)))),join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), single_axiom).
cnf(c_0_3, plain, (join(meet(join(meet(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),join(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))))),X4),meet(join(meet(X1,join(join(meet(X5,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),X6)),join(meet(X2,X3),meet(X3,join(X2,X3))))),meet(X3,join(X1,join(join(meet(X5,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),X6)),join(meet(X2,X3),meet(X3,join(X2,X3))))))),join(join(meet(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),join(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))))),X4)))=join(meet(X2,X3),meet(X3,join(X2,X3)))), inference(pm,[status(thm)],[c_0_2, c_0_2])).
cnf(c_0_4, plain, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(join(meet(X1,join(join(meet(X4,X2),meet(X2,X5)),X2)),meet(join(meet(X6,join(join(meet(X7,X2),meet(X2,X8)),X2)),meet(join(meet(X2,meet(meet(join(X7,join(X2,X8)),join(X9,X2)),X2)),meet(X10,join(X2,meet(meet(join(X7,join(X2,X8)),join(X9,X2)),X2)))),join(X6,join(join(meet(X7,X2),meet(X2,X8)),X2)))),join(X1,join(join(meet(X4,X2),meet(X2,X5)),X2)))),join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_3, c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2])).
cnf(c_0_5, plain, (join(meet(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X2),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X2))),X3),meet(X2,join(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X2),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X2))),X3)))=X2), inference(pm,[status(thm)],[c_0_4, c_0_2])).
cnf(c_0_6, plain, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(join(meet(X1,join(join(meet(X4,X2),meet(X2,X5)),X2)),meet(X2,join(X1,join(join(meet(X4,X2),meet(X2,X5)),X2)))),join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_3, c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5])).
cnf(c_0_7, plain, (join(meet(X1,X2),meet(X1,join(X1,X2)))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_5, c_0_5]), c_0_5])).
cnf(c_0_8, plain, (join(meet(X1,join(join(meet(X2,X3),meet(X3,X4)),X3)),meet(X3,join(X1,join(join(meet(X2,X3),meet(X3,X4)),X3))))=X3), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_6, c_0_6]), c_0_6]), c_0_6]), c_0_7])).
cnf(c_0_9, plain, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), inference(rw,[status(thm)],[c_0_6, c_0_8])).
cnf(c_0_10, plain, (join(meet(X1,join(join(meet(X2,X3),meet(X3,X4)),X3)),meet(join(meet(X3,meet(meet(join(X2,join(X3,X4)),join(X5,X3)),X3)),meet(X6,join(X3,meet(meet(join(X2,join(X3,X4)),join(X5,X3)),X3)))),join(X1,join(join(meet(X2,X3),meet(X3,X4)),X3))))=X3), inference(pm,[status(thm)],[c_0_2, c_0_8])).
cnf(c_0_11, plain, (join(meet(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1)),meet(X5,join(X1,meet(meet(join(X2,join(X1,X3)),join(X4,X1)),X1))))=X1), inference(pm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_12, plain, (join(meet(X1,X2),meet(X2,join(X1,X2)))=X2), inference(pm,[status(thm)],[c_0_9, c_0_11])).
cnf(c_0_13, plain, (meet(meet(join(X1,join(X2,X3)),join(X4,X2)),X2)=X2), inference(pm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_14, plain, (join(meet(X1,join(X2,X2)),meet(X2,join(X1,join(X2,X2))))=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_8, c_0_5]), c_0_5])).
cnf(c_0_15, plain, (join(meet(X1,X1),meet(X2,join(X1,X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_13]), c_0_13])).
cnf(c_0_16, plain, (join(X1,X1)=X1), inference(pm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_17, plain, (meet(meet(join(X1,X2),join(X3,X2)),X2)=X2), inference(pm,[status(thm)],[c_0_13, c_0_16])).
cnf(c_0_18, plain, (meet(X1,X1)=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_16]), c_0_16])).
cnf(c_0_19, plain, (meet(join(X1,X2),X2)=X2), inference(pm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_20, plain, (join(X1,meet(X2,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_16]), c_0_18])).
cnf(c_0_21, plain, (join(meet(meet(X1,X2),meet(X1,join(X1,X2))),meet(meet(X1,X2),X1))=meet(X1,X2)), inference(pm,[status(thm)],[c_0_7, c_0_7])).
cnf(c_0_22, plain, (meet(X1,meet(X2,X1))=meet(X2,X1)), inference(pm,[status(thm)],[c_0_19, c_0_20])).
cnf(c_0_23, plain, (meet(meet(join(X1,X2),join(X3,X1)),X1)=X1), inference(pm,[status(thm)],[c_0_13, c_0_16])).
cnf(c_0_24, plain, (join(meet(X1,meet(X2,X1)),X1)=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_20]), c_0_18])).
cnf(c_0_25, plain, (meet(meet(X1,X2),X2)=meet(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_21, c_0_22]), c_0_20]), c_0_18]), c_0_22]), c_0_16]), c_0_22])).
cnf(c_0_26, plain, (meet(meet(join(X1,X2),X1),X1)=X1), inference(pm,[status(thm)],[c_0_23, c_0_24])).
cnf(c_0_27, plain, (meet(join(X1,join(X2,X3)),X2)=X2), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_13, c_0_16]), c_0_25])).
cnf(c_0_28, plain, (meet(join(X1,X2),X1)=X1), inference(rw,[status(thm)],[c_0_26, c_0_25])).
cnf(c_0_29, plain, (meet(join(X1,X2),meet(X2,X3))=meet(X2,X3)), inference(pm,[status(thm)],[c_0_27, c_0_7])).
cnf(c_0_30, plain, (join(X1,meet(X1,join(join(X1,X2),X1)))=X1), inference(pm,[status(thm)],[c_0_12, c_0_28])).
cnf(c_0_31, plain, (meet(X1,meet(X1,X2))=meet(X1,X2)), inference(pm,[status(thm)],[c_0_28, c_0_7])).
cnf(c_0_32, plain, (meet(meet(X1,join(X2,X1)),X1)=X1), inference(pm,[status(thm)],[c_0_23, c_0_20])).
cnf(c_0_33, plain, (meet(X1,meet(meet(X1,join(X2,X1)),X3))=meet(meet(X1,join(X2,X1)),X3)), inference(pm,[status(thm)],[c_0_29, c_0_12])).
cnf(c_0_34, plain, (join(meet(X1,X2),X2)=X2), inference(rw,[status(thm)],[c_0_24, c_0_22])).
cnf(c_0_35, plain, (meet(X1,join(join(X1,X2),X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_21, c_0_30]), c_0_31]), c_0_18]), c_0_32]), c_0_31]), c_0_32]), c_0_16]), c_0_31])).
cnf(c_0_36, plain, (meet(meet(X1,join(X2,X1)),meet(X1,X3))=meet(X1,X3)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_33, c_0_12]), c_0_32]), c_0_12]), c_0_32])).
cnf(c_0_37, plain, (join(X1,join(join(X1,X2),X1))=join(join(X1,X2),X1)), inference(pm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38, plain, (join(meet(X1,join(X1,join(X2,X1))),X1)=meet(X1,join(X2,X1))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_21, c_0_36]), c_0_32])).
cnf(c_0_39, plain, (meet(X1,join(X2,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_37, c_0_38]), c_0_7]), c_0_38]), c_0_38]), c_0_7])).
cnf(c_0_40, plain, (join(X1,join(X2,X1))=join(X2,X1)), inference(pm,[status(thm)],[c_0_34, c_0_39])).
cnf(c_0_41, plain, (meet(join(X1,join(X2,X3)),meet(X3,X4))=meet(X3,X4)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_29, c_0_29]), c_0_29])).
cnf(c_0_42, plain, (meet(join(X1,join(X2,join(X3,X4))),X3)=X3), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_29, c_0_27]), c_0_27])).
cnf(c_0_43, plain, (join(meet(meet(X1,X2),X1),meet(X1,X2))=meet(X1,X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_40, c_0_21]), c_0_21])).
cnf(c_0_44, plain, (meet(join(X1,join(X2,X3)),X3)=X3), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_29, c_0_19]), c_0_19])).
cnf(c_0_45, plain, (join(meet(meet(X1,join(join(meet(X2,X3),meet(X3,X4)),X3)),meet(X3,join(X1,join(join(meet(X2,X3),meet(X3,X4)),X3)))),meet(meet(X1,join(join(meet(X2,X3),meet(X3,X4)),X3)),X3))=meet(X1,join(join(meet(X2,X3),meet(X3,X4)),X3))), inference(pm,[status(thm)],[c_0_7, c_0_8])).
cnf(c_0_46, plain, (meet(join(X1,join(X2,join(X3,join(X4,join(X5,X6))))),X5)=X5), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_41, c_0_42]), c_0_42])).
cnf(c_0_47, plain, (join(meet(X1,join(X2,join(X3,X1))),X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_43, c_0_44]), c_0_44]), c_0_44])).
cnf(c_0_48, plain, (join(join(meet(X1,X2),meet(X2,X3)),X2)=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_45, c_0_46]), c_0_29]), c_0_46]), c_0_19]), c_0_47]), c_0_46])).
cnf(c_0_49, plain, (join(meet(X1,X2),X1)=X1), inference(pm,[status(thm)],[c_0_48, c_0_43])).
cnf(c_0_50, plain, (meet(meet(X1,X2),X1)=meet(X1,X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_49]), c_0_16])).
cnf(c_0_51, plain, (meet(X1,meet(meet(X2,X1),X3))=meet(meet(X2,X1),X3)), inference(pm,[status(thm)],[c_0_29, c_0_20])).
cnf(c_0_52, plain, (meet(X1,join(X1,X2))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_50, c_0_28]), c_0_28])).
cnf(c_0_53, plain, (join(X1,meet(X1,X2))=X1), inference(pm,[status(thm)],[c_0_20, c_0_50])).
cnf(c_0_54, plain, (join(join(X1,meet(join(X2,X1),X3)),join(X2,X1))=join(X2,X1)), inference(pm,[status(thm)],[c_0_48, c_0_39])).
cnf(c_0_55, plain, (meet(join(X1,X2),meet(X1,X3))=meet(X1,X3)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_51, c_0_52]), c_0_52])).
cnf(c_0_56, plain, (meet(X1,join(X2,join(X1,X3)))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_50, c_0_27]), c_0_27])).
cnf(c_0_57, plain, (join(join(X1,X2),X1)=join(X1,X2)), inference(pm,[status(thm)],[c_0_53, c_0_28])).
cnf(c_0_58, plain, (join(join(X1,meet(X2,X3)),join(X2,X1))=join(X2,X1)), inference(pm,[status(thm)],[c_0_54, c_0_55])).
cnf(c_0_59, plain, (meet(X1,join(join(X1,X2),X3))=X1), inference(pm,[status(thm)],[c_0_56, c_0_57])).
cnf(c_0_60, plain, (join(join(X1,X2),join(X2,X1))=join(X2,X1)), inference(pm,[status(thm)],[c_0_58, c_0_59])).
cnf(c_0_61, plain, (meet(meet(X1,join(X2,meet(X3,X1))),meet(X3,X1))=meet(X3,X1)), inference(pm,[status(thm)],[c_0_23, c_0_12])).
cnf(c_0_62, plain, (join(X1,X2)=join(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_57, c_0_60]), c_0_60]), c_0_60])).
cnf(c_0_63, plain, (meet(meet(X1,X2),meet(X2,X1))=meet(X2,X1)), inference(pm,[status(thm)],[c_0_61, c_0_53])).
cnf(c_0_64, plain, (meet(join(X1,X2),meet(X3,X2))=meet(X3,X2)), inference(pm,[status(thm)],[c_0_27, c_0_12])).
cnf(c_0_65, plain, (join(join(X1,X2),join(X2,meet(X1,X3)))=join(X1,X2)), inference(rw,[status(thm)],[c_0_58, c_0_62])).
cnf(c_0_66, plain, (meet(meet(X1,join(X2,meet(X1,X3))),meet(X1,X3))=meet(X1,X3)), inference(pm,[status(thm)],[c_0_23, c_0_7])).
cnf(c_0_67, plain, (meet(X1,X2)=meet(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_50, c_0_63]), c_0_63]), c_0_63])).
cnf(c_0_68, plain, (meet(X1,meet(X2,meet(X3,X1)))=meet(X2,meet(X3,X1))), inference(pm,[status(thm)],[c_0_64, c_0_20])).
cnf(c_0_69, plain, (join(join(X1,X2),join(X2,meet(X3,X1)))=join(X1,X2)), inference(pm,[status(thm)],[c_0_65, c_0_22])).
cnf(c_0_70, plain, (meet(meet(X1,X2),meet(X1,join(X3,meet(X1,X2))))=meet(X1,X2)), inference(rw,[status(thm)],[c_0_66, c_0_67])).
cnf(c_0_71, plain, (join(X1,meet(X2,meet(X3,X1)))=X1), inference(pm,[status(thm)],[c_0_53, c_0_68])).
cnf(c_0_72, plain, (meet(join(X1,X2),join(X2,meet(X3,X1)))=join(X2,meet(X3,X1))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_39, c_0_69]), c_0_67])).
cnf(c_0_73, plain, (meet(join(X1,X2),meet(X3,X1))=meet(X3,X1)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_68, c_0_52]), c_0_52])).
cnf(c_0_74, plain, (meet(meet(X1,X2),meet(X1,meet(X3,X2)))=meet(X1,meet(X3,X2))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_70, c_0_71]), c_0_67])).
cnf(c_0_75, negated_conjecture, (meet(a,a)!=a|meet(a,b)!=meet(b,a)|meet(meet(a,b),c)!=meet(a,meet(b,c))|join(a,a)!=a|join(a,b)!=join(b,a)|join(join(a,b),c)!=join(a,join(b,c))|meet(a,join(a,b))!=a|join(a,meet(a,b))!=a), prove_normal_axioms).
cnf(c_0_76, plain, (meet(join(X1,X2),join(join(X2,X3),X1))=join(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_72, c_0_52]), c_0_67]), c_0_52])).
cnf(c_0_77, plain, (meet(join(X1,X2),meet(X3,join(X2,X1)))=meet(X3,join(X2,X1))), inference(pm,[status(thm)],[c_0_73, c_0_60])).
cnf(c_0_78, plain, (join(meet(X1,X2),meet(X1,meet(X3,X2)))=meet(X1,X2)), inference(pm,[status(thm)],[c_0_53, c_0_74])).
cnf(c_0_79, negated_conjecture, (meet(meet(a,b),c)!=meet(a,meet(b,c))|join(join(a,b),c)!=join(a,join(b,c))|meet(b,a)!=meet(a,b)|join(b,a)!=join(a,b)|meet(a,join(a,b))!=a|join(a,meet(a,b))!=a|meet(a,a)!=a), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_16])])).
cnf(c_0_80, plain, (meet(join(X1,X2),join(join(X1,X3),X2))=join(X2,X1)), inference(pm,[status(thm)],[c_0_76, c_0_62])).
cnf(c_0_81, plain, (meet(join(X1,X2),join(X1,join(X2,X3)))=join(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_77, c_0_76]), c_0_67]), c_0_76])).
cnf(c_0_82, plain, (join(meet(X1,X2),meet(meet(X3,X1),X2))=meet(X1,X2)), inference(pm,[status(thm)],[c_0_78, c_0_51])).
cnf(c_0_83, negated_conjecture, (meet(meet(a,b),c)!=meet(a,meet(b,c))|join(join(a,b),c)!=join(a,join(b,c))|meet(b,a)!=meet(a,b)|join(b,a)!=join(a,b)|meet(a,join(a,b))!=a|join(a,meet(a,b))!=a), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_18])])).
cnf(c_0_84, plain, (join(X1,join(X2,meet(X3,X1)))=join(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_80, c_0_69]), c_0_67]), c_0_81]), c_0_62])).
cnf(c_0_85, plain, (meet(meet(X1,X2),meet(X2,X3))=meet(meet(X1,X2),X3)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_70, c_0_82]), c_0_74])).
cnf(c_0_86, negated_conjecture, (meet(meet(a,b),c)!=meet(a,meet(b,c))|join(join(a,b),c)!=join(a,join(b,c))|meet(b,a)!=meet(a,b)|join(b,a)!=join(a,b)|meet(a,join(a,b))!=a), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_83, c_0_53])])).
cnf(c_0_87, plain, (join(X1,join(X2,join(X1,X3)))=join(X2,join(X1,X3))), inference(pm,[status(thm)],[c_0_34, c_0_56])).
cnf(c_0_88, plain, (join(join(X1,X2),join(X3,X2))=join(join(X1,X2),X3)), inference(pm,[status(thm)],[c_0_84, c_0_39])).
cnf(c_0_89, plain, (meet(X1,meet(join(X2,X1),X3))=meet(X1,X3)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_85, c_0_39]), c_0_39])).
cnf(c_0_90, plain, (join(meet(X1,X2),meet(X2,X1))=meet(X1,X2)), inference(pm,[status(thm)],[c_0_53, c_0_63])).
cnf(c_0_91, negated_conjecture, (meet(meet(a,b),c)!=meet(a,meet(b,c))|join(join(a,b),c)!=join(a,join(b,c))|meet(b,a)!=meet(a,b)|join(b,a)!=join(a,b)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_52])])).
cnf(c_0_92, plain, (join(X1,join(X2,meet(X1,X3)))=join(X1,X2)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_80, c_0_65]), c_0_67]), c_0_81]), c_0_62])).
cnf(c_0_93, plain, (join(join(X1,X2),join(X3,join(X2,X1)))=join(X3,join(X2,X1))), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_87, c_0_60]), c_0_60])).
cnf(c_0_94, plain, (join(join(X1,X2),X3)=join(join(X3,X2),X1)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_62, c_0_88]), c_0_88])).
cnf(c_0_95, plain, (meet(meet(X1,X2),meet(X3,X2))=meet(meet(X1,X2),X3)), inference(pm,[status(thm)],[c_0_85, c_0_67])).
cnf(c_0_96, plain, (meet(X1,meet(X2,join(X3,X1)))=meet(X1,X2)), inference(pm,[status(thm)],[c_0_89, c_0_67])).
cnf(c_0_97, plain, (meet(meet(X1,X2),meet(X3,meet(X2,X1)))=meet(X3,meet(X2,X1))), inference(pm,[status(thm)],[c_0_64, c_0_90])).
cnf(c_0_98, negated_conjecture, (meet(c,meet(a,b))!=meet(a,meet(b,c))|join(join(a,b),c)!=join(a,join(b,c))|join(b,a)!=join(a,b)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_67]), c_0_67])])).
cnf(c_0_99, plain, (join(X1,join(X2,X3))=join(join(X3,X2),X1)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_92, c_0_80]), c_0_93])).
cnf(c_0_100, plain, (join(join(X1,X2),X3)=join(X1,join(X3,X2))), inference(pm,[status(thm)],[c_0_62, c_0_94])).
cnf(c_0_101, plain, (meet(meet(X1,X2),X3)=meet(meet(X3,X2),X1)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_67, c_0_95]), c_0_95])).
cnf(c_0_102, plain, (meet(meet(X1,X2),meet(X3,X1))=meet(meet(X3,X1),X2)), inference(pm,[status(thm)],[c_0_67, c_0_85])).
cnf(c_0_103, plain, (meet(X1,meet(X2,X3))=meet(meet(X3,X2),X1)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_96, c_0_90]), c_0_97])).
cnf(c_0_104, negated_conjecture, (meet(c,meet(a,b))!=meet(a,meet(b,c))|join(c,join(a,b))!=join(a,join(b,c))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_98, c_0_62]), c_0_62])])).
cnf(c_0_105, plain, (join(X1,join(X2,X3))=join(X3,join(X1,X2))), inference(rw,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_106, plain, (meet(meet(X1,X2),X3)=meet(X1,meet(X3,X2))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_31, c_0_101]), c_0_102]), c_0_103]), c_0_68])).
cnf(c_0_107, negated_conjecture, (meet(c,meet(a,b))!=meet(a,meet(b,c))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_104, c_0_105])])).
cnf(c_0_108, plain, (meet(X1,meet(X2,X3))=meet(X3,meet(X1,X2))), inference(rw,[status(thm)],[c_0_103, c_0_106])).
cnf(c_0_109, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_107, c_0_108])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 110
# Proof object clause steps            : 108
# Proof object formula steps           : 2
# Proof object conjectures             : 10
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 2
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 89
# Proof object simplifying inferences  : 124
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 2
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 2
# Processed clauses                    : 6161
# ...of these trivial                  : 2393
# ...subsumed                          : 3268
# ...remaining for further processing  : 500
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 396
# Generated clauses                    : 158252
# ...of the previous two non-trivial   : 60912
# Contextual simplify-reflections      : 0
# Paramodulations                      : 158252
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 104
#    Positive orientable unit clauses  : 96
#    Positive unorientable unit clauses: 8
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 54099
# ...number of literals in the above   : 54099
# Current number of archived formulas  : 0
# Current number of archived clauses   : 396
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 577
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4631
# BW rewrite match successes           : 2058
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1879777
