# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_3, plain, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(join(meet(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)),meet(join(meet(X2,meet(meet(join(X2,X4),join(X5,X2)),X2)),meet(X6,join(X2,meet(meet(join(X2,X4),join(X5,X2)),X2)))),join(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)))),join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2)).
#
#cnf(i_0_11, plain, (join(meet(join(meet(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),join(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))))),X4),meet(join(meet(X1,join(join(meet(join(meet(X2,X3),meet(X3,join(X2,X3))),X5),meet(X6,join(meet(X2,X3),meet(X3,join(X2,X3))))),join(meet(X2,X3),meet(X3,join(X2,X3))))),meet(X3,join(X1,join(join(meet(join(meet(X2,X3),meet(X3,join(X2,X3))),X5),meet(X6,join(meet(X2,X3),meet(X3,join(X2,X3))))),join(meet(X2,X3),meet(X3,join(X2,X3))))))),join(join(meet(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),join(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))))),X4)))=join(meet(X2,X3),meet(X3,join(X2,X3))))).
#
#cnf(i_0_49, plain, (join(meet(join(meet(X1,X3),meet(X3,join(X1,X3))),X7),meet(join(meet(X1,join(join(meet(X3,X8),meet(X9,X3)),X3)),meet(join(meet(X2,join(join(meet(X3,X4),meet(X5,X3)),X3)),meet(join(meet(X3,meet(meet(join(X3,X4),join(X5,X3)),X3)),meet(X6,join(X3,meet(meet(join(X3,X4),join(X5,X3)),X3)))),join(X2,join(join(meet(X3,X4),meet(X5,X3)),X3)))),join(X1,join(join(meet(X3,X8),meet(X9,X3)),X3)))),join(join(meet(X1,X3),meet(X3,join(X1,X3))),X7)))=X3)).
#
#cnf(i_0_4, negated_conjecture, (meet(b,a)!=meet(a,b))).
#
#cnf(i_0_130, plain, (join(meet(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X2),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X2))),X3),meet(X2,join(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X2),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X2))),X3)))=X2)).
#
#cnf(i_0_191, plain, (join(meet(X2,X3),meet(X2,join(X2,X3)))=X2)).
##
#cnf(i_0_225, plain, (join(meet(meet(X1,X2),meet(X1,join(X1,X2))),meet(meet(X1,X2),X1))=meet(X1,X2))).
#####
#cnf(i_0_286, plain, (join(meet(meet(meet(X1,X2),meet(X1,join(X1,X2))),meet(meet(X1,X2),X1)),meet(meet(meet(X1,X2),meet(X1,join(X1,X2))),meet(X1,X2)))=meet(meet(X1,X2),meet(X1,join(X1,X2))))).
##
#cnf(i_0_179, plain, (join(meet(join(meet(X1,X3),meet(X3,join(X1,X3))),X4),meet(join(meet(X1,join(join(meet(X3,X5),meet(X6,X3)),X3)),meet(X3,join(X1,join(join(meet(X3,X5),meet(X6,X3)),X3)))),join(join(meet(X1,X3),meet(X3,join(X1,X3))),X4)))=X3)).
#
#cnf(i_0_392, plain, (join(meet(X1,join(join(meet(X2,X3),meet(X4,X2)),X2)),meet(X2,join(X1,join(join(meet(X2,X3),meet(X4,X2)),X2))))=X2)).
#
#cnf(i_0_179, plain, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2)).
#
#cnf(i_0_454, plain, (join(meet(meet(X1,join(join(meet(X2,X3),meet(X4,X2)),X2)),meet(X2,join(X1,join(join(meet(X2,X3),meet(X4,X2)),X2)))),meet(meet(X1,join(join(meet(X2,X3),meet(X4,X2)),X2)),X2))=meet(X1,join(join(meet(X2,X3),meet(X4,X2)),X2)))).
#
#cnf(i_0_485, plain, (join(meet(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3))),meet(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),X2))=meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3))).
#
#cnf(i_0_543, plain, (join(meet(join(meet(meet(X1,X2),meet(X1,join(X1,X2))),meet(meet(X1,join(X1,X2)),X1)),X3),meet(meet(X1,join(X1,X2)),join(join(meet(meet(X1,X2),meet(X1,join(X1,X2))),meet(meet(X1,join(X1,X2)),X1)),X3)))=meet(X1,join(X1,X2)))).
###############
#cnf(i_0_478, plain, (join(meet(X1,join(join(meet(X2,X3),meet(X4,X2)),X2)),meet(join(meet(X2,meet(meet(join(X2,X3),join(X4,X2)),X2)),meet(X5,join(X2,meet(meet(join(X2,X3),join(X4,X2)),X2)))),join(X1,join(join(meet(X2,X3),meet(X4,X2)),X2))))=X2)).
#
#cnf(i_0_806, plain, (join(meet(X2,meet(meet(join(X2,X3),join(X4,X2)),X2)),meet(X5,join(X2,meet(meet(join(X2,X3),join(X4,X2)),X2))))=X2)).
#
#cnf(i_0_880, plain, (join(meet(X1,X2),meet(X2,join(X1,X2)))=X2)).
#
#cnf(i_0_948, plain, (meet(meet(join(X1,X2),join(X3,X1)),X1)=X1)).
#
#cnf(i_0_806, plain, (join(meet(X1,X1),meet(X4,join(X1,X1)))=X1)).
#
#cnf(i_0_1008, plain, (join(meet(X1,X1),join(X1,X1))=X1)).
#
#cnf(i_0_1007, plain, (join(join(meet(X1,X2),meet(X3,X1)),X1)=X1)).
#
#cnf(i_0_1032, plain, (meet(meet(join(X1,X2),X1),X1)=X1)).
#
#cnf(i_0_1033, plain, (join(join(meet(X1,X2),X1),X1)=X1)).
##
#cnf(i_0_1049, plain, (meet(meet(X1,meet(X1,X1)),meet(X1,X1))=meet(X1,X1))).
#
#cnf(i_0_1053, plain, (meet(meet(X2,meet(X1,X2)),meet(X1,X2))=meet(X1,X2))).
#
#cnf(i_0_1099, plain, (meet(meet(X1,X1),X1)=X1)).
####
#cnf(i_0_1113, plain, (join(X1,meet(X1,join(meet(X1,X1),X1)))=X1)).
#
#cnf(i_0_1132, plain, (meet(meet(X1,join(X2,X1)),X1)=X1)).
#
#cnf(i_0_1161, plain, (meet(X1,meet(X1,join(X2,X1)))=meet(X1,join(X2,X1)))).
#
#cnf(i_0_1181, plain, (meet(X1,meet(X1,X1))=meet(X1,X1))).
##
#cnf(i_0_1207, plain, (meet(meet(X1,X1),meet(X1,X1))=meet(X1,X1))).
#
#cnf(i_0_1054, plain, (meet(meet(X1,meet(X1,X2)),meet(X1,X2))=meet(X1,X2))).
#
#cnf(i_0_1248, plain, (meet(X1,X1)=X1)).
######
#cnf(i_0_1008, plain, (join(X1,join(X1,X1))=X1)).
#
#cnf(i_0_1264, plain, (meet(join(X1,X1),X1)=X1)).
#
#cnf(i_0_1310, plain, (join(X1,X1)=X1)).
#######
#cnf(i_0_806, plain, (join(X1,meet(X2,X1))=X1)).
######
#cnf(i_0_1335, plain, (join(meet(X1,meet(X2,X1)),X1)=X1)).
#####
#cnf(i_0_1366, plain, (join(meet(meet(X1,meet(X2,X1)),X1),X1)=X1)).
#
#cnf(i_0_1336, plain, (meet(meet(X1,meet(X2,X1)),X1)=meet(X1,meet(X2,X1)))).
#
#cnf(i_0_1418, plain, (join(X1,meet(X1,meet(X2,X1)))=X1)).
#
#cnf(i_0_1419, plain, (join(meet(X1,meet(X1,meet(X2,X1))),X1)=X1)).
##
#cnf(i_0_1452, plain, (join(X1,meet(X1,meet(X1,meet(X2,X1))))=X1)).
##
#cnf(i_0_1348, plain, (meet(meet(meet(X1,X2),X2),meet(X1,X2))=meet(X1,X2))).
#
#cnf(i_0_1548, plain, (meet(X1,meet(join(X1,X2),X1))=meet(join(X1,X2),X1))).
#
#cnf(i_0_1551, plain, (meet(join(X1,X2),X1)=X1)).
#
#cnf(i_0_1625, plain, (meet(X2,meet(X1,X2))=meet(X1,X2))).
#
#cnf(i_0_1419, plain, (join(meet(X2,X1),X1)=X1)).
###
#cnf(i_0_1626, plain, (meet(X1,meet(X1,X2))=meet(X1,X2))).
#
#cnf(i_0_1336, plain, (meet(meet(X2,X1),X1)=meet(X2,X1))).
#####
#cnf(i_0_1603, plain, (join(X1,meet(X1,join(join(X1,X2),X1)))=X1)).
#
#cnf(i_0_1730, plain, (meet(X1,join(join(X1,X2),X1))=X1)).
#
#cnf(i_0_1770, plain, (meet(join(join(X1,X2),X1),X1)=X1)).
#
#cnf(i_0_1644, plain, (join(join(meet(X2,X1),meet(X3,X1)),X1)=X1)).
#
#cnf(i_0_1621, plain, (meet(X1,join(meet(X1,X2),X1))=join(meet(X1,X2),X1))).
#
#cnf(i_0_1868, plain, (meet(join(meet(X1,X2),X1),X1)=X1)).
#
#cnf(i_0_1879, plain, (join(meet(X1,X2),X1)=X1)).
#
#cnf(i_0_1911, plain, (meet(meet(X1,X2),X1)=meet(X1,X2))).
#
#cnf(i_0_1132, plain, (meet(X1,join(X2,X1))=X1)).
#
#cnf(i_0_1946, plain, (join(X1,meet(X1,X2))=X1)).
#
#cnf(i_0_1950, plain, (meet(X1,join(X1,X2))=X1)).
#
#cnf(i_0_1965, plain, (meet(join(X1,X2),X2)=X2)).
#
#cnf(i_0_1929, plain, (join(X1,join(X1,X2))=join(X1,X2))).
#
#cnf(i_0_1959, plain, (join(X1,join(X2,X1))=join(X2,X1))).
#
#cnf(i_0_1964, plain, (join(join(X1,X2),X2)=join(X1,X2))).
#
#cnf(i_0_1984, plain, (join(join(X1,X2),X1)=join(X1,X2))).
#####
#cnf(i_0_1945, plain, (join(join(meet(X1,X2),meet(X2,X3)),X2)=X2)).
#
#cnf(i_0_1948, plain, (join(join(meet(X1,X2),meet(X1,X3)),X1)=X1)).
#
#cnf(i_0_1951, plain, (meet(X1,meet(join(X1,X2),join(X3,X1)))=X1)).
#
#cnf(i_0_2038, plain, (meet(meet(join(X2,X1),join(X3,X1)),X1)=X1)).
#
#cnf(i_0_2072, plain, (meet(meet(join(X1,X2),join(X1,X3)),X1)=X1)).
#
#cnf(i_0_2075, plain, (join(X1,join(meet(X1,X2),meet(X3,X1)))=X1)).
#
#cnf(i_0_2076, plain, (join(X2,join(meet(X1,X2),meet(X3,X2)))=X2)).
#
#cnf(i_0_2089, plain, (join(X2,join(meet(X1,X2),meet(X2,X3)))=X2)).
##
#cnf(i_0_2121, plain, (join(X1,join(meet(X1,X2),meet(X1,X3)))=X1)).
#
#cnf(i_0_2167, plain, (meet(X1,meet(join(X1,X2),join(X1,X3)))=X1)).
#
#cnf(i_0_2184, plain, (meet(X1,meet(join(X2,X1),join(X3,X1)))=X1)).
##
#cnf(i_0_2211, plain, (meet(meet(join(X1,X2),join(X2,X3)),X2)=X2)).
#######
#cnf(i_0_2508, plain, (meet(X1,meet(join(X2,X1),join(X1,X3)))=X1)).
###
#cnf(i_0_1347, plain, (meet(meet(join(meet(X1,X2),X3),X2),meet(X1,X2))=meet(X1,X2))).
#
#cnf(i_0_1604, plain, (join(join(X1,meet(X3,join(X1,X2))),join(X1,X2))=join(X1,X2))).
#
#cnf(i_0_1670, plain, (meet(meet(X2,join(X3,meet(X1,X2))),meet(X1,X2))=meet(X1,X2))).
#
#cnf(i_0_2911, plain, (meet(meet(X1,X2),meet(X2,X1))=meet(X2,X1))).
## meet is commutative
# AC handling enabled dynamically

#cnf(i_0_2920, plain, (meet(X1,X2)=meet(X2,X1))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(single_axiom, axiom, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(join(meet(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)),meet(join(meet(X2,meet(meet(join(X2,X4),join(X5,X2)),X2)),meet(X6,join(X2,meet(meet(join(X2,X4),join(X5,X2)),X2)))),join(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)))),join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), file('./p/LAT093-1.tptp', single_axiom)).
cnf(prove_wal_axioms_2, negated_conjecture, (meet(b,a)!=meet(a,b)), file('./p/LAT093-1.tptp', prove_wal_axioms_2)).
cnf(c_0_2, axiom, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(join(meet(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)),meet(join(meet(X2,meet(meet(join(X2,X4),join(X5,X2)),X2)),meet(X6,join(X2,meet(meet(join(X2,X4),join(X5,X2)),X2)))),join(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)))),join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), single_axiom).
cnf(c_0_3, plain, (join(meet(join(meet(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),join(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))))),X4),meet(join(meet(X1,join(join(meet(join(meet(X2,X3),meet(X3,join(X2,X3))),X5),meet(X6,join(meet(X2,X3),meet(X3,join(X2,X3))))),join(meet(X2,X3),meet(X3,join(X2,X3))))),meet(X3,join(X1,join(join(meet(join(meet(X2,X3),meet(X3,join(X2,X3))),X5),meet(X6,join(meet(X2,X3),meet(X3,join(X2,X3))))),join(meet(X2,X3),meet(X3,join(X2,X3))))))),join(join(meet(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))),meet(join(meet(X2,X3),meet(X3,join(X2,X3))),join(X1,join(meet(X2,X3),meet(X3,join(X2,X3)))))),X4)))=join(meet(X2,X3),meet(X3,join(X2,X3)))), inference(pm,[status(thm)],[c_0_2, c_0_2])).
cnf(c_0_4, plain, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(join(meet(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)),meet(join(meet(X6,join(join(meet(X2,X7),meet(X8,X2)),X2)),meet(join(meet(X2,meet(meet(join(X2,X7),join(X8,X2)),X2)),meet(X9,join(X2,meet(meet(join(X2,X7),join(X8,X2)),X2)))),join(X6,join(join(meet(X2,X7),meet(X8,X2)),X2)))),join(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)))),join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_3, c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2]), c_0_2])).
cnf(c_0_5, plain, (join(meet(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X2),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X2))),X3),meet(X2,join(join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X2),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X2))),X3)))=X2), inference(pm,[status(thm)],[c_0_4, c_0_2])).
cnf(c_0_6, plain, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(join(meet(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)),meet(X2,join(X1,join(join(meet(X2,X4),meet(X5,X2)),X2)))),join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_3, c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5]), c_0_5])).
cnf(c_0_7, plain, (join(meet(X1,X2),meet(X1,join(X1,X2)))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_5, c_0_5]), c_0_5])).
cnf(c_0_8, plain, (join(meet(X1,join(join(meet(X2,X3),meet(X4,X2)),X2)),meet(X2,join(X1,join(join(meet(X2,X3),meet(X4,X2)),X2))))=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_6, c_0_6]), c_0_6]), c_0_6]), c_0_7])).
cnf(c_0_9, plain, (join(meet(join(meet(X1,X2),meet(X2,join(X1,X2))),X3),meet(X2,join(join(meet(X1,X2),meet(X2,join(X1,X2))),X3)))=X2), inference(rw,[status(thm)],[c_0_6, c_0_8])).
cnf(c_0_10, plain, (join(meet(X1,join(join(meet(X2,X3),meet(X4,X2)),X2)),meet(join(meet(X2,meet(meet(join(X2,X3),join(X4,X2)),X2)),meet(X5,join(X2,meet(meet(join(X2,X3),join(X4,X2)),X2)))),join(X1,join(join(meet(X2,X3),meet(X4,X2)),X2))))=X2), inference(pm,[status(thm)],[c_0_2, c_0_8])).
cnf(c_0_11, plain, (join(meet(X1,meet(meet(join(X1,X2),join(X3,X1)),X1)),meet(X4,join(X1,meet(meet(join(X1,X2),join(X3,X1)),X1))))=X1), inference(pm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_12, plain, (join(meet(X1,X2),meet(X2,join(X1,X2)))=X2), inference(pm,[status(thm)],[c_0_9, c_0_11])).
cnf(c_0_13, plain, (meet(meet(join(X1,X2),join(X3,X1)),X1)=X1), inference(pm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_14, plain, (join(meet(X1,X1),meet(X2,join(X1,X1)))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_13]), c_0_13])).
cnf(c_0_15, plain, (join(join(meet(X1,X2),meet(X3,X1)),X1)=X1), inference(pm,[status(thm)],[c_0_8, c_0_14])).
cnf(c_0_16, plain, (meet(meet(join(X1,X2),X1),X1)=X1), inference(pm,[status(thm)],[c_0_13, c_0_15])).
cnf(c_0_17, plain, (meet(meet(X1,meet(X1,X2)),meet(X1,X2))=meet(X1,X2)), inference(pm,[status(thm)],[c_0_16, c_0_7])).
cnf(c_0_18, plain, (meet(meet(X1,meet(X2,X1)),meet(X2,X1))=meet(X2,X1)), inference(pm,[status(thm)],[c_0_16, c_0_12])).
cnf(c_0_19, plain, (meet(X1,X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_17, c_0_13]), c_0_13]), c_0_13]), c_0_13])).
cnf(c_0_20, plain, (join(meet(X1,X1),join(X1,X1))=X1), inference(pm,[status(thm)],[c_0_14, c_0_13])).
cnf(c_0_21, plain, (meet(meet(X1,X1),X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_18, c_0_13]), c_0_13]), c_0_13])).
cnf(c_0_22, plain, (join(join(meet(X1,X2),X1),X1)=X1), inference(pm,[status(thm)],[c_0_15, c_0_13])).
cnf(c_0_23, plain, (meet(join(X1,X1),X1)=X1), inference(pm,[status(thm)],[c_0_13, c_0_19])).
cnf(c_0_24, plain, (join(X1,join(X1,X1))=X1), inference(rw,[status(thm)],[c_0_20, c_0_19])).
cnf(c_0_25, plain, (join(X1,meet(X1,join(meet(X1,X1),X1)))=X1), inference(pm,[status(thm)],[c_0_12, c_0_21])).
cnf(c_0_26, plain, (join(X1,X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_24])).
cnf(c_0_27, plain, (meet(meet(X1,join(X2,X1)),X1)=X1), inference(pm,[status(thm)],[c_0_13, c_0_25])).
cnf(c_0_28, plain, (join(X1,meet(X2,X1))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_19]), c_0_26])).
cnf(c_0_29, plain, (meet(meet(meet(X1,X2),X2),meet(X1,X2))=meet(X1,X2)), inference(pm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_30, plain, (join(meet(meet(X1,X2),meet(X1,join(X1,X2))),meet(meet(X1,X2),X1))=meet(X1,X2)), inference(pm,[status(thm)],[c_0_7, c_0_7])).
cnf(c_0_31, plain, (meet(X1,meet(join(X1,X2),X1))=meet(join(X1,X2),X1)), inference(pm,[status(thm)],[c_0_29, c_0_16])).
cnf(c_0_32, plain, (meet(join(X1,X2),X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_30, c_0_31]), c_0_28]), c_0_19]), c_0_16]), c_0_31]), c_0_16]), c_0_26]), c_0_31])).
cnf(c_0_33, plain, (meet(X1,join(meet(X1,X2),X1))=join(meet(X1,X2),X1)), inference(pm,[status(thm)],[c_0_32, c_0_22])).
cnf(c_0_34, plain, (join(meet(X1,meet(X2,X1)),X1)=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_28]), c_0_19])).
cnf(c_0_35, plain, (meet(meet(X1,meet(X2,X1)),X1)=meet(X1,meet(X2,X1))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_30, c_0_28]), c_0_19]), c_0_26])).
cnf(c_0_36, plain, (meet(join(meet(X1,X2),X1),X1)=X1), inference(pm,[status(thm)],[c_0_27, c_0_33])).
cnf(c_0_37, plain, (join(meet(X1,meet(X1,meet(X2,X1))),X1)=X1), inference(pm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_38, plain, (meet(X1,meet(X2,X1))=meet(X2,X1)), inference(pm,[status(thm)],[c_0_32, c_0_12])).
cnf(c_0_39, plain, (join(meet(X1,X2),X1)=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_36]), c_0_22]), c_0_36]), c_0_26])).
cnf(c_0_40, plain, (join(meet(X1,X2),X2)=X2), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_38]), c_0_38])).
cnf(c_0_41, plain, (meet(meet(X1,X2),X1)=meet(X1,X2)), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_39]), c_0_26])).
cnf(c_0_42, plain, (meet(meet(X1,join(X2,meet(X3,X1))),meet(X3,X1))=meet(X3,X1)), inference(pm,[status(thm)],[c_0_13, c_0_40])).
cnf(c_0_43, plain, (join(X1,meet(X1,X2))=X1), inference(pm,[status(thm)],[c_0_28, c_0_41])).
cnf(c_0_44, plain, (meet(meet(X1,X2),meet(X2,X1))=meet(X2,X1)), inference(pm,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_45, negated_conjecture, (meet(b,a)!=meet(a,b)), prove_wal_axioms_2).
cnf(c_0_46, plain, (meet(X1,X2)=meet(X2,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_41, c_0_44]), c_0_44]), c_0_44])).
cnf(c_0_47, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_46])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 48
# Proof object clause steps            : 46
# Proof object formula steps           : 2
# Proof object conjectures             : 3
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 2
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 38
# Proof object simplifying inferences  : 61
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 2
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 2
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 2
# Processed clauses                    : 151
# ...of these trivial                  : 67
# ...subsumed                          : 0
# ...remaining for further processing  : 84
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 55
# Generated clauses                    : 3037
# ...of the previous two non-trivial   : 1829
# Contextual simplify-reflections      : 0
# Paramodulations                      : 3037
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 29
#    Positive orientable unit clauses  : 28
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 1576
# ...number of literals in the above   : 1576
# Current number of archived formulas  : 0
# Current number of archived clauses   : 55
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 15
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 365
# BW rewrite match successes           : 105
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 441538
