# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_7, plain, (c_Tarski_Opotype_Opset(v_cl,t_a,tc_Product__Type_Ounit)=v_A)).
#
#cnf(i_0_8, plain, (c_Tarski_Opotype_Opset(c_Tarski_Odual(X1,X2),X2,tc_Product__Type_Ounit)=c_Tarski_Opotype_Opset(X1,X2,tc_Product__Type_Ounit))).
#
#cnf(i_0_5, negated_conjecture, (c_lessequals(v_S,v_A,tc_set(t_a)))).
#
# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(cls_conjecture_0, negated_conjecture, (c_lessequals(v_S,c_Tarski_Opotype_Opset(v_cl,t_a,tc_Product__Type_Ounit),tc_set(t_a))), file('./p/LAT262-2.tptp', cls_conjecture_0)).
cnf(cls_Tarski_OA_A_61_61_Apset_Acl_0, axiom, (v_A=c_Tarski_Opotype_Opset(v_cl,t_a,tc_Product__Type_Ounit)), file('./p/LAT262-2.tptp', cls_Tarski_OA_A_61_61_Apset_Acl_0)).
cnf(cls_conjecture_1, negated_conjecture, (~c_lessequals(v_S,c_Tarski_Opotype_Opset(c_Tarski_Odual(v_cl,t_a),t_a,tc_Product__Type_Ounit),tc_set(t_a))), file('./p/LAT262-2.tptp', cls_conjecture_1)).
cnf(cls_Tarski_Opset_A_Idual_Acl_J_A_61_61_Apset_Acl_0, axiom, (c_Tarski_Opotype_Opset(c_Tarski_Odual(X1,X2),X2,tc_Product__Type_Ounit)=c_Tarski_Opotype_Opset(X1,X2,tc_Product__Type_Ounit)), file('./p/LAT262-2.tptp', cls_Tarski_Opset_A_Idual_Acl_J_A_61_61_Apset_Acl_0)).
cnf(c_0_4, negated_conjecture, (c_lessequals(v_S,c_Tarski_Opotype_Opset(v_cl,t_a,tc_Product__Type_Ounit),tc_set(t_a))), cls_conjecture_0).
cnf(c_0_5, axiom, (v_A=c_Tarski_Opotype_Opset(v_cl,t_a,tc_Product__Type_Ounit)), cls_Tarski_OA_A_61_61_Apset_Acl_0).
cnf(c_0_6, negated_conjecture, (~c_lessequals(v_S,c_Tarski_Opotype_Opset(c_Tarski_Odual(v_cl,t_a),t_a,tc_Product__Type_Ounit),tc_set(t_a))), cls_conjecture_1).
cnf(c_0_7, axiom, (c_Tarski_Opotype_Opset(c_Tarski_Odual(X1,X2),X2,tc_Product__Type_Ounit)=c_Tarski_Opotype_Opset(X1,X2,tc_Product__Type_Ounit)), cls_Tarski_Opset_A_Idual_Acl_J_A_61_61_Apset_Acl_0).
cnf(c_0_8, negated_conjecture, (c_lessequals(v_S,v_A,tc_set(t_a))), inference(rw,[status(thm)],[c_0_4, c_0_5])).
cnf(c_0_9, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_6, c_0_7]), c_0_5]), c_0_8])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 10
# Proof object clause steps            : 6
# Proof object formula steps           : 4
# Proof object conjectures             : 6
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 2
# Proof object initial clauses used    : 4
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 0
# Proof object simplifying inferences  : 5
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 4
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 4
# Processed clauses                    : 4
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 3
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 0
# ...of the previous two non-trivial   : 0
# Contextual simplify-reflections      : 0
# Paramodulations                      : 0
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 3
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 0
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 113
