# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_4, plain, (is_a_theorem(X2)|~is_a_theorem(equivalent(X1,X2))|~is_a_theorem(X1))).
#
#cnf(i_0_5, plain, (is_a_theorem(equivalent(equivalent(X1,X2),equivalent(equivalent(X3,X2),equivalent(X1,X3)))))).
#
#cnf(i_0_7, plain, (is_a_theorem(equivalent(equivalent(X1,X2),equivalent(X3,X1)))|~is_a_theorem(equivalent(X3,X2)))).
#
#cnf(i_0_6, negated_conjecture, (~is_a_theorem(equivalent(equivalent(equivalent(a,b),equivalent(c,a)),equivalent(b,c))))).
#
#cnf(i_0_8, plain, (is_a_theorem(equivalent(X1,X2))|~is_a_theorem(equivalent(X2,X3))|~is_a_theorem(equivalent(X1,X3)))).
#
#cnf(i_0_10, plain, (is_a_theorem(X1)|~is_a_theorem(equivalent(X1,X3))|~is_a_theorem(equivalent(X2,X3))|~is_a_theorem(X2))).
#
#cnf(i_0_17, plain, (is_a_theorem(X1)|~is_a_theorem(equivalent(X1,equivalent(equivalent(X2,X3),equivalent(X4,X2))))|~is_a_theorem(equivalent(X4,X3)))).
#
#cnf(i_0_9, negated_conjecture, (~is_a_theorem(equivalent(equivalent(equivalent(a,b),equivalent(c,a)),X1))|~is_a_theorem(equivalent(equivalent(b,c),X1)))).
#
#cnf(i_0_19, plain, (is_a_theorem(X1)|~is_a_theorem(equivalent(X1,X2))|~is_a_theorem(equivalent(X2,X4))|~is_a_theorem(equivalent(X3,X4))|~is_a_theorem(X3))).
#
#cnf(i_0_20, plain, (is_a_theorem(equivalent(X1,X2))|~is_a_theorem(equivalent(X4,equivalent(equivalent(X3,X2),equivalent(X1,X3))))|~is_a_theorem(X4))).
#
#cnf(i_0_22, plain, (is_a_theorem(X1)|~is_a_theorem(equivalent(X3,X2))|~is_a_theorem(equivalent(X2,X4))|~is_a_theorem(equivalent(X1,X4))|~is_a_theorem(X3))).
#
#cnf(i_0_28, negated_conjecture, (~is_a_theorem(equivalent(equivalent(b,c),equivalent(X1,equivalent(a,b))))|~is_a_theorem(equivalent(X1,equivalent(c,a))))).
#
#cnf(i_0_13, plain, (is_a_theorem(equivalent(equivalent(X1,X2),X3))|~is_a_theorem(equivalent(X3,equivalent(equivalent(X4,X2),equivalent(X1,X4)))))).
#
#cnf(i_0_51, plain, (is_a_theorem(equivalent(equivalent(X1,X2),equivalent(X1,X2))))).
#
#cnf(i_0_57, plain, (is_a_theorem(X1)|~is_a_theorem(equivalent(X1,equivalent(X2,X3)))|~is_a_theorem(equivalent(X2,X3)))).
#
#cnf(i_0_68, negated_conjecture, (~is_a_theorem(equivalent(equivalent(b,c),equivalent(equivalent(a,b),equivalent(c,a)))))).
##
#cnf(i_0_66, plain, (is_a_theorem(equivalent(X1,X2))|~is_a_theorem(equivalent(equivalent(X3,X2),equivalent(X1,X3))))).
#
#cnf(i_0_76, plain, (is_a_theorem(equivalent(X1,X1)))).
###
#cnf(i_0_86, plain, (is_a_theorem(X1)|~is_a_theorem(equivalent(X1,X2))|~is_a_theorem(X2))).
#
#cnf(i_0_90, plain, (is_a_theorem(equivalent(X1,X2))|~is_a_theorem(equivalent(X2,X1)))).
#########
#cnf(i_0_100, negated_conjecture, (~is_a_theorem(equivalent(X1,equivalent(equivalent(a,b),equivalent(c,a))))|~is_a_theorem(equivalent(equivalent(b,c),X1)))).
#
#cnf(i_0_92, plain, (is_a_theorem(X1)|~is_a_theorem(equivalent(X2,X3))|~is_a_theorem(equivalent(X3,X1))|~is_a_theorem(X2))).
###
#cnf(i_0_120, negated_conjecture, (~is_a_theorem(equivalent(equivalent(b,c),equivalent(c,b))))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(condensed_detachment, axiom, (is_a_theorem(X2)|~is_a_theorem(equivalent(X1,X2))|~is_a_theorem(X1)), file('./p/LCL022-1.tptp', condensed_detachment)).
cnf(yql, axiom, (is_a_theorem(equivalent(equivalent(X1,X2),equivalent(equivalent(X3,X2),equivalent(X1,X3))))), file('./p/LCL022-1.tptp', yql)).
cnf(prove_ec_1, negated_conjecture, (~is_a_theorem(equivalent(equivalent(equivalent(a,b),equivalent(c,a)),equivalent(b,c)))), file('./p/LCL022-1.tptp', prove_ec_1)).
cnf(c_0_3, axiom, (is_a_theorem(X2)|~is_a_theorem(equivalent(X1,X2))|~is_a_theorem(X1)), condensed_detachment).
cnf(c_0_4, axiom, (is_a_theorem(equivalent(equivalent(X1,X2),equivalent(equivalent(X3,X2),equivalent(X1,X3))))), yql).
cnf(c_0_5, plain, (is_a_theorem(equivalent(equivalent(X1,X2),equivalent(X3,X1)))|~is_a_theorem(equivalent(X3,X2))), inference(pm,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_6, plain, (is_a_theorem(equivalent(X1,X2))|~is_a_theorem(equivalent(X2,X3))|~is_a_theorem(equivalent(X1,X3))), inference(pm,[status(thm)],[c_0_3, c_0_5])).
cnf(c_0_7, plain, (is_a_theorem(X1)|~is_a_theorem(equivalent(X1,X2))|~is_a_theorem(equivalent(X3,X2))|~is_a_theorem(X3)), inference(pm,[status(thm)],[c_0_3, c_0_6])).
cnf(c_0_8, plain, (is_a_theorem(equivalent(equivalent(X1,X2),X3))|~is_a_theorem(equivalent(X3,equivalent(equivalent(X4,X2),equivalent(X1,X4))))), inference(pm,[status(thm)],[c_0_6, c_0_4])).
cnf(c_0_9, plain, (is_a_theorem(equivalent(X1,X2))|~is_a_theorem(equivalent(X3,equivalent(equivalent(X4,X2),equivalent(X1,X4))))|~is_a_theorem(X3)), inference(pm,[status(thm)],[c_0_7, c_0_4])).
cnf(c_0_10, plain, (is_a_theorem(equivalent(equivalent(X1,X2),equivalent(X1,X2)))), inference(pm,[status(thm)],[c_0_8, c_0_4])).
cnf(c_0_11, plain, (is_a_theorem(equivalent(X1,X2))|~is_a_theorem(equivalent(equivalent(X3,X2),equivalent(X1,X3)))), inference(pm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_12, negated_conjecture, (~is_a_theorem(equivalent(equivalent(equivalent(a,b),equivalent(c,a)),equivalent(b,c)))), prove_ec_1).
cnf(c_0_13, plain, (is_a_theorem(equivalent(X1,X1))), inference(pm,[status(thm)],[c_0_11, c_0_10])).
cnf(c_0_14, negated_conjecture, (~is_a_theorem(equivalent(equivalent(equivalent(a,b),equivalent(c,a)),X1))|~is_a_theorem(equivalent(equivalent(b,c),X1))), inference(pm,[status(thm)],[c_0_12, c_0_6])).
cnf(c_0_15, plain, (is_a_theorem(equivalent(X1,X2))|~is_a_theorem(equivalent(X2,X1))), inference(pm,[status(thm)],[c_0_6, c_0_13])).
cnf(c_0_16, negated_conjecture, (~is_a_theorem(equivalent(X1,equivalent(equivalent(a,b),equivalent(c,a))))|~is_a_theorem(equivalent(equivalent(b,c),X1))), inference(pm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_17, negated_conjecture, (~is_a_theorem(equivalent(equivalent(b,c),equivalent(c,b)))), inference(pm,[status(thm)],[c_0_16, c_0_4])).
cnf(c_0_18, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_17, c_0_5]), c_0_13])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 19
# Proof object clause steps            : 16
# Proof object formula steps           : 3
# Proof object conjectures             : 6
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 3
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 13
# Proof object simplifying inferences  : 2
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 3
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 3
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 3
# Processed clauses                    : 36
# ...of these trivial                  : 0
# ...subsumed                          : 13
# ...remaining for further processing  : 23
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 1
# Generated clauses                    : 130
# ...of the previous two non-trivial   : 114
# Contextual simplify-reflections      : 0
# Paramodulations                      : 130
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 21
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 80
# ...number of literals in the above   : 296
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 67
# Rec. Clause-clause subsumption calls : 65
# Non-unit clause-clause subsumptions  : 12
# Unit Clause-clause subsumption calls : 14
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 26
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1386
