# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_6, plain, (implies(truth,X1)=X1)).
#
#cnf(i_0_8, plain, (implies(implies(X1,X2),X2)=implies(implies(X2,X1),X1))).
#
#cnf(i_0_11, plain, (implies(X1,X1)=implies(implies(X1,truth),truth))).
#
#cnf(i_0_10, negated_conjecture, (implies(not(implies(a,b)),not(b))!=truth)).
##
#cnf(i_0_9, plain, (implies(implies(not(X1),not(X2)),implies(X2,X1))=truth)).
#
#cnf(i_0_26, plain, (implies(implies(not(X1),not(truth)),X1)=truth)).
#
#cnf(i_0_7, plain, (implies(implies(X1,X2),implies(implies(X2,X3),implies(X1,X3)))=truth)).
#
#cnf(i_0_45, plain, (implies(X1,implies(implies(X1,X2),X2))=truth)).
#
#cnf(i_0_65, plain, (implies(X1,X1)=truth)).
##
#cnf(i_0_11, plain, (implies(implies(X1,truth),truth)=truth)).
#
#cnf(i_0_102, plain, (implies(X1,truth)=truth)).
#
#cnf(i_0_113, plain, (implies(X2,implies(X1,X2))=truth)).
#####
#cnf(i_0_120, plain, (implies(implies(implies(X2,X1),X1),X1)=implies(X2,X1))).
##
#cnf(i_0_157, plain, (implies(implies(implies(X2,X1),X1),X2)=implies(X1,X2))).
###
#cnf(i_0_121, plain, (implies(implies(implies(X2,X1),X3),implies(X1,X3))=truth)).
#
#cnf(i_0_229, plain, (implies(not(truth),X1)=truth)).
#
#cnf(i_0_227, plain, (implies(not(X2),implies(X2,X1))=truth)).
#
#cnf(i_0_224, plain, (implies(X2,implies(X3,implies(X1,X2)))=truth)).
#
#cnf(i_0_234, plain, (implies(implies(X1,not(truth)),not(truth))=X1)).
#
#cnf(i_0_332, plain, (not(not(truth))=truth)).
#
#cnf(i_0_325, plain, (implies(not(implies(X1,not(truth))),X1)=truth)).
#
#cnf(i_0_365, plain, (implies(X1,implies(not(X1),not(truth)))=truth)).
#
#cnf(i_0_378, plain, (implies(not(X1),not(truth))=X1)).
#
#cnf(i_0_397, plain, (implies(X1,not(truth))=not(X1))).
#
#cnf(i_0_234, plain, (not(not(X1))=X1)).
###############
#cnf(i_0_442, plain, (implies(X1,implies(not(X1),X2))=truth)).
#
#cnf(i_0_425, plain, (implies(not(implies(X1,X2)),not(X2))=truth)).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(wajsberg_2, axiom, (implies(implies(X1,X2),implies(implies(X2,X3),implies(X1,X3)))=truth), file('./p/LCL115-2.tptp', wajsberg_2)).
cnf(wajsberg_1, axiom, (implies(truth,X1)=X1), file('./p/LCL115-2.tptp', wajsberg_1)).
cnf(wajsberg_3, axiom, (implies(implies(X1,X2),X2)=implies(implies(X2,X1),X1)), file('./p/LCL115-2.tptp', wajsberg_3)).
cnf(wajsberg_4, axiom, (implies(implies(not(X1),not(X2)),implies(X2,X1))=truth), file('./p/LCL115-2.tptp', wajsberg_4)).
cnf(prove_mv_39, negated_conjecture, (implies(not(implies(a,b)),not(b))!=truth), file('./p/LCL115-2.tptp', prove_mv_39)).
cnf(c_0_5, axiom, (implies(implies(X1,X2),implies(implies(X2,X3),implies(X1,X3)))=truth), wajsberg_2).
cnf(c_0_6, axiom, (implies(truth,X1)=X1), wajsberg_1).
cnf(c_0_7, axiom, (implies(implies(X1,X2),X2)=implies(implies(X2,X1),X1)), wajsberg_3).
cnf(c_0_8, plain, (implies(X1,implies(implies(X1,X2),X2))=truth), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_5, c_0_6]), c_0_6])).
cnf(c_0_9, plain, (implies(X1,X1)=implies(implies(X1,truth),truth)), inference(pm,[status(thm)],[c_0_7, c_0_6])).
cnf(c_0_10, plain, (implies(X1,X1)=truth), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_6, c_0_8]), c_0_6])).
cnf(c_0_11, plain, (implies(implies(X1,truth),truth)=truth), inference(rw,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_12, plain, (implies(X1,truth)=truth), inference(pm,[status(thm)],[c_0_8, c_0_11])).
cnf(c_0_13, plain, (implies(X1,implies(X2,X1))=truth), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_5, c_0_12]), c_0_6]), c_0_6])).
cnf(c_0_14, axiom, (implies(implies(not(X1),not(X2)),implies(X2,X1))=truth), wajsberg_4).
cnf(c_0_15, plain, (implies(implies(implies(X1,X2),X3),implies(X2,X3))=truth), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_5, c_0_13]), c_0_6])).
cnf(c_0_16, plain, (implies(implies(not(X1),not(truth)),X1)=truth), inference(pm,[status(thm)],[c_0_14, c_0_6])).
cnf(c_0_17, plain, (implies(not(truth),X1)=truth), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_15, c_0_16]), c_0_6])).
cnf(c_0_18, plain, (implies(not(X1),implies(X1,X2))=truth), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_15, c_0_14]), c_0_6])).
cnf(c_0_19, plain, (implies(implies(X1,not(truth)),not(truth))=X1), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_17]), c_0_6])).
cnf(c_0_20, plain, (implies(not(implies(X1,not(truth))),X1)=truth), inference(pm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_21, plain, (implies(X1,implies(not(X1),not(truth)))=truth), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_14, c_0_20]), c_0_6])).
cnf(c_0_22, plain, (implies(not(X1),not(truth))=X1), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_21]), c_0_6]), c_0_16]), c_0_6])).
cnf(c_0_23, plain, (implies(X1,not(truth))=not(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_22]), c_0_17]), c_0_6])).
cnf(c_0_24, negated_conjecture, (implies(not(implies(a,b)),not(b))!=truth), prove_mv_39).
cnf(c_0_25, plain, (implies(not(implies(X1,X2)),not(X2))=truth), inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_15, c_0_23]), c_0_23])).
cnf(c_0_26, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 27
# Proof object clause steps            : 22
# Proof object formula steps           : 5
# Proof object conjectures             : 3
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 1
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 18
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 5
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 5
# Processed clauses                    : 50
# ...of these trivial                  : 22
# ...subsumed                          : 1
# ...remaining for further processing  : 27
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 9
# Generated clauses                    : 495
# ...of the previous two non-trivial   : 182
# Contextual simplify-reflections      : 0
# Paramodulations                      : 495
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 18
#    Positive orientable unit clauses  : 17
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 128
# ...number of literals in the above   : 128
# Current number of archived formulas  : 0
# Current number of archived clauses   : 9
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 41
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3428
