# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_160, negated_conjecture, (actual_world(skc6))).
#
#cnf(i_0_168, negated_conjecture, (group(skc6,skc7))).
#
#cnf(i_0_167, negated_conjecture, (two(skc6,skc7))).
#
#cnf(i_0_93, plain, (~member(X1,X2,X2))).
#
#cnf(i_0_174, negated_conjecture, (event(skc6,skc10))).
#
#cnf(i_0_172, negated_conjecture, (barrel(skc6,skc10))).
#
#cnf(i_0_161, negated_conjecture, (chevy(skc6,skc13))).
#
#cnf(i_0_146, plain, (~multiple(X1,X2)|~singleton(X1,X2))).
#
#cnf(i_0_165, negated_conjecture, (street(skc6,skc11))).
#
#cnf(i_0_162, negated_conjecture, (placename(skc6,skc12))).
#
#cnf(i_0_163, negated_conjecture, (hollywood_placename(skc6,skc12))).
#
#cnf(i_0_145, plain, (~general(X1,X2)|~specific(X1,X2))).
#
#cnf(i_0_164, negated_conjecture, (city(skc6,skc11))).
#
#cnf(i_0_171, negated_conjecture, (old(skc6,skc13))).
#
#cnf(i_0_166, negated_conjecture, (lonely(skc6,skc11))).
#
#cnf(i_0_149, plain, (~nonexistent(X1,X2)|~existent(X1,X2))).
#
#cnf(i_0_169, negated_conjecture, (white(skc6,skc13))).
#
#cnf(i_0_170, negated_conjecture, (dirty(skc6,skc13))).
#
#cnf(i_0_173, negated_conjecture, (present(skc6,skc10))).
#
#cnf(i_0_147, plain, (~nonliving(X1,X2)|~living(X1,X2))).
#
#cnf(i_0_94, plain, (man(X1,X2)|~fellow(X1,X2))).
#
#cnf(i_0_124, plain, (entity(X1,X2)|~object(X1,X2))).
#
#cnf(i_0_111, plain, (thing(X1,X2)|~eventuality(X1,X2))).
#
#cnf(i_0_148, plain, (~nonhuman(X1,X2)|~human(X1,X2))).
#
#cnf(i_0_133, plain, (thing(X1,X2)|~abstraction(X1,X2))).
#
#cnf(i_0_109, plain, (group(X1,X2)|~two(X1,X2))).
#
#cnf(i_0_110, plain, (eventuality(X1,X2)|~state(X1,X2))).
#
#cnf(i_0_150, plain, (~nonliving(X1,X2)|~animate(X1,X2))).
#
#cnf(i_0_115, plain, (event(X1,X2)|~state(X1,X2))).
#
#cnf(i_0_117, plain, (event(X1,X2)|~barrel(X1,X2))).
#
#cnf(i_0_142, plain, (instrumentality(X1,X2)|~furniture(X1,X2))).
#
#cnf(i_0_144, plain, (~unisex(X1,X2)|~male(X1,X2))).
#
#cnf(i_0_129, plain, (artifact(X1,X2)|~way(X1,X2))).
#
#cnf(i_0_123, plain, (object(X1,X2)|~artifact(X1,X2))).
#
#cnf(i_0_139, plain, (object(X1,X2)|~location(X1,X2))).
#
#cnf(i_0_143, plain, (~young(X1,X2)|~old(X1,X2))).
#
#cnf(i_0_137, plain, (placename(X1,X2)|~hollywood_placename(X1,X2))).
#
#cnf(i_0_132, plain, (abstraction(X1,X2)|~relation(X1,X2))).
#
#cnf(i_0_140, plain, (seat(X1,X2)|~frontseat(X1,X2))).
#
#cnf(i_0_154, plain, (skf8(X1,X2)!=skf10(X1,X2)|~two(X2,X1))).
#
#cnf(i_0_141, plain, (furniture(X1,X2)|~seat(X1,X2))).
#
#cnf(i_0_95, plain, (human_person(X1,X2)|~man(X1,X2))).
#
#cnf(i_0_96, plain, (organism(X1,X2)|~human_person(X1,X2))).
#
#cnf(i_0_97, plain, (entity(X1,X2)|~organism(X1,X2))).
#
#cnf(i_0_98, plain, (thing(X1,X2)|~entity(X1,X2))).
#
#cnf(i_0_99, plain, (singleton(X1,X2)|~thing(X1,X2))).
#
#cnf(i_0_100, plain, (specific(X1,X2)|~entity(X1,X2))).
#
#cnf(i_0_112, plain, (specific(X1,X2)|~eventuality(X1,X2))).
#
#cnf(i_0_101, plain, (existent(X1,X2)|~entity(X1,X2))).
#
#cnf(i_0_102, plain, (impartial(X1,X2)|~organism(X1,X2))).
#
#cnf(i_0_126, plain, (impartial(X1,X2)|~object(X1,X2))).
#
#cnf(i_0_103, plain, (living(X1,X2)|~organism(X1,X2))).
#
#cnf(i_0_104, plain, (human(X1,X2)|~human_person(X1,X2))).
#
#cnf(i_0_105, plain, (animate(X1,X2)|~human_person(X1,X2))).
#
#cnf(i_0_106, plain, (male(X1,X2)|~man(X1,X2))).
#
#cnf(i_0_107, plain, (set(X1,X2)|~group(X1,X2))).
#
#cnf(i_0_108, plain, (multiple(X1,X2)|~set(X1,X2))).
#
#cnf(i_0_116, plain, (eventuality(X1,X2)|~event(X1,X2))).
#
#cnf(i_0_113, plain, (nonexistent(X1,X2)|~eventuality(X1,X2))).
#
#cnf(i_0_188, plain, (~set(X1,X2)|~singleton(X1,X2))).
#
#cnf(i_0_114, plain, (unisex(X1,X2)|~eventuality(X1,X2))).
#
#cnf(i_0_127, plain, (unisex(X1,X2)|~object(X1,X2))).
#
#cnf(i_0_136, plain, (unisex(X1,X2)|~abstraction(X1,X2))).
#
#cnf(i_0_190, plain, (~eventuality(X1,X2)|~existent(X1,X2))).
#
#cnf(i_0_118, plain, (car(X1,X2)|~chevy(X1,X2))).
#
#cnf(i_0_119, plain, (vehicle(X1,X2)|~car(X1,X2))).
#
#cnf(i_0_120, plain, (transport(X1,X2)|~vehicle(X1,X2))).
#
#cnf(i_0_191, plain, (~group(X1,X2)|~singleton(X1,X2))).
#
#cnf(i_0_121, plain, (instrumentality(X1,X2)|~transport(X1,X2))).
#
#cnf(i_0_122, plain, (artifact(X1,X2)|~instrumentality(X1,X2))).
#
#cnf(i_0_197, plain, (object(X1,X2)|~instrumentality(X1,X2))).
#
#cnf(i_0_192, plain, (~eventuality(X1,X2)|~male(X1,X2))).
#
#cnf(i_0_198, plain, (object(X1,X2)|~transport(X1,X2))).
#
#cnf(i_0_200, plain, (object(X1,X2)|~vehicle(X1,X2))).
#
#cnf(i_0_201, plain, (object(X1,X2)|~car(X1,X2))).
#
#cnf(i_0_193, plain, (~object(X1,X2)|~male(X1,X2))).
#
#cnf(i_0_125, plain, (nonliving(X1,X2)|~object(X1,X2))).
#
#cnf(i_0_128, plain, (way(X1,X2)|~street(X1,X2))).
#
#cnf(i_0_130, plain, (relname(X1,X2)|~placename(X1,X2))).
#
#cnf(i_0_194, plain, (~abstraction(X1,X2)|~male(X1,X2))).
#
#cnf(i_0_131, plain, (relation(X1,X2)|~relname(X1,X2))).
#
#cnf(i_0_208, plain, (abstraction(X1,X2)|~relname(X1,X2))).
#
#cnf(i_0_134, plain, (nonhuman(X1,X2)|~abstraction(X1,X2))).
#
#cnf(i_0_195, plain, (~eventuality(X1,X2)|~entity(X1,X2))).
#
#cnf(i_0_135, plain, (general(X1,X2)|~abstraction(X1,X2))).
#
#cnf(i_0_138, plain, (location(X1,X2)|~city(X1,X2))).
#
#cnf(i_0_189, plain, (thing(X1,X2)|~event(X1,X2))).
#
#cnf(i_0_196, plain, (~group(X1,X2)|~thing(X1,X2))).
#
#cnf(i_0_202, plain, (object(X1,X2)|~chevy(X1,X2))).
#
#cnf(i_0_206, plain, (artifact(X1,X2)|~street(X1,X2))).
#
#cnf(i_0_209, plain, (abstraction(X1,X2)|~placename(X1,X2))).
#
#cnf(i_0_199, plain, (~eventuality(X1,X2)|~man(X1,X2))).
#
#cnf(i_0_213, plain, (object(X1,X2)|~city(X1,X2))).
#
#cnf(i_0_216, plain, (entity(X1,X2)|~chevy(X1,X2))).
#
#cnf(i_0_217, plain, (object(X1,X2)|~street(X1,X2))).
#
#cnf(i_0_203, plain, (~object(X1,X2)|~man(X1,X2))).
#
#cnf(i_0_218, plain, (thing(X1,X2)|~placename(X1,X2))).
#
#cnf(i_0_220, plain, (entity(X1,X2)|~city(X1,X2))).
#
#cnf(i_0_221, plain, (entity(X1,X2)|~street(X1,X2))).
#
#cnf(i_0_204, plain, (~object(X1,X2)|~living(X1,X2))).
#
#cnf(i_0_175, negated_conjecture, (of(skc6,skc12,skc11))).
#
#cnf(i_0_176, negated_conjecture, (down(skc6,skc10,skc11))).
#
#cnf(i_0_177, negated_conjecture, (in(skc6,skc10,skc11))).
#
#cnf(i_0_205, plain, (~object(X1,X2)|~animate(X1,X2))).
#
#cnf(i_0_178, negated_conjecture, (agent(skc6,skc10,skc13))).
#
#cnf(i_0_151, plain, (X3=X4|~be(X1,X2,X3,X4))).
#
#cnf(i_0_180, negated_conjecture, (fellow(skc6,X1)|~member(skc6,X1,skc7))).
#
#cnf(i_0_207, plain, (~abstraction(X1,X2)|~man(X1,X2))).
#
#cnf(i_0_179, negated_conjecture, (young(skc6,X1)|~member(skc6,X1,skc7))).
#
#cnf(i_0_152, plain, (member(X1,skf10(X2,X1),X2)|~two(X1,X2))).
#
#cnf(i_0_229, negated_conjecture, (fellow(skc6,skf10(skc7,skc6)))).
#
#cnf(i_0_210, plain, (~abstraction(X1,X2)|~human(X1,X2))).
#
#cnf(i_0_230, negated_conjecture, (young(skc6,skf10(skc7,skc6)))).
#
#cnf(i_0_231, negated_conjecture, (man(skc6,skf10(skc7,skc6)))).
#
#cnf(i_0_153, plain, (member(X1,skf8(X2,X1),X2)|~two(X1,X2))).
#
#cnf(i_0_211, plain, (~event(X1,X2)|~entity(X1,X2))).
#
#cnf(i_0_234, negated_conjecture, (fellow(skc6,skf8(skc7,skc6)))).
#
#cnf(i_0_235, negated_conjecture, (young(skc6,skf8(skc7,skc6)))).
#
#cnf(i_0_240, negated_conjecture, (man(skc6,skf8(skc7,skc6)))).
#
#cnf(i_0_212, plain, (~abstraction(X1,X2)|~specific(X1,X2))).
#
#cnf(i_0_182, negated_conjecture, (state(skc6,skf6(X2))|~member(skc6,X1,skc7))).
#
#cnf(i_0_246, negated_conjecture, (state(skc6,skf6(X1)))).
#
#cnf(i_0_248, negated_conjecture, (eventuality(skc6,skf6(X1)))).
#
#cnf(i_0_251, negated_conjecture, (~entity(skc6,skf6(X1)))).
#
#cnf(i_0_249, negated_conjecture, (event(skc6,skf6(X1)))).
#
#cnf(i_0_250, negated_conjecture, (thing(skc6,skf6(X1)))).
#
#cnf(i_0_181, negated_conjecture, (frontseat(skc6,skf5(X2))|~member(skc6,X1,skc7))).
#
#cnf(i_0_252, negated_conjecture, (~man(skc6,skf6(X1)))).
#
#cnf(i_0_259, negated_conjecture, (frontseat(skc6,skf5(X1)))).
#
#cnf(i_0_261, negated_conjecture, (seat(skc6,skf5(X1)))).
#
#cnf(i_0_262, negated_conjecture, (furniture(skc6,skf5(X1)))).
#
#cnf(i_0_253, negated_conjecture, (~organism(skc6,skf6(X1)))).
#
#cnf(i_0_263, negated_conjecture, (instrumentality(skc6,skf5(X1)))).
#
#cnf(i_0_265, negated_conjecture, (object(skc6,skf5(X1)))).
#
#cnf(i_0_266, negated_conjecture, (entity(skc6,skf5(X1)))).
#
#cnf(i_0_254, negated_conjecture, (~chevy(skc6,skf6(X1)))).
#
#cnf(i_0_183, negated_conjecture, (in(skc6,skf5(X2),skf5(X2))|~member(skc6,X1,skc7))).
#
#cnf(i_0_269, negated_conjecture, (in(skc6,skf5(X1),skf5(X1)))).
#
#cnf(i_0_184, negated_conjecture, (be(skc6,skf6(X1),X1,skf5(X1))|~member(skc6,X1,skc7))).
#
#cnf(i_0_255, negated_conjecture, (~city(skc6,skf6(X1)))).
#
#cnf(i_0_271, negated_conjecture, (skf5(X1)=X1|~member(skc6,X1,skc7))).
#
#cnf(i_0_272, negated_conjecture, (skf5(skf8(skc7,skc6))=skf8(skc7,skc6))).
#
#cnf(i_0_277, negated_conjecture, (furniture(skc6,skf8(skc7,skc6)))).
#
#cnf(i_0_256, negated_conjecture, (~street(skc6,skf6(X1)))).
#
#cnf(i_0_278, negated_conjecture, (frontseat(skc6,skf8(skc7,skc6)))).
#
#cnf(i_0_279, negated_conjecture, (seat(skc6,skf8(skc7,skc6)))).
#
#cnf(i_0_280, negated_conjecture, (instrumentality(skc6,skf8(skc7,skc6)))).
#
#cnf(i_0_257, negated_conjecture, (~group(skc6,skf6(X1)))).
#
#cnf(i_0_281, negated_conjecture, (object(skc6,skf8(skc7,skc6)))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(clause89, negated_conjecture, (frontseat(skc6,skf5(X2))|~member(skc6,X1,skc7)), file('./p/NLP141-1.tptp', clause89)).
cnf(clause61, axiom, (member(X1,skf8(X2,X1),X2)|~two(X1,X2)), file('./p/NLP141-1.tptp', clause61)).
cnf(clause75, negated_conjecture, (two(skc6,skc7)), file('./p/NLP141-1.tptp', clause75)).
cnf(clause48, axiom, (seat(X1,X2)|~frontseat(X1,X2)), file('./p/NLP141-1.tptp', clause48)).
cnf(clause49, axiom, (furniture(X1,X2)|~seat(X1,X2)), file('./p/NLP141-1.tptp', clause49)).
cnf(clause31, axiom, (object(X1,X2)|~artifact(X1,X2)), file('./p/NLP141-1.tptp', clause31)).
cnf(clause30, axiom, (artifact(X1,X2)|~instrumentality(X1,X2)), file('./p/NLP141-1.tptp', clause30)).
cnf(clause50, axiom, (instrumentality(X1,X2)|~furniture(X1,X2)), file('./p/NLP141-1.tptp', clause50)).
cnf(clause59, axiom, (X3=X4|~be(X1,X2,X3,X4)), file('./p/NLP141-1.tptp', clause59)).
cnf(clause92, negated_conjecture, (be(skc6,skf6(X1),X1,skf5(X1))|~member(skc6,X1,skc7)), file('./p/NLP141-1.tptp', clause92)).
cnf(clause52, axiom, (~male(X1,X2)|~unisex(X1,X2)), file('./p/NLP141-1.tptp', clause52)).
cnf(clause35, axiom, (unisex(X1,X2)|~object(X1,X2)), file('./p/NLP141-1.tptp', clause35)).
cnf(clause88, negated_conjecture, (fellow(skc6,X1)|~member(skc6,X1,skc7)), file('./p/NLP141-1.tptp', clause88)).
cnf(clause14, axiom, (male(X1,X2)|~man(X1,X2)), file('./p/NLP141-1.tptp', clause14)).
cnf(clause2, axiom, (man(X1,X2)|~fellow(X1,X2)), file('./p/NLP141-1.tptp', clause2)).
cnf(c_0_15, negated_conjecture, (frontseat(skc6,skf5(X2))|~member(skc6,X1,skc7)), clause89).
cnf(c_0_16, axiom, (member(X1,skf8(X2,X1),X2)|~two(X1,X2)), clause61).
cnf(c_0_17, negated_conjecture, (two(skc6,skc7)), clause75).
cnf(c_0_18, axiom, (seat(X1,X2)|~frontseat(X1,X2)), clause48).
cnf(c_0_19, negated_conjecture, (frontseat(skc6,skf5(X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_15, c_0_16]), c_0_17])])).
cnf(c_0_20, axiom, (furniture(X1,X2)|~seat(X1,X2)), clause49).
cnf(c_0_21, negated_conjecture, (seat(skc6,skf5(X1))), inference(pm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_22, axiom, (object(X1,X2)|~artifact(X1,X2)), clause31).
cnf(c_0_23, axiom, (artifact(X1,X2)|~instrumentality(X1,X2)), clause30).
cnf(c_0_24, axiom, (instrumentality(X1,X2)|~furniture(X1,X2)), clause50).
cnf(c_0_25, negated_conjecture, (furniture(skc6,skf5(X1))), inference(pm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26, axiom, (X3=X4|~be(X1,X2,X3,X4)), clause59).
cnf(c_0_27, negated_conjecture, (be(skc6,skf6(X1),X1,skf5(X1))|~member(skc6,X1,skc7)), clause92).
cnf(c_0_28, axiom, (~male(X1,X2)|~unisex(X1,X2)), clause52).
cnf(c_0_29, axiom, (unisex(X1,X2)|~object(X1,X2)), clause35).
cnf(c_0_30, plain, (object(X1,X2)|~instrumentality(X1,X2)), inference(pm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_31, negated_conjecture, (instrumentality(skc6,skf5(X1))), inference(pm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_32, negated_conjecture, (skf5(X1)=X1|~member(skc6,X1,skc7)), inference(pm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_33, negated_conjecture, (fellow(skc6,X1)|~member(skc6,X1,skc7)), clause88).
cnf(c_0_34, plain, (~object(X1,X2)|~male(X1,X2)), inference(pm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_35, axiom, (male(X1,X2)|~man(X1,X2)), clause14).
cnf(c_0_36, negated_conjecture, (object(skc6,skf5(X1))), inference(pm,[status(thm)],[c_0_30, c_0_31])).
cnf(c_0_37, negated_conjecture, (skf5(skf8(skc7,skc6))=skf8(skc7,skc6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_32, c_0_16]), c_0_17])])).
cnf(c_0_38, axiom, (man(X1,X2)|~fellow(X1,X2)), clause2).
cnf(c_0_39, negated_conjecture, (fellow(skc6,skf8(skc7,skc6))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_33, c_0_16]), c_0_17])])).
cnf(c_0_40, plain, (~object(X1,X2)|~man(X1,X2)), inference(pm,[status(thm)],[c_0_34, c_0_35])).
cnf(c_0_41, negated_conjecture, (object(skc6,skf8(skc7,skc6))), inference(pm,[status(thm)],[c_0_36, c_0_37])).
cnf(c_0_42, negated_conjecture, (man(skc6,skf8(skc7,skc6))), inference(pm,[status(thm)],[c_0_38, c_0_39])).
cnf(c_0_43, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_40, c_0_41]), c_0_42])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 44
# Proof object clause steps            : 29
# Proof object formula steps           : 15
# Proof object conjectures             : 19
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 4
# Proof object initial clauses used    : 15
# Proof object initial formulas used   : 15
# Proof object generating inferences   : 14
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 92
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 92
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 92
# Processed clauses                    : 149
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 149
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 0
# Generated clauses                    : 104
# ...of the previous two non-trivial   : 99
# Contextual simplify-reflections      : 0
# Paramodulations                      : 104
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 146
#    Positive orientable unit clauses  : 42
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 96
# Current number of unprocessed clauses: 35
# ...number of literals in the above   : 73
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3
# Clause-clause subsumption calls (NU) : 1659
# Rec. Clause-clause subsumption calls : 1657
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 40
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2785
