# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_4141, hypothesis, (aSet0(slcrc0))).
##
#cnf(i_0_45, plain, (aSet0(szNzAzT0))).
#
#cnf(i_0_4140, hypothesis, (~aElementOf0(X1,slcrc0))).
#
#cnf(i_0_146, hypothesis, (aSet0(xT))).
#
#cnf(i_0_151, hypothesis, (aSet0(xS))).
#
#cnf(i_0_4149, negated_conjecture, (aSet0(esk31_0))).
##
#cnf(i_0_8, plain, (isFinite0(slcrc0))).
#
#cnf(i_0_145, hypothesis, (isFinite0(xT))).
#
#cnf(i_0_44, plain, (isCountable0(szNzAzT0))).
#
#cnf(i_0_6, plain, (X2!=slcrc0|~aElementOf0(X1,X2))).
#
#cnf(i_0_148, hypothesis, (isCountable0(xS))).
#
#cnf(i_0_166, hypothesis, (aFunction0(xc))).
#
#cnf(i_0_4136, hypothesis, (xK=sz00)).
#
#cnf(i_0_10, plain, (~isCountable0(X1)|~isFinite0(X1)|~aSet0(X1))).
#
#cnf(i_0_46, plain, (aElementOf0(sz00,szNzAzT0))).
##
#cnf(i_0_4138, hypothesis, (aSubsetOf0(slcrc0,xS))).
#
#cnf(i_0_4150, hypothesis, (~isFinite0(xS))).
#
#cnf(i_0_149, hypothesis, (aSubsetOf0(xS,szNzAzT0))).
#
#cnf(i_0_4147, negated_conjecture, (aSubsetOf0(esk31_0,xS))).
#
#cnf(i_0_17, plain, (aSubsetOf0(X1,X1)|~aSet0(X1))).
#
#cnf(i_0_4151, plain, (~isFinite0(szNzAzT0))).
#
#cnf(i_0_7, plain, (aSet0(X1)|X1!=slcrc0)).
#
#cnf(i_0_4146, negated_conjecture, (sbrdtbr0(esk31_0)=sz00)).
#
#cnf(i_0_92, plain, (slbdtrb0(sz00)=slcrc0)).
#
#cnf(i_0_4152, plain, (szNzAzT0!=slcrc0)).
#
#cnf(i_0_116, plain, (aSet0(szDzozmdt0(X1))|~aFunction0(X1))).
#
#cnf(i_0_65, plain, (aElement0(sbrdtbr0(X1))|~aSet0(X1))).
#
#cnf(i_0_4153, negated_conjecture, (aElement0(sz00))).
#
#cnf(i_0_52, plain, (szszuzczcdt0(X1)!=X1|~aElementOf0(X1,szNzAzT0))).
#
#cnf(i_0_59, plain, (sdtlseqdt0(X1,X1)|~aElementOf0(X1,szNzAzT0))).
#
#cnf(i_0_15, plain, (aSet0(X1)|~aSubsetOf0(X1,X2)|~aSet0(X2))).
#
#cnf(i_0_3, plain, (aElement0(X2)|~aElementOf0(X2,X1)|~aSet0(X1))).
#
#cnf(i_0_4154, plain, (szszuzczcdt0(sz00)!=sz00)).
#
#cnf(i_0_165, hypothesis, (aSet0(X1)|~aElementOf0(X1,szDzozmdt0(xc)))).
#
#cnf(i_0_150, hypothesis, (aElementOf0(X1,szNzAzT0)|~aElementOf0(X1,xS))).
#
#cnf(i_0_91, plain, (isFinite0(slbdtrb0(X1))|~aElementOf0(X1,szNzAzT0))).
#
#cnf(i_0_11, plain, (X1!=slcrc0|~isCountable0(X1)|~aSet0(X1))).
##
#cnf(i_0_4148, negated_conjecture, (aElementOf0(X1,xS)|~aElementOf0(X1,esk31_0))).
#
#cnf(i_0_4164, negated_conjecture, (aElement0(X1)|~aElementOf0(X1,esk31_0))).
#
#cnf(i_0_4161, hypothesis, (xS!=slcrc0)).
#
#cnf(i_0_54, plain, (sdtlseqdt0(sz00,X1)|~aElementOf0(X1,szNzAzT0))).
#
#cnf(i_0_4165, hypothesis, (aElementOf0(X1,szNzAzT0)|~aElementOf0(X1,esk31_0))).
###
#cnf(i_0_158, hypothesis, (slbdtsldtrb0(xS,sz00)=szDzozmdt0(xc))).
#
#cnf(i_0_4137, hypothesis, (aElementOf0(slcrc0,szDzozmdt0(xc)))).
#
#cnf(i_0_4170, hypothesis, (aElement0(slcrc0)|~aSet0(szDzozmdt0(xc)))).
#
#cnf(i_0_4169, hypothesis, (szDzozmdt0(xc)!=slcrc0)).
#
#cnf(i_0_4173, hypothesis, (aElement0(slcrc0))).
#
#cnf(i_0_4145, negated_conjecture, (aElementOf0(esk31_0,szDzozmdt0(xc)))).
#
#cnf(i_0_4175, negated_conjecture, (aElement0(esk31_0)|~aSet0(szDzozmdt0(xc)))).
##
#cnf(i_0_4178, negated_conjecture, (aElement0(esk31_0))).
#
#cnf(i_0_157, hypothesis, (aSet0(sdtlcdtrc0(xc,szDzozmdt0(xc))))).
#
#cnf(i_0_58, plain, (sdtlseqdt0(X1,szszuzczcdt0(X1))|~aElementOf0(X1,szNzAzT0))).
##
#cnf(i_0_64, plain, (iLess0(X1,szszuzczcdt0(X1))|~aElementOf0(X1,szNzAzT0))).
#
#cnf(i_0_163, hypothesis, (aSubsetOf0(X1,xS)|~aElementOf0(X1,szDzozmdt0(xc)))).
#
#cnf(i_0_100, plain, (sbrdtbr0(slbdtrb0(X1))=X1|~aElementOf0(X1,szNzAzT0))).
##
#cnf(i_0_4182, plain, (sbrdtbr0(slcrc0)=sz00)).
#
#cnf(i_0_48, plain, (aElementOf0(szszuzczcdt0(X1),szNzAzT0)|~aElementOf0(X1,szNzAzT0))).
#
#cnf(i_0_4185, plain, (aElement0(szszuzczcdt0(X1))|~aElementOf0(X1,szNzAzT0))).
##
#cnf(i_0_16, plain, (isFinite0(X2)|~aSubsetOf0(X2,X1)|~isFinite0(X1)|~aSet0(X1))).
##
#cnf(i_0_69, plain, (X1=slcrc0|sbrdtbr0(X1)!=sz00|~aSet0(X1))).
##
#cnf(i_0_4191, negated_conjecture, (slcrc0=esk31_0)).
#
#cnf(i_0_8, plain, (isFinite0(esk31_0))).
#
#cnf(i_0_7, plain, (aSet0(X1)|X1!=esk31_0)).
#
#cnf(i_0_4161, hypothesis, (xS!=esk31_0)).
#
#cnf(i_0_92, plain, (slbdtrb0(sz00)=esk31_0)).
#
#cnf(i_0_69, plain, (X1=esk31_0|sbrdtbr0(X1)!=sz00|~aSet0(X1))).
#
#cnf(i_0_4201, hypothesis, (xT=esk31_0|sbrdtbr0(xT)!=sz00)).
#
#cnf(i_0_4152, plain, (szNzAzT0!=esk31_0)).
#
#cnf(i_0_4204, plain, (sbrdtbr0(szNzAzT0)!=sz00)).
#
#cnf(i_0_4181, plain, (aElement0(X1)|~aElementOf0(X1,szNzAzT0)|~aSet0(slbdtrb0(X1)))).
#
#cnf(i_0_68, plain, (sbrdtbr0(X1)=sz00|X1!=esk31_0|~aSet0(X1))).
##
#cnf(i_0_4214, plain, (sbrdtbr0(X1)=sz00|X1!=esk31_0)).
#
#cnf(i_0_162, hypothesis, (sbrdtbr0(X1)=sz00|~aElementOf0(X1,szDzozmdt0(xc)))).
#
#cnf(i_0_67, plain, (isFinite0(X1)|~aElementOf0(sbrdtbr0(X1),szNzAzT0)|~aSet0(X1))).
##
#cnf(i_0_4222, plain, (isFinite0(X1)|X1!=esk31_0|~aSet0(X1))).
#
#cnf(i_0_4223, hypothesis, (isFinite0(X1)|~aElementOf0(sbrdtbr0(X1),esk31_0)|~aSet0(X1))).
#
#cnf(i_0_4207, plain, (aElement0(X1)|slbdtrb0(X1)!=esk31_0|~aElementOf0(X1,szNzAzT0))).
###
#cnf(i_0_66, plain, (aElementOf0(sbrdtbr0(X1),szNzAzT0)|~isFinite0(X1)|~aSet0(X1))).
###
#cnf(i_0_152, hypothesis, (aSubsetOf0(sdtlcdtrc0(xc,szDzozmdt0(xc)),xT))).
#
#cnf(i_0_4240, hypothesis, (isFinite0(sdtlcdtrc0(xc,szDzozmdt0(xc))))).
#
#cnf(i_0_90, plain, (aSet0(X1)|X1!=slbdtrb0(X2)|~aElementOf0(X2,szNzAzT0))).
#
#cnf(i_0_4140, hypothesis, (~aElementOf0(X1,esk31_0))).
##
#cnf(i_0_4241, plain, (aSet0(slbdtrb0(X1))|~aElementOf0(X1,szNzAzT0))).
#
#cnf(i_0_4249, plain, (aElement0(X1)|~aElementOf0(X1,szNzAzT0))).
#
#cnf(i_0_4169, hypothesis, (szDzozmdt0(xc)!=esk31_0)).
####
#cnf(i_0_4200, hypothesis, (sbrdtbr0(xS)!=sz00)).
##
#cnf(i_0_5, plain, (X1=esk31_0|aElementOf0(esk1_1(X1),X1)|~aSet0(X1))).
#
#cnf(i_0_4261, plain, (aElement0(esk1_1(szNzAzT0)))).
##
#cnf(i_0_4256, hypothesis, (aElementOf0(esk1_1(xS),szNzAzT0))).
#
#cnf(i_0_4264, hypothesis, (aElement0(esk1_1(xS)))).
##
#cnf(i_0_6, plain, (X1!=esk31_0|~aElementOf0(X2,X1))).
#
#cnf(i_0_4262, plain, (aSet0(X1)|X1!=slbdtrb0(esk1_1(szNzAzT0)))).
#
#cnf(i_0_4276, plain, (aSet0(slbdtrb0(esk1_1(szNzAzT0))))).
#
#cnf(i_0_4267, hypothesis, (aSet0(X1)|X1!=slbdtrb0(esk1_1(xS)))).
##
#cnf(i_0_4279, hypothesis, (aSet0(slbdtrb0(esk1_1(xS))))).
#
#cnf(i_0_4263, plain, (X1=esk31_0|aElement0(esk1_1(X1))|~aSet0(X1))).
#
#cnf(i_0_4258, hypothesis, (aSet0(esk1_1(szDzozmdt0(xc)))|~aSet0(szDzozmdt0(xc)))).
##
#cnf(i_0_4216, plain, (sz00=X1|slbdtrb0(X1)!=esk31_0|~aElementOf0(X1,szNzAzT0))).
####
#cnf(i_0_42, plain, (isFinite0(sdtpldt0(X2,X1))|~isFinite0(X2)|~aElement0(X1)|~aSet0(X2))).
#
#cnf(i_0_43, plain, (isFinite0(sdtmndt0(X2,X1))|~isFinite0(X2)|~aElement0(X1)|~aSet0(X2))).
#
#cnf(i_0_40, plain, (isCountable0(sdtpldt0(X2,X1))|~isCountable0(X2)|~aElement0(X1)|~aSet0(X2))).
##
#cnf(i_0_41, plain, (isCountable0(sdtmndt0(X2,X1))|~isCountable0(X2)|~aElement0(X1)|~aSet0(X2))).
#
#cnf(i_0_117, plain, (aElement0(sdtlpdtrp0(X1,X2))|~aFunction0(X1)|~aElementOf0(X2,szDzozmdt0(X1)))).
#
#cnf(i_0_99, plain, (aElementOf0(esk10_1(X1),szNzAzT0)|~aSubsetOf0(X1,szNzAzT0)|~isFinite0(X1))).
##
#cnf(i_0_4291, plain, (aElement0(esk10_1(X1))|~aSubsetOf0(X1,szNzAzT0)|~isFinite0(X1))).
##
#cnf(i_0_28, plain, (aSet0(X1)|X1!=sdtpldt0(X2,X3)|~aElement0(X3)|~aSet0(X2))).
##
#cnf(i_0_4296, plain, (aSet0(sdtpldt0(X1,X2))|~aElement0(X2)|~aSet0(X1))).
#
#cnf(i_0_37, plain, (aSet0(X1)|X1!=sdtmndt0(X2,X3)|~aElement0(X3)|~aSet0(X2))).
#
#cnf(i_0_4299, plain, (aSet0(sdtmndt0(X1,X2))|~aElement0(X2)|~aSet0(X1))).
##
#cnf(i_0_124, plain, (aSet0(X1)|X1!=sdtlbdtrb0(X2,X3)|~aFunction0(X2)|~aElement0(X3))).
#
#cnf(i_0_4302, plain, (aSet0(sdtlbdtrb0(X1,X2))|~aFunction0(X1)|~aElement0(X2))).
#
#cnf(i_0_14, plain, (aElementOf0(X1,X3)|~aSubsetOf0(X2,X3)|~aElementOf0(X1,X2)|~aSet0(X3))).
#
#cnf(i_0_4308, hypothesis, (~aSubsetOf0(X2,esk31_0)|~aElementOf0(X1,X2))).
#####
#cnf(i_0_4315, plain, (aElement0(X1)|~aSubsetOf0(X2,szNzAzT0)|~aElementOf0(X1,X2))).
##
#cnf(i_0_4348, hypothesis, (aElement0(X1)|~aElementOf0(X1,xS))).
#
#cnf(i_0_4332, hypothesis, (~aSubsetOf0(szNzAzT0,esk31_0))).
###
#cnf(i_0_4351, hypothesis, (aElement0(X1)|~aSubsetOf0(X2,xS)|~aElementOf0(X1,X2))).
#
#cnf(i_0_11, plain, (X1!=esk31_0|~isCountable0(X1)|~aSet0(X1))).
###
#cnf(i_0_4333, hypothesis, (X1=esk31_0|~aSubsetOf0(X1,esk31_0)|~aSet0(X1))).
##
#cnf(i_0_4318, plain, (aElement0(X1)|~aSubsetOf0(X3,X2)|~aElementOf0(X1,X3)|~aSet0(X2))).
#########
#cnf(i_0_4377, hypothesis, (aElement0(X1)|~aElementOf0(X1,sdtlcdtrc0(xc,szDzozmdt0(xc))))).
#
#cnf(i_0_4309, hypothesis, (aElementOf0(X1,szNzAzT0)|~aSubsetOf0(X2,xS)|~aElementOf0(X1,X2))).
#
#cnf(i_0_4260, plain, (szszuzczcdt0(esk1_1(szNzAzT0))!=esk1_1(szNzAzT0))).
####
#cnf(i_0_4265, hypothesis, (szszuzczcdt0(esk1_1(xS))!=esk1_1(xS))).
##
#cnf(i_0_4321, hypothesis, (aElementOf0(esk1_1(xS),X1)|~aSubsetOf0(szNzAzT0,X1)|~aSet0(X1))).
#
#cnf(i_0_4245, plain, (aSet0(X1)|X1!=slbdtrb0(szszuzczcdt0(X2))|~aElementOf0(X2,szNzAzT0))).
##
#cnf(i_0_4424, plain, (aSet0(slbdtrb0(szszuzczcdt0(X1)))|~aElementOf0(X1,szNzAzT0))).
#
#cnf(i_0_4283, plain, (esk1_1(szNzAzT0)=sz00|slbdtrb0(esk1_1(szNzAzT0))!=esk31_0)).
#
#cnf(i_0_4284, hypothesis, (esk1_1(xS)=sz00|slbdtrb0(esk1_1(xS))!=esk31_0)).
#
#cnf(i_0_47, plain, (szszuzczcdt0(X1)!=sz00|~aElementOf0(X1,szNzAzT0))).
#
#cnf(i_0_125, plain, (aSubsetOf0(sdtlbdtrb0(X1,X2),szDzozmdt0(X1))|~aFunction0(X1)|~aElement0(X2))).
#######
#cnf(i_0_4429, plain, (szszuzczcdt0(esk1_1(szNzAzT0))!=sz00)).
#
#cnf(i_0_153, hypothesis, (aElementOf0(X1,xT)|~aElementOf0(X1,sdtlcdtrc0(xc,szDzozmdt0(xc))))).
###
#cnf(i_0_4430, hypothesis, (szszuzczcdt0(esk1_1(xS))!=sz00)).
##
#cnf(i_0_4203, plain, (szDzozmdt0(X1)=esk31_0|sbrdtbr0(szDzozmdt0(X1))!=sz00|~aFunction0(X1))).
#
#cnf(i_0_98, plain, (aSubsetOf0(X1,slbdtrb0(esk10_1(X1)))|~aSubsetOf0(X1,szNzAzT0)|~isFinite0(X1))).
##
#cnf(i_0_4244, plain, (aSet0(X1)|X1!=slbdtrb0(sbrdtbr0(X2))|~isFinite0(X2)|~aSet0(X2))).
##
#cnf(i_0_4449, plain, (aSet0(slbdtrb0(sbrdtbr0(X1)))|~isFinite0(X1)|~aSet0(X1))).
##
#cnf(i_0_38, plain, (sdtpldt0(sdtmndt0(X1,X2),X2)=X1|~aElementOf0(X2,X1)|~aSet0(X1))).
##
#cnf(i_0_4306, plain, (isFinite0(X1)|~aSubsetOf0(X2,szNzAzT0)|~aElementOf0(sbrdtbr0(X1),X2)|~aSet0(X1))).
###
#cnf(i_0_4381, hypothesis, (isFinite0(X1)|~aSubsetOf0(X2,xS)|~aElementOf0(sbrdtbr0(X1),X2)|~aSet0(X1))).
####
#cnf(i_0_12, plain, (aSubsetOf0(X2,X1)|~aElementOf0(esk2_2(X1,X2),X1)|~aSet0(X2)|~aSet0(X1))).
#
#cnf(i_0_4259, hypothesis, (aSubsetOf0(esk1_1(szDzozmdt0(xc)),xS)|~aSet0(szDzozmdt0(xc)))).
#
# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(m__, conjecture, ![X1]:(((((aSet0(X1)&![X2]:(aElementOf0(X2,X1)=>aElementOf0(X2,xS)))&aSubsetOf0(X1,xS))&sbrdtbr0(X1)=sz00)&aElementOf0(X1,slbdtsldtrb0(xS,sz00)))=>((aSet0(slcrc0)&~(?[X2]:aElementOf0(X2,slcrc0)))=>sdtlpdtrp0(xc,X1)=sdtlpdtrp0(xc,slcrc0))), file('./p/NUM565+3.tptp', m__)).
fof(mCardEmpty, axiom, ![X1]:(aSet0(X1)=>(sbrdtbr0(X1)=sz00<=>X1=slcrc0)), file('./p/NUM565+3.tptp', mCardEmpty)).
fof(c_0_2, negated_conjecture, ~(![X1]:(((((aSet0(X1)&![X2]:(aElementOf0(X2,X1)=>aElementOf0(X2,xS)))&aSubsetOf0(X1,xS))&sbrdtbr0(X1)=sz00)&aElementOf0(X1,slbdtsldtrb0(xS,sz00)))=>((aSet0(slcrc0)&~(?[X2]:aElementOf0(X2,slcrc0)))=>sdtlpdtrp0(xc,X1)=sdtlpdtrp0(xc,slcrc0)))), inference(assume_negation,[status(cth)],[m__])).
fof(c_0_3, plain, ![X77]:((sbrdtbr0(X77)!=sz00|X77=slcrc0|~aSet0(X77))&(X77!=slcrc0|sbrdtbr0(X77)=sz00|~aSet0(X77))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mCardEmpty])])])).
fof(c_0_4, negated_conjecture, ![X200, X201]:(((((aSet0(esk31_0)&(~aElementOf0(X200,esk31_0)|aElementOf0(X200,xS)))&aSubsetOf0(esk31_0,xS))&sbrdtbr0(esk31_0)=sz00)&aElementOf0(esk31_0,slbdtsldtrb0(xS,sz00)))&((aSet0(slcrc0)&~aElementOf0(X201,slcrc0))&sdtlpdtrp0(xc,esk31_0)!=sdtlpdtrp0(xc,slcrc0))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])])).
cnf(c_0_5, plain, (X1=slcrc0|sbrdtbr0(X1)!=sz00|~aSet0(X1)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6, negated_conjecture, (aSet0(esk31_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, negated_conjecture, (sbrdtbr0(esk31_0)=sz00), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, negated_conjecture, (sdtlpdtrp0(xc,esk31_0)!=sdtlpdtrp0(xc,slcrc0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9, negated_conjecture, (slcrc0=esk31_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_5, c_0_6]), c_0_7])])).
cnf(c_0_10, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_9])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 11
# Proof object clause steps            : 6
# Proof object formula steps           : 5
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 4
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 1
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 80
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 4149
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 4142
# Processed clauses                    : 220
# ...of these trivial                  : 3
# ...subsumed                          : 73
# ...remaining for further processing  : 143
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 17
# Generated clauses                    : 338
# ...of the previous two non-trivial   : 289
# Contextual simplify-reflections      : 0
# Paramodulations                      : 330
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 123
#    Positive orientable unit clauses  : 28
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 14
#    Non-unit-clauses                  : 81
# Current number of unprocessed clauses: 4202
# ...number of literals in the above   : 51351
# Current number of archived formulas  : 0
# Current number of archived clauses   : 20
# Clause-clause subsumption calls (NU) : 483
# Rec. Clause-clause subsumption calls : 362
# Non-unit clause-clause subsumptions  : 27
# Unit Clause-clause subsumption calls : 353
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 330995
