# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_13, negated_conjecture, (aSet0(esk2_0))).
#
#cnf(i_0_8, plain, (isFinite0(slcrc0))).
#
#cnf(i_0_12, negated_conjecture, (isCountable0(esk2_0))).
#
#cnf(i_0_6, plain, (X2!=slcrc0|~aElementOf0(X1,X2))).
#
#cnf(i_0_11, negated_conjecture, (esk2_0=slcrc0)).
#
#cnf(i_0_12, negated_conjecture, (isCountable0(slcrc0))).
#
#cnf(i_0_13, negated_conjecture, (aSet0(slcrc0))).
#
#cnf(i_0_10, plain, (~isCountable0(X1)|~isFinite0(X1)|~aSet0(X1))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(m__, conjecture, ![X1]:((aSet0(X1)&isCountable0(X1))=>X1!=slcrc0), file('./p/NUM531+1.tptp', m__)).
fof(mCountNFin, axiom, ![X1]:((aSet0(X1)&isCountable0(X1))=>~(isFinite0(X1))), file('./p/NUM531+1.tptp', mCountNFin)).
fof(mEmpFin, axiom, isFinite0(slcrc0), file('./p/NUM531+1.tptp', mEmpFin)).
fof(c_0_3, negated_conjecture, ~(![X1]:((aSet0(X1)&isCountable0(X1))=>X1!=slcrc0)), inference(assume_negation,[status(cth)],[m__])).
fof(c_0_4, negated_conjecture, ((aSet0(esk2_0)&isCountable0(esk2_0))&esk2_0=slcrc0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
fof(c_0_5, plain, ![X9]:(~aSet0(X9)|~isCountable0(X9)|~isFinite0(X9)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[mCountNFin])])])).
cnf(c_0_6, negated_conjecture, (isCountable0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, negated_conjecture, (esk2_0=slcrc0), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, negated_conjecture, (aSet0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_9, plain, (~aSet0(X1)|~isCountable0(X1)|~isFinite0(X1)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_10, negated_conjecture, (isCountable0(slcrc0)), inference(rw,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_11, plain, (isFinite0(slcrc0)), inference(split_conjunct,[status(thm)],[mEmpFin])).
cnf(c_0_12, negated_conjecture, (aSet0(slcrc0)), inference(rw,[status(thm)],[c_0_8, c_0_7])).
cnf(c_0_13, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_9, c_0_10]), c_0_11]), c_0_12])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 14
# Proof object clause steps            : 8
# Proof object formula steps           : 6
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 1
# Proof object simplifying inferences  : 5
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 4
# Initial clauses in saturation        : 9
# Processed clauses                    : 8
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 8
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 1
# ...of the previous two non-trivial   : 2
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 6
#    Positive orientable unit clauses  : 4
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 2
# Current number of unprocessed clauses: 3
# ...number of literals in the above   : 8
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 532
