# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_2, plain, (aNaturalNumber0(sz00))).
#
#cnf(i_0_4, plain, (aNaturalNumber0(sz10))).
#
#cnf(i_0_73, hypothesis, (aNaturalNumber0(xn))).
#
#cnf(i_0_3, plain, (sz00!=sz10)).
#
#cnf(i_0_72, hypothesis, (aNaturalNumber0(xm))).
#
#cnf(i_0_71, hypothesis, (aNaturalNumber0(xp))).
#
#cnf(i_0_222, hypothesis, (aNaturalNumber0(xk))).
#
#cnf(i_0_207, hypothesis, (sz00!=xp)).
#
#cnf(i_0_235, hypothesis, (aNaturalNumber0(xr))).
#
#cnf(i_0_202, hypothesis, (aNaturalNumber0(esk9_0))).
#
#cnf(i_0_218, hypothesis, (aNaturalNumber0(esk10_0))).
#
#cnf(i_0_206, hypothesis, (xp!=sz10)).
#
#cnf(i_0_214, hypothesis, (aNaturalNumber0(esk11_0))).
#
#cnf(i_0_234, hypothesis, (aNaturalNumber0(esk12_0))).
#
#cnf(i_0_240, hypothesis, (aNaturalNumber0(esk13_0))).
#
#cnf(i_0_219, hypothesis, (xp!=xn)).
#
#cnf(i_0_238, hypothesis, (aNaturalNumber0(esk14_0))).
#
#cnf(i_0_203, hypothesis, (isPrime0(xp))).
#
#cnf(i_0_227, hypothesis, (isPrime0(xr))).
#
#cnf(i_0_215, hypothesis, (xp!=xm)).
#
#cnf(i_0_216, hypothesis, (sdtlseqdt0(xn,xp))).
#
#cnf(i_0_212, hypothesis, (sdtlseqdt0(xm,xp))).
#
#cnf(i_0_232, hypothesis, (doDivides0(xr,xk))).
#
#cnf(i_0_224, hypothesis, (sz00!=xk)).
#
#cnf(i_0_32, plain, (sdtlseqdt0(X1,X1)|~aNaturalNumber0(X1))).
#
#cnf(i_0_233, hypothesis, (sdtasdt0(xr,esk12_0)=xk)).
#
#cnf(i_0_217, hypothesis, (sdtpldt0(xn,esk10_0)=xp)).
##
#cnf(i_0_213, hypothesis, (sdtpldt0(xm,esk11_0)=xp)).
#
#cnf(i_0_239, hypothesis, (sdtpldt0(xr,esk13_0)=xk)).
#
#cnf(i_0_9, plain, (sdtpldt0(sz00,X1)=X1|~aNaturalNumber0(X1))).
#
#cnf(i_0_223, hypothesis, (xk!=sz10)).
#
#cnf(i_0_13, plain, (sdtasdt0(sz10,X1)=X1|~aNaturalNumber0(X1))).
#
#cnf(i_0_10, plain, (sdtpldt0(X1,sz00)=X1|~aNaturalNumber0(X1))).
#
#cnf(i_0_14, plain, (sdtasdt0(X1,sz10)=X1|~aNaturalNumber0(X1))).
##
#cnf(i_0_241, negated_conjecture, (xk=xp|~sdtlseqdt0(xk,xp))).
#
#cnf(i_0_200, hypothesis, (doDivides0(xp,sdtasdt0(xn,xm)))).
#
#cnf(i_0_236, hypothesis, (doDivides0(xr,sdtasdt0(xn,xm)))).
#
#cnf(i_0_231, hypothesis, (sz00!=xr)).
#
#cnf(i_0_15, plain, (sdtasdt0(sz00,X1)=sz00|~aNaturalNumber0(X1))).
#
#cnf(i_0_16, plain, (sdtasdt0(X1,sz00)=sz00|~aNaturalNumber0(X1))).
#
#cnf(i_0_5, plain, (aNaturalNumber0(sdtpldt0(X1,X2))|~aNaturalNumber0(X2)|~aNaturalNumber0(X1))).
#
#cnf(i_0_230, hypothesis, (xr!=sz10)).
#
#cnf(i_0_6, plain, (aNaturalNumber0(sdtasdt0(X1,X2))|~aNaturalNumber0(X2)|~aNaturalNumber0(X1))).
#
#cnf(i_0_221, hypothesis, (sdtasdt0(xp,xk)=sdtasdt0(xn,xm))).
#
#cnf(i_0_255, hypothesis, (aNaturalNumber0(sdtasdt0(xn,xm)))).
#
#cnf(i_0_208, hypothesis, (~sdtlseqdt0(xp,xn))).
#
#cnf(i_0_201, hypothesis, (sdtasdt0(xp,esk9_0)=sdtasdt0(xn,xm))).
#
#cnf(i_0_237, hypothesis, (sdtasdt0(xr,esk14_0)=sdtasdt0(xn,xm))).
#
#cnf(i_0_36, plain, (sdtlseqdt0(X2,X1)|X1!=X2|~aNaturalNumber0(X2)|~aNaturalNumber0(X1))).
##
#cnf(i_0_220, hypothesis, (sdtsldt0(sdtasdt0(xn,xm),xp)=xk)).
#
#cnf(i_0_35, plain, (sdtlseqdt0(X2,X1)|sdtlseqdt0(X1,X2)|~aNaturalNumber0(X2)|~aNaturalNumber0(X1))).
##
#cnf(i_0_210, hypothesis, (~sdtlseqdt0(xp,xm))).
##
#cnf(i_0_262, negated_conjecture, (xk=xp|sdtlseqdt0(xp,xk))).
###
#cnf(i_0_242, negated_conjecture, (xk=xp|sdtpldt0(xk,X1)!=xp|~aNaturalNumber0(X1))).
#
#cnf(i_0_28, plain, (aNaturalNumber0(esk1_2(X1,X2))|~sdtlseqdt0(X1,X2)|~aNaturalNumber0(X2)|~aNaturalNumber0(X1))).
#
#cnf(i_0_52, plain, (aNaturalNumber0(esk2_2(X1,X2))|~doDivides0(X1,X2)|~aNaturalNumber0(X2)|~aNaturalNumber0(X1))).
#
#cnf(i_0_243, negated_conjecture, (~sdtlseqdt0(xp,xk))).
#
#cnf(i_0_262, negated_conjecture, (xk=xp)).
#
#cnf(i_0_232, hypothesis, (doDivides0(xr,xp))).
#
#cnf(i_0_239, hypothesis, (sdtpldt0(xr,esk13_0)=xp)).
##
#cnf(i_0_233, hypothesis, (sdtasdt0(xr,esk12_0)=xp)).
#
#cnf(i_0_221, hypothesis, (sdtasdt0(xp,xp)=sdtasdt0(xn,xm))).
#
#cnf(i_0_220, hypothesis, (sdtsldt0(sdtasdt0(xn,xm),xp)=xp)).
##
#cnf(i_0_7, plain, (sdtpldt0(X1,X2)=sdtpldt0(X2,X1)|~aNaturalNumber0(X2)|~aNaturalNumber0(X1))).
###
#cnf(i_0_243, negated_conjecture, (~sdtlseqdt0(xp,xp))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(m__, conjecture, (~((?[X1]:(aNaturalNumber0(X1)&sdtpldt0(xp,X1)=xk)|sdtlseqdt0(xp,xk)))=>(xk!=xp&(?[X1]:(aNaturalNumber0(X1)&sdtpldt0(xk,X1)=xp)|sdtlseqdt0(xk,xp)))), file('./p/NUM505+3.tptp', m__)).
fof(mLETotal, axiom, ![X1, X2]:((aNaturalNumber0(X1)&aNaturalNumber0(X2))=>(sdtlseqdt0(X1,X2)|(X2!=X1&sdtlseqdt0(X2,X1)))), file('./p/NUM505+3.tptp', mLETotal)).
fof(m__1837, hypothesis, ((aNaturalNumber0(xn)&aNaturalNumber0(xm))&aNaturalNumber0(xp)), file('./p/NUM505+3.tptp', m__1837)).
fof(m__2306, hypothesis, ((aNaturalNumber0(xk)&sdtasdt0(xn,xm)=sdtasdt0(xp,xk))&xk=sdtsldt0(sdtasdt0(xn,xm),xp)), file('./p/NUM505+3.tptp', m__2306)).
fof(c_0_4, negated_conjecture, ~((~((?[X1]:(aNaturalNumber0(X1)&sdtpldt0(xp,X1)=xk)|sdtlseqdt0(xp,xk)))=>(xk!=xp&(?[X1]:(aNaturalNumber0(X1)&sdtpldt0(xk,X1)=xp)|sdtlseqdt0(xk,xp))))), inference(assume_negation,[status(cth)],[m__])).
fof(c_0_5, negated_conjecture, ![X108, X109]:(((~aNaturalNumber0(X108)|sdtpldt0(xp,X108)!=xk)&~sdtlseqdt0(xp,xk))&((~aNaturalNumber0(X109)|sdtpldt0(xk,X109)!=xp|xk=xp)&(~sdtlseqdt0(xk,xp)|xk=xp))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
fof(c_0_6, plain, ![X49, X50]:((X50!=X49|sdtlseqdt0(X49,X50)|(~aNaturalNumber0(X49)|~aNaturalNumber0(X50)))&(sdtlseqdt0(X50,X49)|sdtlseqdt0(X49,X50)|(~aNaturalNumber0(X49)|~aNaturalNumber0(X50)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[mLETotal])])])).
cnf(c_0_7, negated_conjecture, (xk=xp|~sdtlseqdt0(xk,xp)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, plain, (sdtlseqdt0(X1,X2)|sdtlseqdt0(X2,X1)|~aNaturalNumber0(X2)|~aNaturalNumber0(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, hypothesis, (aNaturalNumber0(xp)), inference(split_conjunct,[status(thm)],[m__1837])).
cnf(c_0_10, hypothesis, (aNaturalNumber0(xk)), inference(split_conjunct,[status(thm)],[m__2306])).
cnf(c_0_11, negated_conjecture, (xk=xp|sdtlseqdt0(xp,xk)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_8]), c_0_9]), c_0_10])])).
cnf(c_0_12, negated_conjecture, (~sdtlseqdt0(xp,xk)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_13, negated_conjecture, (xk=xp), inference(sr,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_14, negated_conjecture, (~sdtlseqdt0(xp,xp)), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_15, plain, (sdtlseqdt0(X2,X1)|X1!=X2|~aNaturalNumber0(X2)|~aNaturalNumber0(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_16, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_14, c_0_15]), c_0_9])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 17
# Proof object clause steps            : 10
# Proof object formula steps           : 7
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 6
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 7
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 50
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 244
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 241
# Processed clauses                    : 76
# ...of these trivial                  : 0
# ...subsumed                          : 11
# ...remaining for further processing  : 65
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 11
# Generated clauses                    : 38
# ...of the previous two non-trivial   : 23
# Contextual simplify-reflections      : 0
# Paramodulations                      : 35
# Factorizations                       : 2
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 53
#    Positive orientable unit clauses  : 29
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 14
# Current number of unprocessed clauses: 187
# ...number of literals in the above   : 1577
# Current number of archived formulas  : 0
# Current number of archived clauses   : 12
# Clause-clause subsumption calls (NU) : 18
# Rec. Clause-clause subsumption calls : 16
# Non-unit clause-clause subsumptions  : 5
# Unit Clause-clause subsumption calls : 77
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 15743
