# Initializing proof state
# Scanning for AC axioms
# intersection is commutative
# AC handling enabled
#
#cnf(i_0_10, plain, (disjoint(X1,X2)|intersect(X1,X2))).
#
#cnf(i_0_13, plain, (intersect(X2,X1)|~intersect(X1,X2))).
#
#cnf(i_0_12, plain, (intersection(X1,X2)=intersection(X2,X1))).
#
#cnf(i_0_11, plain, (~disjoint(X1,X2)|~intersect(X1,X2))).
#
#cnf(i_0_2, plain, (member(X1,X2)|~member(X1,intersection(X3,X2)))).
#
#cnf(i_0_3, plain, (member(X1,X2)|~member(X1,intersection(X2,X3)))).
#
#cnf(i_0_6, plain, (member(X1,X2)|~member(X1,difference(X2,X3)))).
#
#cnf(i_0_5, plain, (~member(X1,difference(X3,X2))|~member(X1,X2))).
####
#cnf(i_0_18, negated_conjecture, (~disjoint(intersection(esk3_0,esk4_0),difference(esk3_0,esk4_0)))).
##
#cnf(i_0_8, plain, (member(esk1_2(X1,X2),X2)|~intersect(X1,X2))).
#
#cnf(i_0_9, plain, (member(esk1_2(X1,X2),X1)|~intersect(X1,X2))).
#
#cnf(i_0_27, plain, (~intersect(X1,difference(X2,X3))|~member(esk1_2(X1,difference(X2,X3)),X3))).
#
#cnf(i_0_7, plain, (intersect(X2,X3)|~member(X1,X3)|~member(X1,X2))).
##
#cnf(i_0_36, plain, (intersect(X1,X2)|~intersect(X3,X1)|~member(esk1_2(X3,X1),X2))).
#
#cnf(i_0_34, plain, (~intersect(X1,difference(X2,X1)))).
#
#cnf(i_0_40, plain, (intersect(X1,X1)|~intersect(X2,X1))).
#####
#cnf(i_0_45, plain, (intersect(X1,X1)|~intersect(X1,X2))).
##
#cnf(i_0_31, plain, (~intersect(difference(X1,X2),X3)|~member(esk1_2(difference(X1,X2),X3),X2))).
####
#cnf(i_0_53, plain, (~intersect(difference(X1,X2),X2))).
##
#cnf(i_0_46, plain, (intersect(X1,X1)|~member(X3,X1)|~member(X3,X2))).
#######
#cnf(i_0_37, plain, (intersect(X1,X2)|~intersect(X1,X3)|~member(esk1_2(X1,X3),X2))).
##
#cnf(i_0_38, plain, (intersect(X1,X2)|~intersect(X3,X2)|~member(esk1_2(X3,X2),X1))).
##
#cnf(i_0_39, plain, (intersect(X1,X2)|~intersect(X2,X3)|~member(esk1_2(X2,X3),X1))).
###
#cnf(i_0_25, negated_conjecture, (intersect(intersection(esk3_0,esk4_0),difference(esk3_0,esk4_0)))).
#
#cnf(i_0_1, plain, (member(X1,intersection(X2,X3))|~member(X1,X3)|~member(X1,X2))).
###
#cnf(i_0_70, plain, (~intersect(difference(X1,intersection(X2,X3)),X4)|~member(esk1_2(difference(X1,intersection(X2,X3)),X4),X3)|~member(esk1_2(difference(X1,intersection(X2,X3)),X4),X2))).
#
#cnf(i_0_59, plain, (intersect(X1,X1)|~intersect(X2,X3)|~member(esk1_2(X2,X3),X1))).
###
#cnf(i_0_87, plain, (~intersect(difference(X1,intersection(X2,X3)),X2)|~member(esk1_2(difference(X1,intersection(X2,X3)),X2),X3))).
####
#cnf(i_0_101, plain, (~intersect(difference(X1,intersection(X2,X2)),X2))).
#####
#cnf(i_0_4, plain, (member(X1,difference(X2,X3))|member(X1,X3)|~member(X1,X2))).
#
#cnf(i_0_69, negated_conjecture, (intersect(intersection(esk3_0,esk4_0),intersection(esk3_0,esk4_0)))).
####
#cnf(i_0_15, plain, (X1=X2|~member(esk2_2(X1,X2),X2)|~member(esk2_2(X1,X2),X1))).
#
#cnf(i_0_26, plain, (member(esk1_2(X1,difference(X2,X3)),X2)|~intersect(X1,difference(X2,X3)))).
#
#cnf(i_0_146, plain, (~intersect(X1,difference(X2,X2)))).
#
#cnf(i_0_141, plain, (intersect(X1,X2)|~intersect(X1,difference(X2,X3)))).
#
#cnf(i_0_144, plain, (intersect(X1,X2)|~intersect(X2,difference(X1,X3)))).
#
#cnf(i_0_153, negated_conjecture, (intersect(intersection(esk3_0,esk4_0),esk3_0))).
#
#cnf(i_0_147, plain, (~intersect(difference(X2,X2),X1))).
#
#cnf(i_0_168, negated_conjecture, (intersect(esk3_0,esk3_0))).
#
#cnf(i_0_154, plain, (intersect(X1,X2)|~intersect(difference(X2,X3),X1))).
#
#cnf(i_0_160, negated_conjecture, (intersect(esk3_0,intersection(esk3_0,esk4_0)))).
##
#cnf(i_0_161, plain, (intersect(X1,X2)|~intersect(difference(X1,X3),X2))).
#####
#cnf(i_0_145, plain, (intersect(X1,X1)|~intersect(X2,difference(X1,X3)))).
#####
#cnf(i_0_202, plain, (intersect(X1,X1)|~intersect(difference(X1,X3),X2))).
#########
#cnf(i_0_104, plain, (~intersect(X2,difference(X1,intersection(X2,X2))))).
###
#cnf(i_0_196, plain, (intersect(X1,X1)|~intersect(X3,difference(difference(X1,X2),X4)))).
#
#cnf(i_0_137, plain, (~intersect(difference(X1,X2),difference(X2,X3)))).
#
#cnf(i_0_208, plain, (intersect(X1,X1)|~intersect(difference(difference(X1,X2),X3),X4))).
###########
#cnf(i_0_191, plain, (~intersect(X2,difference(difference(X1,X1),X3)))).
########
#cnf(i_0_209, plain, (~intersect(difference(difference(X1,X1),X2),X3))).
#
#cnf(i_0_142, plain, (intersect(difference(X1,X2),X1)|~intersect(X3,difference(X1,X2)))).
#
#cnf(i_0_319, negated_conjecture, (intersect(difference(esk3_0,esk4_0),esk3_0))).
#
#cnf(i_0_330, negated_conjecture, (intersect(esk3_0,difference(esk3_0,esk4_0)))).
##########
#cnf(i_0_143, plain, (intersect(X1,difference(X1,X2))|~intersect(X3,difference(X1,X2)))).
##########
#cnf(i_0_157, plain, (intersect(difference(X1,X2),X1)|~intersect(difference(X1,X2),X3))).
###########
#cnf(i_0_164, plain, (intersect(X1,difference(X1,X2))|~intersect(difference(X1,X2),X3))).
##############
#cnf(i_0_249, plain, (~intersect(difference(X2,X3),difference(X1,X2)))).
###########################
#cnf(i_0_155, plain, (intersect(X1,X2)|~member(X4,difference(X2,X3))|~member(X4,X1))).
##
#cnf(i_0_162, plain, (intersect(X1,X2)|~member(X4,difference(X1,X3))|~member(X4,X2))).
####
#cnf(i_0_203, plain, (intersect(X1,X1)|~member(X4,difference(X1,X3))|~member(X4,X2))).
###############
#cnf(i_0_310, plain, (~intersect(X3,difference(difference(X1,X2),X1)))).
######
#cnf(i_0_28, plain, (member(esk1_2(X1,intersection(X2,X3)),X2)|~intersect(X1,intersection(X2,X3)))).
#
#cnf(i_0_578, plain, (intersect(X1,X2)|~intersect(X1,intersection(X2,X3)))).
#
#cnf(i_0_313, plain, (~intersect(X2,difference(X1,intersection(X1,X1))))).
#
#cnf(i_0_581, plain, (intersect(X1,X2)|~intersect(X2,intersection(X1,X3)))).
#
#cnf(i_0_585, plain, (intersect(X1,X2)|~intersect(X1,intersection(X3,X2)))).
#
#cnf(i_0_644, negated_conjecture, (intersect(esk3_0,esk4_0))).
###
#cnf(i_0_591, plain, (intersect(X1,X2)|~intersect(intersection(X2,X3),X1))).
#
#cnf(i_0_623, plain, (intersect(X1,X2)|~intersect(X2,intersection(X3,X1)))).
##
#cnf(i_0_684, negated_conjecture, (intersect(esk4_0,esk3_0))).
##
#cnf(i_0_629, plain, (intersect(X1,X2)|~intersect(intersection(X1,X3),X2))).
####
#cnf(i_0_647, negated_conjecture, (intersect(intersection(esk3_0,esk4_0),esk4_0))).
##
#cnf(i_0_724, negated_conjecture, (intersect(esk4_0,esk4_0))).
#
#cnf(i_0_648, plain, (intersect(X1,X2)|~intersect(intersection(X3,X2),X1))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(intersect_defn, axiom, ![X1, X2]:(intersect(X1,X2)<=>?[X3]:(member(X3,X1)&member(X3,X2))), file('./p/SET629+3.tptp', intersect_defn)).
fof(intersection_defn, axiom, ![X1, X2, X3]:(member(X3,intersection(X1,X2))<=>(member(X3,X1)&member(X3,X2))), file('./p/SET629+3.tptp', intersection_defn)).
fof(commutativity_of_intersection, axiom, ![X1, X2]:intersection(X1,X2)=intersection(X2,X1), file('./p/SET629+3.tptp', commutativity_of_intersection)).
fof(prove_intersection_and_difference_disjoint, conjecture, ![X1, X2]:disjoint(intersection(X1,X2),difference(X1,X2)), file('./p/SET629+3.tptp', prove_intersection_and_difference_disjoint)).
fof(difference_defn, axiom, ![X1, X2, X3]:(member(X3,difference(X1,X2))<=>(member(X3,X1)&~(member(X3,X2)))), file('./p/SET629+3.tptp', difference_defn)).
fof(symmetry_of_intersect, axiom, ![X1, X2]:(intersect(X1,X2)=>intersect(X2,X1)), file('./p/SET629+3.tptp', symmetry_of_intersect)).
fof(disjoint_defn, axiom, ![X1, X2]:(disjoint(X1,X2)<=>~(intersect(X1,X2))), file('./p/SET629+3.tptp', disjoint_defn)).
fof(c_0_7, plain, ![X10, X11, X13, X14, X15]:(((member(esk1_2(X10,X11),X10)|~intersect(X10,X11))&(member(esk1_2(X10,X11),X11)|~intersect(X10,X11)))&(~member(X15,X13)|~member(X15,X14)|intersect(X13,X14))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[intersect_defn])])])])])])).
fof(c_0_8, plain, ![X4, X5, X6]:(((member(X6,X4)|~member(X6,intersection(X4,X5)))&(member(X6,X5)|~member(X6,intersection(X4,X5))))&(~member(X6,X4)|~member(X6,X5)|member(X6,intersection(X4,X5)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[intersection_defn])])])).
cnf(c_0_9, plain, (intersect(X2,X3)|~member(X1,X2)|~member(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, plain, (member(esk1_2(X1,X2),X1)|~intersect(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (member(X1,X2)|~member(X1,intersection(X2,X3))), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (member(esk1_2(X1,X2),X2)|~intersect(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_13, plain, (intersect(X1,X2)|~intersect(X1,X3)|~member(esk1_2(X1,X3),X2)), inference(pm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_14, plain, (member(esk1_2(X1,intersection(X2,X3)),X2)|~intersect(X1,intersection(X2,X3))), inference(pm,[status(thm)],[c_0_11, c_0_12])).
fof(c_0_15, plain, ![X18, X19]:intersection(X18,X19)=intersection(X19,X18), inference(variable_rename,[status(thm)],[commutativity_of_intersection])).
fof(c_0_16, negated_conjecture, ~(![X1, X2]:disjoint(intersection(X1,X2),difference(X1,X2))), inference(assume_negation,[status(cth)],[prove_intersection_and_difference_disjoint])).
fof(c_0_17, plain, ![X7, X8, X9]:(((member(X9,X7)|~member(X9,difference(X7,X8)))&(~member(X9,X8)|~member(X9,difference(X7,X8))))&(~member(X9,X7)|member(X9,X8)|member(X9,difference(X7,X8)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[difference_defn])])])])).
cnf(c_0_18, plain, (intersect(X1,X2)|~intersect(X1,intersection(X2,X3))), inference(pm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_19, plain, (intersection(X1,X2)=intersection(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_20, plain, ![X20, X21]:(~intersect(X20,X21)|intersect(X21,X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[symmetry_of_intersect])])).
fof(c_0_21, negated_conjecture, ~disjoint(intersection(esk3_0,esk4_0),difference(esk3_0,esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_16])])])).
fof(c_0_22, plain, ![X16, X17]:((~disjoint(X16,X17)|~intersect(X16,X17))&(intersect(X16,X17)|disjoint(X16,X17))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[disjoint_defn])])])).
cnf(c_0_23, plain, (~member(X1,X2)|~member(X1,difference(X3,X2))), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_24, plain, (intersect(X1,X2)|~intersect(X1,intersection(X3,X2))), inference(pm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_25, plain, (intersect(X2,X1)|~intersect(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_26, negated_conjecture, (~disjoint(intersection(esk3_0,esk4_0),difference(esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_27, plain, (intersect(X1,X2)|disjoint(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_22])).
cnf(c_0_28, plain, (~intersect(difference(X1,X2),X3)|~member(esk1_2(difference(X1,X2),X3),X2)), inference(pm,[status(thm)],[c_0_23, c_0_10])).
cnf(c_0_29, plain, (intersect(X1,X2)|~intersect(intersection(X3,X2),X1)), inference(pm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_30, negated_conjecture, (intersect(intersection(esk3_0,esk4_0),difference(esk3_0,esk4_0))), inference(pm,[status(thm)],[c_0_26, c_0_27])).
cnf(c_0_31, plain, (~intersect(difference(X1,X2),X2)), inference(pm,[status(thm)],[c_0_28, c_0_12])).
cnf(c_0_32, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_29, c_0_30]), c_0_31]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 33
# Proof object clause steps            : 18
# Proof object formula steps           : 15
# Proof object conjectures             : 6
# Proof object clause conjectures      : 3
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 1
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 18
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 16
# Processed clauses                    : 246
# ...of these trivial                  : 1
# ...subsumed                          : 171
# ...remaining for further processing  : 74
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 732
# ...of the previous two non-trivial   : 682
# Contextual simplify-reflections      : 0
# Paramodulations                      : 732
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 74
#    Positive orientable unit clauses  : 11
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 13
#    Non-unit-clauses                  : 49
# Current number of unprocessed clauses: 435
# ...number of literals in the above   : 1012
# Current number of archived formulas  : 0
# Current number of archived clauses   : 0
# Clause-clause subsumption calls (NU) : 1076
# Rec. Clause-clause subsumption calls : 1072
# Non-unit clause-clause subsumptions  : 133
# Unit Clause-clause subsumption calls : 178
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5510
