# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_2, negated_conjecture, (element(X1,esk2_1(X1))|element(X1,esk1_0))).
#
#cnf(i_0_1, negated_conjecture, (element(esk2_1(X1),X1)|element(X1,esk1_0))).
#
#cnf(i_0_3, negated_conjecture, (~element(X1,esk1_0)|~element(X2,X1)|~element(X1,X2))).
#
#cnf(i_0_4, negated_conjecture, (element(X1,esk1_0)|~element(esk2_1(X1),esk1_0)|~element(X1,esk2_1(X1)))).
#
#cnf(i_0_7, negated_conjecture, (element(X1,esk1_0)|~element(esk2_1(X1),esk1_0))).
#
#cnf(i_0_9, negated_conjecture, (element(esk1_0,esk1_0))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(thm25, conjecture, ~(?[X1]:![X2]:(element(X2,X1)<=>~(?[X3]:(element(X2,X3)&element(X3,X2))))), file('./p/SET786+1.tptp', thm25)).
fof(c_0_1, negated_conjecture, ~(~(?[X1]:![X2]:(element(X2,X1)<=>~(?[X3]:(element(X2,X3)&element(X3,X2)))))), inference(assume_negation,[status(cth)],[thm25])).
fof(c_0_2, negated_conjecture, ![X5, X6]:((~element(X5,esk1_0)|(~element(X5,X6)|~element(X6,X5)))&((element(X5,esk2_1(X5))|element(X5,esk1_0))&(element(esk2_1(X5),X5)|element(X5,esk1_0)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
cnf(c_0_3, negated_conjecture, (~element(X1,esk1_0)|~element(X1,X2)|~element(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_4, negated_conjecture, (element(esk2_1(X1),X1)|element(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_5, negated_conjecture, (element(X1,esk1_0)|~element(esk2_1(X1),esk1_0)|~element(X1,esk2_1(X1))), inference(pm,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_6, negated_conjecture, (element(X1,esk2_1(X1))|element(X1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_7, negated_conjecture, (element(X1,esk1_0)|~element(esk2_1(X1),esk1_0)), inference(pm,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_8, negated_conjecture, (element(esk1_0,esk1_0)), inference(pm,[status(thm)],[c_0_7, c_0_4])).
cnf(c_0_9, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_3, c_0_8]), c_0_8])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 10
# Proof object clause steps            : 7
# Proof object formula steps           : 3
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 3
# Proof object initial formulas used   : 1
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 2
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 1
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 3
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 3
# Processed clauses                    : 6
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 6
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 0
# Generated clauses                    : 9
# ...of the previous two non-trivial   : 5
# Contextual simplify-reflections      : 0
# Paramodulations                      : 9
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 5
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 4
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 5
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 3
# Rec. Clause-clause subsumption calls : 3
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 260
