# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_9, plain, (c_in(X1,c_insert(X1,X2,X3),X3))).
#
#cnf(i_0_11, plain, (c_lessequals(X1,X2,tc_set(X3))|~c_in(c_Main_OsubsetI__1(X1,X2,X3),X2,X3))).
#
#cnf(i_0_10, plain, (c_in(c_Main_OsubsetI__1(X1,X2,X3),X1,X3)|c_lessequals(X1,X2,tc_set(X3)))).
#
#cnf(i_0_7, negated_conjecture, (~c_lessequals(v_S,c_insert(X1,c_emptyset,tc_set(t_a)),tc_set(tc_set(t_a))))).
#
#cnf(i_0_13, plain, (c_lessequals(X1,X1,tc_set(X2)))).
#
#cnf(i_0_12, plain, (X1=X2|~c_lessequals(X2,X1,tc_set(X3))|~c_lessequals(X1,X2,tc_set(X3)))).
#
#cnf(i_0_8, negated_conjecture, (c_lessequals(X1,X2,tc_set(t_a))|~c_in(X2,v_S,tc_set(t_a))|~c_in(X1,v_S,tc_set(t_a)))).
#
#cnf(i_0_14, negated_conjecture, (X1=X2|~c_in(X2,v_S,tc_set(t_a))|~c_in(X1,v_S,tc_set(t_a))|~c_lessequals(X2,X1,tc_set(t_a)))).
#
#cnf(i_0_16, negated_conjecture, (X1=X2|~c_in(X2,v_S,tc_set(t_a))|~c_in(X1,v_S,tc_set(t_a)))).
#
#cnf(i_0_23, negated_conjecture, (c_lessequals(X1,X1,X3)|~c_in(tc_set(X2),v_S,tc_set(t_a))|~c_in(X3,v_S,tc_set(t_a)))).
##
#cnf(i_0_19, negated_conjecture, (~c_in(c_insert(X1,c_emptyset,tc_set(t_a)),v_S,tc_set(t_a))|~c_in(X2,v_S,tc_set(t_a))|~c_lessequals(v_S,X2,tc_set(tc_set(t_a))))).
#
#cnf(i_0_20, negated_conjecture, (c_in(X1,X4,X3)|~c_in(c_insert(X1,X2,X3),v_S,tc_set(t_a))|~c_in(X4,v_S,tc_set(t_a)))).
####
#cnf(i_0_37, negated_conjecture, (c_lessequals(X1,X2,tc_set(X3))|~c_in(c_insert(c_Main_OsubsetI__1(X1,X2,X3),X4,X3),v_S,tc_set(t_a))|~c_in(X2,v_S,tc_set(t_a)))).
#
#cnf(i_0_31, negated_conjecture, (c_Main_OsubsetI__1(v_S,X1,tc_set(t_a))=X2|c_lessequals(v_S,X1,tc_set(tc_set(t_a)))|~c_in(X2,v_S,tc_set(t_a)))).
#
#cnf(i_0_59, negated_conjecture, (c_lessequals(v_S,X1,tc_set(tc_set(t_a)))|~c_in(X2,v_S,tc_set(t_a))|~c_in(X2,X1,tc_set(t_a)))).
#
#cnf(i_0_65, negated_conjecture, (~c_in(X2,c_insert(X1,c_emptyset,tc_set(t_a)),tc_set(t_a))|~c_in(X2,v_S,tc_set(t_a)))).
#
#cnf(i_0_68, negated_conjecture, (c_lessequals(v_S,c_insert(X1,X2,tc_set(t_a)),tc_set(tc_set(t_a)))|~c_in(X1,v_S,tc_set(t_a)))).
#
#cnf(i_0_63, negated_conjecture, (v_S=X1|~c_in(X2,v_S,tc_set(t_a))|~c_in(X2,X1,tc_set(t_a))|~c_lessequals(X1,v_S,tc_set(tc_set(t_a))))).
##
#cnf(i_0_76, negated_conjecture, (~c_in(X1,v_S,tc_set(t_a)))).
#
#cnf(i_0_86, negated_conjecture, (c_lessequals(v_S,X1,tc_set(tc_set(t_a))))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(cls_Set_Osubset__antisym_0, axiom, (X2=X1|~c_lessequals(X1,X2,tc_set(X3))|~c_lessequals(X2,X1,tc_set(X3))), file('./p/SET839-2.tptp', cls_Set_Osubset__antisym_0)).
cnf(cls_conjecture_1, negated_conjecture, (c_lessequals(X1,X2,tc_set(t_a))|~c_in(X2,v_S,tc_set(t_a))|~c_in(X1,v_S,tc_set(t_a))), file('./p/SET839-2.tptp', cls_conjecture_1)).
cnf(cls_Set_OsubsetI_0, axiom, (c_in(c_Main_OsubsetI__1(X1,X2,X3),X1,X3)|c_lessequals(X1,X2,tc_set(X3))), file('./p/SET839-2.tptp', cls_Set_OsubsetI_0)).
cnf(cls_Set_OsubsetI_1, axiom, (c_lessequals(X1,X2,tc_set(X3))|~c_in(c_Main_OsubsetI__1(X1,X2,X3),X2,X3)), file('./p/SET839-2.tptp', cls_Set_OsubsetI_1)).
cnf(cls_conjecture_0, negated_conjecture, (~c_lessequals(v_S,c_insert(X1,c_emptyset,tc_set(t_a)),tc_set(tc_set(t_a)))), file('./p/SET839-2.tptp', cls_conjecture_0)).
cnf(cls_Set_OinsertCI_1, axiom, (c_in(X1,c_insert(X1,X2,X3),X3)), file('./p/SET839-2.tptp', cls_Set_OinsertCI_1)).
cnf(c_0_6, axiom, (X2=X1|~c_lessequals(X1,X2,tc_set(X3))|~c_lessequals(X2,X1,tc_set(X3))), cls_Set_Osubset__antisym_0).
cnf(c_0_7, negated_conjecture, (c_lessequals(X1,X2,tc_set(t_a))|~c_in(X2,v_S,tc_set(t_a))|~c_in(X1,v_S,tc_set(t_a))), cls_conjecture_1).
cnf(c_0_8, negated_conjecture, (X1=X2|~c_in(X2,v_S,tc_set(t_a))|~c_in(X1,v_S,tc_set(t_a))|~c_lessequals(X2,X1,tc_set(t_a))), inference(pm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_9, negated_conjecture, (X1=X2|~c_in(X2,v_S,tc_set(t_a))|~c_in(X1,v_S,tc_set(t_a))), inference(pm,[status(thm)],[c_0_8, c_0_7])).
cnf(c_0_10, axiom, (c_in(c_Main_OsubsetI__1(X1,X2,X3),X1,X3)|c_lessequals(X1,X2,tc_set(X3))), cls_Set_OsubsetI_0).
cnf(c_0_11, axiom, (c_lessequals(X1,X2,tc_set(X3))|~c_in(c_Main_OsubsetI__1(X1,X2,X3),X2,X3)), cls_Set_OsubsetI_1).
cnf(c_0_12, negated_conjecture, (c_Main_OsubsetI__1(v_S,X1,tc_set(t_a))=X2|c_lessequals(v_S,X1,tc_set(tc_set(t_a)))|~c_in(X2,v_S,tc_set(t_a))), inference(pm,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_13, negated_conjecture, (~c_lessequals(v_S,c_insert(X1,c_emptyset,tc_set(t_a)),tc_set(tc_set(t_a)))), cls_conjecture_0).
cnf(c_0_14, negated_conjecture, (c_lessequals(v_S,X1,tc_set(tc_set(t_a)))|~c_in(X2,v_S,tc_set(t_a))|~c_in(X2,X1,tc_set(t_a))), inference(pm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_15, negated_conjecture, (~c_in(X1,c_insert(X2,c_emptyset,tc_set(t_a)),tc_set(t_a))|~c_in(X1,v_S,tc_set(t_a))), inference(pm,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_16, axiom, (c_in(X1,c_insert(X1,X2,X3),X3)), cls_Set_OinsertCI_1).
cnf(c_0_17, negated_conjecture, (~c_in(X1,v_S,tc_set(t_a))), inference(pm,[status(thm)],[c_0_15, c_0_16])).
cnf(c_0_18, negated_conjecture, (c_lessequals(v_S,X1,tc_set(tc_set(t_a)))), inference(pm,[status(thm)],[c_0_17, c_0_10])).
cnf(c_0_19, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_18])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 20
# Proof object clause steps            : 14
# Proof object formula steps           : 6
# Proof object conjectures             : 12
# Proof object clause conjectures      : 10
# Proof object formula conjectures     : 2
# Proof object initial clauses used    : 6
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 2
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 6
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 6
# Processed clauses                    : 25
# ...of these trivial                  : 0
# ...subsumed                          : 5
# ...remaining for further processing  : 20
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 4
# Generated clauses                    : 76
# ...of the previous two non-trivial   : 74
# Contextual simplify-reflections      : 0
# Paramodulations                      : 76
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 14
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 10
# Current number of unprocessed clauses: 54
# ...number of literals in the above   : 215
# Current number of archived formulas  : 0
# Current number of archived clauses   : 6
# Clause-clause subsumption calls (NU) : 16
# Rec. Clause-clause subsumption calls : 15
# Non-unit clause-clause subsumptions  : 6
# Unit Clause-clause subsumption calls : 14
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 6
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1031
