# Initializing proof state
# Scanning for AC axioms
# unordered_pair is commutative
# AC handling enabled
#
#cnf(i_0_16, plain, (empty(esk4_0))).
#
#cnf(i_0_18, plain, (subset(X1,X1))).
#
#cnf(i_0_2, plain, (unordered_pair(X1,X2)=unordered_pair(X2,X1))).
#
#cnf(i_0_17, plain, (~empty(esk5_0))).
#
#cnf(i_0_13, plain, (subset(X1,X2)|~in(esk3_2(X1,X2),X2))).
#
#cnf(i_0_14, plain, (subset(X1,X2)|in(esk3_2(X1,X2),X1))).
#
#cnf(i_0_15, plain, (in(X3,X2)|~subset(X1,X2)|~in(X3,X1))).
#
#cnf(i_0_19, negated_conjecture, (esk8_0!=esk6_0)).
#
#cnf(i_0_6, plain, (X1=X3|X2!=singleton(X3)|~in(X1,X2))).
#
#cnf(i_0_26, plain, (X1=X2|~in(X1,singleton(X2)))).
#
#cnf(i_0_5, plain, (in(X1,X3)|X3!=singleton(X2)|X1!=X2)).
#
#cnf(i_0_1, plain, (~in(X2,X1)|~in(X1,X2))).
#
#cnf(i_0_44, plain, (in(X1,X2)|X2!=singleton(X1))).
#
#cnf(i_0_43, plain, (in(X1,singleton(X2))|X1!=X2)).
#
#cnf(i_0_58, plain, (X1=X2|singleton(X2)!=singleton(X1))).
#
#cnf(i_0_59, plain, (X2!=singleton(X1)|~in(X2,X1))).
#
#cnf(i_0_51, plain, (subset(X1,X2)|~in(X1,esk3_2(X1,X2)))).
##
#cnf(i_0_76, plain, (subset(X1,X2)|esk3_2(X1,X2)!=singleton(X1))).
###
#cnf(i_0_57, plain, (subset(X1,X2)|singleton(esk3_2(X1,X2))!=X2)).
#
#cnf(i_0_20, negated_conjecture, (subset(unordered_pair(esk6_0,esk7_0),singleton(esk8_0)))).
#
#cnf(i_0_64, plain, (X1!=X2|~in(singleton(X2),X1))).
#
#cnf(i_0_42, plain, (X1=X2|~subset(X3,singleton(X2))|~in(X1,X3))).
##
#cnf(i_0_62, plain, (in(X1,X2)|X3!=singleton(X1)|~subset(X3,X2))).
###
#cnf(i_0_141, plain, (in(X1,X2)|~subset(singleton(X1),X2))).
#
#cnf(i_0_161, plain, (X1=X2|~subset(singleton(X1),singleton(X2)))).
#
#cnf(i_0_162, plain, (~subset(singleton(X1),X2)|~in(X2,X1))).
#
#cnf(i_0_158, plain, (subset(X1,X2)|~subset(singleton(esk3_2(X1,X2)),X2))).
#
#cnf(i_0_160, plain, (subset(X1,X2)|~subset(singleton(X1),esk3_2(X1,X2)))).
###
#cnf(i_0_67, plain, (in(X1,X2)|X1!=X3|~subset(singleton(X3),X2))).
###
#cnf(i_0_77, plain, (X1!=singleton(singleton(X2))|X1!=X2)).
#
#cnf(i_0_63, plain, (subset(X1,singleton(X2))|esk3_2(X1,singleton(X2))!=X2)).
#
#cnf(i_0_10, plain, (in(X1,X3)|X3!=unordered_pair(X4,X2)|X1!=X2)).
#
#cnf(i_0_233, plain, (in(X1,unordered_pair(X2,X3))|X1!=X3)).
#
#cnf(i_0_216, plain, (singleton(singleton(X1))!=X1)).
#
#cnf(i_0_265, plain, (in(X1,unordered_pair(X3,X2))|X1!=X3)).
###
#cnf(i_0_254, plain, (X1!=X3|~in(unordered_pair(X2,X3),X1))).
##
#cnf(i_0_234, plain, (in(X1,X2)|X2!=unordered_pair(X3,X1))).
#
#cnf(i_0_306, plain, (in(X1,unordered_pair(X2,X1)))).
#
#cnf(i_0_325, plain, (unordered_pair(X1,singleton(X2))!=X2)).
#
#cnf(i_0_336, plain, (in(X1,unordered_pair(X1,X2)))).
##
#cnf(i_0_323, plain, (in(X1,X2)|unordered_pair(X1,X3)!=X2)).
#
#cnf(i_0_327, plain, (singleton(unordered_pair(X2,X1))!=X1)).
####
#cnf(i_0_346, plain, (unordered_pair(singleton(X2),X1)!=X2)).
#
#cnf(i_0_330, plain, (in(X1,X2)|~subset(unordered_pair(X3,X1),X2))).
#
#cnf(i_0_407, negated_conjecture, (in(esk7_0,singleton(esk8_0)))).
#
#cnf(i_0_413, negated_conjecture, (esk8_0=esk7_0)).
#
#cnf(i_0_19, negated_conjecture, (esk7_0!=esk6_0)).
#
#cnf(i_0_407, negated_conjecture, (in(esk7_0,singleton(esk7_0)))).
#
#cnf(i_0_353, plain, (in(X1,X2)|~subset(unordered_pair(X1,X3),X2))).
#######
#cnf(i_0_311, plain, (X1=X2|singleton(X2)!=unordered_pair(X3,X1))).
##
#cnf(i_0_352, plain, (X1=X2|unordered_pair(X1,X3)!=singleton(X2))).
#
#cnf(i_0_350, plain, (singleton(unordered_pair(X1,X2))!=X1)).
#########
#cnf(i_0_331, plain, (X1=X2|~subset(unordered_pair(X3,X1),singleton(X2)))).
#
#cnf(i_0_354, plain, (X1=X2|~subset(unordered_pair(X1,X3),singleton(X2)))).
########
#cnf(i_0_307, plain, (subset(X1,X2)|unordered_pair(X3,esk3_2(X1,X2))!=X2)).
#
#cnf(i_0_362, plain, (subset(X1,X2)|unordered_pair(esk3_2(X1,X2),X3)!=X2)).
#######
#cnf(i_0_20, negated_conjecture, (subset(unordered_pair(esk6_0,esk7_0),singleton(esk7_0)))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=unordered_pair(X1,X2)<=>![X4]:(in(X4,X3)<=>(X4=X1|X4=X2))), file('./p/SET885+1.tptp', d2_tarski)).
fof(d3_tarski, axiom, ![X1, X2]:(subset(X1,X2)<=>![X3]:(in(X3,X1)=>in(X3,X2))), file('./p/SET885+1.tptp', d3_tarski)).
fof(t26_zfmisc_1, conjecture, ![X1, X2, X3]:(subset(unordered_pair(X1,X2),singleton(X3))=>X1=X3), file('./p/SET885+1.tptp', t26_zfmisc_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=singleton(X1)<=>![X3]:(in(X3,X2)<=>X3=X1)), file('./p/SET885+1.tptp', d1_tarski)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:unordered_pair(X1,X2)=unordered_pair(X2,X1), file('./p/SET885+1.tptp', commutativity_k2_tarski)).
fof(c_0_5, plain, ![X16, X17, X18, X19, X20, X21, X22, X23]:(((~in(X19,X18)|(X19=X16|X19=X17)|X18!=unordered_pair(X16,X17))&((X20!=X16|in(X20,X18)|X18!=unordered_pair(X16,X17))&(X20!=X17|in(X20,X18)|X18!=unordered_pair(X16,X17))))&(((esk2_3(X21,X22,X23)!=X21|~in(esk2_3(X21,X22,X23),X23)|X23=unordered_pair(X21,X22))&(esk2_3(X21,X22,X23)!=X22|~in(esk2_3(X21,X22,X23),X23)|X23=unordered_pair(X21,X22)))&(in(esk2_3(X21,X22,X23),X23)|(esk2_3(X21,X22,X23)=X21|esk2_3(X21,X22,X23)=X22)|X23=unordered_pair(X21,X22)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_6, plain, (in(X1,X3)|X1!=X2|X3!=unordered_pair(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_7, plain, ![X25, X26, X27, X28, X29]:((~subset(X25,X26)|(~in(X27,X25)|in(X27,X26)))&((in(esk3_2(X28,X29),X28)|subset(X28,X29))&(~in(esk3_2(X28,X29),X29)|subset(X28,X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_8, plain, (in(X1,X2)|X2!=unordered_pair(X3,X1)), inference(er,[status(thm)],[c_0_6])).
fof(c_0_9, negated_conjecture, ~(![X1, X2, X3]:(subset(unordered_pair(X1,X2),singleton(X3))=>X1=X3)), inference(assume_negation,[status(cth)],[t26_zfmisc_1])).
fof(c_0_10, plain, ![X9, X10, X11, X12, X13, X14]:(((~in(X11,X10)|X11=X9|X10!=singleton(X9))&(X12!=X9|in(X12,X10)|X10!=singleton(X9)))&((~in(esk1_2(X13,X14),X14)|esk1_2(X13,X14)!=X13|X14=singleton(X13))&(in(esk1_2(X13,X14),X14)|esk1_2(X13,X14)=X13|X14=singleton(X13)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_11, plain, (in(X3,X2)|~subset(X1,X2)|~in(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_12, plain, (in(X1,unordered_pair(X2,X1))), inference(er,[status(thm)],[c_0_8])).
fof(c_0_13, negated_conjecture, (subset(unordered_pair(esk6_0,esk7_0),singleton(esk8_0))&esk6_0!=esk8_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_9])])])).
cnf(c_0_14, plain, (X1=X3|~in(X1,X2)|X2!=singleton(X3)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_15, plain, ![X7, X8]:unordered_pair(X7,X8)=unordered_pair(X8,X7), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_16, plain, (in(X1,X2)|~subset(unordered_pair(X3,X1),X2)), inference(pm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_17, negated_conjecture, (subset(unordered_pair(esk6_0,esk7_0),singleton(esk8_0))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (X1=X2|~in(X1,singleton(X2))), inference(er,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (unordered_pair(X1,X2)=unordered_pair(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, negated_conjecture, (in(esk7_0,singleton(esk8_0))), inference(pm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_21, plain, (X1=X2|~subset(X3,singleton(X2))|~in(X1,X3)), inference(pm,[status(thm)],[c_0_18, c_0_11])).
cnf(c_0_22, plain, (in(X1,unordered_pair(X1,X2))), inference(pm,[status(thm)],[c_0_12, c_0_19])).
cnf(c_0_23, negated_conjecture, (esk8_0=esk7_0), inference(pm,[status(thm)],[c_0_18, c_0_20])).
cnf(c_0_24, negated_conjecture, (esk6_0!=esk8_0), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_25, plain, (X1=X2|~subset(unordered_pair(X1,X3),singleton(X2))), inference(pm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_26, negated_conjecture, (subset(unordered_pair(esk6_0,esk7_0),singleton(esk7_0))), inference(rw,[status(thm)],[c_0_17, c_0_23])).
cnf(c_0_27, negated_conjecture, (esk7_0!=esk6_0), inference(rw,[status(thm)],[c_0_24, c_0_23])).
cnf(c_0_28, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_25, c_0_26]), c_0_27]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 29
# Proof object clause steps            : 18
# Proof object formula steps           : 11
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 6
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 3
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 20
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 20
# Processed clauses                    : 102
# ...of these trivial                  : 1
# ...subsumed                          : 44
# ...remaining for further processing  : 57
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 484
# ...of the previous two non-trivial   : 473
# Contextual simplify-reflections      : 0
# Paramodulations                      : 473
# Factorizations                       : 0
# Equation resolutions                 : 11
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 54
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 387
# ...number of literals in the above   : 1326
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3
# Clause-clause subsumption calls (NU) : 176
# Rec. Clause-clause subsumption calls : 176
# Non-unit clause-clause subsumptions  : 36
# Unit Clause-clause subsumption calls : 132
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3305
