# Initializing proof state
# Scanning for AC axioms
# unordered_pair is commutative
# AC handling enabled
#
#cnf(i_0_5, plain, (empty(esk1_0))).
#
#cnf(i_0_2, plain, (unordered_pair(X1,X2)=unordered_pair(X2,X1))).
#
#cnf(i_0_13, plain, (X1=X2|~in(esk9_2(X1,X2),X2)|~in(esk9_2(X1,X2),X1))).
#
#cnf(i_0_6, plain, (~empty(esk2_0))).
#
#cnf(i_0_12, plain, (X1=X2|in(esk9_2(X1,X2),X2)|in(esk9_2(X1,X2),X1))).
#
#cnf(i_0_11, negated_conjecture, (unordered_pair(singleton(esk5_1(X1)),unordered_pair(esk5_1(X1),esk6_1(X1)))=X1|~in(X1,esk3_0))).
#
#cnf(i_0_10, negated_conjecture, (unordered_pair(singleton(esk7_1(X1)),unordered_pair(esk8_1(X1),esk7_1(X1)))=X1|~in(X1,esk4_0))).
#
#cnf(i_0_7, negated_conjecture, (esk4_0!=esk3_0)).
#
#cnf(i_0_8, negated_conjecture, (in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk3_0)|~in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk4_0))).
#
#cnf(i_0_9, negated_conjecture, (in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk4_0)|~in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk3_0))).
#
#cnf(i_0_20, negated_conjecture, (in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk3_0)|~in(unordered_pair(singleton(X1),unordered_pair(X2,X1)),esk4_0))).
#
#cnf(i_0_1, plain, (~in(X2,X1)|~in(X1,X2))).
#
#cnf(i_0_34, plain, (X1=X2|in(esk9_2(X1,X2),X2)|~in(X1,esk9_2(X1,X2)))).
##
#cnf(i_0_42, plain, (X1=X2|~in(X2,esk9_2(X1,X2))|~in(X1,esk9_2(X1,X2)))).
#
#cnf(i_0_4, plain, (~empty(unordered_pair(singleton(X1),unordered_pair(X1,X2))))).
##
#cnf(i_0_35, plain, (X2=X1|in(esk9_2(X2,X1),X2)|~in(X1,esk9_2(X2,X1)))).
##
#cnf(i_0_49, negated_conjecture, (~empty(X1)|~in(X1,esk3_0))).
####
#cnf(i_0_47, plain, (~empty(unordered_pair(singleton(X1),unordered_pair(X2,X1))))).
##
#cnf(i_0_54, negated_conjecture, (esk3_0=X1|~empty(esk9_2(esk3_0,X1))|~in(X1,esk9_2(esk3_0,X1)))).
#
#cnf(i_0_56, negated_conjecture, (X1=esk3_0|~empty(esk9_2(X1,esk3_0))|~in(X1,esk9_2(X1,esk3_0)))).
#
#cnf(i_0_62, negated_conjecture, (~empty(X1)|~in(X1,esk4_0))).
#
#cnf(i_0_64, negated_conjecture, (esk4_0=X1|~empty(esk9_2(esk4_0,X1))|~in(X1,esk9_2(esk4_0,X1)))).
#
#cnf(i_0_66, negated_conjecture, (X1=esk4_0|~empty(esk9_2(X1,esk4_0))|~in(X1,esk9_2(X1,esk4_0)))).
#
#cnf(i_0_53, negated_conjecture, (esk3_0=X1|in(esk9_2(esk3_0,X1),X1)|~empty(esk9_2(esk3_0,X1)))).
#
#cnf(i_0_68, negated_conjecture, (~empty(esk9_2(esk3_0,esk4_0)))).
###
#cnf(i_0_55, negated_conjecture, (X1=esk3_0|in(esk9_2(X1,esk3_0),X1)|~empty(esk9_2(X1,esk3_0)))).
#
#cnf(i_0_72, negated_conjecture, (~empty(esk9_2(esk4_0,esk3_0)))).
###
#cnf(i_0_63, negated_conjecture, (esk4_0=X1|in(esk9_2(esk4_0,X1),X1)|~empty(esk9_2(esk4_0,X1)))).
####
#cnf(i_0_65, negated_conjecture, (X1=esk4_0|in(esk9_2(X1,esk4_0),X1)|~empty(esk9_2(X1,esk4_0)))).
############
#cnf(i_0_32, negated_conjecture, (in(unordered_pair(singleton(X1),unordered_pair(X2,X1)),esk3_0)|~in(unordered_pair(singleton(X1),unordered_pair(X2,X1)),esk4_0))).
##
#cnf(i_0_88, negated_conjecture, (in(unordered_pair(singleton(esk7_1(X1)),unordered_pair(esk8_1(X1),esk7_1(X1))),esk3_0)|~in(X1,esk4_0))).
#
#cnf(i_0_93, negated_conjecture, (in(X1,esk3_0)|~in(X1,esk4_0))).
##
#cnf(i_0_45, negated_conjecture, (~empty(unordered_pair(X1,singleton(singleton(esk7_1(X1)))))|~in(X1,esk4_0))).
##
#cnf(i_0_97, negated_conjecture, (esk4_0=X1|in(esk9_2(esk4_0,X1),esk3_0)|~in(X1,esk9_2(esk4_0,X1)))).
##
#cnf(i_0_46, negated_conjecture, (~empty(unordered_pair(X1,singleton(singleton(esk5_1(X1)))))|~in(X1,esk3_0))).
#
#cnf(i_0_101, negated_conjecture, (esk4_0=X1|~in(esk3_0,esk9_2(esk4_0,X1))|~in(X1,esk9_2(esk4_0,X1)))).
##
#cnf(i_0_99, negated_conjecture, (X1=esk4_0|in(esk9_2(X1,esk4_0),esk3_0)|~in(X1,esk9_2(X1,esk4_0)))).
#
#cnf(i_0_58, negated_conjecture, (~empty(unordered_pair(X1,singleton(unordered_pair(esk8_1(X1),esk7_1(X1)))))|~in(X1,esk4_0))).
##
#cnf(i_0_104, negated_conjecture, (X1=esk4_0|~in(esk3_0,esk9_2(X1,esk4_0))|~in(X1,esk9_2(X1,esk4_0)))).
##
#cnf(i_0_59, negated_conjecture, (~empty(unordered_pair(X1,singleton(unordered_pair(esk5_1(X1),esk6_1(X1)))))|~in(X1,esk3_0))).
#
#cnf(i_0_96, negated_conjecture, (esk4_0=X1|in(esk9_2(esk4_0,X1),esk3_0)|in(esk9_2(esk4_0,X1),X1))).
#
#cnf(i_0_107, negated_conjecture, (in(esk9_2(esk4_0,esk3_0),esk3_0))).
####
#cnf(i_0_112, negated_conjecture, (esk4_0=X1|in(esk9_2(esk4_0,X1),X1)|~in(esk3_0,esk9_2(esk4_0,X1)))).
##
#cnf(i_0_115, negated_conjecture, (~in(esk3_0,esk9_2(esk4_0,esk3_0)))).
###
#cnf(i_0_121, negated_conjecture, (esk4_0=X1|~in(esk9_2(esk4_0,X1),esk4_0)|~in(esk3_0,esk9_2(esk4_0,X1)))).
########
#cnf(i_0_98, negated_conjecture, (X1=esk4_0|in(esk9_2(X1,esk4_0),esk3_0)|in(esk9_2(X1,esk4_0),X1))).
##
#cnf(i_0_125, negated_conjecture, (in(esk9_2(esk3_0,esk4_0),esk3_0))).
####
#cnf(i_0_130, negated_conjecture, (X1=esk4_0|in(esk9_2(X1,esk4_0),X1)|~in(esk3_0,esk9_2(X1,esk4_0)))).
###
#cnf(i_0_133, negated_conjecture, (~in(esk3_0,esk9_2(esk3_0,esk4_0)))).
##
#cnf(i_0_139, negated_conjecture, (X1=esk4_0|~in(esk9_2(X1,esk4_0),esk4_0)|~in(esk3_0,esk9_2(X1,esk4_0)))).
#########
#cnf(i_0_26, negated_conjecture, (in(unordered_pair(singleton(X1),unordered_pair(X2,X1)),esk4_0)|~in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk3_0))).
##
#cnf(i_0_148, negated_conjecture, (in(unordered_pair(singleton(X1),unordered_pair(X2,X1)),esk4_0)|~in(unordered_pair(singleton(X1),unordered_pair(X2,X1)),esk3_0))).
######
#cnf(i_0_146, negated_conjecture, (in(unordered_pair(singleton(X1),unordered_pair(X2,X1)),esk3_0)|~in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk3_0))).
###
#cnf(i_0_157, negated_conjecture, (in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk3_0)|~in(unordered_pair(singleton(X1),unordered_pair(X2,X1)),esk3_0))).
###
#cnf(i_0_160, negated_conjecture, (in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk4_0)|~in(unordered_pair(singleton(X1),unordered_pair(X2,X1)),esk3_0))).
##########
#cnf(i_0_28, negated_conjecture, (in(X1,esk4_0)|~in(unordered_pair(singleton(esk5_1(X1)),unordered_pair(esk5_1(X1),esk6_1(X1))),esk3_0)|~in(X1,esk3_0))).
##
#cnf(i_0_191, negated_conjecture, (in(X1,esk4_0)|~in(X1,esk3_0))).
#
#cnf(i_0_192, negated_conjecture, (esk4_0=X1|~in(esk9_2(esk4_0,X1),esk3_0)|~in(esk9_2(esk4_0,X1),X1))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t112_zfmisc_1, conjecture, ![X1, X2]:(((![X3]:~((in(X3,X1)&![X4, X5]:X3!=ordered_pair(X4,X5)))&![X3]:~((in(X3,X2)&![X4, X5]:X3!=ordered_pair(X4,X5))))&![X3, X4]:(in(ordered_pair(X3,X4),X1)<=>in(ordered_pair(X3,X4),X2)))=>X1=X2), file('./p/SET959+1.tptp', t112_zfmisc_1)).
fof(d5_tarski, axiom, ![X1, X2]:ordered_pair(X1,X2)=unordered_pair(unordered_pair(X1,X2),singleton(X1)), file('./p/SET959+1.tptp', d5_tarski)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:unordered_pair(X1,X2)=unordered_pair(X2,X1), file('./p/SET959+1.tptp', commutativity_k2_tarski)).
fof(t2_tarski, axiom, ![X1, X2]:(![X3]:(in(X3,X1)<=>in(X3,X2))=>X1=X2), file('./p/SET959+1.tptp', t2_tarski)).
fof(c_0_4, negated_conjecture, ~(![X1, X2]:(((![X3]:~((in(X3,X1)&![X4, X5]:X3!=ordered_pair(X4,X5)))&![X3]:~((in(X3,X2)&![X4, X5]:X3!=ordered_pair(X4,X5))))&![X3, X4]:(in(ordered_pair(X3,X4),X1)<=>in(ordered_pair(X3,X4),X2)))=>X1=X2)), inference(assume_negation,[status(cth)],[t112_zfmisc_1])).
fof(c_0_5, negated_conjecture, ![X18, X21, X24, X25]:((((~in(X18,esk3_0)|X18=ordered_pair(esk5_1(X18),esk6_1(X18)))&(~in(X21,esk4_0)|X21=ordered_pair(esk7_1(X21),esk8_1(X21))))&((~in(ordered_pair(X24,X25),esk3_0)|in(ordered_pair(X24,X25),esk4_0))&(~in(ordered_pair(X24,X25),esk4_0)|in(ordered_pair(X24,X25),esk3_0))))&esk3_0!=esk4_0), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])])).
fof(c_0_6, plain, ![X10, X11]:ordered_pair(X10,X11)=unordered_pair(unordered_pair(X10,X11),singleton(X10)), inference(variable_rename,[status(thm)],[d5_tarski])).
cnf(c_0_7, negated_conjecture, (in(ordered_pair(X1,X2),esk3_0)|~in(ordered_pair(X1,X2),esk4_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_8, plain, (ordered_pair(X1,X2)=unordered_pair(unordered_pair(X1,X2),singleton(X1))), inference(split_conjunct,[status(thm)],[c_0_6])).
fof(c_0_9, plain, ![X8, X9]:unordered_pair(X8,X9)=unordered_pair(X9,X8), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_10, negated_conjecture, (in(unordered_pair(unordered_pair(X1,X2),singleton(X1)),esk3_0)|~in(unordered_pair(unordered_pair(X1,X2),singleton(X1)),esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_8]), c_0_8])).
cnf(c_0_11, plain, (unordered_pair(X1,X2)=unordered_pair(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_12, negated_conjecture, (in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk3_0)|~in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_11]), c_0_11])).
cnf(c_0_13, negated_conjecture, (X1=ordered_pair(esk7_1(X1),esk8_1(X1))|~in(X1,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_14, negated_conjecture, (in(ordered_pair(X1,X2),esk4_0)|~in(ordered_pair(X1,X2),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15, negated_conjecture, (X1=ordered_pair(esk5_1(X1),esk6_1(X1))|~in(X1,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_16, negated_conjecture, (in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk3_0)|~in(unordered_pair(singleton(X1),unordered_pair(X2,X1)),esk4_0)), inference(pm,[status(thm)],[c_0_12, c_0_11])).
cnf(c_0_17, negated_conjecture, (X1=unordered_pair(unordered_pair(esk7_1(X1),esk8_1(X1)),singleton(esk7_1(X1)))|~in(X1,esk4_0)), inference(rw,[status(thm)],[c_0_13, c_0_8])).
cnf(c_0_18, negated_conjecture, (in(unordered_pair(unordered_pair(X1,X2),singleton(X1)),esk4_0)|~in(unordered_pair(unordered_pair(X1,X2),singleton(X1)),esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_8]), c_0_8])).
cnf(c_0_19, negated_conjecture, (X1=unordered_pair(unordered_pair(esk5_1(X1),esk6_1(X1)),singleton(esk5_1(X1)))|~in(X1,esk3_0)), inference(rw,[status(thm)],[c_0_15, c_0_8])).
cnf(c_0_20, negated_conjecture, (in(unordered_pair(singleton(X1),unordered_pair(X2,X1)),esk3_0)|~in(unordered_pair(singleton(X1),unordered_pair(X2,X1)),esk4_0)), inference(pm,[status(thm)],[c_0_16, c_0_11])).
cnf(c_0_21, negated_conjecture, (unordered_pair(singleton(esk7_1(X1)),unordered_pair(esk8_1(X1),esk7_1(X1)))=X1|~in(X1,esk4_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_11]), c_0_11])).
cnf(c_0_22, negated_conjecture, (in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk4_0)|~in(unordered_pair(singleton(X1),unordered_pair(X1,X2)),esk3_0)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_18, c_0_11]), c_0_11])).
cnf(c_0_23, negated_conjecture, (unordered_pair(singleton(esk5_1(X1)),unordered_pair(esk5_1(X1),esk6_1(X1)))=X1|~in(X1,esk3_0)), inference(rw,[status(thm)],[c_0_19, c_0_11])).
cnf(c_0_24, negated_conjecture, (in(unordered_pair(singleton(esk7_1(X1)),unordered_pair(esk8_1(X1),esk7_1(X1))),esk3_0)|~in(X1,esk4_0)), inference(pm,[status(thm)],[c_0_20, c_0_21])).
fof(c_0_25, plain, ![X26, X27]:((~in(esk9_2(X26,X27),X26)|~in(esk9_2(X26,X27),X27)|X26=X27)&(in(esk9_2(X26,X27),X26)|in(esk9_2(X26,X27),X27)|X26=X27)), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_tarski])])])])).
cnf(c_0_26, negated_conjecture, (in(X1,esk4_0)|~in(unordered_pair(singleton(esk5_1(X1)),unordered_pair(esk5_1(X1),esk6_1(X1))),esk3_0)|~in(X1,esk3_0)), inference(pm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_27, negated_conjecture, (in(X1,esk3_0)|~in(X1,esk4_0)), inference(pm,[status(thm)],[c_0_24, c_0_21])).
cnf(c_0_28, plain, (in(esk9_2(X1,X2),X1)|in(esk9_2(X1,X2),X2)|X1=X2), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_29, plain, (X1=X2|~in(esk9_2(X1,X2),X1)|~in(esk9_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_25])).
cnf(c_0_30, negated_conjecture, (in(X1,esk4_0)|~in(X1,esk3_0)), inference(pm,[status(thm)],[c_0_26, c_0_23])).
cnf(c_0_31, negated_conjecture, (esk4_0=X1|in(esk9_2(esk4_0,X1),esk3_0)|in(esk9_2(esk4_0,X1),X1)), inference(pm,[status(thm)],[c_0_27, c_0_28])).
cnf(c_0_32, negated_conjecture, (esk3_0!=esk4_0), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_33, negated_conjecture, (esk4_0=X1|~in(esk9_2(esk4_0,X1),esk3_0)|~in(esk9_2(esk4_0,X1),X1)), inference(pm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_34, negated_conjecture, (in(esk9_2(esk4_0,esk3_0),esk3_0)), inference(sr,[status(thm)],[inference(ef,[status(thm)],[c_0_31]), c_0_32])).
cnf(c_0_35, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_33, c_0_34]), c_0_34])]), c_0_32]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 36
# Proof object clause steps            : 27
# Proof object formula steps           : 9
# Proof object conjectures             : 26
# Proof object clause conjectures      : 23
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 10
# Proof object simplifying inferences  : 17
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 13
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 12
# Processed clauses                    : 138
# ...of these trivial                  : 2
# ...subsumed                          : 78
# ...remaining for further processing  : 58
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 0
# Generated clauses                    : 200
# ...of the previous two non-trivial   : 163
# Contextual simplify-reflections      : 0
# Paramodulations                      : 194
# Factorizations                       : 6
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 53
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 8
#    Non-unit-clauses                  : 41
# Current number of unprocessed clauses: 34
# ...number of literals in the above   : 80
# Current number of archived formulas  : 0
# Current number of archived clauses   : 6
# Clause-clause subsumption calls (NU) : 298
# Rec. Clause-clause subsumption calls : 292
# Non-unit clause-clause subsumptions  : 62
# Unit Clause-clause subsumption calls : 25
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 3023
