# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_3, plain, (empty(esk1_0))).
#
#cnf(i_0_1, plain, (subset(X1,X1))).
#
#cnf(i_0_8, plain, (in(X1,X2)|~subset(singleton(X1),X2))).
#
#cnf(i_0_4, plain, (~empty(esk2_0))).
#
#cnf(i_0_7, plain, (subset(singleton(X1),X2)|~in(X1,X2))).
#
#cnf(i_0_5, negated_conjecture, (in(esk3_0,esk4_0)|subset(singleton(esk3_0),esk4_0))).
#
#cnf(i_0_10, negated_conjecture, (subset(singleton(esk3_0),esk4_0))).
#
#cnf(i_0_2, plain, (~in(X2,X1)|~in(X1,X2))).
#
#cnf(i_0_11, plain, (~in(X1,X2)|~subset(singleton(X2),X1))).
##
#cnf(i_0_13, plain, (~subset(singleton(X2),X1)|~subset(singleton(X1),X2))).
#
#cnf(i_0_6, negated_conjecture, (~in(esk3_0,esk4_0))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t37_zfmisc_1, conjecture, ![X1, X2]:(subset(singleton(X1),X2)<=>in(X1,X2)), file('./p/SEU158+3.tptp', t37_zfmisc_1)).
fof(l2_zfmisc_1, axiom, ![X1, X2]:(subset(singleton(X1),X2)<=>in(X1,X2)), file('./p/SEU158+3.tptp', l2_zfmisc_1)).
fof(c_0_2, negated_conjecture, ~(![X1, X2]:(subset(singleton(X1),X2)<=>in(X1,X2))), inference(assume_negation,[status(cth)],[t37_zfmisc_1])).
fof(c_0_3, plain, ![X10, X11]:((~subset(singleton(X10),X11)|in(X10,X11))&(~in(X10,X11)|subset(singleton(X10),X11))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l2_zfmisc_1])])).
fof(c_0_4, negated_conjecture, ((~subset(singleton(esk3_0),esk4_0)|~in(esk3_0,esk4_0))&(subset(singleton(esk3_0),esk4_0)|in(esk3_0,esk4_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_2])])])).
cnf(c_0_5, plain, (subset(singleton(X1),X2)|~in(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_6, negated_conjecture, (subset(singleton(esk3_0),esk4_0)|in(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, negated_conjecture, (~subset(singleton(esk3_0),esk4_0)|~in(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_8, negated_conjecture, (subset(singleton(esk3_0),esk4_0)), inference(pm,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_9, negated_conjecture, (~in(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_7, c_0_8])])).
cnf(c_0_10, plain, (in(X1,X2)|~subset(singleton(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_3])).
cnf(c_0_11, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_9, c_0_10]), c_0_8])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 12
# Proof object clause steps            : 7
# Proof object formula steps           : 5
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 4
# Proof object initial formulas used   : 2
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 8
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 8
# Processed clauses                    : 12
# ...of these trivial                  : 0
# ...subsumed                          : 1
# ...remaining for further processing  : 11
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 6
# ...of the previous two non-trivial   : 4
# Contextual simplify-reflections      : 0
# Paramodulations                      : 6
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 10
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 5
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 1
# Rec. Clause-clause subsumption calls : 1
# Non-unit clause-clause subsumptions  : 1
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 434
