# Initializing proof state
# Scanning for AC axioms
# unordered_pair is commutative
# AC handling enabled
#
#cnf(i_0_1, plain, (unordered_pair(X1,X2)=unordered_pair(X2,X1))).
#
#cnf(i_0_6, negated_conjecture, (unordered_pair(esk2_0,esk3_0)=singleton(esk1_0))).
#
#cnf(i_0_4, plain, (X1=X2|singleton(X1)!=unordered_pair(X2,X3))).
#
#cnf(i_0_5, negated_conjecture, (esk3_0!=esk2_0)).
#
#cnf(i_0_8, plain, (X1=X2|unordered_pair(X3,X2)!=singleton(X1))).
####
#cnf(i_0_7, negated_conjecture, (X1=esk2_0|singleton(esk1_0)!=singleton(X1))).
#
#cnf(i_0_13, negated_conjecture, (esk2_0=esk1_0)).
#
#cnf(i_0_6, negated_conjecture, (unordered_pair(esk1_0,esk3_0)=singleton(esk1_0))).
#
#cnf(i_0_5, negated_conjecture, (esk3_0!=esk1_0)).
#
#cnf(i_0_7, negated_conjecture, (X1=esk1_0|singleton(esk1_0)!=singleton(X1))).
##
#cnf(i_0_15, negated_conjecture, (X1=esk3_0|singleton(esk1_0)!=singleton(X1))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t9_zfmisc_1, conjecture, ![X1, X2, X3]:(singleton(X1)=unordered_pair(X2,X3)=>X2=X3), file('./p/SEU150+1.tptp', t9_zfmisc_1)).
fof(t8_zfmisc_1, axiom, ![X1, X2, X3]:(singleton(X1)=unordered_pair(X2,X3)=>X1=X2), file('./p/SEU150+1.tptp', t8_zfmisc_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:unordered_pair(X1,X2)=unordered_pair(X2,X1), file('./p/SEU150+1.tptp', commutativity_k2_tarski)).
fof(c_0_3, negated_conjecture, ~(![X1, X2, X3]:(singleton(X1)=unordered_pair(X2,X3)=>X2=X3)), inference(assume_negation,[status(cth)],[t9_zfmisc_1])).
fof(c_0_4, plain, ![X6, X7, X8]:(singleton(X6)!=unordered_pair(X7,X8)|X6=X7), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_zfmisc_1])])).
fof(c_0_5, negated_conjecture, (singleton(esk1_0)=unordered_pair(esk2_0,esk3_0)&esk2_0!=esk3_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_3])])])).
cnf(c_0_6, plain, (X1=X2|singleton(X1)!=unordered_pair(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_7, negated_conjecture, (singleton(esk1_0)=unordered_pair(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_5])).
fof(c_0_8, plain, ![X4, X5]:unordered_pair(X4,X5)=unordered_pair(X5,X4), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_9, negated_conjecture, (X1=esk2_0|singleton(esk1_0)!=singleton(X1)), inference(pm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_10, plain, (unordered_pair(X1,X2)=unordered_pair(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, negated_conjecture, (esk2_0=esk1_0), inference(er,[status(thm)],[c_0_9])).
cnf(c_0_12, plain, (X1=X2|unordered_pair(X3,X2)!=singleton(X1)), inference(pm,[status(thm)],[c_0_6, c_0_10])).
cnf(c_0_13, negated_conjecture, (unordered_pair(esk1_0,esk3_0)=singleton(esk1_0)), inference(rw,[status(thm)],[c_0_7, c_0_11])).
cnf(c_0_14, negated_conjecture, (esk2_0!=esk3_0), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_15, negated_conjecture, (X1=esk3_0|singleton(esk1_0)!=singleton(X1)), inference(pm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16, negated_conjecture, (esk3_0!=esk1_0), inference(rw,[status(thm)],[c_0_14, c_0_11])).
cnf(c_0_17, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(er,[status(thm)],[c_0_15]), c_0_16]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 18
# Proof object clause steps            : 11
# Proof object formula steps           : 7
# Proof object conjectures             : 11
# Proof object clause conjectures      : 8
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 4
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 3
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 6
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 4
# Processed clauses                    : 15
# ...of these trivial                  : 0
# ...subsumed                          : 4
# ...remaining for further processing  : 11
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 11
# ...of the previous two non-trivial   : 12
# Contextual simplify-reflections      : 0
# Paramodulations                      : 8
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 8
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 4
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3
# Clause-clause subsumption calls (NU) : 8
# Rec. Clause-clause subsumption calls : 8
# Non-unit clause-clause subsumptions  : 4
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 238
