# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_2, plain, (icext(uri_owl_Thing,X1))).
#
#cnf(i_0_17, plain, (iext(uri_rdf_first,esk4_0,uri_ex_w))).
#
#cnf(i_0_16, plain, (iext(uri_rdf_rest,esk4_0,uri_rdf_nil))).
#
#cnf(i_0_4, plain, (~icext(uri_owl_Nothing,X1))).
#
#cnf(i_0_18, plain, (iext(uri_owl_oneOf,esk3_0,esk4_0))).
#
#cnf(i_0_19, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,esk3_0))).
#
#cnf(i_0_14, plain, (ic(X1)|~iext(uri_owl_equivalentClass,X2,X1))).
#
#cnf(i_0_20, plain, (ic(esk3_0))).
#
#cnf(i_0_15, plain, (ic(X1)|~iext(uri_owl_equivalentClass,X1,X2))).
#
#cnf(i_0_21, plain, (ic(uri_owl_Thing))).
#
#cnf(i_0_12, plain, (icext(X3,X2)|~iext(uri_owl_equivalentClass,X3,X1)|~icext(X1,X2))).
#
#cnf(i_0_22, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X2)|~icext(X2,X1))).
#
#cnf(i_0_24, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X1,uri_owl_Thing))).
#
#cnf(i_0_13, plain, (icext(X3,X2)|~iext(uri_owl_equivalentClass,X1,X3)|~icext(X1,X2))).
#
#cnf(i_0_38, plain, (icext(esk3_0,X1))).
#
#cnf(i_0_26, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,uri_owl_Thing))).
#
#cnf(i_0_35, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X1))).
#
#cnf(i_0_40, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X1,esk3_0))).
#
#cnf(i_0_41, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,esk3_0,X1))).
##
#cnf(i_0_30, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X3,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X1,X3))).
#
#cnf(i_0_37, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X3,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X3,X1))).
#
#cnf(i_0_44, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X3)|~iext(uri_owl_equivalentClass,X1,X3))).
#
#cnf(i_0_39, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,esk3_0))).
##
#cnf(i_0_45, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X3)|~iext(uri_owl_equivalentClass,X3,X1))).
##
#cnf(i_0_42, plain, (~iext(uri_owl_equivalentClass,uri_owl_Thing,uri_owl_Nothing))).
#
#cnf(i_0_49, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X3,esk3_0)|~iext(uri_owl_equivalentClass,X1,X3))).
##
#cnf(i_0_50, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X3,esk3_0)|~iext(uri_owl_equivalentClass,X3,X1))).
###
#cnf(i_0_54, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,esk3_0,X3)|~iext(uri_owl_equivalentClass,X1,X3))).
#
#cnf(i_0_55, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,esk3_0,X3)|~iext(uri_owl_equivalentClass,X3,X1))).
#
#cnf(i_0_52, plain, (~iext(uri_owl_equivalentClass,esk3_0,uri_owl_Nothing))).
#
#cnf(i_0_23, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X1,X3)|~iext(uri_owl_equivalentClass,X3,X4)|~icext(X4,X2))).
###
#cnf(i_0_31, plain, (~iext(uri_owl_equivalentClass,X2,uri_owl_Nothing)|~icext(X2,X1))).
##
#cnf(i_0_33, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X1,X3)|~iext(uri_owl_equivalentClass,X4,X3)|~icext(X4,X2))).
#####
#cnf(i_0_34, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X3,X1)|~iext(uri_owl_equivalentClass,X4,X3)|~icext(X4,X2))).
#####
#cnf(i_0_43, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X1))).
#
#cnf(i_0_36, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X3,X1)|~iext(uri_owl_equivalentClass,X3,X4)|~icext(X4,X2))).
###
#cnf(i_0_53, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,esk3_0,X1))).
##
#cnf(i_0_9, plain, (ic(X1)|~iext(uri_rdf_rest,X2,uri_rdf_nil)|~iext(uri_rdf_first,X2,X3)|~iext(uri_owl_oneOf,X1,X2))).
#
#cnf(i_0_58, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X4,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X1,X3)|~iext(uri_owl_equivalentClass,X3,X4))).
#
#cnf(i_0_61, plain, (~iext(uri_owl_equivalentClass,X2,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X2,uri_owl_Nothing))).
##
#cnf(i_0_59, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X4,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X3,X1)|~iext(uri_owl_equivalentClass,X3,X4))).
###
#cnf(i_0_63, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X4,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X1,X3)|~iext(uri_owl_equivalentClass,X4,X3))).
#
#cnf(i_0_64, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X4,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X3,X1)|~iext(uri_owl_equivalentClass,X4,X3))).
#
#cnf(i_0_67, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X4)|~iext(uri_owl_equivalentClass,X1,X3)|~iext(uri_owl_equivalentClass,X3,X4))).
#
#cnf(i_0_82, plain, (~iext(uri_owl_equivalentClass,X2,esk3_0)|~iext(uri_owl_equivalentClass,X2,uri_owl_Nothing))).
###
#cnf(i_0_68, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X4)|~iext(uri_owl_equivalentClass,X3,X1)|~iext(uri_owl_equivalentClass,X3,X4))).
####
#cnf(i_0_73, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X4)|~iext(uri_owl_equivalentClass,X1,X3)|~iext(uri_owl_equivalentClass,X4,X3))).
###
#cnf(i_0_74, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X4)|~iext(uri_owl_equivalentClass,X3,X1)|~iext(uri_owl_equivalentClass,X4,X3))).
###
#cnf(i_0_78, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X4,esk3_0)|~iext(uri_owl_equivalentClass,X1,X3)|~iext(uri_owl_equivalentClass,X3,X4))).
###
#cnf(i_0_29, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,X1,uri_owl_Thing))).
#
#cnf(i_0_79, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X4,esk3_0)|~iext(uri_owl_equivalentClass,X3,X1)|~iext(uri_owl_equivalentClass,X3,X4))).
###
#cnf(i_0_48, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,X1,esk3_0))).
#
#cnf(i_0_84, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X4,esk3_0)|~iext(uri_owl_equivalentClass,X1,X3)|~iext(uri_owl_equivalentClass,X4,X3))).
####
#cnf(i_0_85, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,X4,esk3_0)|~iext(uri_owl_equivalentClass,X3,X1)|~iext(uri_owl_equivalentClass,X4,X3))).
###
#cnf(i_0_71, plain, (~iext(uri_owl_equivalentClass,uri_owl_Thing,X2)|~iext(uri_owl_equivalentClass,X2,uri_owl_Nothing))).
#
#cnf(i_0_90, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,esk3_0,X4)|~iext(uri_owl_equivalentClass,X1,X3)|~iext(uri_owl_equivalentClass,X3,X4))).
##
#cnf(i_0_91, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,esk3_0,X4)|~iext(uri_owl_equivalentClass,X3,X1)|~iext(uri_owl_equivalentClass,X3,X4))).
###
#cnf(i_0_95, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,esk3_0,X4)|~iext(uri_owl_equivalentClass,X1,X3)|~iext(uri_owl_equivalentClass,X4,X3))).
###
#cnf(i_0_96, plain, (icext(X1,X2)|~iext(uri_owl_equivalentClass,esk3_0,X4)|~iext(uri_owl_equivalentClass,X3,X1)|~iext(uri_owl_equivalentClass,X4,X3))).
###
#cnf(i_0_93, plain, (~iext(uri_owl_equivalentClass,esk3_0,X2)|~iext(uri_owl_equivalentClass,X2,uri_owl_Nothing))).
############
#cnf(i_0_27, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,X1,X3)|~icext(X3,X2))).
#########
#cnf(i_0_11, plain, (iext(uri_owl_equivalentClass,X1,X2)|~ic(X2)|~ic(X1)|~icext(X2,esk2_2(X1,X2))|~icext(X1,esk2_2(X1,X2)))).
#
#cnf(i_0_440, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~icext(X1,esk2_2(esk3_0,X1)))).
#
#cnf(i_0_512, plain, (iext(uri_owl_equivalentClass,esk3_0,esk3_0))).
#
#cnf(i_0_32, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,X3,X1)|~icext(X3,X2))).
#
#cnf(i_0_513, plain, (iext(uri_owl_equivalentClass,esk3_0,uri_owl_Thing))).
###############
#cnf(i_0_123, plain, (~iext(uri_owl_equivalentClass,X1,uri_owl_Nothing)|~iext(uri_owl_equivalentClass,X1,X3)|~icext(X3,X2))).
############
#cnf(i_0_124, plain, (~iext(uri_owl_equivalentClass,X1,uri_owl_Nothing)|~iext(uri_owl_equivalentClass,X3,X1)|~icext(X3,X2))).
##########
#cnf(i_0_521, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X1))).
##################
#cnf(i_0_57, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,X3,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X1,X3))).
##
#cnf(i_0_517, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X1,uri_owl_Thing))).
#########
#cnf(i_0_524, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X1,esk3_0))).
##########
#cnf(i_0_441, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~icext(X1,esk2_2(uri_owl_Thing,X1)))).
#
#cnf(i_0_873, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,uri_owl_Thing))).
##
#cnf(i_0_62, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,X3,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X3,X1))).
####
#cnf(i_0_66, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X3)|~iext(uri_owl_equivalentClass,X1,X3))).
####
#cnf(i_0_72, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X3)|~iext(uri_owl_equivalentClass,X3,X1))).
####
#cnf(i_0_77, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,X3,esk3_0)|~iext(uri_owl_equivalentClass,X1,X3))).
########
#cnf(i_0_83, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,X3,esk3_0)|~iext(uri_owl_equivalentClass,X3,X1))).
####
#cnf(i_0_89, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,esk3_0,X3)|~iext(uri_owl_equivalentClass,X1,X3))).
####
#cnf(i_0_94, plain, (~iext(uri_owl_equivalentClass,uri_owl_Nothing,X1)|~iext(uri_owl_equivalentClass,esk3_0,X3)|~iext(uri_owl_equivalentClass,X3,X1))).
###
#cnf(i_0_885, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,esk3_0,X1))).
####################################
#cnf(i_0_877, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X1,uri_owl_Thing))).
#############
#cnf(i_0_127, plain, (~iext(uri_owl_equivalentClass,X1,uri_owl_Nothing)|~iext(uri_owl_equivalentClass,X3,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X1,X3))).
#####
#cnf(i_0_884, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X1,esk3_0))).
############################
#cnf(i_0_476, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~icext(X1,esk2_2(X1,esk3_0)))).
#
#cnf(i_0_1260, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~iext(uri_owl_equivalentClass,X1,uri_owl_Thing))).
#
#cnf(i_0_1264, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X1))).
############
#cnf(i_0_1268, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~iext(uri_owl_equivalentClass,esk3_0,X1))).
######################
#cnf(i_0_477, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~icext(X1,esk2_2(X1,uri_owl_Thing)))).
#
#cnf(i_0_1499, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~iext(uri_owl_equivalentClass,X1,esk3_0))).
#####################
#cnf(i_0_1496, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X1))).
############
#cnf(i_0_1500, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~iext(uri_owl_equivalentClass,esk3_0,X1))).
#################################################################################################################################################################################################################################################
#cnf(i_0_129, plain, (~iext(uri_owl_equivalentClass,X1,uri_owl_Nothing)|~iext(uri_owl_equivalentClass,X3,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X3,X1))).
###
#cnf(i_0_514, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X1,X2)|~icext(X2,esk2_2(esk3_0,X1)))).
####
#cnf(i_0_515, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,X1)|~icext(X2,esk2_2(esk3_0,X1)))).
###
#cnf(i_0_874, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X1,X2)|~icext(X2,esk2_2(uri_owl_Thing,X1)))).
####
#cnf(i_0_875, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,X1)|~icext(X2,esk2_2(uri_owl_Thing,X1)))).
###
#cnf(i_0_1257, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~iext(uri_owl_equivalentClass,X1,X2)|~icext(X2,esk2_2(X1,esk3_0)))).
##
#cnf(i_0_1258, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,X1)|~icext(X2,esk2_2(X1,esk3_0)))).
####
#cnf(i_0_1489, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~iext(uri_owl_equivalentClass,X1,X2)|~icext(X2,esk2_2(X1,uri_owl_Thing)))).
###
#cnf(i_0_1490, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,X1)|~icext(X2,esk2_2(X1,uri_owl_Thing)))).
####
#cnf(i_0_518, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X1,X2))).
##############
#cnf(i_0_132, plain, (~iext(uri_owl_equivalentClass,X1,uri_owl_Nothing)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X3)|~iext(uri_owl_equivalentClass,X1,X3))).
#####
#cnf(i_0_523, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X2,X1))).
#####################
#cnf(i_0_526, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X2)|~iext(uri_owl_equivalentClass,X1,X2))).
####################
#cnf(i_0_527, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X2)|~iext(uri_owl_equivalentClass,X2,X1))).
##
#cnf(i_0_133, plain, (~iext(uri_owl_equivalentClass,X1,uri_owl_Nothing)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X3)|~iext(uri_owl_equivalentClass,X3,X1))).
#################
#cnf(i_0_528, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,esk3_0)|~iext(uri_owl_equivalentClass,X1,X2))).
####################
#cnf(i_0_529, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,esk3_0)|~iext(uri_owl_equivalentClass,X2,X1))).
######################
#cnf(i_0_530, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,esk3_0,X2)|~iext(uri_owl_equivalentClass,X1,X2))).
#############
#cnf(i_0_134, plain, (~iext(uri_owl_equivalentClass,X1,uri_owl_Nothing)|~iext(uri_owl_equivalentClass,X3,esk3_0)|~iext(uri_owl_equivalentClass,X1,X3))).
#########
#cnf(i_0_531, plain, (iext(uri_owl_equivalentClass,esk3_0,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,esk3_0,X2)|~iext(uri_owl_equivalentClass,X2,X1))).
##############
#cnf(i_0_878, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X1,X2))).
##################
#cnf(i_0_883, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X2,X1))).
################
#cnf(i_0_886, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X2)|~iext(uri_owl_equivalentClass,X1,X2))).
###################
#cnf(i_0_135, plain, (~iext(uri_owl_equivalentClass,X1,uri_owl_Nothing)|~iext(uri_owl_equivalentClass,X3,esk3_0)|~iext(uri_owl_equivalentClass,X3,X1))).
###
#cnf(i_0_887, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X2)|~iext(uri_owl_equivalentClass,X2,X1))).
################
#cnf(i_0_888, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,esk3_0)|~iext(uri_owl_equivalentClass,X1,X2))).
###################
#cnf(i_0_889, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,esk3_0)|~iext(uri_owl_equivalentClass,X2,X1))).
################
#cnf(i_0_890, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,esk3_0,X2)|~iext(uri_owl_equivalentClass,X1,X2))).
################
#cnf(i_0_891, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~ic(X1)|~iext(uri_owl_equivalentClass,esk3_0,X2)|~iext(uri_owl_equivalentClass,X2,X1))).
#############
#cnf(i_0_1261, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X1,X2))).
###################
#cnf(i_0_1266, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X2,X1))).
#####################
#cnf(i_0_1269, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X2)|~iext(uri_owl_equivalentClass,X1,X2))).
######################
#cnf(i_0_1270, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X2)|~iext(uri_owl_equivalentClass,X2,X1))).
####################
#cnf(i_0_1271, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,esk3_0)|~iext(uri_owl_equivalentClass,X1,X2))).
####################
#cnf(i_0_1272, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,esk3_0)|~iext(uri_owl_equivalentClass,X2,X1))).
######################
#cnf(i_0_1273, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~iext(uri_owl_equivalentClass,esk3_0,X2)|~iext(uri_owl_equivalentClass,X1,X2))).
######################
#cnf(i_0_1274, plain, (iext(uri_owl_equivalentClass,X1,esk3_0)|~ic(X1)|~iext(uri_owl_equivalentClass,esk3_0,X2)|~iext(uri_owl_equivalentClass,X2,X1))).
##############
#cnf(i_0_1493, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X1,X2))).
###############
#cnf(i_0_1498, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,uri_owl_Thing)|~iext(uri_owl_equivalentClass,X2,X1))).
#####################
#cnf(i_0_1501, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X2)|~iext(uri_owl_equivalentClass,X1,X2))).
###################
#cnf(i_0_1502, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X2)|~iext(uri_owl_equivalentClass,X2,X1))).
####################
#cnf(i_0_1503, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,esk3_0)|~iext(uri_owl_equivalentClass,X1,X2))).
###############
#cnf(i_0_1504, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,esk3_0)|~iext(uri_owl_equivalentClass,X2,X1))).
#####################
#cnf(i_0_1505, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~iext(uri_owl_equivalentClass,esk3_0,X2)|~iext(uri_owl_equivalentClass,X1,X2))).
################
#cnf(i_0_1506, plain, (iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~ic(X1)|~iext(uri_owl_equivalentClass,esk3_0,X2)|~iext(uri_owl_equivalentClass,X2,X1))).
##################################################################################################################################################
#cnf(i_0_136, plain, (~iext(uri_owl_equivalentClass,X1,uri_owl_Nothing)|~iext(uri_owl_equivalentClass,esk3_0,X3)|~iext(uri_owl_equivalentClass,X1,X3))).
########################
#cnf(i_0_137, plain, (~iext(uri_owl_equivalentClass,X1,uri_owl_Nothing)|~iext(uri_owl_equivalentClass,esk3_0,X3)|~iext(uri_owl_equivalentClass,X3,X1))).
#########################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################################
#cnf(i_0_445, plain, (iext(uri_owl_equivalentClass,X1,X2)|~ic(X2)|~ic(X1)|~iext(uri_owl_equivalentClass,X1,uri_owl_Thing)|~icext(X2,esk2_2(X1,X2)))).
##
#cnf(i_0_449, plain, (iext(uri_owl_equivalentClass,X1,X2)|~ic(X2)|~ic(X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X1)|~icext(X2,esk2_2(X1,X2)))).
####
#cnf(i_0_452, plain, (iext(uri_owl_equivalentClass,X1,X2)|~ic(X2)|~ic(X1)|~iext(uri_owl_equivalentClass,X1,esk3_0)|~icext(X2,esk2_2(X1,X2)))).
###
#cnf(i_0_453, plain, (iext(uri_owl_equivalentClass,X1,X2)|~ic(X2)|~ic(X1)|~iext(uri_owl_equivalentClass,esk3_0,X1)|~icext(X2,esk2_2(X1,X2)))).
####
#cnf(i_0_481, plain, (iext(uri_owl_equivalentClass,X1,X2)|~ic(X2)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,uri_owl_Thing)|~icext(X1,esk2_2(X1,X2)))).
####
#cnf(i_0_485, plain, (iext(uri_owl_equivalentClass,X1,X2)|~ic(X2)|~ic(X1)|~iext(uri_owl_equivalentClass,uri_owl_Thing,X2)|~icext(X1,esk2_2(X1,X2)))).
###
#cnf(i_0_488, plain, (iext(uri_owl_equivalentClass,X1,X2)|~ic(X2)|~ic(X1)|~iext(uri_owl_equivalentClass,X2,esk3_0)|~icext(X1,esk2_2(X1,X2)))).
####
#cnf(i_0_489, plain, (iext(uri_owl_equivalentClass,X1,X2)|~ic(X2)|~ic(X1)|~iext(uri_owl_equivalentClass,esk3_0,X2)|~icext(X1,esk2_2(X1,X2)))).
#########################################################################################################################
#cnf(i_0_8, plain, (X2=X3|~iext(uri_rdf_rest,X4,uri_rdf_nil)|~iext(uri_rdf_first,X4,X3)|~iext(uri_owl_oneOf,X1,X4)|~icext(X1,X2))).
#
#cnf(i_0_8115, plain, (X1=X2|~iext(uri_rdf_first,esk4_0,X2))).
#
#cnf(i_0_8505, plain, (~iext(uri_rdf_first,esk4_0,uri_owl_Nothing)|~icext(X2,X1))).
#
#cnf(i_0_8644, plain, (X1=uri_ex_w)).
#
#cnf(i_0_8685, plain, (X2=X1)).
#
#cnf(i_0_21, plain, (ic(uri_ex_w))).
#
#cnf(i_0_8505, plain, (~icext(X1,X2))).
#
#cnf(i_0_8694, plain, (ic(X1))).
##
# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
fof(owl_class_thing_ext, axiom, ![X1]:(icext(uri_owl_Thing,X1)<=>ir(X1)), file('./p/SWB031+2.tptp', owl_class_thing_ext)).
fof(simple_ir, axiom, ![X1]:ir(X1), file('./p/SWB031+2.tptp', simple_ir)).
fof(owl_eqdis_equivalentclass, axiom, ![X5, X6]:(iext(uri_owl_equivalentClass,X5,X6)<=>((ic(X5)&ic(X6))&![X1]:(icext(X5,X1)<=>icext(X6,X1)))), file('./p/SWB031+2.tptp', owl_eqdis_equivalentclass)).
fof(testcase_premise_fullish_031_Large_Universe, axiom, ?[X7, X8]:(((iext(uri_owl_equivalentClass,uri_owl_Thing,X7)&iext(uri_owl_oneOf,X7,X8))&iext(uri_rdf_first,X8,uri_ex_w))&iext(uri_rdf_rest,X8,uri_rdf_nil)), file('./p/SWB031+2.tptp', testcase_premise_fullish_031_Large_Universe)).
fof(owl_enum_class_001, axiom, ![X2, X3, X4]:((iext(uri_rdf_first,X3,X4)&iext(uri_rdf_rest,X3,uri_rdf_nil))=>(iext(uri_owl_oneOf,X2,X3)<=>(ic(X2)&![X1]:(icext(X2,X1)<=>X1=X4)))), file('./p/SWB031+2.tptp', owl_enum_class_001)).
fof(owl_class_nothing_ext, axiom, ![X1]:~(icext(uri_owl_Nothing,X1)), file('./p/SWB031+2.tptp', owl_class_nothing_ext)).
fof(c_0_6, plain, ![X10]:((~icext(uri_owl_Thing,X10)|ir(X10))&(~ir(X10)|icext(uri_owl_Thing,X10))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[owl_class_thing_ext])])).
fof(c_0_7, plain, ![X9]:ir(X9), inference(variable_rename,[status(thm)],[simple_ir])).
fof(c_0_8, plain, ![X18, X19, X20, X21, X22, X23]:((((ic(X18)|~iext(uri_owl_equivalentClass,X18,X19))&(ic(X19)|~iext(uri_owl_equivalentClass,X18,X19)))&((~icext(X18,X20)|icext(X19,X20)|~iext(uri_owl_equivalentClass,X18,X19))&(~icext(X19,X21)|icext(X18,X21)|~iext(uri_owl_equivalentClass,X18,X19))))&((~icext(X22,esk2_2(X22,X23))|~icext(X23,esk2_2(X22,X23))|(~ic(X22)|~ic(X23))|iext(uri_owl_equivalentClass,X22,X23))&(icext(X22,esk2_2(X22,X23))|icext(X23,esk2_2(X22,X23))|(~ic(X22)|~ic(X23))|iext(uri_owl_equivalentClass,X22,X23)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[owl_eqdis_equivalentclass])])])])])])).
fof(c_0_9, plain, (((iext(uri_owl_equivalentClass,uri_owl_Thing,esk3_0)&iext(uri_owl_oneOf,esk3_0,esk4_0))&iext(uri_rdf_first,esk4_0,uri_ex_w))&iext(uri_rdf_rest,esk4_0,uri_rdf_nil)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[testcase_premise_fullish_031_Large_Universe])])).
cnf(c_0_10, plain, (icext(uri_owl_Thing,X1)|~ir(X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_11, plain, (ir(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
fof(c_0_12, plain, ![X12, X13, X14, X15, X16]:(((ic(X12)|~iext(uri_owl_oneOf,X12,X13)|(~iext(uri_rdf_first,X13,X14)|~iext(uri_rdf_rest,X13,uri_rdf_nil)))&((~icext(X12,X15)|X15=X14|~iext(uri_owl_oneOf,X12,X13)|(~iext(uri_rdf_first,X13,X14)|~iext(uri_rdf_rest,X13,uri_rdf_nil)))&(X16!=X14|icext(X12,X16)|~iext(uri_owl_oneOf,X12,X13)|(~iext(uri_rdf_first,X13,X14)|~iext(uri_rdf_rest,X13,uri_rdf_nil)))))&((~icext(X12,esk1_3(X12,X13,X14))|esk1_3(X12,X13,X14)!=X14|~ic(X12)|iext(uri_owl_oneOf,X12,X13)|(~iext(uri_rdf_first,X13,X14)|~iext(uri_rdf_rest,X13,uri_rdf_nil)))&(icext(X12,esk1_3(X12,X13,X14))|esk1_3(X12,X13,X14)=X14|~ic(X12)|iext(uri_owl_oneOf,X12,X13)|(~iext(uri_rdf_first,X13,X14)|~iext(uri_rdf_rest,X13,uri_rdf_nil))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[owl_enum_class_001])])])])])])).
cnf(c_0_13, plain, (icext(X3,X2)|~icext(X1,X2)|~iext(uri_owl_equivalentClass,X1,X3)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_14, plain, (iext(uri_owl_equivalentClass,uri_owl_Thing,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_15, plain, (icext(uri_owl_Thing,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_11])])).
fof(c_0_16, plain, ![X11]:~icext(uri_owl_Nothing,X11), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[owl_class_nothing_ext])])).
cnf(c_0_17, plain, (X2=X3|~icext(X1,X2)|~iext(uri_owl_oneOf,X1,X4)|~iext(uri_rdf_first,X4,X3)|~iext(uri_rdf_rest,X4,uri_rdf_nil)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (iext(uri_owl_oneOf,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_19, plain, (iext(uri_rdf_rest,esk4_0,uri_rdf_nil)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_20, plain, (icext(esk3_0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_13, c_0_14]), c_0_15])])).
cnf(c_0_21, plain, (~icext(uri_owl_Nothing,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_22, plain, (X1=X2|~iext(uri_rdf_first,esk4_0,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_20])])).
cnf(c_0_23, plain, (iext(uri_rdf_first,esk4_0,uri_ex_w)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_24, plain, (~iext(uri_rdf_first,esk4_0,uri_owl_Nothing)|~icext(X1,X2)), inference(pm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_25, plain, (X1=uri_ex_w), inference(pm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_26, plain, (~icext(X1,X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_24, c_0_25]), c_0_23])])).
cnf(c_0_27, plain, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_25]), c_0_26]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 28
# Proof object clause steps            : 16
# Proof object formula steps           : 12
# Proof object conjectures             : 0
# Proof object clause conjectures      : 0
# Proof object formula conjectures     : 0
# Proof object initial clauses used    : 9
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 4
# Proof object simplifying inferences  : 12
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 19
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 17
# Processed clauses                    : 2363
# ...of these trivial                  : 1
# ...subsumed                          : 2205
# ...remaining for further processing  : 156
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 150
# Generated clauses                    : 8676
# ...of the previous two non-trivial   : 6967
# Contextual simplify-reflections      : 0
# Paramodulations                      : 8676
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 4
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 4619
# ...number of literals in the above   : 23887
# Current number of archived formulas  : 0
# Current number of archived clauses   : 153
# Clause-clause subsumption calls (NU) : 36954
# Rec. Clause-clause subsumption calls : 7753
# Non-unit clause-clause subsumptions  : 2088
# Unit Clause-clause subsumption calls : 58
# Rewrite failures with RHS unbound    : 21
# BW rewrite match attempts            : 37
# BW rewrite match successes           : 23
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 56729
