# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_11, plain, (iext(uri_owl_propertyDisjointWith,uri_ex_notInstanceOf,uri_rdf_type))).
#
#cnf(i_0_9, plain, (iext(uri_ex_notInstanceOf,uri_ex_u,uri_ex_c))).
#
#cnf(i_0_10, plain, (iext(uri_rdf_type,uri_ex_w,uri_ex_c))).
#
#cnf(i_0_8, negated_conjecture, (~iext(uri_owl_differentFrom,uri_ex_w,uri_ex_u))).
#
#cnf(i_0_6, plain, (ip(X1)|~iext(uri_owl_propertyDisjointWith,X2,X1))).
#
#cnf(i_0_12, plain, (ip(uri_rdf_type))).
#
#cnf(i_0_7, plain, (ip(X1)|~iext(uri_owl_propertyDisjointWith,X1,X2))).
#
#cnf(i_0_2, plain, (X1!=X2|~iext(uri_owl_differentFrom,X1,X2))).
#
#cnf(i_0_13, plain, (ip(uri_ex_notInstanceOf))).
#
#cnf(i_0_1, plain, (X1=X2|iext(uri_owl_differentFrom,X1,X2))).
#
#cnf(i_0_14, negated_conjecture, (uri_ex_w=uri_ex_u)).
#
#cnf(i_0_8, negated_conjecture, (~iext(uri_owl_differentFrom,uri_ex_u,uri_ex_u))).
#
#cnf(i_0_10, plain, (iext(uri_rdf_type,uri_ex_u,uri_ex_c))).
#
#cnf(i_0_3, plain, (iext(X1,esk1_2(X2,X1),esk2_2(X2,X1))|iext(uri_owl_propertyDisjointWith,X2,X1)|~ip(X2)|~ip(X1))).
#
#cnf(i_0_16, plain, (ip(esk2_2(X1,uri_owl_propertyDisjointWith))|iext(uri_owl_propertyDisjointWith,X1,uri_owl_propertyDisjointWith)|~ip(uri_owl_propertyDisjointWith)|~ip(X1))).
#
#cnf(i_0_5, plain, (~iext(uri_owl_propertyDisjointWith,X1,X4)|~iext(X4,X2,X3)|~iext(X1,X2,X3))).
#
#cnf(i_0_22, plain, (X2=X3|~iext(uri_owl_propertyDisjointWith,uri_owl_differentFrom,X1)|~iext(X1,X2,X3))).
#
#cnf(i_0_32, plain, (X1=X2|~iext(uri_owl_propertyDisjointWith,uri_owl_differentFrom,uri_owl_differentFrom))).
#
#cnf(i_0_41, plain, (ip(X1)|~iext(uri_owl_propertyDisjointWith,uri_owl_differentFrom,uri_owl_differentFrom))).
#
#cnf(i_0_29, plain, (~iext(uri_rdf_type,X1,X2)|~iext(uri_ex_notInstanceOf,X1,X2))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(testcase_conclusion_fullish_017_Built_in_Based_Definitions, conjecture, iext(uri_owl_differentFrom,uri_ex_w,uri_ex_u), file('./p/SWB017+2.tptp', testcase_conclusion_fullish_017_Built_in_Based_Definitions)).
fof(owl_eqdis_differentfrom, axiom, ![X1, X2]:(iext(uri_owl_differentFrom,X1,X2)<=>X1!=X2), file('./p/SWB017+2.tptp', owl_eqdis_differentfrom)).
fof(owl_eqdis_propertydisjointwith, axiom, ![X3, X4]:(iext(uri_owl_propertyDisjointWith,X3,X4)<=>((ip(X3)&ip(X4))&![X1, X2]:~((iext(X3,X1,X2)&iext(X4,X1,X2))))), file('./p/SWB017+2.tptp', owl_eqdis_propertydisjointwith)).
fof(testcase_premise_fullish_017_Built_in_Based_Definitions, axiom, ((iext(uri_owl_propertyDisjointWith,uri_ex_notInstanceOf,uri_rdf_type)&iext(uri_rdf_type,uri_ex_w,uri_ex_c))&iext(uri_ex_notInstanceOf,uri_ex_u,uri_ex_c)), file('./p/SWB017+2.tptp', testcase_premise_fullish_017_Built_in_Based_Definitions)).
fof(c_0_4, negated_conjecture, ~(iext(uri_owl_differentFrom,uri_ex_w,uri_ex_u)), inference(assume_negation,[status(cth)],[testcase_conclusion_fullish_017_Built_in_Based_Definitions])).
fof(c_0_5, negated_conjecture, ~iext(uri_owl_differentFrom,uri_ex_w,uri_ex_u), inference(fof_simplification,[status(thm)],[c_0_4])).
fof(c_0_6, plain, ![X5, X6]:((~iext(uri_owl_differentFrom,X5,X6)|X5!=X6)&(X5=X6|iext(uri_owl_differentFrom,X5,X6))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[owl_eqdis_differentfrom])])).
fof(c_0_7, plain, ![X7, X8, X9, X10, X11, X12]:((((ip(X7)|~iext(uri_owl_propertyDisjointWith,X7,X8))&(ip(X8)|~iext(uri_owl_propertyDisjointWith,X7,X8)))&(~iext(X7,X9,X10)|~iext(X8,X9,X10)|~iext(uri_owl_propertyDisjointWith,X7,X8)))&((iext(X11,esk1_2(X11,X12),esk2_2(X11,X12))|(~ip(X11)|~ip(X12))|iext(uri_owl_propertyDisjointWith,X11,X12))&(iext(X12,esk1_2(X11,X12),esk2_2(X11,X12))|(~ip(X11)|~ip(X12))|iext(uri_owl_propertyDisjointWith,X11,X12)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[owl_eqdis_propertydisjointwith])])])])])])).
cnf(c_0_8, negated_conjecture, (~iext(uri_owl_differentFrom,uri_ex_w,uri_ex_u)), inference(split_conjunct,[status(thm)],[c_0_5])).
cnf(c_0_9, plain, (X1=X2|iext(uri_owl_differentFrom,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_10, plain, (~iext(X1,X2,X3)|~iext(X4,X2,X3)|~iext(uri_owl_propertyDisjointWith,X1,X4)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_11, plain, (iext(uri_owl_propertyDisjointWith,uri_ex_notInstanceOf,uri_rdf_type)), inference(split_conjunct,[status(thm)],[testcase_premise_fullish_017_Built_in_Based_Definitions])).
cnf(c_0_12, plain, (iext(uri_rdf_type,uri_ex_w,uri_ex_c)), inference(split_conjunct,[status(thm)],[testcase_premise_fullish_017_Built_in_Based_Definitions])).
cnf(c_0_13, negated_conjecture, (uri_ex_w=uri_ex_u), inference(pm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_14, plain, (~iext(uri_rdf_type,X1,X2)|~iext(uri_ex_notInstanceOf,X1,X2)), inference(pm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_15, plain, (iext(uri_rdf_type,uri_ex_u,uri_ex_c)), inference(rw,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_16, plain, (iext(uri_ex_notInstanceOf,uri_ex_u,uri_ex_c)), inference(split_conjunct,[status(thm)],[testcase_premise_fullish_017_Built_in_Based_Definitions])).
cnf(c_0_17, plain, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_14, c_0_15]), c_0_16])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 18
# Proof object clause steps            : 10
# Proof object formula steps           : 8
# Proof object conjectures             : 5
# Proof object clause conjectures      : 2
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 6
# Proof object initial formulas used   : 4
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 3
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 4
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 11
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 11
# Processed clauses                    : 20
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 20
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 91
# ...of the previous two non-trivial   : 89
# Contextual simplify-reflections      : 0
# Paramodulations                      : 91
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 18
#    Positive orientable unit clauses  : 6
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 11
# Current number of unprocessed clauses: 77
# ...number of literals in the above   : 244
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 10
# Rec. Clause-clause subsumption calls : 9
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 1
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 906
