# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_5, plain, (ssItem(esk1_0))).
#
#cnf(i_0_4, plain, (ssItem(esk2_0))).
#
#cnf(i_0_84, plain, (ssList(nil))).
#
#cnf(i_0_115, plain, (~singletonP(nil))).
#
#cnf(i_0_203, negated_conjecture, (ssList(esk48_0))).
#
#cnf(i_0_202, negated_conjecture, (ssList(esk49_0))).
#
#cnf(i_0_201, negated_conjecture, (ssList(esk50_0))).
#
#cnf(i_0_3, plain, (esk2_0!=esk1_0)).
#
#cnf(i_0_200, negated_conjecture, (ssList(esk51_0))).
#
#cnf(i_0_141, plain, (cyclefreeP(nil))).
#
#cnf(i_0_143, plain, (totalorderP(nil))).
#
#cnf(i_0_183, plain, (~lt(X1,X1)|~ssItem(X1))).
#
#cnf(i_0_145, plain, (strictorderP(nil))).
#
#cnf(i_0_147, plain, (totalorderedP(nil))).
#
#cnf(i_0_154, plain, (strictorderedP(nil))).
#
#cnf(i_0_114, plain, (~memberP(nil,X1)|~ssItem(X1))).
#
#cnf(i_0_161, plain, (duplicatefreeP(nil))).
#
#cnf(i_0_163, plain, (equalelemsP(nil))).
#
#cnf(i_0_198, negated_conjecture, (esk50_0=esk48_0)).
#
#cnf(i_0_195, negated_conjecture, (nil!=esk48_0|nil!=esk49_0)).
#
#cnf(i_0_199, negated_conjecture, (esk49_0=esk51_0)).
#
#cnf(i_0_118, plain, (frontsegP(X1,X1)|~ssList(X1))).
#
#cnf(i_0_128, plain, (rearsegP(X1,X1)|~ssList(X1))).
#
#cnf(i_0_195, negated_conjecture, (nil!=esk48_0|nil!=esk51_0)).
#
#cnf(i_0_135, plain, (segmentP(X1,X1)|~ssList(X1))).
#
#cnf(i_0_101, plain, (leq(X1,X1)|~ssItem(X1))).
#
#cnf(i_0_182, plain, (geq(X1,X1)|~ssItem(X1))).
#
#cnf(i_0_85, plain, (cons(X2,X1)!=X1|~ssList(X1)|~ssItem(X2))).
#
#cnf(i_0_123, plain, (frontsegP(X1,nil)|~ssList(X1))).
#
#cnf(i_0_130, plain, (rearsegP(X1,nil)|~ssList(X1))).
#
#cnf(i_0_137, plain, (segmentP(X1,nil)|~ssList(X1))).
#
#cnf(i_0_91, plain, (cons(X2,X1)!=nil|~ssList(X1)|~ssItem(X2))).
#
#cnf(i_0_196, negated_conjecture, (nil=esk48_0|nil!=esk51_0)).
#
#cnf(i_0_194, negated_conjecture, (ssList(esk52_0)|~neq(esk51_0,nil))).
#
#cnf(i_0_177, plain, (app(X1,nil)=X1|~ssList(X1))).
#
#cnf(i_0_2, plain, (X1!=X2|~neq(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_98, plain, (app(nil,X1)=X1|~ssList(X1))).
#
#cnf(i_0_140, plain, (cyclefreeP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_142, plain, (totalorderP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_188, plain, (X1!=X2|~lt(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_144, plain, (strictorderP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_146, plain, (totalorderedP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_153, plain, (strictorderedP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_82, plain, (X1!=X2|~ssList(X2)|~ssList(X1)|~neq(X1,X2))).
#
#cnf(i_0_160, plain, (duplicatefreeP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_162, plain, (equalelemsP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_12, plain, (ssItem(esk5_1(X1))|~singletonP(X1)|~ssList(X1))).
#
#cnf(i_0_104, plain, (~lt(X2,X1)|~lt(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_30, plain, (cyclefreeP(X1)|ssItem(esk10_1(X1))|~ssList(X1))).
#
#cnf(i_0_29, plain, (cyclefreeP(X1)|ssItem(esk11_1(X1))|~ssList(X1))).
#
#cnf(i_0_28, plain, (cyclefreeP(X1)|ssList(esk12_1(X1))|~ssList(X1))).
#
#cnf(i_0_189, plain, (~gt(X2,X1)|~gt(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_27, plain, (cyclefreeP(X1)|ssList(esk13_1(X1))|~ssList(X1))).
#
#cnf(i_0_26, plain, (cyclefreeP(X1)|ssList(esk14_1(X1))|~ssList(X1))).
#
#cnf(i_0_39, plain, (totalorderP(X1)|ssItem(esk15_1(X1))|~ssList(X1))).
#
#cnf(i_0_197, negated_conjecture, (~segmentP(esk48_0,X1)|~segmentP(esk51_0,X1)|~ssList(X1)|~neq(X1,nil))).
#
#cnf(i_0_38, plain, (totalorderP(X1)|ssItem(esk16_1(X1))|~ssList(X1))).
#
#cnf(i_0_37, plain, (totalorderP(X1)|ssList(esk17_1(X1))|~ssList(X1))).
#
#cnf(i_0_36, plain, (totalorderP(X1)|ssList(esk18_1(X1))|~ssList(X1))).
#
#cnf(i_0_73, plain, (app(app(X4,cons(X2,X5)),cons(X3,X6))!=X1|X2!=X3|~duplicatefreeP(X1)|~ssList(X6)|~ssList(X5)|~ssList(X4)|~ssList(X1)|~ssItem(X3)|~ssItem(X2))).
#
#cnf(i_0_35, plain, (totalorderP(X1)|ssList(esk19_1(X1))|~ssList(X1))).
#
#cnf(i_0_48, plain, (strictorderP(X1)|ssItem(esk20_1(X1))|~ssList(X1))).
#
#cnf(i_0_47, plain, (strictorderP(X1)|ssItem(esk21_1(X1))|~ssList(X1))).
#
#cnf(i_0_206, plain, (app(cons(X1,X2),cons(X3,X4))!=X5|X1!=X3|~duplicatefreeP(X5)|~ssList(cons(X1,X2))|~ssList(X4)|~ssList(X2)|~ssList(X5)|~ssItem(X3)|~ssItem(X1))).
#
#cnf(i_0_46, plain, (strictorderP(X1)|ssList(esk22_1(X1))|~ssList(X1))).
#
#cnf(i_0_45, plain, (strictorderP(X1)|ssList(esk23_1(X1))|~ssList(X1))).
#
#cnf(i_0_44, plain, (strictorderP(X1)|ssList(esk24_1(X1))|~ssList(X1))).
#
#cnf(i_0_31, plain, (app(app(X4,cons(X2,X5)),cons(X3,X6))!=X1|~leq(X3,X2)|~leq(X2,X3)|~cyclefreeP(X1)|~ssList(X6)|~ssList(X5)|~ssList(X4)|~ssList(X1)|~ssItem(X3)|~ssItem(X2))).
#
#cnf(i_0_56, plain, (totalorderedP(X1)|ssItem(esk25_1(X1))|~ssList(X1))).
#
#cnf(i_0_55, plain, (totalorderedP(X1)|ssItem(esk26_1(X1))|~ssList(X1))).
#
#cnf(i_0_54, plain, (totalorderedP(X1)|ssList(esk27_1(X1))|~ssList(X1))).
#
#cnf(i_0_211, plain, (app(cons(X1,X2),cons(X3,X4))!=X5|~leq(X3,X1)|~leq(X1,X3)|~cyclefreeP(X5)|~ssList(cons(X1,X2))|~ssList(X4)|~ssList(X2)|~ssList(X5)|~ssItem(X3)|~ssItem(X1))).
#
#cnf(i_0_53, plain, (totalorderedP(X1)|ssList(esk28_1(X1))|~ssList(X1))).
#
#cnf(i_0_52, plain, (totalorderedP(X1)|ssList(esk29_1(X1))|~ssList(X1))).
#
#cnf(i_0_64, plain, (strictorderedP(X1)|ssItem(esk30_1(X1))|~ssList(X1))).
#
#cnf(i_0_207, plain, (X1!=X2|~duplicatefreeP(app(cons(X1,X3),cons(X2,X4)))|~ssList(app(cons(X1,X3),cons(X2,X4)))|~ssList(cons(X1,X3))|~ssList(X4)|~ssList(X3)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_63, plain, (strictorderedP(X1)|ssItem(esk31_1(X1))|~ssList(X1))).
#
#cnf(i_0_62, plain, (strictorderedP(X1)|ssList(esk32_1(X1))|~ssList(X1))).
#
#cnf(i_0_61, plain, (strictorderedP(X1)|ssList(esk33_1(X1))|~ssList(X1))).
#
#cnf(i_0_204, plain, (X1!=X2|~duplicatefreeP(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(app(app(X3,cons(X1,X4)),cons(X2,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_60, plain, (strictorderedP(X1)|ssList(esk34_1(X1))|~ssList(X1))).
#
#cnf(i_0_72, plain, (duplicatefreeP(X1)|ssItem(esk35_1(X1))|~ssList(X1))).
#
#cnf(i_0_71, plain, (duplicatefreeP(X1)|ssItem(esk36_1(X1))|~ssList(X1))).
#
#cnf(i_0_208, plain, (app(cons(X1,X2),cons(X3,X4))!=cons(X5,nil)|X1!=X3|~ssList(cons(X5,nil))|~ssList(cons(X1,X2))|~ssList(X4)|~ssList(X2)|~ssItem(X3)|~ssItem(X1)|~ssItem(X5))).
#
#cnf(i_0_70, plain, (duplicatefreeP(X1)|ssList(esk37_1(X1))|~ssList(X1))).
#
#cnf(i_0_69, plain, (duplicatefreeP(X1)|ssList(esk38_1(X1))|~ssList(X1))).
#
#cnf(i_0_68, plain, (duplicatefreeP(X1)|ssList(esk39_1(X1))|~ssList(X1))).
#
#cnf(i_0_205, plain, (app(app(X1,cons(X2,X3)),cons(X4,X5))!=cons(X6,nil)|X2!=X4|~ssList(cons(X6,nil))|~ssList(X5)|~ssList(X3)|~ssList(X1)|~ssItem(X4)|~ssItem(X2)|~ssItem(X6))).
#
#cnf(i_0_79, plain, (equalelemsP(X1)|ssItem(esk40_1(X1))|~ssList(X1))).
#
#cnf(i_0_78, plain, (equalelemsP(X1)|ssItem(esk41_1(X1))|~ssList(X1))).
#
#cnf(i_0_77, plain, (equalelemsP(X1)|ssList(esk42_1(X1))|~ssList(X1))).
##
#cnf(i_0_76, plain, (equalelemsP(X1)|ssList(esk43_1(X1))|~ssList(X1))).
#
#cnf(i_0_191, negated_conjecture, (segmentP(esk48_0,esk52_0)|~neq(esk51_0,nil))).
#
#cnf(i_0_192, negated_conjecture, (segmentP(esk51_0,esk52_0)|~neq(esk51_0,nil))).
#
#cnf(i_0_214, plain, (X1!=X2|~duplicatefreeP(app(cons(X1,X3),cons(X2,X4)))|~ssList(app(app(nil,cons(X1,X3)),cons(X2,X4)))|~ssList(cons(X1,X3))|~ssList(X4)|~ssList(X3)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_193, negated_conjecture, (neq(esk52_0,nil)|~neq(esk51_0,nil))).
#
#cnf(i_0_83, plain, (ssList(cons(X2,X1))|~ssList(X1)|~ssItem(X2))).
#
#cnf(i_0_96, plain, (ssList(app(X1,X2))|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_218, negated_conjecture, (esk52_0!=nil|~ssList(esk52_0)|~neq(esk51_0,nil))).
#
#cnf(i_0_74, plain, (equalelemsP(X1)|esk41_1(X1)!=esk40_1(X1)|~ssList(X1))).
#
#cnf(i_0_125, plain, (nil=X1|~frontsegP(nil,X1)|~ssList(X1))).
#
#cnf(i_0_132, plain, (nil=X1|~rearsegP(nil,X1)|~ssList(X1))).
#
#cnf(i_0_217, negated_conjecture, (esk52_0!=nil|~neq(esk51_0,nil)|~ssItem(nil)|~ssItem(esk52_0))).
#
#cnf(i_0_139, plain, (nil=X1|~segmentP(nil,X1)|~ssList(X1))).
#
#cnf(i_0_124, plain, (frontsegP(nil,X1)|nil!=X1|~ssList(X1))).
#
#cnf(i_0_131, plain, (rearsegP(nil,X1)|nil!=X1|~ssList(X1))).
#
#cnf(i_0_219, negated_conjecture, (~segmentP(esk48_0,esk52_0)|~segmentP(esk51_0,esk52_0)|~ssList(esk52_0)|~neq(esk51_0,nil))).
#
#cnf(i_0_138, plain, (segmentP(nil,X1)|nil!=X1|~ssList(X1))).
#
#cnf(i_0_92, plain, (nil=X1|ssItem(hd(X1))|~ssList(X1))).
#
#cnf(i_0_89, plain, (nil=X1|ssItem(esk45_1(X1))|~ssList(X1))).
#
#cnf(i_0_221, plain, (app(cons(X1,X2),cons(X3,X4))!=cons(X5,nil)|X1!=X3|~ssList(cons(X1,X2))|~ssList(X4)|~ssList(X2)|~ssItem(X3)|~ssItem(X1)|~ssItem(X5))).
#
#cnf(i_0_165, plain, (nil=X1|ssItem(esk46_1(X1))|~ssList(X1))).
#
#cnf(i_0_94, plain, (nil=X1|ssList(tl(X1))|~ssList(X1))).
#
#cnf(i_0_90, plain, (nil=X1|ssList(esk44_1(X1))|~ssList(X1))).
#
#cnf(i_0_220, plain, (app(app(X1,cons(X2,X3)),cons(X4,X5))!=cons(X6,nil)|X2!=X4|~ssList(X5)|~ssList(X3)|~ssList(X1)|~ssItem(X4)|~ssItem(X2)|~ssItem(X6))).
#
#cnf(i_0_167, plain, (nil=X1|ssList(esk47_1(X1))|~ssList(X1))).
#
#cnf(i_0_66, plain, (esk35_1(X1)=esk36_1(X1)|duplicatefreeP(X1)|~ssList(X1))).
###
#cnf(i_0_1, plain, (X1=X2|neq(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_81, plain, (X1=X2|neq(X1,X2)|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_239, negated_conjecture, (nil=esk51_0|ssList(esk52_0))).
##
#cnf(i_0_240, negated_conjecture, (nil=esk51_0|segmentP(esk48_0,esk52_0))).
#
#cnf(i_0_241, negated_conjecture, (nil=esk51_0|segmentP(esk51_0,esk52_0))).
#
#cnf(i_0_243, negated_conjecture, (nil=esk51_0|esk52_0!=nil|~ssList(esk52_0))).
#
#cnf(i_0_222, plain, (X1!=X2|~duplicatefreeP(app(cons(X1,X3),cons(X2,X4)))|~ssList(app(nil,cons(X1,X3)))|~ssList(cons(X1,X3))|~ssList(cons(X2,X4))|~ssList(X4)|~ssList(X3)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_248, negated_conjecture, (nil=esk51_0|esk52_0!=nil)).
#
#cnf(i_0_103, plain, (leq(X2,X1)|~geq(X1,X2)|~ssItem(X2)|~ssItem(X1))).
##
#cnf(i_0_212, plain, (~leq(X1,X2)|~leq(X2,X1)|~cyclefreeP(app(cons(X2,X3),cons(X1,X4)))|~ssList(app(cons(X2,X3),cons(X1,X4)))|~ssList(cons(X2,X3))|~ssList(X4)|~ssList(X3)|~ssItem(X1)|~ssItem(X2))).
#
#cnf(i_0_187, plain, (leq(X1,X2)|~lt(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_107, plain, (lt(X2,X1)|~gt(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_102, plain, (geq(X2,X1)|~leq(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_209, plain, (~leq(X1,X2)|~leq(X2,X1)|~cyclefreeP(app(app(X3,cons(X2,X4)),cons(X1,X5)))|~ssList(app(app(X3,cons(X2,X4)),cons(X1,X5)))|~ssList(X5)|~ssList(X4)|~ssList(X3)|~ssItem(X1)|~ssItem(X2))).
#
#cnf(i_0_106, plain, (gt(X2,X1)|~lt(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_33, plain, (totalorderP(X1)|~leq(esk15_1(X1),esk16_1(X1))|~ssList(X1))).
#
#cnf(i_0_32, plain, (totalorderP(X1)|~leq(esk16_1(X1),esk15_1(X1))|~ssList(X1))).
#
#cnf(i_0_252, plain, (~gt(X2,X1)|~lt(X2,X1)|~ssItem(X1)|~ssItem(X2))).
#
#cnf(i_0_42, plain, (strictorderP(X1)|~lt(esk20_1(X1),esk21_1(X1))|~ssList(X1))).
#
#cnf(i_0_41, plain, (strictorderP(X1)|~lt(esk21_1(X1),esk20_1(X1))|~ssList(X1))).
#
#cnf(i_0_50, plain, (totalorderedP(X1)|~leq(esk25_1(X1),esk26_1(X1))|~ssList(X1))).
##
#cnf(i_0_58, plain, (strictorderedP(X1)|~lt(esk30_1(X1),esk31_1(X1))|~ssList(X1))).
#
#cnf(i_0_164, plain, (hd(X1)=esk46_1(X1)|nil=X1|~ssList(X1))).
###
#cnf(i_0_166, plain, (tl(X1)=esk47_1(X1)|nil=X1|~ssList(X1))).
##
#cnf(i_0_93, plain, (hd(cons(X2,X1))=X2|~ssList(X1)|~ssItem(X2))).
#
#cnf(i_0_213, plain, (app(cons(X1,X2),cons(X3,X4))!=cons(X5,nil)|~leq(X3,X1)|~leq(X1,X3)|~ssList(cons(X5,nil))|~ssList(cons(X1,X2))|~ssList(X4)|~ssList(X2)|~ssItem(X3)|~ssItem(X1)|~ssItem(X5))).
#
#cnf(i_0_95, plain, (tl(cons(X2,X1))=X1|~ssList(X1)|~ssItem(X2))).
#
#cnf(i_0_24, plain, (leq(esk10_1(X1),esk11_1(X1))|cyclefreeP(X1)|~ssList(X1))).
#
#cnf(i_0_23, plain, (leq(esk11_1(X1),esk10_1(X1))|cyclefreeP(X1)|~ssList(X1))).
#
#cnf(i_0_260, plain, (app(cons(X1,X2),cons(X3,X4))!=cons(X5,nil)|~leq(X3,X1)|~leq(X1,X3)|~ssList(cons(X1,X2))|~ssList(X4)|~ssList(X2)|~ssItem(X3)|~ssItem(X1)|~ssItem(X5))).
#
#cnf(i_0_11, plain, (cons(esk5_1(X1),nil)=X1|~singletonP(X1)|~ssList(X1))).
#
#cnf(i_0_294, plain, (cyclefreeP(X1)|~singletonP(X1)|~ssList(X1)|~ssItem(esk5_1(X1)))).
#
#cnf(i_0_304, plain, (cyclefreeP(X1)|~singletonP(X1)|~ssList(X1))).
#
#cnf(i_0_275, plain, (X1!=nil|~singletonP(X1)|~ssList(X1)|~ssItem(esk5_1(X1)))).
#
#cnf(i_0_295, plain, (totalorderP(X1)|~singletonP(X1)|~ssList(X1)|~ssItem(esk5_1(X1)))).
#
#cnf(i_0_310, plain, (totalorderP(X1)|~singletonP(X1)|~ssList(X1))).
#
#cnf(i_0_296, plain, (strictorderP(X1)|~singletonP(X1)|~ssList(X1)|~ssItem(esk5_1(X1)))).
#
#cnf(i_0_309, plain, (X1!=nil|~singletonP(X1)|~ssList(X1))).
#
#cnf(i_0_311, plain, (strictorderP(X1)|~singletonP(X1)|~ssList(X1))).
##
#cnf(i_0_297, plain, (totalorderedP(X1)|~singletonP(X1)|~ssList(X1)|~ssItem(esk5_1(X1)))).
##
#cnf(i_0_341, plain, (totalorderedP(X1)|~singletonP(X1)|~ssList(X1))).
#
#cnf(i_0_298, plain, (strictorderedP(X1)|~singletonP(X1)|~ssList(X1)|~ssItem(esk5_1(X1)))).
#
#cnf(i_0_342, plain, (strictorderedP(X1)|~singletonP(X1)|~ssList(X1))).
#
#cnf(i_0_313, negated_conjecture, (nil!=esk51_0|~singletonP(esk51_0))).
#
#cnf(i_0_299, plain, (duplicatefreeP(X1)|~singletonP(X1)|~ssList(X1)|~ssItem(esk5_1(X1)))).
#
#cnf(i_0_343, plain, (duplicatefreeP(X1)|~singletonP(X1)|~ssList(X1))).
#
#cnf(i_0_300, plain, (equalelemsP(X1)|~singletonP(X1)|~ssList(X1)|~ssItem(esk5_1(X1)))).
#
#cnf(i_0_314, negated_conjecture, (nil!=esk48_0|~singletonP(esk48_0))).
#
#cnf(i_0_349, plain, (equalelemsP(X1)|~singletonP(X1)|~ssList(X1))).
#
#cnf(i_0_9, plain, (ssList(esk3_2(X1,X2))|~memberP(X1,X2)|~ssList(X1)|~ssItem(X2))).
#
#cnf(i_0_8, plain, (ssList(esk4_2(X1,X2))|~memberP(X1,X2)|~ssList(X1)|~ssItem(X2))).
##
#cnf(i_0_15, plain, (ssList(esk6_2(X1,X2))|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_18, plain, (ssList(esk7_2(X1,X2))|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_22, plain, (ssList(esk8_2(X1,X2))|~segmentP(X1,X2)|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_339, plain, (app(X1,X2)!=nil|~singletonP(app(X1,X2))|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_21, plain, (ssList(esk9_2(X1,X2))|~segmentP(X1,X2)|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_317, plain, (equalelemsP(X1)|esk43_1(X1)!=nil|~singletonP(esk43_1(X1))|~ssList(X1))).
#
#cnf(i_0_318, plain, (equalelemsP(X1)|esk42_1(X1)!=nil|~singletonP(esk42_1(X1))|~ssList(X1))).
#
#cnf(i_0_355, plain, (app(nil,X1)!=nil|~singletonP(X1)|~ssList(X1))).
#
#cnf(i_0_319, plain, (duplicatefreeP(X1)|esk39_1(X1)!=nil|~singletonP(esk39_1(X1))|~ssList(X1))).
#
#cnf(i_0_320, plain, (duplicatefreeP(X1)|esk38_1(X1)!=nil|~singletonP(esk38_1(X1))|~ssList(X1))).
#
#cnf(i_0_321, plain, (duplicatefreeP(X1)|esk37_1(X1)!=nil|~singletonP(esk37_1(X1))|~ssList(X1))).
##
#cnf(i_0_322, plain, (strictorderedP(X1)|esk34_1(X1)!=nil|~singletonP(esk34_1(X1))|~ssList(X1))).
#
#cnf(i_0_323, plain, (strictorderedP(X1)|esk33_1(X1)!=nil|~singletonP(esk33_1(X1))|~ssList(X1))).
#
#cnf(i_0_324, plain, (strictorderedP(X1)|esk32_1(X1)!=nil|~singletonP(esk32_1(X1))|~ssList(X1))).
##
#cnf(i_0_325, plain, (totalorderedP(X1)|esk29_1(X1)!=nil|~singletonP(esk29_1(X1))|~ssList(X1))).
#
#cnf(i_0_326, plain, (totalorderedP(X1)|esk28_1(X1)!=nil|~singletonP(esk28_1(X1))|~ssList(X1))).
#
#cnf(i_0_327, plain, (totalorderedP(X1)|esk27_1(X1)!=nil|~singletonP(esk27_1(X1))|~ssList(X1))).
#
#cnf(i_0_356, plain, (app(X1,nil)!=nil|~singletonP(X1)|~ssList(X1))).
#
#cnf(i_0_328, plain, (strictorderP(X1)|esk24_1(X1)!=nil|~singletonP(esk24_1(X1))|~ssList(X1))).
#
#cnf(i_0_329, plain, (strictorderP(X1)|esk23_1(X1)!=nil|~singletonP(esk23_1(X1))|~ssList(X1))).
#
#cnf(i_0_330, plain, (strictorderP(X1)|esk22_1(X1)!=nil|~singletonP(esk22_1(X1))|~ssList(X1))).
##
#cnf(i_0_331, plain, (totalorderP(X1)|esk19_1(X1)!=nil|~singletonP(esk19_1(X1))|~ssList(X1))).
#
#cnf(i_0_332, plain, (totalorderP(X1)|esk18_1(X1)!=nil|~singletonP(esk18_1(X1))|~ssList(X1))).
#
#cnf(i_0_333, plain, (totalorderP(X1)|esk17_1(X1)!=nil|~singletonP(esk17_1(X1))|~ssList(X1))).
##
#cnf(i_0_334, plain, (cyclefreeP(X1)|esk14_1(X1)!=nil|~singletonP(esk14_1(X1))|~ssList(X1))).
#
#cnf(i_0_335, plain, (cyclefreeP(X1)|esk13_1(X1)!=nil|~singletonP(esk13_1(X1))|~ssList(X1))).
#
#cnf(i_0_336, plain, (cyclefreeP(X1)|esk12_1(X1)!=nil|~singletonP(esk12_1(X1))|~ssList(X1))).
##
#cnf(i_0_149, plain, (totalorderedP(cons(X2,X1))|nil!=X1|~ssList(X1)|~ssItem(X2))).
##
#cnf(i_0_156, plain, (strictorderedP(cons(X2,X1))|nil!=X1|~ssList(X1)|~ssItem(X2))).
#
#cnf(i_0_350, plain, (esk3_2(X1,X2)!=nil|~singletonP(esk3_2(X1,X2))|~memberP(X1,X2)|~ssList(X1)|~ssItem(X2))).
##
#cnf(i_0_10, plain, (singletonP(X2)|cons(X1,nil)!=X2|~ssList(X2)|~ssItem(X1))).
#
#cnf(i_0_279, plain, (tl(X1)=nil|~singletonP(X1)|~ssList(X1)|~ssItem(esk5_1(X1)))).
#
#cnf(i_0_351, plain, (esk4_2(X1,X2)!=nil|~singletonP(esk4_2(X1,X2))|~memberP(X1,X2)|~ssList(X1)|~ssItem(X2))).
#
#cnf(i_0_397, plain, (tl(X1)=nil|~singletonP(X1)|~ssList(X1))).
#
#cnf(i_0_399, plain, (esk47_1(X1)=nil|nil=X1|~singletonP(X1)|~ssList(X1))).
##
#cnf(i_0_352, plain, (esk6_2(X1,X2)!=nil|~frontsegP(X1,X2)|~singletonP(esk6_2(X1,X2))|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_176, plain, (nil=X1|app(X2,X1)!=nil|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_175, plain, (nil=X1|app(X1,X2)!=nil|~ssList(X2)|~ssList(X1))).
##
#cnf(i_0_353, plain, (esk7_2(X1,X2)!=nil|~rearsegP(X1,X2)|~singletonP(esk7_2(X1,X2))|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_245, negated_conjecture, (X1=nil|~segmentP(esk48_0,X1)|~segmentP(esk51_0,X1)|~ssList(X1))).
#
#cnf(i_0_404, negated_conjecture, (nil=esk51_0|~segmentP(esk48_0,esk51_0))).
#
#cnf(i_0_364, plain, (singletonP(cons(X1,nil))|~ssList(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_354, plain, (esk8_2(X1,X2)!=nil|~segmentP(X1,X2)|~singletonP(esk8_2(X1,X2))|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_407, plain, (singletonP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_99, plain, (X1=X2|~leq(X2,X1)|~leq(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_180, plain, (X1=X2|~geq(X2,X1)|~geq(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_357, plain, (esk9_2(X1,X2)!=nil|~segmentP(X1,X2)|~singletonP(esk9_2(X1,X2))|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_411, plain, (X1=X2|~geq(X2,X1)|~leq(X2,X1)|~ssItem(X2)|~ssItem(X1))).
##
#cnf(i_0_117, plain, (X1=X2|~frontsegP(X2,X1)|~frontsegP(X1,X2)|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_348, plain, (app(cons(X1,X2),cons(X3,X4))!=X5|X1!=X3|~singletonP(X5)|~ssList(cons(X1,X2))|~ssList(X4)|~ssList(X2)|~ssList(X5)|~ssItem(X3)|~ssItem(X1))).
###
#cnf(i_0_127, plain, (X1=X2|~rearsegP(X2,X1)|~rearsegP(X1,X2)|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_347, plain, (app(app(X1,cons(X2,X3)),cons(X4,X5))!=X6|X2!=X4|~singletonP(X6)|~ssList(X5)|~ssList(X3)|~ssList(X1)|~ssList(X6)|~ssItem(X4)|~ssItem(X2))).
###
#cnf(i_0_134, plain, (X1=X2|~segmentP(X2,X1)|~segmentP(X1,X2)|~ssList(X2)|~ssList(X1))).
####
#cnf(i_0_278, plain, (esk5_1(X1)=hd(X1)|~singletonP(X1)|~ssList(X1)|~ssItem(esk5_1(X1)))).
#
#cnf(i_0_288, plain, (app(cons(X1,X2),X3)!=X4|X1!=esk5_1(X3)|~duplicatefreeP(X4)|~singletonP(X3)|~ssList(cons(X1,X2))|~ssList(X2)|~ssList(X4)|~ssList(X3)|~ssItem(esk5_1(X3))|~ssItem(X1))).
#
#cnf(i_0_495, plain, (esk5_1(X1)=hd(X1)|~singletonP(X1)|~ssList(X1))).
#
#cnf(i_0_501, plain, (ssItem(hd(X1))|~singletonP(X1)|~ssList(X1))).
###
#cnf(i_0_502, plain, (cons(hd(X1),nil)=X1|~singletonP(X1)|~ssList(X1))).
###########
#cnf(i_0_450, plain, (app(cons(X1,X2),X3)!=X4|X1!=esk5_1(X3)|~singletonP(X4)|~singletonP(X3)|~ssList(cons(X1,X2))|~ssList(X2)|~ssList(X4)|~ssList(X3)|~ssItem(esk5_1(X3))|~ssItem(X1))).
#####
#cnf(i_0_185, plain, (X1=X2|lt(X1,X2)|~leq(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_596, plain, (X1=X2|~lt(X2,X1)|~leq(X1,X2)|~ssItem(X1)|~ssItem(X2))).
#####
#cnf(i_0_315, plain, (nil=X1|esk47_1(X1)!=nil|~singletonP(esk47_1(X1))|~ssList(X1))).
#
#cnf(i_0_535, plain, (app(cons(X1,X2),X3)!=X4|X1!=hd(X3)|~duplicatefreeP(X4)|~singletonP(X3)|~ssList(cons(X1,X2))|~ssList(X2)|~ssList(X4)|~ssList(X3)|~ssItem(hd(X3))|~ssItem(X1))).
#
#cnf(i_0_316, plain, (nil=X1|esk44_1(X1)!=nil|~singletonP(esk44_1(X1))|~ssList(X1))).
#
#cnf(i_0_337, plain, (nil=X1|tl(X1)!=nil|~singletonP(tl(X1))|~ssList(X1))).
#
#cnf(i_0_607, plain, (nil=X1|tl(X1)!=nil|~singletonP(esk47_1(X1))|~ssList(X1))).
##
#cnf(i_0_169, plain, (cons(hd(X1),tl(X1))=X1|nil=X1|~ssList(X1))).
#
#cnf(i_0_88, plain, (cons(esk45_1(X1),esk44_1(X1))=X1|nil=X1|~ssList(X1))).
#
#cnf(i_0_649, plain, (cons(hd(X1),esk47_1(X1))=X1|nil=X1|~ssList(X1))).
#
#cnf(i_0_539, plain, (app(cons(X1,X2),X3)!=X4|X1!=hd(X3)|~singletonP(X4)|~singletonP(X3)|~ssList(cons(X1,X2))|~ssList(X2)|~ssList(X4)|~ssList(X3)|~ssItem(hd(X3))|~ssItem(X1))).
#
#cnf(i_0_652, plain, (cons(esk46_1(X1),tl(X1))=X1|nil=X1|~ssList(X1))).
#
#cnf(i_0_734, plain, (cons(esk46_1(X1),esk47_1(X1))=X1|nil=X1|~ssList(X1))).
#####
#cnf(i_0_172, plain, (app(cons(X2,nil),X1)=cons(X2,X1)|~ssList(X1)|~ssItem(X2))).
##
#cnf(i_0_151, plain, (nil=X1|totalorderedP(X1)|~totalorderedP(cons(X2,X1))|~ssList(X1)|~ssItem(X2))).
#
#cnf(i_0_158, plain, (nil=X1|strictorderedP(X1)|~strictorderedP(cons(X2,X1))|~ssList(X1)|~ssItem(X2))).
#
#cnf(i_0_112, plain, (memberP(cons(X2,X3),X1)|X1!=X2|~ssList(X3)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_867, plain, (app(cons(X1,nil),X2)!=nil|~singletonP(cons(X1,X2))|~ssList(cons(X1,nil))|~ssList(X2)|~ssItem(X1))).
#
#cnf(i_0_13, plain, (frontsegP(X3,X2)|app(X2,X1)!=X3|~ssList(X3)|~ssList(X2)|~ssList(X1))).
####
#cnf(i_0_908, plain, (frontsegP(app(X1,X2),X1)|~ssList(app(X1,X2))|~ssList(X1)|~ssList(X2))).
####
#cnf(i_0_16, plain, (rearsegP(X3,X2)|app(X1,X2)!=X3|~ssList(X3)|~ssList(X2)|~ssList(X1))).
###
#cnf(i_0_904, plain, (app(X1,X2)!=nil|~singletonP(cons(esk5_1(X1),X2))|~singletonP(X1)|~ssList(cons(esk5_1(X1),nil))|~ssList(X2)|~ssList(X1)|~ssItem(esk5_1(X1)))).
#
#cnf(i_0_919, plain, (rearsegP(app(X1,X2),X2)|~ssList(app(X1,X2))|~ssList(X2)|~ssList(X1))).
###
#cnf(i_0_926, plain, (app(X1,X2)!=nil|~singletonP(cons(esk5_1(X1),X2))|~singletonP(X1)|~ssList(X1)|~ssList(X2)|~ssItem(esk5_1(X1)))).
#
#cnf(i_0_405, negated_conjecture, (esk52_0=nil|nil=esk51_0|~segmentP(esk48_0,esk52_0)|~ssList(esk52_0))).
#
#cnf(i_0_937, negated_conjecture, (esk52_0=nil|nil=esk51_0|~ssList(esk52_0))).
#
#cnf(i_0_938, negated_conjecture, (esk52_0=nil|nil=esk51_0)).
##
#cnf(i_0_943, negated_conjecture, (nil=esk51_0)).
#
#cnf(i_0_163, plain, (equalelemsP(esk51_0))).
#
#cnf(i_0_161, plain, (duplicatefreeP(esk51_0))).
#
#cnf(i_0_313, negated_conjecture, (~singletonP(esk51_0))).
#
#cnf(i_0_154, plain, (strictorderedP(esk51_0))).
#
#cnf(i_0_147, plain, (totalorderedP(esk51_0))).
#
#cnf(i_0_145, plain, (strictorderP(esk51_0))).
##
#cnf(i_0_143, plain, (totalorderP(esk51_0))).
#
#cnf(i_0_141, plain, (cyclefreeP(esk51_0))).
#
#cnf(i_0_196, negated_conjecture, (esk51_0=esk48_0)).
#
#cnf(i_0_313, negated_conjecture, (~singletonP(esk48_0))).
#
#cnf(i_0_141, plain, (cyclefreeP(esk48_0))).
#
#cnf(i_0_143, plain, (totalorderP(esk48_0))).
#
#cnf(i_0_145, plain, (strictorderP(esk48_0))).
#
# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(co1, conjecture, ![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(X2!=X4|X1!=X3|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5))|(nil!=X3&nil=X4)|(nil=X2&nil=X1)|(![X6]:(ssList(X6)=>(~(neq(X6,nil))|~(segmentP(X4,X6))|~(segmentP(X3,X6))))&neq(X4,nil))))))), file('./p/SWC069+1.tptp', co1)).
fof(ax15, axiom, ![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>(neq(X1,X2)<=>X1!=X2))), file('./p/SWC069+1.tptp', ax15)).
fof(ax17, axiom, ssList(nil), file('./p/SWC069+1.tptp', ax17)).
fof(c_0_3, negated_conjecture, ~(![X1]:(ssList(X1)=>![X2]:(ssList(X2)=>![X3]:(ssList(X3)=>![X4]:(ssList(X4)=>(X2!=X4|X1!=X3|?[X5]:(((ssList(X5)&neq(X5,nil))&segmentP(X2,X5))&segmentP(X1,X5))|(nil!=X3&nil=X4)|(nil=X2&nil=X1)|(![X6]:(ssList(X6)=>(~(neq(X6,nil))|~(segmentP(X4,X6))|~(segmentP(X3,X6))))&neq(X4,nil)))))))), inference(assume_negation,[status(cth)],[co1])).
fof(c_0_4, negated_conjecture, ![X255]:(ssList(esk48_0)&(ssList(esk49_0)&(ssList(esk50_0)&(ssList(esk51_0)&(((((esk49_0=esk51_0&esk48_0=esk50_0)&(~ssList(X255)|~neq(X255,nil)|~segmentP(esk49_0,X255)|~segmentP(esk48_0,X255)))&(nil=esk50_0|nil!=esk51_0))&(nil!=esk49_0|nil!=esk48_0))&((ssList(esk52_0)|~neq(esk51_0,nil))&(((neq(esk52_0,nil)|~neq(esk51_0,nil))&(segmentP(esk51_0,esk52_0)|~neq(esk51_0,nil)))&(segmentP(esk50_0,esk52_0)|~neq(esk51_0,nil))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_3])])])])])])).
cnf(c_0_5, negated_conjecture, (~ssList(X1)|~neq(X1,nil)|~segmentP(esk49_0,X1)|~segmentP(esk48_0,X1)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_6, negated_conjecture, (esk49_0=esk51_0), inference(split_conjunct,[status(thm)],[c_0_4])).
fof(c_0_7, plain, ![X109, X110]:((~neq(X109,X110)|X109!=X110|~ssList(X110)|~ssList(X109))&(X109=X110|neq(X109,X110)|~ssList(X110)|~ssList(X109))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[ax15])])])])).
cnf(c_0_8, negated_conjecture, (~segmentP(esk48_0,X1)|~segmentP(esk51_0,X1)|~ssList(X1)|~neq(X1,nil)), inference(rw,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_9, plain, (X1=X2|neq(X1,X2)|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, plain, (ssList(nil)), inference(split_conjunct,[status(thm)],[ax17])).
cnf(c_0_11, negated_conjecture, (segmentP(esk51_0,esk52_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_12, negated_conjecture, (ssList(esk51_0)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_13, negated_conjecture, (segmentP(esk50_0,esk52_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_14, negated_conjecture, (esk48_0=esk50_0), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_15, plain, (~neq(X1,X2)|X1!=X2|~ssList(X2)|~ssList(X1)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_16, negated_conjecture, (neq(esk52_0,nil)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_17, negated_conjecture, (X1=nil|~segmentP(esk48_0,X1)|~segmentP(esk51_0,X1)|~ssList(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_8, c_0_9]), c_0_10])])).
cnf(c_0_18, negated_conjecture, (nil=esk51_0|segmentP(esk51_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_11, c_0_9]), c_0_10]), c_0_12])])).
cnf(c_0_19, negated_conjecture, (segmentP(esk48_0,esk52_0)|~neq(esk51_0,nil)), inference(rw,[status(thm)],[c_0_13, c_0_14])).
cnf(c_0_20, negated_conjecture, (esk52_0!=nil|~ssList(esk52_0)|~neq(esk51_0,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_15, c_0_16]), c_0_10])])).
cnf(c_0_21, negated_conjecture, (ssList(esk52_0)|~neq(esk51_0,nil)), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_22, negated_conjecture, (esk52_0=nil|nil=esk51_0|~segmentP(esk48_0,esk52_0)|~ssList(esk52_0)), inference(pm,[status(thm)],[c_0_17, c_0_18])).
cnf(c_0_23, negated_conjecture, (nil=esk51_0|segmentP(esk48_0,esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_19, c_0_9]), c_0_10]), c_0_12])])).
cnf(c_0_24, negated_conjecture, (nil=esk51_0|esk52_0!=nil|~ssList(esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_20, c_0_9]), c_0_10]), c_0_12])])).
cnf(c_0_25, negated_conjecture, (nil=esk51_0|ssList(esk52_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_21, c_0_9]), c_0_10]), c_0_12])])).
cnf(c_0_26, negated_conjecture, (esk52_0=nil|nil=esk51_0|~ssList(esk52_0)), inference(pm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_27, negated_conjecture, (nil=esk50_0|nil!=esk51_0), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_28, negated_conjecture, (nil=esk51_0|esk52_0!=nil), inference(pm,[status(thm)],[c_0_24, c_0_25])).
cnf(c_0_29, negated_conjecture, (esk52_0=nil|nil=esk51_0), inference(pm,[status(thm)],[c_0_26, c_0_25])).
cnf(c_0_30, negated_conjecture, (nil!=esk49_0|nil!=esk48_0), inference(split_conjunct,[status(thm)],[c_0_4])).
cnf(c_0_31, negated_conjecture, (nil=esk48_0|nil!=esk51_0), inference(rw,[status(thm)],[c_0_27, c_0_14])).
cnf(c_0_32, negated_conjecture, (nil=esk51_0), inference(pm,[status(thm)],[c_0_28, c_0_29])).
cnf(c_0_33, negated_conjecture, (nil!=esk48_0|nil!=esk51_0), inference(rw,[status(thm)],[c_0_30, c_0_6])).
cnf(c_0_34, negated_conjecture, (esk51_0=esk48_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_32]), c_0_32])])).
cnf(c_0_35, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_32]), c_0_32])]), c_0_34])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 36
# Proof object clause steps            : 30
# Proof object formula steps           : 6
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 13
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 11
# Proof object simplifying inferences  : 28
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 96
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 203
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 201
# Processed clauses                    : 336
# ...of these trivial                  : 0
# ...subsumed                          : 67
# ...remaining for further processing  : 268
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 18
# Backward-rewritten                   : 129
# Generated clauses                    : 740
# ...of the previous two non-trivial   : 816
# Contextual simplify-reflections      : 0
# Paramodulations                      : 727
# Factorizations                       : 0
# Equation resolutions                 : 13
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 121
#    Positive orientable unit clauses  : 8
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 111
# Current number of unprocessed clauses: 681
# ...number of literals in the above   : 5550
# Current number of archived formulas  : 0
# Current number of archived clauses   : 147
# Clause-clause subsumption calls (NU) : 2396
# Rec. Clause-clause subsumption calls : 951
# Non-unit clause-clause subsumptions  : 77
# Unit Clause-clause subsumption calls : 588
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 27285
