# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_199, plain, (equalelemsP(nil))).
#
#cnf(i_0_200, plain, (duplicatefreeP(nil))).
#
#cnf(i_0_201, plain, (strictorderedP(nil))).
#
#cnf(i_0_209, plain, (~singletonP(nil))).
#
#cnf(i_0_202, plain, (totalorderedP(nil))).
#
#cnf(i_0_203, plain, (strictorderP(nil))).
#
#cnf(i_0_204, plain, (totalorderP(nil))).
#
#cnf(i_0_252, plain, (skac2!=skac3)).
#
#cnf(i_0_205, plain, (cyclefreeP(nil))).
#
#cnf(i_0_206, plain, (ssList(nil))).
#
#cnf(i_0_384, negated_conjecture, (ssList(sk1))).
#
#cnf(i_0_261, plain, (~lt(X1,X1)|~ssItem(X1))).
#
#cnf(i_0_385, negated_conjecture, (ssList(sk2))).
#
#cnf(i_0_386, negated_conjecture, (ssList(sk3))).
#
#cnf(i_0_387, negated_conjecture, (ssList(sk4))).
#
#cnf(i_0_269, plain, (~memberP(nil,X1)|~ssItem(X1))).
#
#cnf(i_0_207, plain, (ssItem(skac3))).
#
#cnf(i_0_208, plain, (ssItem(skac2))).
#
#cnf(i_0_389, negated_conjecture, (sk3=sk1)).
#
#cnf(i_0_297, plain, (cons(X1,X2)!=X2|~ssItem(X1)|~ssList(X2))).
#
#cnf(i_0_388, negated_conjecture, (sk2=sk4)).
#
#cnf(i_0_211, plain, (ssList(skaf82(X1)))).
#
#cnf(i_0_212, plain, (ssList(skaf81(X1)))).
#
#cnf(i_0_296, plain, (cons(X1,X2)!=nil|~ssItem(X1)|~ssList(X2))).
#
#cnf(i_0_213, plain, (ssList(skaf80(X1)))).
#
#cnf(i_0_216, plain, (ssList(skaf77(X1)))).
#
#cnf(i_0_217, plain, (ssList(skaf76(X1)))).
#
#cnf(i_0_313, plain, (X1!=X2|~neq(X1,X2)|~ssList(X2)|~ssList(X1))).
#
#cnf(i_0_218, plain, (ssList(skaf75(X1)))).
#
#cnf(i_0_220, plain, (ssList(skaf73(X1)))).
#
#cnf(i_0_221, plain, (ssList(skaf72(X1)))).
#
#cnf(i_0_309, plain, (X1!=X2|~lt(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_222, plain, (ssList(skaf71(X1)))).
#
#cnf(i_0_225, plain, (ssList(skaf68(X1)))).
#
#cnf(i_0_226, plain, (ssList(skaf67(X1)))).
#
#cnf(i_0_315, plain, (X1!=X2|~neq(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_227, plain, (ssList(skaf66(X1)))).
#
#cnf(i_0_230, plain, (ssList(skaf63(X1)))).
#
#cnf(i_0_231, plain, (ssList(skaf62(X1)))).
#
#cnf(i_0_312, plain, (~lt(X2,X1)|~lt(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_232, plain, (ssList(skaf61(X1)))).
#
#cnf(i_0_235, plain, (ssList(skaf58(X1)))).
#
#cnf(i_0_236, plain, (ssList(skaf57(X1)))).
#
#cnf(i_0_308, plain, (~gt(X2,X1)|~gt(X1,X2)|~ssItem(X2)|~ssItem(X1))).
#
#cnf(i_0_237, plain, (ssList(skaf56(X1)))).
#
#cnf(i_0_240, plain, (ssList(skaf53(X1)))).
#
#cnf(i_0_241, plain, (ssList(skaf52(X1)))).
#
#cnf(i_0_391, negated_conjecture, (~neq(X1,nil)|~segmentP(sk1,X1)|~segmentP(sk4,X1)|~ssList(X1))).
#
#cnf(i_0_242, plain, (ssList(skaf51(X1)))).
#
#cnf(i_0_210, plain, (ssItem(skaf83(X1)))).
#
#cnf(i_0_214, plain, (ssItem(skaf79(X1)))).
#
#cnf(i_0_377, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=X5|~ssItem(X2)|~ssList(X5)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(X5))).
#
#cnf(i_0_215, plain, (ssItem(skaf78(X1)))).
#
#cnf(i_0_219, plain, (ssItem(skaf74(X1)))).
#
#cnf(i_0_223, plain, (ssItem(skaf70(X1)))).
#
#cnf(i_0_398, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf51(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf51(X5)))).
#
#cnf(i_0_224, plain, (ssItem(skaf69(X1)))).
#
#cnf(i_0_228, plain, (ssItem(skaf65(X1)))).
#
#cnf(i_0_229, plain, (ssItem(skaf64(X1)))).
#
#cnf(i_0_399, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf52(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf52(X5)))).
#
#cnf(i_0_233, plain, (ssItem(skaf60(X1)))).
#
#cnf(i_0_234, plain, (ssItem(skaf59(X1)))).
#
#cnf(i_0_238, plain, (ssItem(skaf55(X1)))).
#
#cnf(i_0_400, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf53(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf53(X5)))).
#
#cnf(i_0_239, plain, (ssItem(skaf54(X1)))).
#
#cnf(i_0_243, plain, (ssItem(skaf50(X1)))).
#
#cnf(i_0_244, plain, (ssItem(skaf49(X1)))).
#
#cnf(i_0_401, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf56(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf56(X5)))).
#
#cnf(i_0_245, plain, (ssItem(skaf44(X1)))).
#
#cnf(i_0_390, negated_conjecture, (neq(sk4,nil))).
#
#cnf(i_0_246, plain, (ssList(skaf48(X1,X2)))).
#
#cnf(i_0_419, negated_conjecture, (nil!=sk4)).
#
#cnf(i_0_247, plain, (ssList(skaf47(X1,X2)))).
#
#cnf(i_0_248, plain, (ssList(skaf46(X1,X2)))).
#
#cnf(i_0_249, plain, (ssList(skaf45(X1,X2)))).
##
#cnf(i_0_250, plain, (ssList(skaf43(X1,X2)))).
#
#cnf(i_0_251, plain, (ssList(skaf42(X1,X2)))).
#
#cnf(i_0_253, plain, (geq(X1,X1)|~ssItem(X1))).
#
#cnf(i_0_402, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf57(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf57(X5)))).
#
#cnf(i_0_255, plain, (segmentP(X1,X1)|~ssList(X1))).
#
#cnf(i_0_257, plain, (rearsegP(X1,X1)|~ssList(X1))).
#
#cnf(i_0_259, plain, (frontsegP(X1,X1)|~ssList(X1))).
#
#cnf(i_0_427, negated_conjecture, (~segmentP(sk1,sk4))).
#
#cnf(i_0_260, plain, (leq(X1,X1)|~ssItem(X1))).
#
#cnf(i_0_254, plain, (segmentP(X1,nil)|~ssList(X1))).
#
#cnf(i_0_256, plain, (rearsegP(X1,nil)|~ssList(X1))).
#
#cnf(i_0_428, negated_conjecture, (~neq(nil,nil)|~segmentP(sk1,nil))).
#
#cnf(i_0_258, plain, (frontsegP(X1,nil)|~ssList(X1))).
#
#cnf(i_0_270, plain, (ssItem(X2)|duplicatefreeP(X1)|~ssList(X1))).
#
#cnf(i_0_436, plain, (ssItem(X1)|duplicatefreeP(skaf51(X2)))).
#
#cnf(i_0_429, negated_conjecture, (~neq(nil,nil))).
#
#cnf(i_0_437, plain, (ssItem(X1)|duplicatefreeP(skaf52(X2)))).
#
#cnf(i_0_438, plain, (ssItem(X1)|duplicatefreeP(skaf53(X2)))).
#
#cnf(i_0_439, plain, (ssItem(X1)|duplicatefreeP(skaf56(X2)))).
#
#cnf(i_0_403, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf58(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf58(X5)))).
#
#cnf(i_0_440, plain, (ssItem(X1)|duplicatefreeP(skaf57(X2)))).
#
#cnf(i_0_441, plain, (ssItem(X1)|duplicatefreeP(skaf58(X2)))).
#
#cnf(i_0_442, plain, (ssItem(X1)|duplicatefreeP(skaf61(X2)))).
#
#cnf(i_0_404, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf61(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf61(X5)))).
#
#cnf(i_0_443, plain, (ssItem(X1)|duplicatefreeP(skaf62(X2)))).
#
#cnf(i_0_444, plain, (ssItem(X1)|duplicatefreeP(skaf63(X2)))).
#
#cnf(i_0_445, plain, (ssItem(X1)|duplicatefreeP(skaf66(X2)))).
#
#cnf(i_0_405, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf62(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf62(X5)))).
#
#cnf(i_0_446, plain, (ssItem(X1)|duplicatefreeP(skaf67(X2)))).
#
#cnf(i_0_447, plain, (ssItem(X1)|duplicatefreeP(skaf68(X2)))).
#
#cnf(i_0_448, plain, (ssItem(X1)|duplicatefreeP(skaf71(X2)))).
#
#cnf(i_0_406, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf63(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf63(X5)))).
#
#cnf(i_0_449, plain, (ssItem(X1)|duplicatefreeP(skaf72(X2)))).
#
#cnf(i_0_450, plain, (ssItem(X1)|duplicatefreeP(skaf73(X2)))).
#
#cnf(i_0_451, plain, (ssItem(X1)|duplicatefreeP(skaf75(X2)))).
#
#cnf(i_0_407, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf66(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf66(X5)))).
#
#cnf(i_0_452, plain, (ssItem(X1)|duplicatefreeP(skaf76(X2)))).
#
#cnf(i_0_453, plain, (ssItem(X1)|duplicatefreeP(skaf77(X2)))).
#
#cnf(i_0_454, plain, (ssItem(X1)|duplicatefreeP(skaf80(X2)))).
#
#cnf(i_0_408, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf67(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf67(X5)))).
#
#cnf(i_0_455, plain, (ssItem(X1)|duplicatefreeP(skaf81(X2)))).
#
#cnf(i_0_456, plain, (ssItem(X1)|duplicatefreeP(skaf82(X2)))).
##
#cnf(i_0_409, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf68(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf68(X5)))).
#
#cnf(i_0_430, plain, (ssItem(X1)|duplicatefreeP(skaf42(X2,X3)))).
#
#cnf(i_0_431, plain, (ssItem(X1)|duplicatefreeP(skaf43(X2,X3)))).
#
#cnf(i_0_432, plain, (ssItem(X1)|duplicatefreeP(skaf45(X2,X3)))).
#
#cnf(i_0_410, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf71(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf71(X5)))).
#
#cnf(i_0_433, plain, (ssItem(X1)|duplicatefreeP(skaf46(X2,X3)))).
#
#cnf(i_0_434, plain, (ssItem(X1)|duplicatefreeP(skaf47(X2,X3)))).
#
#cnf(i_0_435, plain, (ssItem(X1)|duplicatefreeP(skaf48(X2,X3)))).
#
#cnf(i_0_411, plain, (app(app(X1,cons(X2,X3)),cons(X2,X4))!=skaf72(X5)|~ssItem(X2)|~ssList(X4)|~ssList(X3)|~ssList(X1)|~duplicatefreeP(skaf72(X5)))).
#
#cnf(i_0_393, negated_conjecture, (ssList(sk5))).
#
#cnf(i_0_271, plain, (app(X1,nil)=X1|~ssList(X1))).
#
#cnf(i_0_272, plain, (app(nil,X1)=X1|~ssList(X1))).
#
#cnf(i_0_471, plain, (app(cons(X1,X2),cons(X1,X3))!=X4|~ssItem(X1)|~ssList(cons(X1,X2))|~ssList(X4)|~ssList(X3)|~ssList(X2)|~duplicatefreeP(X4))).
#
#cnf(i_0_262, plain, (equalelemsP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_263, plain, (duplicatefreeP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_264, plain, (strictorderedP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_472, plain, (app(cons(X1,X2),cons(X1,X3))!=skaf51(X4)|~ssItem(X1)|~ssList(cons(X1,X2))|~ssList(X3)|~ssList(X2)|~duplicatefreeP(skaf51(X4)))).
#
#cnf(i_0_265, plain, (totalorderedP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_266, plain, (strictorderP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_267, plain, (totalorderP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_473, plain, (app(cons(X1,X2),cons(X1,X3))!=skaf52(X4)|~ssItem(X1)|~ssList(cons(X1,X2))|~ssList(X3)|~ssList(X2)|~duplicatefreeP(skaf52(X4)))).
#
#cnf(i_0_268, plain, (cyclefreeP(cons(X1,nil))|~ssItem(X1))).
#
#cnf(i_0_396, negated_conjecture, (segmentP(sk1,sk5))).
#
#cnf(i_0_395, negated_conjecture, (segmentP(sk4,sk5))).
#
#cnf(i_0_516, negated_conjecture, (~neq(sk5,nil))).
#
# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(co1_7, negated_conjecture, (neq(sk2,nil)), file('./p/SWC082-1.tptp', co1_7)).
cnf(co1_5, negated_conjecture, (sk2=sk4), file('./p/SWC082-1.tptp', co1_5)).
cnf(co1_8, negated_conjecture, (~ssList(X1)|~neq(X1,nil)|~segmentP(sk2,X1)|~segmentP(sk1,X1)), file('./p/SWC082-1.tptp', co1_8)).
cnf(co1_12, negated_conjecture, (segmentP(sk4,sk5)|~neq(sk4,nil)), file('./p/SWC082-1.tptp', co1_12)).
cnf(co1_13, negated_conjecture, (segmentP(sk3,sk5)|~neq(sk4,nil)), file('./p/SWC082-1.tptp', co1_13)).
cnf(co1_6, negated_conjecture, (sk1=sk3), file('./p/SWC082-1.tptp', co1_6)).
cnf(co1_10, negated_conjecture, (ssList(sk5)|~neq(sk4,nil)), file('./p/SWC082-1.tptp', co1_10)).
cnf(co1_11, negated_conjecture, (neq(sk5,nil)|~neq(sk4,nil)), file('./p/SWC082-1.tptp', co1_11)).
cnf(c_0_8, negated_conjecture, (neq(sk2,nil)), co1_7).
cnf(c_0_9, negated_conjecture, (sk2=sk4), co1_5).
cnf(c_0_10, negated_conjecture, (~ssList(X1)|~neq(X1,nil)|~segmentP(sk2,X1)|~segmentP(sk1,X1)), co1_8).
cnf(c_0_11, negated_conjecture, (segmentP(sk4,sk5)|~neq(sk4,nil)), co1_12).
cnf(c_0_12, negated_conjecture, (neq(sk4,nil)), inference(rw,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_13, negated_conjecture, (segmentP(sk3,sk5)|~neq(sk4,nil)), co1_13).
cnf(c_0_14, negated_conjecture, (sk1=sk3), co1_6).
cnf(c_0_15, negated_conjecture, (ssList(sk5)|~neq(sk4,nil)), co1_10).
cnf(c_0_16, negated_conjecture, (~neq(X1,nil)|~segmentP(sk1,X1)|~segmentP(sk4,X1)|~ssList(X1)), inference(rw,[status(thm)],[c_0_10, c_0_9])).
cnf(c_0_17, negated_conjecture, (segmentP(sk4,sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_12])])).
cnf(c_0_18, negated_conjecture, (segmentP(sk1,sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_13, c_0_14]), c_0_12])])).
cnf(c_0_19, negated_conjecture, (ssList(sk5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_15, c_0_12])])).
cnf(c_0_20, negated_conjecture, (neq(sk5,nil)|~neq(sk4,nil)), co1_11).
cnf(c_0_21, negated_conjecture, (~neq(sk5,nil)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_16, c_0_17]), c_0_18]), c_0_19])])).
cnf(c_0_22, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_20, c_0_12])]), c_0_21]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 23
# Proof object clause steps            : 15
# Proof object formula steps           : 8
# Proof object conjectures             : 23
# Proof object clause conjectures      : 15
# Proof object formula conjectures     : 8
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 1
# Proof object simplifying inferences  : 15
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 198
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 198
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 198
# Processed clauses                    : 145
# ...of these trivial                  : 0
# ...subsumed                          : 2
# ...remaining for further processing  : 142
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 120
# ...of the previous two non-trivial   : 120
# Contextual simplify-reflections      : 0
# Paramodulations                      : 118
# Factorizations                       : 0
# Equation resolutions                 : 2
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 140
#    Positive orientable unit clauses  : 60
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 74
# Current number of unprocessed clauses: 173
# ...number of literals in the above   : 896
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 626
# Rec. Clause-clause subsumption calls : 358
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 26
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 8525
