# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_134, plain, (true)).
#
#cnf(i_0_139, negated_conjecture, (epred2_0)).
#
#cnf(i_0_4, plain, (leq(X1,X1))).
#
#cnf(i_0_3, plain, (~gt(X1,X1))).
#
#cnf(i_0_146, plain, (gt(n0,tptp_minus_1))).
#
#cnf(i_0_153, plain, (gt(n1,n0))).
#
#cnf(i_0_147, plain, (gt(n1,tptp_minus_1))).
#
#cnf(i_0_135, plain, (def!=use)).
#
#cnf(i_0_154, plain, (gt(n2,n0))).
#
#cnf(i_0_148, plain, (gt(n2,tptp_minus_1))).
#
#cnf(i_0_159, plain, (gt(n2,n1))).
#
#cnf(i_0_136, negated_conjecture, (a_select2(muold_init,esk24_0)!=init)).
#
#cnf(i_0_155, plain, (gt(n3,n0))).
#
#cnf(i_0_149, plain, (gt(n3,tptp_minus_1))).
#
#cnf(i_0_160, plain, (gt(n3,n1))).
#
#cnf(i_0_164, plain, (gt(n3,n2))).
#
#cnf(i_0_150, plain, (gt(n4,n0))).
#
#cnf(i_0_143, plain, (gt(n4,tptp_minus_1))).
#
#cnf(i_0_156, plain, (gt(n4,n1))).
#
#cnf(i_0_161, plain, (gt(n4,n2))).
#
#cnf(i_0_165, plain, (gt(n4,n3))).
#
#cnf(i_0_151, plain, (gt(n5,n0))).
#
#cnf(i_0_144, plain, (gt(n5,tptp_minus_1))).
#
#cnf(i_0_157, plain, (gt(n5,n1))).
#
#cnf(i_0_162, plain, (gt(n5,n2))).
#
#cnf(i_0_166, plain, (gt(n5,n3))).
#
#cnf(i_0_140, plain, (gt(n5,n4))).
#
#cnf(i_0_152, plain, (gt(n135299,n0))).
#
#cnf(i_0_145, plain, (gt(n135299,tptp_minus_1))).
#
#cnf(i_0_158, plain, (gt(n135299,n1))).
#
#cnf(i_0_163, plain, (gt(n135299,n2))).
#
#cnf(i_0_167, plain, (gt(n135299,n3))).
#
#cnf(i_0_141, plain, (gt(n135299,n4))).
#
#cnf(i_0_142, plain, (gt(n135299,n5))).
#
#cnf(i_0_138, negated_conjecture, (leq(n0,esk24_0))).
#
#cnf(i_0_137, negated_conjecture, (leq(esk24_0,n4))).
#
#cnf(i_0_7, plain, (gt(X2,X1)|~lt(X1,X2))).
#
#cnf(i_0_10, plain, (leq(X2,X1)|~gt(X1,X2))).
#
#cnf(i_0_212, plain, (leq(n5,n135299))).
#
#cnf(i_0_213, plain, (leq(n4,n135299))).
#
#cnf(i_0_214, plain, (leq(n3,n135299))).
#
#cnf(i_0_215, plain, (leq(n2,n135299))).
#
#cnf(i_0_216, plain, (leq(n1,n135299))).
#
#cnf(i_0_217, plain, (leq(tptp_minus_1,n135299))).
#
#cnf(i_0_218, plain, (leq(n0,n135299))).
#
#cnf(i_0_219, plain, (leq(n4,n5))).
#
#cnf(i_0_220, plain, (leq(n3,n5))).
#
#cnf(i_0_221, plain, (leq(n2,n5))).
#
#cnf(i_0_222, plain, (leq(n1,n5))).
#
#cnf(i_0_223, plain, (leq(tptp_minus_1,n5))).
#
#cnf(i_0_224, plain, (leq(n0,n5))).
#
#cnf(i_0_225, plain, (leq(n3,n4))).
#
#cnf(i_0_226, plain, (leq(n2,n4))).
#
#cnf(i_0_227, plain, (leq(n1,n4))).
#
#cnf(i_0_228, plain, (leq(tptp_minus_1,n4))).
#
#cnf(i_0_229, plain, (leq(n0,n4))).
#
#cnf(i_0_230, plain, (leq(n2,n3))).
#
#cnf(i_0_231, plain, (leq(n1,n3))).
#
#cnf(i_0_232, plain, (leq(tptp_minus_1,n3))).
#
#cnf(i_0_233, plain, (leq(n0,n3))).
#
#cnf(i_0_234, plain, (leq(n1,n2))).
#
#cnf(i_0_235, plain, (leq(tptp_minus_1,n2))).
#
#cnf(i_0_236, plain, (leq(n0,n2))).
#
#cnf(i_0_237, plain, (leq(tptp_minus_1,n1))).
#
#cnf(i_0_238, plain, (leq(n0,n1))).
#
#cnf(i_0_239, plain, (leq(tptp_minus_1,n0))).
#
#cnf(i_0_9, plain, (leq(X2,X1)|~geq(X1,X2))).
#
#cnf(i_0_6, plain, (lt(X2,X1)|~gt(X1,X2))).
#
#cnf(i_0_8, plain, (geq(X2,X1)|~leq(X1,X2))).
#
#cnf(i_0_209, plain, (gt(n1,loopcounter))).
#
#cnf(i_0_242, plain, (leq(loopcounter,n1))).
#
#cnf(i_0_210, plain, (leq(n1,loopcounter))).
#
#cnf(i_0_211, plain, (leq(tptp_float_0_001,pv76))).
#
#cnf(i_0_14, plain, (gt(plus(X1,n1),X1))).
#
#cnf(i_0_243, plain, (leq(X1,plus(X1,n1)))).
#
#cnf(i_0_176, plain, (plus(n0,n1)=n1)).
#
#cnf(i_0_104, plain, (plus(tptp_minus_1,n1)=n0)).
#
#cnf(i_0_103, plain, (sum(n0,tptp_minus_1,X1)=tptp_float_0_0)).
#
#cnf(i_0_102, plain, (tptp_float_0_0=n0)).
#
#cnf(i_0_103, plain, (sum(n0,tptp_minus_1,X1)=n0)).
#
#cnf(i_0_129, plain, (a_select2(tptp_update2(X1,X2,X3),X2)=X3)).
#
#cnf(i_0_11, plain, (X1=X2|gt(X2,X1)|~leq(X1,X2))).
#
#cnf(i_0_13, plain, (gt(X2,X1)|~leq(X1,minus(X2,n1)))).
#
#cnf(i_0_249, plain, (gt(X1,minus(X1,n1)))).
#
#cnf(i_0_250, plain, (leq(minus(X1,n1),X1))).
#
#cnf(i_0_120, plain, (gt(X2,X1)|~leq(plus(X1,n1),X2))).
#
#cnf(i_0_252, plain, (gt(X1,tptp_minus_1)|~leq(n0,X1))).
#
#cnf(i_0_257, plain, (~leq(n0,tptp_minus_1))).
#
#cnf(i_0_253, plain, (gt(X1,n0)|~leq(n1,X1))).
#
#cnf(i_0_258, plain, (leq(n0,X1)|~leq(n1,X1))).
#
#cnf(i_0_263, plain, (leq(n0,loopcounter))).
#
#cnf(i_0_259, plain, (~leq(n1,n0))).
#
#cnf(i_0_256, plain, (leq(tptp_minus_1,X1)|~leq(n0,X1))).
#
#cnf(i_0_260, plain, (leq(n0,plus(n1,n1)))).
#
#cnf(i_0_16, plain, (leq(X2,X1)|~gt(plus(X1,n1),X2))).
#
#cnf(i_0_269, plain, (leq(X1,tptp_minus_1)|~gt(n0,X1))).
#
#cnf(i_0_270, plain, (leq(X1,n0)|~gt(n1,X1))).
#
#cnf(i_0_285, plain, (leq(loopcounter,n0))).
#
#cnf(i_0_276, plain, (leq(minus(n0,n1),tptp_minus_1))).
##
#cnf(i_0_281, plain, (leq(minus(n1,n1),n0))).
#
#cnf(i_0_121, plain, (leq(n0,X2)|~leq(minus(X1,X2),X1))).
#
#cnf(i_0_117, plain, (plus(minus(X1,n1),n1)=X1)).
#
#cnf(i_0_116, plain, (minus(plus(X1,n1),n1)=X1)).
#
#cnf(i_0_297, plain, (minus(n0,n1)=tptp_minus_1)).
#
#cnf(i_0_298, plain, (minus(n1,n1)=n0)).
#
#cnf(i_0_301, plain, (gt(n0,X1)|~leq(X1,tptp_minus_1))).
##
#cnf(i_0_309, plain, (leq(X1,n0)|~leq(X1,tptp_minus_1))).
#
#cnf(i_0_306, plain, (gt(n1,X1)|~leq(X1,n0))).
#
#cnf(i_0_315, plain, (leq(X1,n1)|~leq(X1,n0))).
##
#cnf(i_0_312, plain, (leq(minus(tptp_minus_1,n1),n0))).
#
#cnf(i_0_1, plain, (X1=X2|gt(X2,X1)|gt(X1,X2))).
#
#cnf(i_0_326, plain, (X2=X1|leq(X1,X2)|gt(X1,X2))).
##
#cnf(i_0_334, plain, (X1=X2|leq(X1,X2)|leq(X2,X1))).
#
#cnf(i_0_18, plain, (leq(uniform_int_rnd(X2,X1),X1)|~leq(n0,X1))).
#
#cnf(i_0_2, plain, (gt(X1,X3)|~gt(X2,X3)|~gt(X1,X2))).
#
#cnf(i_0_356, plain, (~gt(X2,X1)|~gt(X1,X2))).
#
#cnf(i_0_394, plain, (gt(X1,n5)|~gt(X1,n135299))).
#
#cnf(i_0_395, plain, (gt(X1,n4)|~gt(X1,n135299))).
#
#cnf(i_0_396, plain, (gt(X1,n3)|~gt(X1,n135299))).
#
#cnf(i_0_433, plain, (~gt(n5,n135299))).
#
#cnf(i_0_397, plain, (gt(X1,n2)|~gt(X1,n135299))).
#
#cnf(i_0_398, plain, (gt(X1,n1)|~gt(X1,n135299))).
#
#cnf(i_0_399, plain, (gt(X1,tptp_minus_1)|~gt(X1,n135299))).
#
#cnf(i_0_434, plain, (~gt(n4,n135299))).
#
#cnf(i_0_400, plain, (gt(X1,n0)|~gt(X1,n135299))).
#
#cnf(i_0_401, plain, (gt(X1,n4)|~gt(X1,n5))).
#
#cnf(i_0_402, plain, (gt(X1,n3)|~gt(X1,n5))).
#
#cnf(i_0_435, plain, (~gt(n3,n135299))).
#
#cnf(i_0_403, plain, (gt(X1,n2)|~gt(X1,n5))).
#
#cnf(i_0_404, plain, (gt(X1,n1)|~gt(X1,n5))).
#
#cnf(i_0_405, plain, (gt(X1,tptp_minus_1)|~gt(X1,n5))).
#
#cnf(i_0_436, plain, (~gt(n2,n135299))).
#
#cnf(i_0_406, plain, (gt(X1,n0)|~gt(X1,n5))).
#
#cnf(i_0_407, plain, (gt(X1,n3)|~gt(X1,n4))).
#
#cnf(i_0_408, plain, (gt(X1,n2)|~gt(X1,n4))).
#
#cnf(i_0_437, plain, (~gt(n1,n135299))).
#
#cnf(i_0_409, plain, (gt(X1,n1)|~gt(X1,n4))).
#
#cnf(i_0_410, plain, (gt(X1,tptp_minus_1)|~gt(X1,n4))).
#
#cnf(i_0_411, plain, (gt(X1,n0)|~gt(X1,n4))).
#
#cnf(i_0_438, plain, (~gt(tptp_minus_1,n135299))).
#
#cnf(i_0_412, plain, (gt(X1,n2)|~gt(X1,n3))).
#
#cnf(i_0_413, plain, (gt(X1,n1)|~gt(X1,n3))).
#
#cnf(i_0_414, plain, (gt(X1,tptp_minus_1)|~gt(X1,n3))).
#
#cnf(i_0_439, plain, (~gt(n0,n135299))).
#
#cnf(i_0_415, plain, (gt(X1,n0)|~gt(X1,n3))).
#
#cnf(i_0_416, plain, (gt(X1,n1)|~gt(X1,n2))).
#
#cnf(i_0_417, plain, (gt(X1,tptp_minus_1)|~gt(X1,n2))).
#
#cnf(i_0_440, plain, (~gt(n4,n5))).
#
#cnf(i_0_418, plain, (gt(X1,n0)|~gt(X1,n2))).
#
#cnf(i_0_422, plain, (gt(X1,n0)|~gt(X1,n1))).
#
#cnf(i_0_358, plain, (gt(n135299,X1)|~gt(n5,X1))).
#
#cnf(i_0_441, plain, (~gt(n3,n5))).
#
#cnf(i_0_811, plain, (leq(X1,n135299)|~gt(n5,X1))).
#
#cnf(i_0_359, plain, (gt(n135299,X1)|~gt(n4,X1))).
#
#cnf(i_0_860, plain, (leq(X1,n135299)|~gt(n4,X1))).
#
#cnf(i_0_442, plain, (~gt(n2,n5))).
#
#cnf(i_0_360, plain, (gt(n135299,X1)|~gt(n3,X1))).
#
#cnf(i_0_908, plain, (leq(X1,n135299)|~gt(n3,X1))).
#
#cnf(i_0_361, plain, (gt(n135299,X1)|~gt(n2,X1))).
#
#cnf(i_0_443, plain, (~gt(n1,n5))).
#
#cnf(i_0_952, plain, (leq(X1,n135299)|~gt(n2,X1))).
#
#cnf(i_0_362, plain, (gt(n135299,X1)|~gt(n1,X1))).
#
#cnf(i_0_999, plain, (leq(X1,n135299)|~gt(n1,X1))).
#
#cnf(i_0_444, plain, (~gt(tptp_minus_1,n5))).
#
#cnf(i_0_1012, plain, (leq(loopcounter,n135299))).
#
#cnf(i_0_363, plain, (gt(n135299,X1)|~gt(tptp_minus_1,X1))).
#
#cnf(i_0_1046, plain, (leq(X1,n135299)|~gt(tptp_minus_1,X1))).
#
#cnf(i_0_445, plain, (~gt(n0,n5))).
#
#cnf(i_0_364, plain, (gt(n135299,X1)|~gt(n0,X1))).
#
#cnf(i_0_1090, plain, (leq(X1,n135299)|~gt(n0,X1))).
#
#cnf(i_0_365, plain, (gt(n5,X1)|~gt(n4,X1))).
#
#cnf(i_0_446, plain, (~gt(n3,n4))).
#
#cnf(i_0_1132, plain, (leq(X1,n5)|~gt(n4,X1))).
##
#cnf(i_0_366, plain, (gt(n5,X1)|~gt(n3,X1))).
#
#cnf(i_0_447, plain, (~gt(n2,n4))).
#
#cnf(i_0_1182, plain, (leq(X1,n5)|~gt(n3,X1))).
##
#cnf(i_0_367, plain, (gt(n5,X1)|~gt(n2,X1))).
#
#cnf(i_0_448, plain, (~gt(n1,n4))).
#
#cnf(i_0_1231, plain, (leq(X1,n5)|~gt(n2,X1))).
##
#cnf(i_0_368, plain, (gt(n5,X1)|~gt(n1,X1))).
#
#cnf(i_0_449, plain, (~gt(tptp_minus_1,n4))).
#
#cnf(i_0_1280, plain, (leq(X1,n5)|~gt(n1,X1))).
#
#cnf(i_0_1298, plain, (leq(loopcounter,n5))).
##
#cnf(i_0_450, plain, (~gt(n0,n4))).
#
#cnf(i_0_369, plain, (gt(n5,X1)|~gt(tptp_minus_1,X1))).
#
#cnf(i_0_1333, plain, (leq(X1,n5)|~gt(tptp_minus_1,X1))).
##
#cnf(i_0_451, plain, (~gt(n2,n3))).
#
#cnf(i_0_370, plain, (gt(n5,X1)|~gt(n0,X1))).
#
#cnf(i_0_1378, plain, (leq(X1,n5)|~gt(n0,X1))).
##
#cnf(i_0_452, plain, (~gt(n1,n3))).
#
#cnf(i_0_371, plain, (gt(n4,X1)|~gt(n3,X1))).
#
#cnf(i_0_1426, plain, (leq(X1,n4)|~gt(n3,X1))).
##
#cnf(i_0_453, plain, (~gt(tptp_minus_1,n3))).
##
#cnf(i_0_372, plain, (gt(n4,X1)|~gt(n2,X1))).
#
#cnf(i_0_1477, plain, (leq(X1,n4)|~gt(n2,X1))).
#
#cnf(i_0_454, plain, (~gt(n0,n3))).
###
#cnf(i_0_373, plain, (gt(n4,X1)|~gt(n1,X1))).
#
#cnf(i_0_455, plain, (~gt(n1,n2))).
#
#cnf(i_0_1528, plain, (leq(X1,n4)|~gt(n1,X1))).
#
#cnf(i_0_1549, plain, (leq(loopcounter,n4))).
##
#cnf(i_0_456, plain, (~gt(tptp_minus_1,n2))).
##
#cnf(i_0_374, plain, (gt(n4,X1)|~gt(tptp_minus_1,X1))).
#
#cnf(i_0_1583, plain, (leq(X1,n4)|~gt(tptp_minus_1,X1))).
#
#cnf(i_0_457, plain, (~gt(n0,n2))).
###
#cnf(i_0_375, plain, (gt(n4,X1)|~gt(n0,X1))).
#
#cnf(i_0_459, plain, (~gt(n0,n1))).
#
#cnf(i_0_1631, plain, (leq(X1,n4)|~gt(n0,X1))).
####
#cnf(i_0_376, plain, (gt(n3,X1)|~gt(n2,X1))).
#
#cnf(i_0_1681, plain, (leq(X1,n3)|~gt(n2,X1))).
#####
#cnf(i_0_377, plain, (gt(n3,X1)|~gt(n1,X1))).
##
#cnf(i_0_1730, plain, (leq(X1,n3)|~gt(n1,X1))).
#
#cnf(i_0_1749, plain, (leq(loopcounter,n3))).
#####
#cnf(i_0_378, plain, (gt(n3,X1)|~gt(tptp_minus_1,X1))).
##
#cnf(i_0_1780, plain, (leq(X1,n3)|~gt(tptp_minus_1,X1))).
#####
#cnf(i_0_379, plain, (gt(n3,X1)|~gt(n0,X1))).
#
#cnf(i_0_1826, plain, (leq(X1,n3)|~gt(n0,X1))).
######
#cnf(i_0_380, plain, (gt(n2,X1)|~gt(n1,X1))).
#
#cnf(i_0_1874, plain, (leq(X1,n2)|~gt(n1,X1))).
#
#cnf(i_0_1895, plain, (leq(loopcounter,n2))).
#######
#cnf(i_0_381, plain, (gt(n2,X1)|~gt(tptp_minus_1,X1))).
#
#cnf(i_0_1926, plain, (leq(X1,n2)|~gt(tptp_minus_1,X1))).
#######
#cnf(i_0_382, plain, (gt(n2,X1)|~gt(n0,X1))).
#
#cnf(i_0_1974, plain, (leq(X1,n2)|~gt(n0,X1))).
#######
#cnf(i_0_384, plain, (gt(n1,X1)|~gt(n0,X1))).
#
#cnf(i_0_2024, plain, (leq(X1,n1)|~gt(n0,X1))).
##########
#cnf(i_0_2049, plain, (leq(X1,n1)|~leq(X1,tptp_minus_1))).
#
#cnf(i_0_420, plain, (gt(X1,loopcounter)|~gt(X1,n1))).
#
#cnf(i_0_2086, plain, (leq(loopcounter,X1)|~gt(X1,n1))).
##
#cnf(i_0_421, plain, (gt(X1,tptp_minus_1)|~gt(X1,n1))).
#
#cnf(i_0_2160, plain, (leq(tptp_minus_1,X1)|~gt(X1,n1))).
#
#cnf(i_0_424, plain, (gt(X1,tptp_minus_1)|~gt(X1,n0))).
##
#cnf(i_0_2234, plain, (leq(tptp_minus_1,X1)|~gt(X1,n0))).
#
#cnf(i_0_1015, plain, (leq(X1,n135299)|~leq(X1,n0))).
#
#cnf(i_0_1104, plain, (leq(X1,n135299)|~leq(X1,tptp_minus_1))).
##
#cnf(i_0_1301, plain, (leq(X1,n5)|~leq(X1,n0))).
#
#cnf(i_0_1394, plain, (leq(X1,n5)|~leq(X1,tptp_minus_1))).
#
#cnf(i_0_1552, plain, (leq(X1,n4)|~leq(X1,n0))).
##
#cnf(i_0_1653, plain, (leq(X1,n4)|~leq(X1,tptp_minus_1))).
#
#cnf(i_0_1752, plain, (leq(X1,n3)|~leq(X1,n0))).
#
#cnf(i_0_1846, plain, (leq(X1,n3)|~leq(X1,tptp_minus_1))).
##
#cnf(i_0_1898, plain, (leq(X1,n2)|~leq(X1,n0))).
#
#cnf(i_0_1996, plain, (leq(X1,n2)|~leq(X1,tptp_minus_1))).
#
#cnf(i_0_2279, plain, (leq(tptp_minus_1,X1)|~leq(n1,X1))).
##
#cnf(i_0_530, plain, (n135299=n5|~leq(n135299,n5))).
#
#cnf(i_0_557, plain, (n135299=n4|~leq(n135299,n4))).
#
#cnf(i_0_586, plain, (n135299=n3|~leq(n135299,n3))).
##
#cnf(i_0_616, plain, (n135299=n2|~leq(n135299,n2))).
#
#cnf(i_0_648, plain, (n135299=n1|~leq(n135299,n1))).
#
#cnf(i_0_682, plain, (n135299=tptp_minus_1|~leq(n135299,tptp_minus_1))).
##
#cnf(i_0_718, plain, (n135299=n0|~leq(n135299,n0))).
#
#cnf(i_0_757, plain, (n5=n4|~leq(n5,n4))).
#
#cnf(i_0_817, plain, (n5=n3|~leq(n5,n3))).
##
#cnf(i_0_879, plain, (n5=n2|~leq(n5,n2))).
#
#cnf(i_0_958, plain, (n5=n1|~leq(n5,n1))).
#
#cnf(i_0_1017, plain, (n5=tptp_minus_1|~leq(n5,tptp_minus_1))).
#
#cnf(i_0_541, plain, (~leq(n135299,n0))).
#
#cnf(i_0_1060, plain, (n5=n0|~leq(n5,n0))).
#
#cnf(i_0_1140, plain, (n4=n3|~leq(n4,n3))).
#
#cnf(i_0_1190, plain, (n4=n2|~leq(n4,n2))).
#
#cnf(i_0_561, plain, (~leq(n135299,tptp_minus_1))).
#
#cnf(i_0_1239, plain, (n4=n1|~leq(n4,n1))).
#
#cnf(i_0_1288, plain, (n4=tptp_minus_1|~leq(n4,tptp_minus_1))).
#
#cnf(i_0_1303, plain, (n4=n0|~leq(n4,n0))).
#
#cnf(i_0_599, plain, (~leq(n5,n0))).
#
#cnf(i_0_1349, plain, (n3=n2|~leq(n3,n2))).
#
#cnf(i_0_1396, plain, (n3=n1|~leq(n3,n1))).
#
#cnf(i_0_1448, plain, (n3=tptp_minus_1|~leq(n3,tptp_minus_1))).
#
#cnf(i_0_621, plain, (~leq(n5,tptp_minus_1))).
#
#cnf(i_0_1498, plain, (n3=n0|~leq(n3,n0))).
#
#cnf(i_0_1538, plain, (n2=n1|~leq(n2,n1))).
#
#cnf(i_0_1554, plain, (n2=tptp_minus_1|~leq(n2,tptp_minus_1))).
##
#cnf(i_0_1601, plain, (n2=n0|~leq(n2,n0))).
###
#cnf(i_0_654, plain, (~leq(n4,n0))).
####
#cnf(i_0_675, plain, (~leq(n4,tptp_minus_1))).
####
#cnf(i_0_699, plain, (~leq(n3,n0))).
########
#cnf(i_0_726, plain, (~leq(n3,tptp_minus_1))).
####
#cnf(i_0_737, plain, (~leq(n2,n0))).
####
#cnf(i_0_765, plain, (~leq(n2,tptp_minus_1))).
####
#cnf(i_0_778, plain, (~leq(n1,tptp_minus_1))).
#####
#cnf(i_0_467, plain, (X1=X2|~leq(X1,X2)|~gt(X1,X2))).
#
#cnf(i_0_2459, plain, (loopcounter=n1)).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(leq_succ_gt_equiv, axiom, ![X1, X2]:(leq(X1,X2)<=>gt(succ(X2),X1)), file('./p/SWV181+1.tptp', leq_succ_gt_equiv)).
fof(succ_plus_1_r, axiom, ![X1]:plus(X1,n1)=succ(X1), file('./p/SWV181+1.tptp', succ_plus_1_r)).
fof(cl5_nebula_init_0081, conjecture, ((((((((leq(tptp_float_0_001,pv76)&leq(n1,loopcounter))&gt(n1,loopcounter))&![X14]:((leq(n0,X14)&leq(X14,n135299))=>![X18]:((leq(n0,X18)&leq(X18,n4))=>a_select3(q_init,X14,X18)=init)))&![X4]:((leq(n0,X4)&leq(X4,n4))=>a_select3(center_init,X4,n0)=init))&(gt(loopcounter,n0)=>![X20]:((leq(n0,X20)&leq(X20,n4))=>a_select2(mu_init,X20)=init)))&(gt(loopcounter,n0)=>![X21]:((leq(n0,X21)&leq(X21,n4))=>a_select2(rho_init,X21)=init)))&(gt(loopcounter,n0)=>![X22]:((leq(n0,X22)&leq(X22,n4))=>a_select2(sigma_init,X22)=init)))=>![X28]:((leq(n0,X28)&leq(X28,n4))=>a_select2(muold_init,X28)=init)), file('./p/SWV181+1.tptp', cl5_nebula_init_0081)).
fof(irreflexivity_gt, axiom, ![X1]:~(gt(X1,X1)), file('./p/SWV181+1.tptp', irreflexivity_gt)).
fof(transitivity_gt, axiom, ![X1, X2, X3]:((gt(X1,X2)&gt(X2,X3))=>gt(X1,X3)), file('./p/SWV181+1.tptp', transitivity_gt)).
fof(leq_succ_gt, axiom, ![X1, X2]:(leq(succ(X1),X2)=>gt(X2,X1)), file('./p/SWV181+1.tptp', leq_succ_gt)).
fof(successor_1, axiom, succ(n0)=n1, file('./p/SWV181+1.tptp', successor_1)).
fof(leq_gt2, axiom, ![X1, X2]:((leq(X1,X2)&X1!=X2)=>gt(X2,X1)), file('./p/SWV181+1.tptp', leq_gt2)).
fof(c_0_8, plain, (epred2_0<=>(((((((leq(tptp_float_0_001,pv76)&leq(n1,loopcounter))&gt(n1,loopcounter))&![X14]:((leq(n0,X14)&leq(X14,n135299))=>![X18]:((leq(n0,X18)&leq(X18,n4))=>a_select3(q_init,X14,X18)=init)))&![X4]:((leq(n0,X4)&leq(X4,n4))=>a_select3(center_init,X4,n0)=init))&(gt(loopcounter,n0)=>![X20]:((leq(n0,X20)&leq(X20,n4))=>a_select2(mu_init,X20)=init)))&(gt(loopcounter,n0)=>![X21]:((leq(n0,X21)&leq(X21,n4))=>a_select2(rho_init,X21)=init)))&(gt(loopcounter,n0)=>![X22]:((leq(n0,X22)&leq(X22,n4))=>a_select2(sigma_init,X22)=init)))), introduced(definition)).
fof(c_0_9, plain, ![X52, X53]:((~leq(X52,X53)|gt(succ(X53),X52))&(~gt(succ(X53),X52)|leq(X52,X53))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[leq_succ_gt_equiv])])).
fof(c_0_10, plain, ![X136]:plus(X136,n1)=succ(X136), inference(variable_rename,[status(thm)],[succ_plus_1_r])).
fof(c_0_11, plain, (epred2_0=>(((((((leq(tptp_float_0_001,pv76)&leq(n1,loopcounter))&gt(n1,loopcounter))&![X14]:((leq(n0,X14)&leq(X14,n135299))=>![X18]:((leq(n0,X18)&leq(X18,n4))=>a_select3(q_init,X14,X18)=init)))&![X4]:((leq(n0,X4)&leq(X4,n4))=>a_select3(center_init,X4,n0)=init))&(gt(loopcounter,n0)=>![X20]:((leq(n0,X20)&leq(X20,n4))=>a_select2(mu_init,X20)=init)))&(gt(loopcounter,n0)=>![X21]:((leq(n0,X21)&leq(X21,n4))=>a_select2(rho_init,X21)=init)))&(gt(loopcounter,n0)=>![X22]:((leq(n0,X22)&leq(X22,n4))=>a_select2(sigma_init,X22)=init)))), inference(split_equiv,[status(thm)],[c_0_8])).
fof(c_0_12, negated_conjecture, ~((epred2_0=>![X28]:((leq(n0,X28)&leq(X28,n4))=>a_select2(muold_init,X28)=init))), inference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[cl5_nebula_init_0081]), c_0_8])).
fof(c_0_13, plain, ![X34]:~gt(X34,X34), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[irreflexivity_gt])])).
fof(c_0_14, plain, ![X31, X32, X33]:(~gt(X31,X32)|~gt(X32,X33)|gt(X31,X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[transitivity_gt])])).
fof(c_0_15, plain, ![X151, X152]:(~leq(succ(X151),X152)|gt(X152,X151)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[leq_succ_gt])])).
cnf(c_0_16, plain, (leq(X2,X1)|~gt(succ(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_9])).
cnf(c_0_17, plain, (plus(X1,n1)=succ(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_18, plain, (succ(n0)=n1), inference(split_conjunct,[status(thm)],[successor_1])).
fof(c_0_19, plain, ![X202, X203, X204, X205, X206, X207]:((((((((leq(tptp_float_0_001,pv76)|~epred2_0)&(leq(n1,loopcounter)|~epred2_0))&(gt(n1,loopcounter)|~epred2_0))&(~leq(n0,X202)|~leq(X202,n135299)|(~leq(n0,X203)|~leq(X203,n4)|a_select3(q_init,X202,X203)=init)|~epred2_0))&(~leq(n0,X204)|~leq(X204,n4)|a_select3(center_init,X204,n0)=init|~epred2_0))&(~gt(loopcounter,n0)|(~leq(n0,X205)|~leq(X205,n4)|a_select2(mu_init,X205)=init)|~epred2_0))&(~gt(loopcounter,n0)|(~leq(n0,X206)|~leq(X206,n4)|a_select2(rho_init,X206)=init)|~epred2_0))&(~gt(loopcounter,n0)|(~leq(n0,X207)|~leq(X207,n4)|a_select2(sigma_init,X207)=init)|~epred2_0)), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_11])])])])).
fof(c_0_20, negated_conjecture, (epred2_0&((leq(n0,esk24_0)&leq(esk24_0,n4))&a_select2(muold_init,esk24_0)!=init)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_12])])])).
cnf(c_0_21, plain, (~gt(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_22, plain, (gt(X1,X3)|~gt(X1,X2)|~gt(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_23, plain, ![X45, X46]:(~leq(X45,X46)|X45=X46|gt(X46,X45)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[leq_gt2])])).
cnf(c_0_24, plain, (gt(X2,X1)|~leq(succ(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_25, plain, (leq(X2,X1)|~gt(plus(X1,n1),X2)), inference(rw,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_26, plain, (plus(n0,n1)=n1), inference(rw,[status(thm)],[c_0_18, c_0_17])).
cnf(c_0_27, plain, (gt(n1,loopcounter)|~epred2_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_28, negated_conjecture, (epred2_0), inference(split_conjunct,[status(thm)],[c_0_20])).
cnf(c_0_29, plain, (~gt(X1,X2)|~gt(X2,X1)), inference(pm,[status(thm)],[c_0_21, c_0_22])).
cnf(c_0_30, plain, (X1=X2|gt(X2,X1)|~leq(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_23])).
cnf(c_0_31, plain, (leq(n1,loopcounter)|~epred2_0), inference(split_conjunct,[status(thm)],[c_0_19])).
cnf(c_0_32, plain, (gt(X2,X1)|~leq(plus(X1,n1),X2)), inference(rw,[status(thm)],[c_0_24, c_0_17])).
cnf(c_0_33, plain, (leq(X1,n0)|~gt(n1,X1)), inference(pm,[status(thm)],[c_0_25, c_0_26])).
cnf(c_0_34, plain, (gt(n1,loopcounter)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28])])).
cnf(c_0_35, plain, (X1=X2|~leq(X1,X2)|~gt(X1,X2)), inference(pm,[status(thm)],[c_0_29, c_0_30])).
cnf(c_0_36, plain, (leq(n1,loopcounter)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_31, c_0_28])])).
cnf(c_0_37, plain, (gt(X1,n0)|~leq(n1,X1)), inference(pm,[status(thm)],[c_0_32, c_0_26])).
cnf(c_0_38, plain, (leq(loopcounter,n0)), inference(pm,[status(thm)],[c_0_33, c_0_34])).
cnf(c_0_39, plain, (loopcounter=n1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_35, c_0_34]), c_0_36])])).
cnf(c_0_40, plain, (~leq(n1,n0)), inference(pm,[status(thm)],[c_0_21, c_0_37])).
cnf(c_0_41, plain, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_39]), c_0_40]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 42
# Proof object clause steps            : 23
# Proof object formula steps           : 19
# Proof object conjectures             : 4
# Proof object clause conjectures      : 1
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 10
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 7
# Proof object simplifying inferences  : 11
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 92
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 211
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 209
# Processed clauses                    : 386
# ...of these trivial                  : 0
# ...subsumed                          : 105
# ...remaining for further processing  : 281
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 15
# Generated clauses                    : 2265
# ...of the previous two non-trivial   : 1156
# Contextual simplify-reflections      : 0
# Paramodulations                      : 2261
# Factorizations                       : 4
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 266
#    Positive orientable unit clauses  : 78
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 42
#    Non-unit-clauses                  : 146
# Current number of unprocessed clauses: 978
# ...number of literals in the above   : 2633
# Current number of archived formulas  : 0
# Current number of archived clauses   : 17
# Clause-clause subsumption calls (NU) : 2313
# Rec. Clause-clause subsumption calls : 2278
# Non-unit clause-clause subsumptions  : 47
# Unit Clause-clause subsumption calls : 895
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 4
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 21308
