# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_134, plain, (true)).
#
#cnf(i_0_4, plain, (leq(X1,X1))).
#
#cnf(i_0_149, plain, (gt(n0,tptp_minus_1))).
#
#cnf(i_0_3, plain, (~gt(X1,X1))).
#
#cnf(i_0_155, plain, (gt(n1,n0))).
#
#cnf(i_0_150, plain, (gt(n1,tptp_minus_1))).
#
#cnf(i_0_156, plain, (gt(n2,n0))).
#
#cnf(i_0_135, plain, (def!=use)).
#
#cnf(i_0_151, plain, (gt(n2,tptp_minus_1))).
#
#cnf(i_0_160, plain, (gt(n2,n1))).
#
#cnf(i_0_157, plain, (gt(n3,n0))).
#
#cnf(i_0_136, negated_conjecture, (~leq(n0,pv1325))).
#
#cnf(i_0_152, plain, (gt(n3,tptp_minus_1))).
#
#cnf(i_0_161, plain, (gt(n3,n1))).
#
#cnf(i_0_164, plain, (gt(n3,n2))).
#
#cnf(i_0_153, plain, (gt(n4,n0))).
#
#cnf(i_0_147, plain, (gt(n4,tptp_minus_1))).
#
#cnf(i_0_158, plain, (gt(n4,n1))).
#
#cnf(i_0_162, plain, (gt(n4,n2))).
#
#cnf(i_0_165, plain, (gt(n4,n3))).
#
#cnf(i_0_154, plain, (gt(n5,n0))).
#
#cnf(i_0_148, plain, (gt(n5,tptp_minus_1))).
#
#cnf(i_0_159, plain, (gt(n5,n1))).
#
#cnf(i_0_163, plain, (gt(n5,n2))).
#
#cnf(i_0_166, plain, (gt(n5,n3))).
#
#cnf(i_0_146, plain, (gt(n5,n4))).
#
#cnf(i_0_145, negated_conjecture, (leq(n0,s_best7))).
#
#cnf(i_0_144, negated_conjecture, (leq(n0,s_sworst7))).
#
#cnf(i_0_143, negated_conjecture, (leq(n0,s_worst7))).
#
#cnf(i_0_142, negated_conjecture, (leq(n2,pv1325))).
#
#cnf(i_0_141, negated_conjecture, (leq(s_best7,n3))).
#
#cnf(i_0_140, negated_conjecture, (leq(s_sworst7,n3))).
#
#cnf(i_0_139, negated_conjecture, (leq(s_worst7,n3))).
#
#cnf(i_0_138, negated_conjecture, (leq(pv1325,n3))).
#
#cnf(i_0_7, plain, (gt(X2,X1)|~lt(X1,X2))).
#
#cnf(i_0_10, plain, (leq(X2,X1)|~gt(X1,X2))).
#
#cnf(i_0_203, plain, (leq(n4,n5))).
#
#cnf(i_0_204, plain, (leq(n3,n5))).
#
#cnf(i_0_205, plain, (leq(n2,n5))).
#
#cnf(i_0_206, plain, (leq(n1,n5))).
#
#cnf(i_0_207, plain, (leq(tptp_minus_1,n5))).
#
#cnf(i_0_208, plain, (leq(n0,n5))).
#
#cnf(i_0_209, plain, (leq(n3,n4))).
#
#cnf(i_0_210, plain, (leq(n2,n4))).
#
#cnf(i_0_211, plain, (leq(n1,n4))).
#
#cnf(i_0_212, plain, (leq(tptp_minus_1,n4))).
#
#cnf(i_0_213, plain, (leq(n0,n4))).
#
#cnf(i_0_214, plain, (leq(n2,n3))).
#
#cnf(i_0_215, plain, (leq(n1,n3))).
#
#cnf(i_0_216, plain, (leq(tptp_minus_1,n3))).
#
#cnf(i_0_217, plain, (leq(n0,n3))).
#
#cnf(i_0_218, plain, (leq(n1,n2))).
#
#cnf(i_0_219, plain, (leq(tptp_minus_1,n2))).
#
#cnf(i_0_220, plain, (leq(n0,n2))).
#
#cnf(i_0_221, plain, (leq(tptp_minus_1,n1))).
#
#cnf(i_0_222, plain, (leq(n0,n1))).
#
#cnf(i_0_223, plain, (leq(tptp_minus_1,n0))).
#
#cnf(i_0_9, plain, (leq(X2,X1)|~geq(X1,X2))).
#
#cnf(i_0_6, plain, (lt(X2,X1)|~gt(X1,X2))).
#
#cnf(i_0_8, plain, (geq(X2,X1)|~leq(X1,X2))).
#
#cnf(i_0_14, plain, (gt(plus(X1,n1),X1))).
#
#cnf(i_0_226, plain, (leq(X1,plus(X1,n1)))).
#
#cnf(i_0_175, plain, (plus(n0,n1)=n1)).
#
#cnf(i_0_104, plain, (plus(tptp_minus_1,n1)=n0)).
#
#cnf(i_0_103, plain, (sum(n0,tptp_minus_1,X1)=tptp_float_0_0)).
#
#cnf(i_0_102, plain, (tptp_float_0_0=n0)).
#
#cnf(i_0_103, plain, (sum(n0,tptp_minus_1,X1)=n0)).
#
#cnf(i_0_129, plain, (a_select2(tptp_update2(X1,X2,X3),X2)=X3)).
#
#cnf(i_0_11, plain, (X1=X2|gt(X2,X1)|~leq(X1,X2))).
#
#cnf(i_0_13, plain, (gt(X2,X1)|~leq(X1,minus(X2,n1)))).
#
#cnf(i_0_232, plain, (gt(X1,minus(X1,n1)))).
#
#cnf(i_0_233, plain, (leq(minus(X1,n1),X1))).
#
#cnf(i_0_120, plain, (gt(X2,X1)|~leq(plus(X1,n1),X2))).
#
#cnf(i_0_235, plain, (gt(X1,tptp_minus_1)|~leq(n0,X1))).
#
#cnf(i_0_236, plain, (gt(X1,n0)|~leq(n1,X1))).
#
#cnf(i_0_240, plain, (~leq(n0,tptp_minus_1))).
#
#cnf(i_0_241, plain, (leq(n0,X1)|~leq(n1,X1))).
#
#cnf(i_0_239, plain, (leq(tptp_minus_1,X1)|~leq(n0,X1))).
#
#cnf(i_0_243, plain, (leq(n0,plus(n1,n1)))).
#
#cnf(i_0_242, plain, (~leq(n1,n0))).
#
#cnf(i_0_16, plain, (leq(X2,X1)|~gt(plus(X1,n1),X2))).
#
#cnf(i_0_250, plain, (leq(X1,tptp_minus_1)|~gt(n0,X1))).
#
#cnf(i_0_251, plain, (leq(X1,n0)|~gt(n1,X1))).
##
#cnf(i_0_257, plain, (leq(minus(n0,n1),tptp_minus_1))).
#
#cnf(i_0_262, plain, (leq(minus(n1,n1),n0))).
#
#cnf(i_0_121, plain, (leq(n0,X2)|~leq(minus(X1,X2),X1))).
#
#cnf(i_0_117, plain, (plus(minus(X1,n1),n1)=X1)).
#
#cnf(i_0_116, plain, (minus(plus(X1,n1),n1)=X1)).
#
#cnf(i_0_277, plain, (minus(n0,n1)=tptp_minus_1)).
#
#cnf(i_0_278, plain, (minus(n1,n1)=n0)).
#
#cnf(i_0_281, plain, (gt(n0,X1)|~leq(X1,tptp_minus_1))).
#
#cnf(i_0_289, plain, (leq(X1,n0)|~leq(X1,tptp_minus_1))).
#
#cnf(i_0_286, plain, (gt(n1,X1)|~leq(X1,n0))).
#
#cnf(i_0_295, plain, (leq(X1,n1)|~leq(X1,n0))).
##
#cnf(i_0_292, plain, (leq(minus(tptp_minus_1,n1),n0))).
#
#cnf(i_0_1, plain, (X1=X2|gt(X2,X1)|gt(X1,X2))).
#
#cnf(i_0_306, plain, (X2=X1|leq(X1,X2)|gt(X1,X2))).
###
#cnf(i_0_314, plain, (X1=X2|leq(X1,X2)|leq(X2,X1))).
#
#cnf(i_0_18, plain, (leq(uniform_int_rnd(X2,X1),X1)|~leq(n0,X1))).
#
#cnf(i_0_106, plain, (plus(n1,X1)=plus(X1,n1))).
#
#cnf(i_0_333, plain, (gt(plus(n1,X1),X1))).
#
#cnf(i_0_334, plain, (leq(X1,plus(n1,X1)))).
##
#cnf(i_0_104, plain, (plus(n1,tptp_minus_1)=n0)).
#
#cnf(i_0_335, plain, (gt(X1,X2)|~leq(plus(n1,X2),X1))).
##
#cnf(i_0_336, plain, (leq(X1,X2)|~gt(plus(n1,X2),X1))).
####
#cnf(i_0_117, plain, (plus(n1,minus(X1,n1))=X1)).
#
#cnf(i_0_332, plain, (minus(plus(n1,X1),n1)=X1)).
#
#cnf(i_0_2, plain, (gt(X1,X3)|~gt(X2,X3)|~gt(X1,X2))).
#
#cnf(i_0_426, plain, (gt(X1,n4)|~gt(X1,n5))).
#
#cnf(i_0_427, plain, (gt(X1,n3)|~gt(X1,n5))).
#
#cnf(i_0_394, plain, (~gt(X2,X1)|~gt(X1,X2))).
#
#cnf(i_0_428, plain, (gt(X1,n2)|~gt(X1,n5))).
#
#cnf(i_0_429, plain, (gt(X1,n1)|~gt(X1,n5))).
#
#cnf(i_0_430, plain, (gt(X1,tptp_minus_1)|~gt(X1,n5))).
#
#cnf(i_0_476, plain, (~gt(n4,n5))).
#
#cnf(i_0_431, plain, (gt(X1,n0)|~gt(X1,n5))).
#
#cnf(i_0_432, plain, (gt(X1,n3)|~gt(X1,n4))).
#
#cnf(i_0_433, plain, (gt(X1,n2)|~gt(X1,n4))).
#
#cnf(i_0_477, plain, (~gt(n3,n5))).
#
#cnf(i_0_434, plain, (gt(X1,n1)|~gt(X1,n4))).
#
#cnf(i_0_435, plain, (gt(X1,tptp_minus_1)|~gt(X1,n4))).
#
#cnf(i_0_436, plain, (gt(X1,n0)|~gt(X1,n4))).
#
#cnf(i_0_478, plain, (~gt(n2,n5))).
#
#cnf(i_0_437, plain, (gt(X1,n2)|~gt(X1,n3))).
#
#cnf(i_0_438, plain, (gt(X1,n1)|~gt(X1,n3))).
#
#cnf(i_0_439, plain, (gt(X1,tptp_minus_1)|~gt(X1,n3))).
#
#cnf(i_0_479, plain, (~gt(n1,n5))).
#
#cnf(i_0_440, plain, (gt(X1,n0)|~gt(X1,n3))).
#
#cnf(i_0_441, plain, (gt(X1,n1)|~gt(X1,n2))).
#
#cnf(i_0_442, plain, (gt(X1,tptp_minus_1)|~gt(X1,n2))).
#
#cnf(i_0_480, plain, (~gt(tptp_minus_1,n5))).
#
#cnf(i_0_443, plain, (gt(X1,n0)|~gt(X1,n2))).
#
#cnf(i_0_447, plain, (gt(X1,n0)|~gt(X1,n1))).
#
#cnf(i_0_396, plain, (gt(n5,X1)|~gt(n4,X1))).
#
#cnf(i_0_481, plain, (~gt(n0,n5))).
#
#cnf(i_0_751, plain, (leq(X1,n5)|~gt(n4,X1))).
#
#cnf(i_0_397, plain, (gt(n5,X1)|~gt(n3,X1))).
#
#cnf(i_0_793, plain, (leq(X1,n5)|~gt(n3,X1))).
#
#cnf(i_0_482, plain, (~gt(n3,n4))).
#
#cnf(i_0_398, plain, (gt(n5,X1)|~gt(n2,X1))).
#
#cnf(i_0_833, plain, (leq(X1,n5)|~gt(n2,X1))).
#
#cnf(i_0_399, plain, (gt(n5,X1)|~gt(n1,X1))).
#
#cnf(i_0_483, plain, (~gt(n2,n4))).
#
#cnf(i_0_869, plain, (leq(X1,n5)|~gt(n1,X1))).
#
#cnf(i_0_400, plain, (gt(n5,X1)|~gt(tptp_minus_1,X1))).
#
#cnf(i_0_908, plain, (leq(X1,n5)|~gt(tptp_minus_1,X1))).
#
#cnf(i_0_484, plain, (~gt(n1,n4))).
#
#cnf(i_0_401, plain, (gt(n5,X1)|~gt(n0,X1))).
#
#cnf(i_0_945, plain, (leq(X1,n5)|~gt(n0,X1))).
#
#cnf(i_0_402, plain, (gt(n4,X1)|~gt(n3,X1))).
#
#cnf(i_0_485, plain, (~gt(tptp_minus_1,n4))).
#
#cnf(i_0_980, plain, (leq(X1,n4)|~gt(n3,X1))).
##
#cnf(i_0_403, plain, (gt(n4,X1)|~gt(n2,X1))).
#
#cnf(i_0_486, plain, (~gt(n0,n4))).
#
#cnf(i_0_1022, plain, (leq(X1,n4)|~gt(n2,X1))).
##
#cnf(i_0_404, plain, (gt(n4,X1)|~gt(n1,X1))).
#
#cnf(i_0_487, plain, (~gt(n2,n3))).
#
#cnf(i_0_1064, plain, (leq(X1,n4)|~gt(n1,X1))).
##
#cnf(i_0_405, plain, (gt(n4,X1)|~gt(tptp_minus_1,X1))).
#
#cnf(i_0_488, plain, (~gt(n1,n3))).
#
#cnf(i_0_1105, plain, (leq(X1,n4)|~gt(tptp_minus_1,X1))).
##
#cnf(i_0_406, plain, (gt(n4,X1)|~gt(n0,X1))).
#
#cnf(i_0_489, plain, (~gt(tptp_minus_1,n3))).
#
#cnf(i_0_1144, plain, (leq(X1,n4)|~gt(n0,X1))).
##
#cnf(i_0_407, plain, (gt(n3,X1)|~gt(n2,X1))).
#
#cnf(i_0_490, plain, (~gt(n0,n3))).
#
#cnf(i_0_1184, plain, (leq(X1,n3)|~gt(n2,X1))).
###
#cnf(i_0_491, plain, (~gt(n1,n2))).
#
#cnf(i_0_408, plain, (gt(n3,X1)|~gt(n1,X1))).
#
#cnf(i_0_1232, plain, (leq(X1,n3)|~gt(n1,X1))).
##
#cnf(i_0_492, plain, (~gt(tptp_minus_1,n2))).
##
#cnf(i_0_409, plain, (gt(n3,X1)|~gt(tptp_minus_1,X1))).
#
#cnf(i_0_1275, plain, (leq(X1,n3)|~gt(tptp_minus_1,X1))).
#
#cnf(i_0_493, plain, (~gt(n0,n2))).
###
#cnf(i_0_410, plain, (gt(n3,X1)|~gt(n0,X1))).
#
#cnf(i_0_496, plain, (~gt(n0,n1))).
#
#cnf(i_0_1316, plain, (leq(X1,n3)|~gt(n0,X1))).
####
#cnf(i_0_411, plain, (gt(n2,X1)|~gt(n1,X1))).
#
#cnf(i_0_1359, plain, (leq(X1,n2)|~gt(n1,X1))).
#####
#cnf(i_0_412, plain, (gt(n2,X1)|~gt(tptp_minus_1,X1))).
##
#cnf(i_0_1401, plain, (leq(X1,n2)|~gt(tptp_minus_1,X1))).
#####
#cnf(i_0_413, plain, (gt(n2,X1)|~gt(n0,X1))).
#
#cnf(i_0_1440, plain, (leq(X1,n2)|~gt(n0,X1))).
######
#cnf(i_0_416, plain, (gt(n1,X1)|~gt(n0,X1))).
#
#cnf(i_0_1481, plain, (leq(X1,n1)|~gt(n0,X1))).
########
#cnf(i_0_1504, plain, (leq(X1,n1)|~leq(X1,tptp_minus_1))).
#
#cnf(i_0_446, plain, (gt(X1,tptp_minus_1)|~gt(X1,n1))).
##
#cnf(i_0_1535, plain, (leq(tptp_minus_1,X1)|~gt(X1,n1))).
#
#cnf(i_0_449, plain, (gt(X1,tptp_minus_1)|~gt(X1,n0))).
#
#cnf(i_0_1596, plain, (leq(tptp_minus_1,X1)|~gt(X1,n0))).
##
#cnf(i_0_887, plain, (leq(X1,n5)|~leq(X1,n0))).
#
#cnf(i_0_959, plain, (leq(X1,n5)|~leq(X1,tptp_minus_1))).
#
#cnf(i_0_1084, plain, (leq(X1,n4)|~leq(X1,n0))).
##
#cnf(i_0_1163, plain, (leq(X1,n4)|~leq(X1,tptp_minus_1))).
#
#cnf(i_0_1251, plain, (leq(X1,n3)|~leq(X1,n0))).
#
#cnf(i_0_1338, plain, (leq(X1,n3)|~leq(X1,tptp_minus_1))).
##
#cnf(i_0_1380, plain, (leq(X1,n2)|~leq(X1,n0))).
#
#cnf(i_0_1460, plain, (leq(X1,n2)|~leq(X1,tptp_minus_1))).
#
#cnf(i_0_1634, plain, (leq(tptp_minus_1,X1)|~leq(n1,X1))).
##
#cnf(i_0_564, plain, (n5=n4|~leq(n5,n4))).
#
#cnf(i_0_596, plain, (n5=n3|~leq(n5,n3))).
#
#cnf(i_0_629, plain, (n5=n2|~leq(n5,n2))).
##
#cnf(i_0_665, plain, (n5=n1|~leq(n5,n1))).
#
#cnf(i_0_704, plain, (n5=tptp_minus_1|~leq(n5,tptp_minus_1))).
#
#cnf(i_0_757, plain, (n5=n0|~leq(n5,n0))).
##
#cnf(i_0_811, plain, (n4=n3|~leq(n4,n3))).
#
#cnf(i_0_875, plain, (n4=n2|~leq(n4,n2))).
#
#cnf(i_0_922, plain, (n4=n1|~leq(n4,n1))).
##
#cnf(i_0_988, plain, (n4=tptp_minus_1|~leq(n4,tptp_minus_1))).
#
#cnf(i_0_1030, plain, (n4=n0|~leq(n4,n0))).
#
#cnf(i_0_1072, plain, (n3=n2|~leq(n3,n2))).
##
#cnf(i_0_1113, plain, (n3=n1|~leq(n3,n1))).
#
#cnf(i_0_1152, plain, (n3=tptp_minus_1|~leq(n3,tptp_minus_1))).
#
#cnf(i_0_1194, plain, (n3=n0|~leq(n3,n0))).
##
#cnf(i_0_1209, plain, (n2=n1|~leq(n2,n1))).
#
#cnf(i_0_1253, plain, (n2=tptp_minus_1|~leq(n2,tptp_minus_1))).
#
#cnf(i_0_1293, plain, (n2=n0|~leq(n2,n0))).
###
#cnf(i_0_1672, plain, (n5=n4|~leq(n5,n0))).
#
#cnf(i_0_1673, plain, (n5=n4|~leq(n5,tptp_minus_1))).
#
#cnf(i_0_547, plain, (~leq(n5,n0))).
##
#cnf(i_0_1677, plain, (n5=n3|~leq(n5,tptp_minus_1))).
#
#cnf(i_0_1680, plain, (n5=n2|~leq(n5,tptp_minus_1))).
#
#cnf(i_0_569, plain, (~leq(n5,tptp_minus_1))).
###
#cnf(i_0_1691, plain, (n4=n3|~leq(n4,n0))).
#
#cnf(i_0_601, plain, (~leq(n4,n0))).
#
#cnf(i_0_1692, plain, (n4=n3|~leq(n4,tptp_minus_1))).
#
#cnf(i_0_1695, plain, (n4=n2|~leq(n4,tptp_minus_1))).
##
#cnf(i_0_622, plain, (~leq(n4,tptp_minus_1))).
##
#cnf(i_0_1706, plain, (n3=n2|~leq(n3,tptp_minus_1))).
#
#cnf(i_0_1707, plain, (n3=n2|~leq(n3,n0))).
#
#cnf(i_0_646, plain, (~leq(n3,n0))).
##
#cnf(i_0_1717, plain, (n2=n1|~leq(n2,n0))).
#
#cnf(i_0_504, plain, (X1=X2|~leq(X1,X2)|~gt(X1,X2))).
#####
#cnf(i_0_673, plain, (~leq(n3,tptp_minus_1))).
####
#cnf(i_0_684, plain, (~leq(n2,n0))).
####
#cnf(i_0_712, plain, (~leq(n2,tptp_minus_1))).
####
#cnf(i_0_725, plain, (~leq(n1,tptp_minus_1))).
##############
#cnf(i_0_458, plain, (gt(plus(n1,n5),n4))).
###
#cnf(i_0_467, plain, (gt(plus(n1,n5),n3))).
##
#cnf(i_0_536, plain, (gt(plus(n1,n5),n2))).
###
#cnf(i_0_545, plain, (gt(plus(n1,n5),n1))).
###
#cnf(i_0_554, plain, (gt(plus(n1,n5),tptp_minus_1))).
##
#cnf(i_0_566, plain, (gt(plus(n1,n5),n0))).
###
#cnf(i_0_576, plain, (gt(plus(n1,n4),n3))).
###
#cnf(i_0_586, plain, (gt(plus(n1,n4),n2))).
##
#cnf(i_0_599, plain, (gt(plus(n1,n4),n1))).
###
#cnf(i_0_609, plain, (gt(plus(n1,n4),tptp_minus_1))).
###
#cnf(i_0_619, plain, (gt(plus(n1,n4),n0))).
##
#cnf(i_0_633, plain, (gt(plus(n1,n3),n2))).
###
#cnf(i_0_644, plain, (gt(plus(n1,n3),n1))).
###
#cnf(i_0_655, plain, (gt(plus(n1,n3),tptp_minus_1))).
##
#cnf(i_0_670, plain, (gt(plus(n1,n3),n0))).
###
#cnf(i_0_682, plain, (gt(plus(n1,n2),n1))).
###
#cnf(i_0_694, plain, (gt(plus(n1,n2),tptp_minus_1))).
##
#cnf(i_0_709, plain, (gt(plus(n1,n2),n0))).
###
#cnf(i_0_722, plain, (gt(plus(n1,n1),n0))).
###
#cnf(i_0_760, plain, (leq(minus(n4,n1),n5))).
#
#cnf(i_0_798, plain, (leq(minus(n3,n1),n5))).
#
#cnf(i_0_838, plain, (leq(minus(n2,n1),n5))).
##
#cnf(i_0_913, plain, (leq(minus(tptp_minus_1,n1),n5))).
#
#cnf(i_0_990, plain, (leq(minus(n3,n1),n4))).
#
#cnf(i_0_1033, plain, (leq(minus(n2,n1),n4))).
##
#cnf(i_0_1116, plain, (leq(minus(tptp_minus_1,n1),n4))).
#
#cnf(i_0_1197, plain, (leq(minus(n2,n1),n3))).
#
#cnf(i_0_1284, plain, (leq(minus(tptp_minus_1,n1),n3))).
##
#cnf(i_0_1412, plain, (leq(minus(tptp_minus_1,n1),n2))).
#
#cnf(i_0_1505, plain, (leq(minus(tptp_minus_1,n1),n1))).
#
#cnf(i_0_1562, plain, (leq(tptp_minus_1,plus(n1,n1)))).
###
#cnf(i_0_1814, plain, (leq(n4,plus(n1,n5)))).
########
#cnf(i_0_1826, plain, (leq(n3,plus(n1,n5)))).
#######
#cnf(i_0_1837, plain, (leq(n2,plus(n1,n5)))).
#####
#cnf(i_0_1847, plain, (leq(n1,plus(n1,n5)))).
###
#cnf(i_0_1852, plain, (leq(tptp_minus_1,plus(n1,n5)))).
##
#cnf(i_0_1863, plain, (leq(n0,plus(n1,n5)))).
###
#cnf(i_0_1871, plain, (leq(n3,plus(n1,n4)))).
#######
#cnf(i_0_1882, plain, (leq(n2,plus(n1,n4)))).
#####
#cnf(i_0_1892, plain, (leq(n1,plus(n1,n4)))).
###
#cnf(i_0_1897, plain, (leq(tptp_minus_1,plus(n1,n4)))).
##
#cnf(i_0_1908, plain, (leq(n0,plus(n1,n4)))).
###
#cnf(i_0_1916, plain, (leq(n2,plus(n1,n3)))).
#####
#cnf(i_0_1926, plain, (leq(n1,plus(n1,n3)))).
###
#cnf(i_0_1931, plain, (leq(tptp_minus_1,plus(n1,n3)))).
###
#cnf(i_0_1942, plain, (leq(n0,plus(n1,n3)))).
##
#cnf(i_0_1950, plain, (leq(n1,plus(n1,n2)))).
###
#cnf(i_0_1955, plain, (leq(tptp_minus_1,plus(n1,n2)))).
###
#cnf(i_0_1966, plain, (leq(n0,plus(n1,n2)))).
############
#cnf(i_0_1804, plain, (X1=X2|~leq(X1,X2)|~leq(X2,X1))).
#
#cnf(i_0_1989, negated_conjecture, (pv1325=n3|~leq(n3,pv1325))).
#
#cnf(i_0_1990, negated_conjecture, (s_worst7=n3|~leq(n3,s_worst7))).
##
#cnf(i_0_1991, negated_conjecture, (s_sworst7=n3|~leq(n3,s_sworst7))).
#
#cnf(i_0_1992, negated_conjecture, (s_best7=n3|~leq(n3,s_best7))).
###############
#cnf(i_0_419, plain, (gt(X1,X2)|~gt(minus(X1,n1),X2))).
#
#cnf(i_0_2058, plain, (gt(n0,X1)|~gt(tptp_minus_1,X1))).
###
#cnf(i_0_2069, plain, (gt(n0,minus(tptp_minus_1,n1)))).
#
#cnf(i_0_444, plain, (gt(X1,X2)|~gt(X1,plus(n1,X2)))).
###
#cnf(i_0_445, plain, (gt(X1,X2)|~gt(X1,plus(X2,n1)))).
##########################
#cnf(i_0_5, plain, (leq(X1,X3)|~leq(X2,X3)|~leq(X1,X2))).
##
#cnf(i_0_2242, negated_conjecture, (leq(X1,n3)|~leq(X1,pv1325))).
#
#cnf(i_0_2243, negated_conjecture, (leq(X1,n3)|~leq(X1,s_worst7))).
#
#cnf(i_0_2333, negated_conjecture, (leq(uniform_int_rnd(X1,s_worst7),n3))).
##
#cnf(i_0_2244, negated_conjecture, (leq(X1,n3)|~leq(X1,s_sworst7))).
#
#cnf(i_0_2346, negated_conjecture, (leq(uniform_int_rnd(X1,s_sworst7),n3))).
#
#cnf(i_0_2245, negated_conjecture, (leq(X1,n3)|~leq(X1,s_best7))).
##
#cnf(i_0_2359, negated_conjecture, (leq(uniform_int_rnd(X1,s_best7),n3))).
####
#cnf(i_0_2206, negated_conjecture, (leq(pv1325,X1)|~leq(n3,X1))).
#
#cnf(i_0_2207, negated_conjecture, (leq(s_worst7,X1)|~leq(n3,X1))).
#
#cnf(i_0_2208, negated_conjecture, (leq(s_sworst7,X1)|~leq(n3,X1))).
##
#cnf(i_0_2209, negated_conjecture, (leq(s_best7,X1)|~leq(n3,X1))).
####
#cnf(i_0_2259, plain, (leq(X1,n5)|~leq(X1,n4))).
#
#cnf(i_0_2261, plain, (leq(X1,n5)|~leq(X1,n3))).
#
#cnf(i_0_2262, plain, (leq(X1,n4)|~leq(X1,n3))).
##
#cnf(i_0_2265, negated_conjecture, (leq(X1,pv1325)|~leq(X1,n2))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(leq_succ_gt, axiom, ![X1, X2]:(leq(succ(X1),X2)=>gt(X2,X1)), file('./p/SWV138+1.tptp', leq_succ_gt)).
fof(succ_plus_1_r, axiom, ![X1]:plus(X1,n1)=succ(X1), file('./p/SWV138+1.tptp', succ_plus_1_r)).
fof(successor_1, axiom, succ(n0)=n1, file('./p/SWV138+1.tptp', successor_1)).
fof(gauss_array_0008, conjecture, (((((((((leq(n0,s_best7)&leq(n0,s_sworst7))&leq(n0,s_worst7))&leq(n2,pv1325))&leq(s_best7,n3))&leq(s_sworst7,n3))&leq(s_worst7,n3))&leq(pv1325,n3))&gt(a_select2(s_values7,pv1325),a_select2(s_values7,s_best7)))=>leq(n0,pv1325)), file('./p/SWV138+1.tptp', gauss_array_0008)).
fof(leq_gt1, axiom, ![X1, X2]:(gt(X2,X1)=>leq(X1,X2)), file('./p/SWV138+1.tptp', leq_gt1)).
fof(transitivity_leq, axiom, ![X1, X2, X3]:((leq(X1,X2)&leq(X2,X3))=>leq(X1,X3)), file('./p/SWV138+1.tptp', transitivity_leq)).
fof(gt_2_1, axiom, gt(n2,n1), file('./p/SWV138+1.tptp', gt_2_1)).
fof(c_0_7, plain, ![X150, X151]:(~leq(succ(X150),X151)|gt(X151,X150)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[leq_succ_gt])])).
fof(c_0_8, plain, ![X135]:plus(X135,n1)=succ(X135), inference(variable_rename,[status(thm)],[succ_plus_1_r])).
cnf(c_0_9, plain, (gt(X2,X1)|~leq(succ(X1),X2)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, plain, (plus(X1,n1)=succ(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_11, plain, (succ(n0)=n1), inference(split_conjunct,[status(thm)],[successor_1])).
fof(c_0_12, negated_conjecture, ~((((((((((leq(n0,s_best7)&leq(n0,s_sworst7))&leq(n0,s_worst7))&leq(n2,pv1325))&leq(s_best7,n3))&leq(s_sworst7,n3))&leq(s_worst7,n3))&leq(pv1325,n3))&gt(a_select2(s_values7,pv1325),a_select2(s_values7,s_best7)))=>leq(n0,pv1325))), inference(assume_negation,[status(cth)],[gauss_array_0008])).
fof(c_0_13, plain, ![X42, X43]:(~gt(X43,X42)|leq(X42,X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[leq_gt1])])).
cnf(c_0_14, plain, (gt(X2,X1)|~leq(plus(X1,n1),X2)), inference(rw,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_15, plain, (plus(n0,n1)=n1), inference(rw,[status(thm)],[c_0_11, c_0_10])).
fof(c_0_16, plain, ![X35, X36, X37]:(~leq(X35,X36)|~leq(X36,X37)|leq(X35,X37)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[transitivity_leq])])).
fof(c_0_17, negated_conjecture, (((((((((leq(n0,s_best7)&leq(n0,s_sworst7))&leq(n0,s_worst7))&leq(n2,pv1325))&leq(s_best7,n3))&leq(s_sworst7,n3))&leq(s_worst7,n3))&leq(pv1325,n3))&gt(a_select2(s_values7,pv1325),a_select2(s_values7,s_best7)))&~leq(n0,pv1325)), inference(fof_nnf,[status(thm)],[c_0_12])).
cnf(c_0_18, plain, (leq(X2,X1)|~gt(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_19, plain, (gt(X1,n0)|~leq(n1,X1)), inference(pm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_20, plain, (leq(X1,X3)|~leq(X1,X2)|~leq(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_21, negated_conjecture, (leq(n2,pv1325)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_22, plain, (gt(n2,n1)), inference(split_conjunct,[status(thm)],[gt_2_1])).
cnf(c_0_23, plain, (leq(n0,X1)|~leq(n1,X1)), inference(pm,[status(thm)],[c_0_18, c_0_19])).
cnf(c_0_24, negated_conjecture, (leq(X1,pv1325)|~leq(X1,n2)), inference(pm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_25, plain, (leq(n1,n2)), inference(pm,[status(thm)],[c_0_18, c_0_22])).
cnf(c_0_26, negated_conjecture, (~leq(n0,pv1325)), inference(split_conjunct,[status(thm)],[c_0_17])).
cnf(c_0_27, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_23, c_0_24]), c_0_25])]), c_0_26]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 28
# Proof object clause steps            : 15
# Proof object formula steps           : 13
# Proof object conjectures             : 7
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 5
# Proof object simplifying inferences  : 5
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 85
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 202
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 200
# Processed clauses                    : 553
# ...of these trivial                  : 60
# ...subsumed                          : 181
# ...remaining for further processing  : 312
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 5
# Generated clauses                    : 2246
# ...of the previous two non-trivial   : 1455
# Contextual simplify-reflections      : 0
# Paramodulations                      : 2242
# Factorizations                       : 4
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 307
#    Positive orientable unit clauses  : 123
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 33
#    Non-unit-clauses                  : 150
# Current number of unprocessed clauses: 1098
# ...number of literals in the above   : 3043
# Current number of archived formulas  : 0
# Current number of archived clauses   : 7
# Clause-clause subsumption calls (NU) : 3284
# Rec. Clause-clause subsumption calls : 3186
# Non-unit clause-clause subsumptions  : 60
# Unit Clause-clause subsumption calls : 1389
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 11
# BW rewrite match successes           : 10
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 20780
