# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_134, plain, (true)).
#
#cnf(i_0_4, plain, (leq(X1,X1))).
#
#cnf(i_0_148, plain, (gt(n0,tptp_minus_1))).
#
#cnf(i_0_3, plain, (~gt(X1,X1))).
#
#cnf(i_0_155, plain, (gt(n1,n0))).
#
#cnf(i_0_149, plain, (gt(n1,tptp_minus_1))).
#
#cnf(i_0_156, plain, (gt(n2,n0))).
#
#cnf(i_0_135, plain, (def!=use)).
#
#cnf(i_0_150, plain, (gt(n2,tptp_minus_1))).
#
#cnf(i_0_161, plain, (gt(n2,n1))).
#
#cnf(i_0_157, plain, (gt(n3,n0))).
#
#cnf(i_0_238, plain, (n0!=X1|n1!=X2|~epred2_2(X2,X1))).
#
#cnf(i_0_151, plain, (gt(n3,tptp_minus_1))).
#
#cnf(i_0_162, plain, (gt(n3,n1))).
#
#cnf(i_0_166, plain, (gt(n3,n2))).
#
#cnf(i_0_243, plain, (n0!=X1|n1!=X2|~epred2_2(X1,X2))).
#
#cnf(i_0_152, plain, (gt(n4,n0))).
#
#cnf(i_0_145, plain, (gt(n4,tptp_minus_1))).
#
#cnf(i_0_158, plain, (gt(n4,n1))).
#
#cnf(i_0_237, plain, (n0!=X1|n2!=X2|~epred2_2(X2,X1))).
#
#cnf(i_0_163, plain, (gt(n4,n2))).
#
#cnf(i_0_167, plain, (gt(n4,n3))).
#
#cnf(i_0_153, plain, (gt(n5,n0))).
#
#cnf(i_0_242, plain, (n0!=X1|n2!=X2|~epred2_2(X1,X2))).
#
#cnf(i_0_146, plain, (gt(n5,tptp_minus_1))).
#
#cnf(i_0_159, plain, (gt(n5,n1))).
#
#cnf(i_0_164, plain, (gt(n5,n2))).
#
#cnf(i_0_236, plain, (n0!=X1|n3!=X2|~epred2_2(X2,X1))).
#
#cnf(i_0_168, plain, (gt(n5,n3))).
#
#cnf(i_0_142, plain, (gt(n5,n4))).
#
#cnf(i_0_154, plain, (gt(n400,n0))).
#
#cnf(i_0_241, plain, (n0!=X1|n3!=X2|~epred2_2(X1,X2))).
#
#cnf(i_0_147, plain, (gt(n400,tptp_minus_1))).
#
#cnf(i_0_160, plain, (gt(n400,n1))).
#
#cnf(i_0_165, plain, (gt(n400,n2))).
#
#cnf(i_0_235, plain, (n0!=X1|n4!=X2|~epred2_2(X2,X1))).
#
#cnf(i_0_169, plain, (gt(n400,n3))).
#
#cnf(i_0_143, plain, (gt(n400,n4))).
#
#cnf(i_0_144, plain, (gt(n400,n5))).
#
#cnf(i_0_240, plain, (n0!=X1|n4!=X2|~epred2_2(X1,X2))).
#
#cnf(i_0_141, negated_conjecture, (leq(n0,esk24_0))).
#
#cnf(i_0_140, negated_conjecture, (leq(n0,esk25_0))).
#
#cnf(i_0_139, negated_conjecture, (leq(esk24_0,n5))).
#
#cnf(i_0_234, plain, (n0!=X1|n5!=X2|~epred2_2(X2,X1))).
#
#cnf(i_0_138, negated_conjecture, (leq(esk25_0,n5))).
#
#cnf(i_0_137, negated_conjecture, (epred2_2(esk24_0,esk25_0))).
#
#cnf(i_0_7, plain, (gt(X2,X1)|~lt(X1,X2))).
#
#cnf(i_0_244, plain, (esk24_0!=n0|esk25_0!=n1)).
#
#cnf(i_0_10, plain, (leq(X2,X1)|~gt(X1,X2))).
#
#cnf(i_0_253, plain, (leq(n5,n400))).
#
#cnf(i_0_254, plain, (leq(n4,n400))).
#
#cnf(i_0_245, plain, (esk24_0!=n0|esk25_0!=n2)).
#
#cnf(i_0_255, plain, (leq(n3,n400))).
#
#cnf(i_0_256, plain, (leq(n2,n400))).
#
#cnf(i_0_257, plain, (leq(n1,n400))).
#
#cnf(i_0_246, plain, (esk24_0!=n0|esk25_0!=n3)).
#
#cnf(i_0_258, plain, (leq(tptp_minus_1,n400))).
#
#cnf(i_0_259, plain, (leq(n0,n400))).
#
#cnf(i_0_260, plain, (leq(n4,n5))).
#
#cnf(i_0_247, plain, (esk24_0!=n0|esk25_0!=n4)).
#
#cnf(i_0_261, plain, (leq(n3,n5))).
#
#cnf(i_0_262, plain, (leq(n2,n5))).
#
#cnf(i_0_263, plain, (leq(n1,n5))).
#
#cnf(i_0_248, plain, (esk25_0!=n0|esk24_0!=n1)).
#
#cnf(i_0_264, plain, (leq(tptp_minus_1,n5))).
#
#cnf(i_0_265, plain, (leq(n0,n5))).
#
#cnf(i_0_266, plain, (leq(n3,n4))).
#
#cnf(i_0_249, plain, (esk25_0!=n0|esk24_0!=n2)).
#
#cnf(i_0_267, plain, (leq(n2,n4))).
#
#cnf(i_0_268, plain, (leq(n1,n4))).
#
#cnf(i_0_269, plain, (leq(tptp_minus_1,n4))).
#
#cnf(i_0_250, plain, (esk25_0!=n0|esk24_0!=n3)).
#
#cnf(i_0_270, plain, (leq(n0,n4))).
#
#cnf(i_0_271, plain, (leq(n2,n3))).
#
#cnf(i_0_272, plain, (leq(n1,n3))).
#
#cnf(i_0_251, plain, (esk25_0!=n0|esk24_0!=n4)).
#
#cnf(i_0_273, plain, (leq(tptp_minus_1,n3))).
#
#cnf(i_0_274, plain, (leq(n0,n3))).
#
#cnf(i_0_275, plain, (leq(n1,n2))).
#
#cnf(i_0_252, plain, (esk25_0!=n0|esk24_0!=n5)).
#
#cnf(i_0_276, plain, (leq(tptp_minus_1,n2))).
#
#cnf(i_0_277, plain, (leq(n0,n2))).
#
#cnf(i_0_278, plain, (leq(tptp_minus_1,n1))).
#
#cnf(i_0_239, plain, (n0!=X1|n5!=X2|~epred2_2(X1,X2))).
#
#cnf(i_0_279, plain, (leq(n0,n1))).
#
#cnf(i_0_280, plain, (leq(tptp_minus_1,n0))).
#
#cnf(i_0_9, plain, (leq(X2,X1)|~geq(X1,X2))).
#
#cnf(i_0_281, negated_conjecture, (esk24_0!=n0|esk25_0!=n5)).
#
#cnf(i_0_6, plain, (lt(X2,X1)|~gt(X1,X2))).
#
#cnf(i_0_8, plain, (geq(X2,X1)|~leq(X1,X2))).
#
#cnf(i_0_207, plain, (n0=X1|~epred2_2(X2,X1))).
#
#cnf(i_0_233, plain, (n1!=X2|n1!=X1|~epred2_2(X1,X2))).
#
#cnf(i_0_284, negated_conjecture, (esk25_0=n0)).
#
#cnf(i_0_137, negated_conjecture, (epred2_2(esk24_0,n0))).
#
#cnf(i_0_208, plain, (n0=X1|~epred2_2(X1,X2))).
#
#cnf(i_0_252, plain, (esk24_0!=n5)).
#
#cnf(i_0_297, negated_conjecture, (esk24_0=n0)).
#
#cnf(i_0_137, negated_conjecture, (epred2_2(n0,n0))).
#
#cnf(i_0_206, plain, (n3=X1|~epred2_2(X1,X2))).
#
#cnf(i_0_251, plain, (n4!=n0)).
#
#cnf(i_0_309, negated_conjecture, (n3=n0)).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(quaternion_ds1_symm_0041, conjecture, ![X14, X18]:((((leq(n0,X14)&leq(n0,X18))&leq(X14,n5))&leq(X18,n5))=>((((((((((((((((((((((((((((((((((((((~((n0=X14&n1=X18))&~((n0=X14&n2=X18)))&~((n0=X14&n3=X18)))&~((n0=X14&n4=X18)))&~((n0=X14&n5=X18)))&~((n0=X18&n1=X14)))&~((n0=X18&n2=X14)))&~((n0=X18&n3=X14)))&~((n0=X18&n4=X14)))&~((n0=X18&n5=X14)))&~((n1=X14&n1=X18)))&~((n1=X14&n2=X18)))&~((n1=X14&n3=X18)))&~((n1=X14&n4=X18)))&~((n1=X14&n5=X18)))&~((n1=X18&n2=X14)))&~((n1=X18&n3=X14)))&~((n1=X18&n4=X14)))&~((n1=X18&n5=X14)))&~((n2=X14&n2=X18)))&~((n2=X14&n3=X18)))&~((n2=X14&n4=X18)))&~((n2=X14&n5=X18)))&~((n2=X18&n3=X14)))&~((n2=X18&n4=X14)))&~((n2=X18&n5=X14)))&~((n3=X14&n3=X18)))&~((n3=X14&n4=X18)))&~((n3=X14&n5=X18)))&~((n3=X18&n4=X14)))&~((n3=X18&n5=X14)))&~((n4=X14&n4=X18)))&~((n4=X14&n5=X18)))&~((n4=X18&n5=X14)))&~((n5=X14&n5=X18)))&n0=X14)&n0=X18)&n3=X14)=>n0=times(divide(n1,n400),a_select2(sigma,n0)))), file('./p/SWV212+1.tptp', quaternion_ds1_symm_0041)).
fof(irreflexivity_gt, axiom, ![X1]:~(gt(X1,X1)), file('./p/SWV212+1.tptp', irreflexivity_gt)).
fof(gt_3_0, axiom, gt(n3,n0), file('./p/SWV212+1.tptp', gt_3_0)).
fof(c_0_3, plain, ![X18, X14]:(epred2_2(X14,X18)<=>(((((((((((((((((((((((((((((((((((((~((n0=X14&n1=X18))&~((n0=X14&n2=X18)))&~((n0=X14&n3=X18)))&~((n0=X14&n4=X18)))&~((n0=X14&n5=X18)))&~((n0=X18&n1=X14)))&~((n0=X18&n2=X14)))&~((n0=X18&n3=X14)))&~((n0=X18&n4=X14)))&~((n0=X18&n5=X14)))&~((n1=X14&n1=X18)))&~((n1=X14&n2=X18)))&~((n1=X14&n3=X18)))&~((n1=X14&n4=X18)))&~((n1=X14&n5=X18)))&~((n1=X18&n2=X14)))&~((n1=X18&n3=X14)))&~((n1=X18&n4=X14)))&~((n1=X18&n5=X14)))&~((n2=X14&n2=X18)))&~((n2=X14&n3=X18)))&~((n2=X14&n4=X18)))&~((n2=X14&n5=X18)))&~((n2=X18&n3=X14)))&~((n2=X18&n4=X14)))&~((n2=X18&n5=X14)))&~((n3=X14&n3=X18)))&~((n3=X14&n4=X18)))&~((n3=X14&n5=X18)))&~((n3=X18&n4=X14)))&~((n3=X18&n5=X14)))&~((n4=X14&n4=X18)))&~((n4=X14&n5=X18)))&~((n4=X18&n5=X14)))&~((n5=X14&n5=X18)))&n0=X14)&n0=X18)&n3=X14)), introduced(definition)).
fof(c_0_4, plain, ![X18, X14]:(epred2_2(X14,X18)=>(((((((((((((((((((((((((((((((((((((~((n0=X14&n1=X18))&~((n0=X14&n2=X18)))&~((n0=X14&n3=X18)))&~((n0=X14&n4=X18)))&~((n0=X14&n5=X18)))&~((n0=X18&n1=X14)))&~((n0=X18&n2=X14)))&~((n0=X18&n3=X14)))&~((n0=X18&n4=X14)))&~((n0=X18&n5=X14)))&~((n1=X14&n1=X18)))&~((n1=X14&n2=X18)))&~((n1=X14&n3=X18)))&~((n1=X14&n4=X18)))&~((n1=X14&n5=X18)))&~((n1=X18&n2=X14)))&~((n1=X18&n3=X14)))&~((n1=X18&n4=X14)))&~((n1=X18&n5=X14)))&~((n2=X14&n2=X18)))&~((n2=X14&n3=X18)))&~((n2=X14&n4=X18)))&~((n2=X14&n5=X18)))&~((n2=X18&n3=X14)))&~((n2=X18&n4=X14)))&~((n2=X18&n5=X14)))&~((n3=X14&n3=X18)))&~((n3=X14&n4=X18)))&~((n3=X14&n5=X18)))&~((n3=X18&n4=X14)))&~((n3=X18&n5=X14)))&~((n4=X14&n4=X18)))&~((n4=X14&n5=X18)))&~((n4=X18&n5=X14)))&~((n5=X14&n5=X18)))&n0=X14)&n0=X18)&n3=X14)), inference(split_equiv,[status(thm)],[c_0_3])).
fof(c_0_5, negated_conjecture, ~(![X14, X18]:((((leq(n0,X14)&leq(n0,X18))&leq(X14,n5))&leq(X18,n5))=>(epred2_2(X14,X18)=>n0=times(divide(n1,n400),a_select2(sigma,n0))))), inference(apply_def,[status(thm)],[inference(assume_negation,[status(cth)],[quaternion_ds1_symm_0041]), c_0_3])).
fof(c_0_6, plain, ![X202, X203]:((((((((((((((((((((((((((((((((((((((n0!=X203|n1!=X202|~epred2_2(X203,X202))&(n0!=X203|n2!=X202|~epred2_2(X203,X202)))&(n0!=X203|n3!=X202|~epred2_2(X203,X202)))&(n0!=X203|n4!=X202|~epred2_2(X203,X202)))&(n0!=X203|n5!=X202|~epred2_2(X203,X202)))&(n0!=X202|n1!=X203|~epred2_2(X203,X202)))&(n0!=X202|n2!=X203|~epred2_2(X203,X202)))&(n0!=X202|n3!=X203|~epred2_2(X203,X202)))&(n0!=X202|n4!=X203|~epred2_2(X203,X202)))&(n0!=X202|n5!=X203|~epred2_2(X203,X202)))&(n1!=X203|n1!=X202|~epred2_2(X203,X202)))&(n1!=X203|n2!=X202|~epred2_2(X203,X202)))&(n1!=X203|n3!=X202|~epred2_2(X203,X202)))&(n1!=X203|n4!=X202|~epred2_2(X203,X202)))&(n1!=X203|n5!=X202|~epred2_2(X203,X202)))&(n1!=X202|n2!=X203|~epred2_2(X203,X202)))&(n1!=X202|n3!=X203|~epred2_2(X203,X202)))&(n1!=X202|n4!=X203|~epred2_2(X203,X202)))&(n1!=X202|n5!=X203|~epred2_2(X203,X202)))&(n2!=X203|n2!=X202|~epred2_2(X203,X202)))&(n2!=X203|n3!=X202|~epred2_2(X203,X202)))&(n2!=X203|n4!=X202|~epred2_2(X203,X202)))&(n2!=X203|n5!=X202|~epred2_2(X203,X202)))&(n2!=X202|n3!=X203|~epred2_2(X203,X202)))&(n2!=X202|n4!=X203|~epred2_2(X203,X202)))&(n2!=X202|n5!=X203|~epred2_2(X203,X202)))&(n3!=X203|n3!=X202|~epred2_2(X203,X202)))&(n3!=X203|n4!=X202|~epred2_2(X203,X202)))&(n3!=X203|n5!=X202|~epred2_2(X203,X202)))&(n3!=X202|n4!=X203|~epred2_2(X203,X202)))&(n3!=X202|n5!=X203|~epred2_2(X203,X202)))&(n4!=X203|n4!=X202|~epred2_2(X203,X202)))&(n4!=X203|n5!=X202|~epred2_2(X203,X202)))&(n4!=X202|n5!=X203|~epred2_2(X203,X202)))&(n5!=X203|n5!=X202|~epred2_2(X203,X202)))&(n0=X203|~epred2_2(X203,X202)))&(n0=X202|~epred2_2(X203,X202)))&(n3=X203|~epred2_2(X203,X202))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_4])])])).
fof(c_0_7, negated_conjecture, ((((leq(n0,esk24_0)&leq(n0,esk25_0))&leq(esk24_0,n5))&leq(esk25_0,n5))&(epred2_2(esk24_0,esk25_0)&n0!=times(divide(n1,n400),a_select2(sigma,n0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_5])])])).
cnf(c_0_8, plain, (n0=X1|~epred2_2(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_9, negated_conjecture, (epred2_2(esk24_0,esk25_0)), inference(split_conjunct,[status(thm)],[c_0_7])).
cnf(c_0_10, negated_conjecture, (esk25_0=n0), inference(pm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_11, plain, (n0=X1|~epred2_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_12, negated_conjecture, (epred2_2(esk24_0,n0)), inference(rw,[status(thm)],[c_0_9, c_0_10])).
cnf(c_0_13, negated_conjecture, (esk24_0=n0), inference(pm,[status(thm)],[c_0_11, c_0_12])).
cnf(c_0_14, plain, (n3=X1|~epred2_2(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_6])).
cnf(c_0_15, negated_conjecture, (epred2_2(n0,n0)), inference(rw,[status(thm)],[c_0_12, c_0_13])).
fof(c_0_16, plain, ![X33]:~gt(X33,X33), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[irreflexivity_gt])])).
cnf(c_0_17, plain, (gt(n3,n0)), inference(split_conjunct,[status(thm)],[gt_3_0])).
cnf(c_0_18, negated_conjecture, (n3=n0), inference(pm,[status(thm)],[c_0_14, c_0_15])).
cnf(c_0_19, plain, (~gt(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_20, plain, ($false), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_17, c_0_18]), c_0_19]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 21
# Proof object clause steps            : 12
# Proof object formula steps           : 9
# Proof object conjectures             : 9
# Proof object clause conjectures      : 6
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 6
# Proof object initial formulas used   : 3
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 4
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 92
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 243
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 241
# Processed clauses                    : 101
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 101
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 34
# Generated clauses                    : 66
# ...of the previous two non-trivial   : 84
# Contextual simplify-reflections      : 0
# Paramodulations                      : 66
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 67
#    Positive orientable unit clauses  : 48
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 16
# Current number of unprocessed clauses: 217
# ...number of literals in the above   : 883
# Current number of archived formulas  : 0
# Current number of archived clauses   : 36
# Clause-clause subsumption calls (NU) : 158
# Rec. Clause-clause subsumption calls : 94
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 41
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 14978
