# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_30, negated_conjecture, (X1=k|element(X1,k)|X1!=m)).
#
#cnf(i_0_31, negated_conjecture, (X1=k|element(X1,k)|X1!=n)).
#
#cnf(i_0_21, negated_conjecture, (X1=m|f(X1)!=X1|~element(X1,m))).
#
#cnf(i_0_17, negated_conjecture, (m!=n)).
#
#cnf(i_0_26, negated_conjecture, (X1=n|element(X1,n)|g(X1)!=X1)).
#
#cnf(i_0_20, negated_conjecture, (X1=m|f(X1)!=m|~element(X1,m))).
#
#cnf(i_0_25, negated_conjecture, (X1=n|element(X1,n)|g(X1)!=n)).
#
#cnf(i_0_19, negated_conjecture, (m!=k)).
#
#cnf(i_0_22, negated_conjecture, (X1=m|element(X1,f(X1))|~element(X1,m))).
#
#cnf(i_0_23, negated_conjecture, (X1=m|element(f(X1),X1)|~element(X1,m))).
#
#cnf(i_0_27, negated_conjecture, (X1=n|element(X1,g(X1))|element(X1,n))).
#
#cnf(i_0_18, negated_conjecture, (k!=n)).
#
#cnf(i_0_28, negated_conjecture, (X1=n|element(g(X1),X1)|element(X1,n))).
#
#cnf(i_0_32, negated_conjecture, (X1=m|X1=n|X1=k|~element(X1,k))).
#
#cnf(i_0_33, negated_conjecture, (g(m)=m|element(m,n)|f(g(m))!=m)).
#
#cnf(i_0_34, negated_conjecture, (g(m)=m|element(m,n)|f(g(m))!=g(m))).
#
#cnf(i_0_29, negated_conjecture, (X2=n|X1=n|X1=X2|~element(X1,n)|~element(X2,X1)|~element(X1,X2))).
#
#cnf(i_0_41, negated_conjecture, (X1=n|X1=k|X1!=m|~element(X1,n)|~element(k,X1))).
#
#cnf(i_0_53, negated_conjecture, (g(m)=m|f(g(m))!=m|~element(k,m))).
#
#cnf(i_0_54, negated_conjecture, (g(m)=m|f(g(m))!=g(m)|~element(k,m))).
#
#cnf(i_0_56, negated_conjecture, (X1=n|X1=k|g(X1)!=X1|X1!=m|~element(k,X1))).
#
#cnf(i_0_55, negated_conjecture, (X1=n|X1=k|g(X1)!=n|X1!=m|~element(k,X1))).
#
#cnf(i_0_45, negated_conjecture, (X1=n|k=X1|X1!=m|~element(k,n)|~element(k,X1))).
#
#cnf(i_0_63, negated_conjecture, (X1=n|k=X1|g(k)!=n|X1!=m|~element(k,X1))).
#
#cnf(i_0_64, negated_conjecture, (X1=n|k=X1|g(k)!=k|X1!=m|~element(k,X1))).
#
#cnf(i_0_51, negated_conjecture, (X1=n|X2=n|X2=X1|g(X2)!=X2|~element(X1,X2)|~element(X2,X1))).
###
#cnf(i_0_50, negated_conjecture, (X1=n|X2=n|X2=X1|g(X2)!=n|~element(X1,X2)|~element(X2,X1))).
###
#cnf(i_0_24, negated_conjecture, (X2=m|X1=m|X1=X2|element(X1,m)|~element(X2,X1)|~element(X1,X2))).
#
#cnf(i_0_105, negated_conjecture, (X1=m|k=X1|element(k,m)|X1!=n|~element(k,X1))).
#
#cnf(i_0_110, negated_conjecture, (X1=m|k=X1|X1!=n|~element(k,n)|~element(k,X1))).
#
#cnf(i_0_108, negated_conjecture, (X1=m|k=X1|f(k)!=m|X1!=n|~element(k,X1))).
#
#cnf(i_0_121, negated_conjecture, (g(k)!=n|~element(k,n))).
#
#cnf(i_0_109, negated_conjecture, (X1=m|k=X1|f(k)!=k|X1!=n|~element(k,X1))).
#
#cnf(i_0_111, negated_conjecture, (X1=m|k=X1|g(m)!=n|X1!=n|~element(k,X1))).
#
#cnf(i_0_112, negated_conjecture, (X1=m|k=X1|g(m)!=m|X1!=n|~element(k,X1))).
#
#cnf(i_0_129, negated_conjecture, (g(k)!=n)).
##
#cnf(i_0_114, negated_conjecture, (X1=m|k=X1|g(k)!=k|X1!=n|~element(k,X1))).
##
#cnf(i_0_146, negated_conjecture, (g(k)!=k)).
##
#cnf(i_0_99, negated_conjecture, (X1=m|X2=m|X1=X2|f(X1)!=X1|~element(X2,X1)|~element(X1,X2))).
#
#cnf(i_0_150, negated_conjecture, (X1=m|X1=k|f(X1)!=X1|X1!=n|~element(k,X1))).
##
#cnf(i_0_152, negated_conjecture, (X1=m|X1=n|f(X1)!=X1|g(X1)!=X1|~element(n,X1))).
#
#cnf(i_0_151, negated_conjecture, (X1=m|X1=n|g(X1)!=n|f(X1)!=X1|~element(n,X1))).
###
#cnf(i_0_158, negated_conjecture, (X1=m|n=X1|f(n)!=n|g(X1)!=X1|~element(n,X1))).
#
#cnf(i_0_85, negated_conjecture, (X1=m|k=X1|~element(k,n)|~element(X1,k)|~element(k,X1))).
###
#cnf(i_0_86, negated_conjecture, (X1=m|k=X1|g(m)!=n|~element(X1,k)|~element(k,X1))).
##
#cnf(i_0_87, negated_conjecture, (X1=m|k=X1|g(m)!=m|~element(X1,k)|~element(k,X1))).
########
#cnf(i_0_107, negated_conjecture, (X1=m|n=X1|element(n,m)|g(X1)!=X1|~element(n,X1))).
##
#cnf(i_0_192, negated_conjecture, (X1=m|n=X1|f(n)!=m|g(X1)!=X1|~element(n,X1))).
##
#cnf(i_0_157, negated_conjecture, (X1=m|n=X1|f(n)!=n|g(X1)!=n|~element(n,X1))).
##
#cnf(i_0_98, negated_conjecture, (X1=m|X2=m|X1=X2|f(X1)!=m|~element(X2,X1)|~element(X1,X2))).
#
#cnf(i_0_204, negated_conjecture, (X1=m|X1=k|f(X1)!=m|X1!=n|~element(k,X1))).
#
#cnf(i_0_206, negated_conjecture, (X1=m|X1=n|f(X1)!=m|g(X1)!=X1|~element(n,X1))).
###
#cnf(i_0_205, negated_conjecture, (X1=m|X1=n|f(X1)!=m|g(X1)!=n|~element(n,X1))).
###
#cnf(i_0_211, negated_conjecture, (X1=m|n=X1|f(n)!=m|g(X1)!=n|~element(n,X1))).
#
#cnf(i_0_106, negated_conjecture, (X1=m|n=X1|element(n,m)|g(X1)!=n|~element(n,X1))).
####
#cnf(i_0_36, negated_conjecture, (f(k)=m|f(k)=n|f(k)=k|~element(k,m))).
#
#cnf(i_0_35, negated_conjecture, (g(k)=m|element(k,n))).
##
#cnf(i_0_242, negated_conjecture, (g(k)=m|~element(k,n))).
#
#cnf(i_0_254, negated_conjecture, (g(k)=m)).
#
#cnf(i_0_255, negated_conjecture, (element(k,m)|element(k,n))).
##
#cnf(i_0_256, negated_conjecture, (element(m,k)|element(k,n))).
#
#cnf(i_0_257, negated_conjecture, (element(k,n)|f(k)!=m)).
#
#cnf(i_0_258, negated_conjecture, (element(k,n)|f(k)!=k)).
##
#cnf(i_0_260, negated_conjecture, (element(k,n)|g(m)!=n)).
#
#cnf(i_0_261, negated_conjecture, (element(k,n)|g(m)!=m)).
#
#cnf(i_0_268, negated_conjecture, (element(k,n)|~element(m,n)|~element(k,m))).
#
#cnf(i_0_281, negated_conjecture, (f(k)!=m)).
####
#cnf(i_0_298, negated_conjecture, (f(k)!=k)).
####
#cnf(i_0_321, negated_conjecture, (g(m)!=n)).
##
#cnf(i_0_328, negated_conjecture, (element(n,m)|g(m)!=m|~element(n,k))).
#
#cnf(i_0_264, negated_conjecture, (g(m)=m|element(k,n)|f(g(m))!=m)).
#
#cnf(i_0_341, negated_conjecture, (g(m)!=m)).
#
#cnf(i_0_364, negated_conjecture, (f(g(m))!=m|~element(k,n))).
###
#cnf(i_0_378, negated_conjecture, (f(g(m))!=m)).
#########
#cnf(i_0_263, negated_conjecture, (element(k,n)|f(g(m))!=g(m))).
#
#cnf(i_0_379, negated_conjecture, (element(n,m)|f(g(m))!=g(m)|~element(n,k))).
#########################
#cnf(i_0_262, negated_conjecture, (f(k)=n|element(k,n))).
##
#cnf(i_0_412, negated_conjecture, (f(k)=n|~element(k,n))).
#
#cnf(i_0_427, negated_conjecture, (f(k)=n)).
####
#cnf(i_0_430, negated_conjecture, (element(k,n)|~element(k,m))).
#
#cnf(i_0_433, negated_conjecture, (element(k,n))).
##
#cnf(i_0_431, negated_conjecture, (element(n,k)|~element(k,m))).
#
#cnf(i_0_435, negated_conjecture, (element(n,m)|~element(n,k))).
#
#cnf(i_0_110, negated_conjecture, (X1=m|k=X1|X1!=n|~element(k,X1))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(c_13, negated_conjecture, (X1=n|X2=n|X2=X1|~element(X1,n)|~element(X1,X2)|~element(X2,X1)), file('./p/SYN013-1.tptp', c_13)).
cnf(c_14, negated_conjecture, (X1=k|element(X1,k)|X1!=m), file('./p/SYN013-1.tptp', c_14)).
cnf(c_2, negated_conjecture, (n!=k), file('./p/SYN013-1.tptp', c_2)).
cnf(c_8, negated_conjecture, (X1=m|element(X1,m)|X2=m|X2=X1|~element(X1,X2)|~element(X2,X1)), file('./p/SYN013-1.tptp', c_8)).
cnf(c_15, negated_conjecture, (X1=k|element(X1,k)|X1!=n), file('./p/SYN013-1.tptp', c_15)).
cnf(c_3, negated_conjecture, (k!=m), file('./p/SYN013-1.tptp', c_3)).
cnf(c_1, negated_conjecture, (m!=n), file('./p/SYN013-1.tptp', c_1)).
cnf(c_10, negated_conjecture, (X1=n|element(X1,n)|g(X1)!=X1), file('./p/SYN013-1.tptp', c_10)).
cnf(c_9, negated_conjecture, (X1=n|element(X1,n)|g(X1)!=n), file('./p/SYN013-1.tptp', c_9)).
cnf(c_16, negated_conjecture, (X1=k|X1=m|X1=n|~element(X1,k)), file('./p/SYN013-1.tptp', c_16)).
cnf(c_12, negated_conjecture, (X1=n|element(X1,n)|element(g(X1),X1)), file('./p/SYN013-1.tptp', c_12)).
cnf(c_11, negated_conjecture, (X1=n|element(X1,n)|element(X1,g(X1))), file('./p/SYN013-1.tptp', c_11)).
cnf(c_4, negated_conjecture, (X1=m|~element(X1,m)|f(X1)!=m), file('./p/SYN013-1.tptp', c_4)).
cnf(c_5, negated_conjecture, (X1=m|~element(X1,m)|f(X1)!=X1), file('./p/SYN013-1.tptp', c_5)).
cnf(c_7, negated_conjecture, (X1=m|element(f(X1),X1)|~element(X1,m)), file('./p/SYN013-1.tptp', c_7)).
cnf(c_6, negated_conjecture, (X1=m|element(X1,f(X1))|~element(X1,m)), file('./p/SYN013-1.tptp', c_6)).
cnf(c_0_16, negated_conjecture, (X1=n|X2=n|X2=X1|~element(X1,n)|~element(X1,X2)|~element(X2,X1)), c_13).
cnf(c_0_17, negated_conjecture, (X1=k|element(X1,k)|X1!=m), c_14).
cnf(c_0_18, negated_conjecture, (n!=k), c_2).
cnf(c_0_19, negated_conjecture, (X1=m|element(X1,m)|X2=m|X2=X1|~element(X1,X2)|~element(X2,X1)), c_8).
cnf(c_0_20, negated_conjecture, (X1=k|element(X1,k)|X1!=n), c_15).
cnf(c_0_21, negated_conjecture, (k!=m), c_3).
cnf(c_0_22, negated_conjecture, (X1=n|k=X1|X1!=m|~element(k,n)|~element(k,X1)), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_16, c_0_17]), c_0_18])).
cnf(c_0_23, negated_conjecture, (X1=m|k=X1|element(k,m)|X1!=n|~element(k,X1)), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_19, c_0_20]), c_0_21])).
cnf(c_0_24, negated_conjecture, (m!=n), c_1).
cnf(c_0_25, negated_conjecture, (X1=n|element(X1,n)|g(X1)!=X1), c_10).
cnf(c_0_26, negated_conjecture, (X1=m|k=X1|X1!=n|~element(k,n)|~element(k,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_22, c_0_23]), c_0_24]), c_0_21])).
cnf(c_0_27, negated_conjecture, (X1=n|element(X1,n)|g(X1)!=n), c_9).
cnf(c_0_28, negated_conjecture, (X1=n|k=X1|g(k)!=k|X1!=m|~element(k,X1)), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_22, c_0_25]), c_0_18])).
cnf(c_0_29, negated_conjecture, (g(k)!=n|~element(k,n)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_26, c_0_27]), c_0_24]), c_0_18])).
cnf(c_0_30, negated_conjecture, (X1=m|k=X1|g(k)!=k|X1!=n|~element(k,X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_28, c_0_23]), c_0_24]), c_0_21])).
cnf(c_0_31, negated_conjecture, (X1=k|X1=m|X1=n|~element(X1,k)), c_16).
cnf(c_0_32, negated_conjecture, (X1=n|element(X1,n)|element(g(X1),X1)), c_12).
cnf(c_0_33, negated_conjecture, (g(k)!=n), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_29, c_0_27]), c_0_18])).
cnf(c_0_34, negated_conjecture, (g(k)!=k), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_30, c_0_25]), c_0_24]), c_0_18])).
cnf(c_0_35, negated_conjecture, (g(k)=m|element(k,n)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_31, c_0_32]), c_0_18]), c_0_33]), c_0_34])).
cnf(c_0_36, negated_conjecture, (g(k)=m|~element(k,n)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_26, c_0_35]), c_0_24]), c_0_18])).
cnf(c_0_37, negated_conjecture, (X1=n|element(X1,n)|element(X1,g(X1))), c_11).
cnf(c_0_38, negated_conjecture, (g(k)=m), inference(pm,[status(thm)],[c_0_36, c_0_35])).
cnf(c_0_39, negated_conjecture, (X1=m|~element(X1,m)|f(X1)!=m), c_4).
cnf(c_0_40, negated_conjecture, (element(k,m)|element(k,n)), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_37, c_0_38]), c_0_18])).
cnf(c_0_41, negated_conjecture, (X1=m|~element(X1,m)|f(X1)!=X1), c_5).
cnf(c_0_42, negated_conjecture, (X1=m|element(f(X1),X1)|~element(X1,m)), c_7).
cnf(c_0_43, negated_conjecture, (X1=m|k=X1|f(k)!=m|X1!=n|~element(k,X1)), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_39, c_0_23]), c_0_21])).
cnf(c_0_44, negated_conjecture, (element(k,n)|f(k)!=m), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_39, c_0_40]), c_0_21])).
cnf(c_0_45, negated_conjecture, (X1=m|k=X1|f(k)!=k|X1!=n|~element(k,X1)), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_41, c_0_23]), c_0_21])).
cnf(c_0_46, negated_conjecture, (element(k,n)|f(k)!=k), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_41, c_0_40]), c_0_21])).
cnf(c_0_47, negated_conjecture, (f(k)=m|f(k)=n|f(k)=k|~element(k,m)), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_31, c_0_42]), c_0_21])).
cnf(c_0_48, negated_conjecture, (f(k)!=m), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_43, c_0_44]), c_0_24]), c_0_18])).
cnf(c_0_49, negated_conjecture, (f(k)!=k), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_45, c_0_46]), c_0_24]), c_0_18])).
cnf(c_0_50, negated_conjecture, (f(k)=n|element(k,n)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_47, c_0_40]), c_0_48]), c_0_49])).
cnf(c_0_51, negated_conjecture, (f(k)=n|~element(k,n)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_26, c_0_50]), c_0_24]), c_0_18])).
cnf(c_0_52, negated_conjecture, (X1=m|element(X1,f(X1))|~element(X1,m)), c_6).
cnf(c_0_53, negated_conjecture, (f(k)=n), inference(pm,[status(thm)],[c_0_51, c_0_50])).
cnf(c_0_54, negated_conjecture, (element(k,n)|~element(k,m)), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_52, c_0_53]), c_0_21])).
cnf(c_0_55, negated_conjecture, (element(k,n)), inference(pm,[status(thm)],[c_0_54, c_0_40])).
cnf(c_0_56, negated_conjecture, (X1=m|k=X1|X1!=n|~element(k,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_26, c_0_55])])).
cnf(c_0_57, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_56, c_0_55]), c_0_24]), c_0_18]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 58
# Proof object clause steps            : 42
# Proof object formula steps           : 16
# Proof object conjectures             : 58
# Proof object clause conjectures      : 42
# Proof object formula conjectures     : 16
# Proof object initial clauses used    : 16
# Proof object initial formulas used   : 16
# Proof object generating inferences   : 25
# Proof object simplifying inferences  : 36
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 16
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 16
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 16
# Processed clauses                    : 163
# ...of these trivial                  : 2
# ...subsumed                          : 77
# ...remaining for further processing  : 84
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 12
# Backward-rewritten                   : 24
# Generated clauses                    : 423
# ...of the previous two non-trivial   : 391
# Contextual simplify-reflections      : 0
# Paramodulations                      : 423
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 48
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 39
# Current number of unprocessed clauses: 205
# ...number of literals in the above   : 784
# Current number of archived formulas  : 0
# Current number of archived clauses   : 36
# Clause-clause subsumption calls (NU) : 125
# Rec. Clause-clause subsumption calls : 64
# Non-unit clause-clause subsumptions  : 25
# Unit Clause-clause subsumption calls : 54
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2999
