# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_9, plain, (q|~r)).
#
#cnf(i_0_8, plain, (p|~r)).
#
#cnf(i_0_7, plain, (r|~q)).
#
#cnf(i_0_11, negated_conjecture, (~p|~q)).
#
#cnf(i_0_14, plain, (p|~q)).
#
#cnf(i_0_12, negated_conjecture, (p|q)).
#
#cnf(i_0_16, negated_conjecture, (p)).
#
#cnf(i_0_11, negated_conjecture, (~q)).
#
#cnf(i_0_10, plain, (r)).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(clause_2, axiom, (p|~r), file('./p/SYN044-1.tptp', clause_2)).
cnf(clause_1, axiom, (r|~q), file('./p/SYN044-1.tptp', clause_1)).
cnf(clause_6, negated_conjecture, (p|q), file('./p/SYN044-1.tptp', clause_6)).
cnf(clause_5, negated_conjecture, (~p|~q), file('./p/SYN044-1.tptp', clause_5)).
cnf(clause_4, axiom, (q|r|~p), file('./p/SYN044-1.tptp', clause_4)).
cnf(clause_3, axiom, (q|~r), file('./p/SYN044-1.tptp', clause_3)).
cnf(c_0_6, axiom, (p|~r), clause_2).
cnf(c_0_7, axiom, (r|~q), clause_1).
cnf(c_0_8, plain, (p|~q), inference(pm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_9, negated_conjecture, (p|q), clause_6).
cnf(c_0_10, negated_conjecture, (~p|~q), clause_5).
cnf(c_0_11, negated_conjecture, (p), inference(pm,[status(thm)],[c_0_8, c_0_9])).
cnf(c_0_12, axiom, (q|r|~p), clause_4).
cnf(c_0_13, negated_conjecture, (~q), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_11])])).
cnf(c_0_14, axiom, (q|~r), clause_3).
cnf(c_0_15, plain, (r), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_12, c_0_11])]), c_0_13])).
cnf(c_0_16, plain, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_14, c_0_15])]), c_0_13]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 17
# Proof object clause steps            : 11
# Proof object formula steps           : 6
# Proof object conjectures             : 6
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 2
# Proof object initial clauses used    : 6
# Proof object initial formulas used   : 6
# Proof object generating inferences   : 2
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 6
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 6
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 6
# Processed clauses                    : 9
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 9
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 5
# Generated clauses                    : 4
# ...of the previous two non-trivial   : 3
# Contextual simplify-reflections      : 0
# Paramodulations                      : 4
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 3
#    Positive orientable unit clauses  : 2
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 6
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 2
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 118
