# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_72, plain, (p_2_2(a))).
#
#cnf(i_0_70, plain, (p_1_1(a))).
#
#cnf(i_0_73, plain, (q_2_2(a))).
#
#cnf(i_0_51, negated_conjecture, (~p_1_3(X1))).
#
#cnf(i_0_71, plain, (q_1_1(a))).
#
#cnf(i_0_74, plain, (p_3_3(a))).
#
#cnf(i_0_75, plain, (q_3_3(a))).
#
#cnf(i_0_97, plain, (~sym_p_2_2(a))).
#
#cnf(i_0_66, plain, (p_2_2(X1)|~p_1_2(X1))).
#
#cnf(i_0_64, plain, (p_1_1(X1)|~p_1_2(X1))).
#
#cnf(i_0_68, plain, (q_2_2(X1)|~q_1_2(X1))).
#
#cnf(i_0_95, plain, (~sym_p_1_1(a))).
#
#cnf(i_0_65, plain, (q_1_1(X1)|~q_1_2(X1))).
#
#cnf(i_0_67, plain, (p_3_3(X1)|~p_2_3(X1))).
#
#cnf(i_0_69, plain, (q_3_3(X1)|~q_2_3(X1))).
#
#cnf(i_0_98, plain, (~sym_q_2_2(a))).
#
#cnf(i_0_91, plain, (sym_p_1_2(X1)|~sym_p_2_2(X1))).
#
#cnf(i_0_89, plain, (sym_p_1_2(X1)|~sym_p_1_1(X1))).
#
#cnf(i_0_92, plain, (sym_p_2_3(X1)|~sym_p_3_3(X1))).
#
#cnf(i_0_96, plain, (~sym_q_1_1(a))).
#
#cnf(i_0_94, plain, (sym_q_2_3(X1)|~sym_q_3_3(X1))).
#
#cnf(i_0_93, plain, (sym_q_1_2(X1)|~sym_q_2_2(X1))).
#
#cnf(i_0_90, plain, (sym_q_1_2(X1)|~sym_q_1_1(X1))).
#
#cnf(i_0_99, plain, (~sym_p_3_3(a))).
#
#cnf(i_0_79, plain, (sym_p_2_3(X1)|sym_p_1_2(X1))).
#
#cnf(i_0_80, plain, (sym_q_1_2(X1)|sym_q_2_3(X1))).
#
#cnf(i_0_61, plain, (q_1_3(X1)|~q_2_3(X1)|~p_1_2(X1))).
#
#cnf(i_0_100, plain, (~sym_q_3_3(a))).
#
#cnf(i_0_60, plain, (q_1_3(X1)|~q_1_2(X1)|~p_2_3(X1))).
#
#cnf(i_0_54, plain, (~p_2_3(X1)|~p_1_2(X1))).
#
#cnf(i_0_55, plain, (~q_1_2(X1)|~q_2_3(X1))).
#
#cnf(i_0_86, plain, (sym_q_2_3(X1)|sym_p_1_2(X1)|~sym_q_1_3(X1))).
#
#cnf(i_0_52, plain, (p_1_2(X1)|~p_1_1(X1)|~p_2_2(X1))).
#
#cnf(i_0_53, plain, (p_1_2(X1)|~q_1_1(X1)|~q_2_2(X1))).
#
#cnf(i_0_103, plain, (p_2_2(X1)|~q_1_1(X1)|~q_2_2(X1))).
#
#cnf(i_0_77, plain, (sym_p_1_1(X1)|sym_p_2_2(X1)|~sym_p_1_2(X1))).
#
#cnf(i_0_104, plain, (p_1_1(X1)|~q_1_1(X1)|~q_2_2(X1))).
#
#cnf(i_0_56, plain, (p_2_3(X1)|~p_3_3(X1)|~p_2_2(X1))).
#
#cnf(i_0_57, plain, (p_2_3(X1)|~q_3_3(X1)|~q_2_2(X1))).
#
#cnf(i_0_106, plain, (~sym_p_1_2(a))).
#
#cnf(i_0_112, plain, (p_3_3(X1)|~q_3_3(X1)|~q_2_2(X1))).
#
#cnf(i_0_63, plain, (q_2_3(X1)|~q_3_3(X1)|~p_2_2(X1))).
#
#cnf(i_0_62, plain, (q_2_3(X1)|~p_3_3(X1)|~q_2_2(X1))).
#
#cnf(i_0_110, plain, (~p_3_3(X1)|~p_2_2(X1)|~p_1_2(X1))).
#
#cnf(i_0_58, plain, (q_1_2(X1)|~q_1_1(X1)|~p_2_2(X1))).
#
#cnf(i_0_124, plain, (q_2_2(X1)|~q_1_1(X1)|~p_2_2(X1))).
#
#cnf(i_0_59, plain, (q_1_2(X1)|~q_2_2(X1)|~p_1_1(X1))).
#
#cnf(i_0_113, plain, (~q_3_3(X1)|~q_2_2(X1)|~p_1_2(X1))).
#
#cnf(i_0_119, plain, (q_3_3(X1)|~p_3_3(X1)|~q_2_2(X1))).
#
#cnf(i_0_128, plain, (q_1_1(X1)|~q_2_2(X1)|~p_1_1(X1))).
#
#cnf(i_0_133, plain, (p_2_2(X1)|~q_2_2(X1)|~p_1_1(X1))).
#
#cnf(i_0_117, plain, (~q_3_3(X1)|~q_1_2(X1)|~p_2_2(X1))).
#
#cnf(i_0_85, plain, (sym_q_1_2(X1)|sym_p_2_3(X1)|~sym_q_1_3(X1))).
#
#cnf(i_0_83, plain, (sym_q_1_1(X1)|sym_p_2_2(X1)|~sym_q_1_2(X1))).
#
#cnf(i_0_81, plain, (sym_p_3_3(X1)|sym_p_2_2(X1)|~sym_p_2_3(X1))).
#
#cnf(i_0_138, plain, (~sym_q_1_2(a))).
#
#cnf(i_0_142, plain, (sym_q_2_3(a))).
#
#cnf(i_0_88, plain, (sym_q_3_3(X1)|sym_p_2_2(X1)|~sym_q_2_3(X1))).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(sym_m_t2_2, axiom, (~sym_q_1_1(a)), file('./p/SYN099-1.003.tptp', sym_m_t2_2)).
cnf(sym_m_s2_type12_1, axiom, (sym_p_2_2(X1)|sym_q_1_1(X1)|~sym_q_1_2(X1)), file('./p/SYN099-1.003.tptp', sym_m_s2_type12_1)).
cnf(sym_m_t2_3, axiom, (~sym_p_2_2(a)), file('./p/SYN099-1.003.tptp', sym_m_t2_3)).
cnf(sym_m_s2_type11_4, axiom, (sym_q_2_3(X1)|sym_q_1_2(X1)|~sym_p_1_3(X1)), file('./p/SYN099-1.003.tptp', sym_m_s2_type11_4)).
cnf(sym_m_s2_goal_1, axiom, (sym_p_1_3(X1)), file('./p/SYN099-1.003.tptp', sym_m_s2_goal_1)).
cnf(sym_m_t2_6, axiom, (~sym_q_3_3(a)), file('./p/SYN099-1.003.tptp', sym_m_t2_6)).
cnf(sym_m_s2_type12_6, axiom, (sym_q_3_3(X1)|sym_p_2_2(X1)|~sym_q_2_3(X1)), file('./p/SYN099-1.003.tptp', sym_m_s2_type12_6)).
cnf(c_0_7, axiom, (~sym_q_1_1(a)), sym_m_t2_2).
cnf(c_0_8, axiom, (sym_p_2_2(X1)|sym_q_1_1(X1)|~sym_q_1_2(X1)), sym_m_s2_type12_1).
cnf(c_0_9, axiom, (~sym_p_2_2(a)), sym_m_t2_3).
cnf(c_0_10, axiom, (sym_q_2_3(X1)|sym_q_1_2(X1)|~sym_p_1_3(X1)), sym_m_s2_type11_4).
cnf(c_0_11, axiom, (sym_p_1_3(X1)), sym_m_s2_goal_1).
cnf(c_0_12, plain, (~sym_q_1_2(a)), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_8]), c_0_9])).
cnf(c_0_13, plain, (sym_q_2_3(X1)|sym_q_1_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_10, c_0_11])])).
cnf(c_0_14, axiom, (~sym_q_3_3(a)), sym_m_t2_6).
cnf(c_0_15, axiom, (sym_q_3_3(X1)|sym_p_2_2(X1)|~sym_q_2_3(X1)), sym_m_s2_type12_6).
cnf(c_0_16, plain, (sym_q_2_3(a)), inference(pm,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_17, plain, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_14, c_0_15]), c_0_16])]), c_0_9]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 18
# Proof object clause steps            : 11
# Proof object formula steps           : 7
# Proof object conjectures             : 0
# Proof object clause conjectures      : 0
# Proof object formula conjectures     : 0
# Proof object initial clauses used    : 7
# Proof object initial formulas used   : 7
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 6
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 50
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 50
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 49
# Processed clauses                    : 58
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 58
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 44
# ...of the previous two non-trivial   : 25
# Contextual simplify-reflections      : 0
# Paramodulations                      : 44
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 58
#    Positive orientable unit clauses  : 7
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 9
#    Non-unit-clauses                  : 42
# Current number of unprocessed clauses: 16
# ...number of literals in the above   : 53
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 80
# Rec. Clause-clause subsumption calls : 79
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 13
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1494
