# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_50, negated_conjecture, (p_5_1)).
#
#cnf(i_0_51, negated_conjecture, (p_5_2)).
#
#cnf(i_0_52, negated_conjecture, (p_5_3)).
#
#cnf(i_0_28, negated_conjecture, (~p_0)).
#
#cnf(i_0_53, negated_conjecture, (p_5_4)).
#
#cnf(i_0_54, negated_conjecture, (p_5_5)).
#
#cnf(i_0_29, negated_conjecture, (~p_1_5|~p_1_4|~p_1_3|~p_1_2|~p_1_1)).
#
#cnf(i_0_30, negated_conjecture, (p_1_1|~p_2_5|~p_2_4|~p_2_3|~p_2_2|~p_2_1)).
#
#cnf(i_0_31, negated_conjecture, (p_1_2|~p_2_5|~p_2_4|~p_2_3|~p_2_2|~p_2_1)).
#
#cnf(i_0_32, negated_conjecture, (p_1_3|~p_2_5|~p_2_4|~p_2_3|~p_2_2|~p_2_1)).
#
#cnf(i_0_33, negated_conjecture, (p_1_4|~p_2_5|~p_2_4|~p_2_3|~p_2_2|~p_2_1)).
#
#cnf(i_0_34, negated_conjecture, (p_1_5|~p_2_5|~p_2_4|~p_2_3|~p_2_2|~p_2_1)).
#
#cnf(i_0_35, negated_conjecture, (p_2_1|~p_3_5|~p_3_4|~p_3_3|~p_3_2|~p_3_1)).
#
#cnf(i_0_36, negated_conjecture, (p_2_2|~p_3_5|~p_3_4|~p_3_3|~p_3_2|~p_3_1)).
#
#cnf(i_0_37, negated_conjecture, (p_2_3|~p_3_5|~p_3_4|~p_3_3|~p_3_2|~p_3_1)).
#
#cnf(i_0_38, negated_conjecture, (p_2_4|~p_3_5|~p_3_4|~p_3_3|~p_3_2|~p_3_1)).
#
#cnf(i_0_39, negated_conjecture, (p_2_5|~p_3_5|~p_3_4|~p_3_3|~p_3_2|~p_3_1)).
#
#cnf(i_0_40, negated_conjecture, (p_3_1|~p_4_5|~p_4_4|~p_4_3|~p_4_2|~p_4_1)).
#
#cnf(i_0_41, negated_conjecture, (p_3_2|~p_4_5|~p_4_4|~p_4_3|~p_4_2|~p_4_1)).
#
#cnf(i_0_42, negated_conjecture, (p_3_3|~p_4_5|~p_4_4|~p_4_3|~p_4_2|~p_4_1)).
#
#cnf(i_0_43, negated_conjecture, (p_3_4|~p_4_5|~p_4_4|~p_4_3|~p_4_2|~p_4_1)).
#
#cnf(i_0_44, negated_conjecture, (p_3_5|~p_4_5|~p_4_4|~p_4_3|~p_4_2|~p_4_1)).
#
#cnf(i_0_45, negated_conjecture, (p_4_1)).
#
#cnf(i_0_44, negated_conjecture, (p_3_5|~p_4_5|~p_4_4|~p_4_3|~p_4_2)).
#
#cnf(i_0_43, negated_conjecture, (p_3_4|~p_4_5|~p_4_4|~p_4_3|~p_4_2)).
#
#cnf(i_0_42, negated_conjecture, (p_3_3|~p_4_5|~p_4_4|~p_4_3|~p_4_2)).
#
#cnf(i_0_41, negated_conjecture, (p_3_2|~p_4_5|~p_4_4|~p_4_3|~p_4_2)).
#
#cnf(i_0_40, negated_conjecture, (p_3_1|~p_4_5|~p_4_4|~p_4_3|~p_4_2)).
#
#cnf(i_0_46, negated_conjecture, (p_4_2)).
#
#cnf(i_0_40, negated_conjecture, (p_3_1|~p_4_5|~p_4_4|~p_4_3)).
#
#cnf(i_0_41, negated_conjecture, (p_3_2|~p_4_5|~p_4_4|~p_4_3)).
#
#cnf(i_0_42, negated_conjecture, (p_3_3|~p_4_5|~p_4_4|~p_4_3)).
#
#cnf(i_0_44, negated_conjecture, (p_3_5|~p_4_5|~p_4_4|~p_4_3)).
#
#cnf(i_0_43, negated_conjecture, (p_3_4|~p_4_5|~p_4_4|~p_4_3)).
#
#cnf(i_0_47, negated_conjecture, (p_4_3)).
#
#cnf(i_0_43, negated_conjecture, (p_3_4|~p_4_5|~p_4_4)).
#
#cnf(i_0_44, negated_conjecture, (p_3_5|~p_4_5|~p_4_4)).
#
#cnf(i_0_42, negated_conjecture, (p_3_3|~p_4_5|~p_4_4)).
#
#cnf(i_0_41, negated_conjecture, (p_3_2|~p_4_5|~p_4_4)).
#
#cnf(i_0_40, negated_conjecture, (p_3_1|~p_4_5|~p_4_4)).
#
#cnf(i_0_48, negated_conjecture, (p_4_4)).
#
#cnf(i_0_40, negated_conjecture, (p_3_1|~p_4_5)).
#
#cnf(i_0_41, negated_conjecture, (p_3_2|~p_4_5)).
#
#cnf(i_0_42, negated_conjecture, (p_3_3|~p_4_5)).
#
#cnf(i_0_44, negated_conjecture, (p_3_5|~p_4_5)).
#
#cnf(i_0_43, negated_conjecture, (p_3_4|~p_4_5)).
#
#cnf(i_0_49, negated_conjecture, (p_4_5)).
#
#cnf(i_0_43, negated_conjecture, (p_3_4)).
#
#cnf(i_0_44, negated_conjecture, (p_3_5)).
#
#cnf(i_0_42, negated_conjecture, (p_3_3)).
#
#cnf(i_0_41, negated_conjecture, (p_3_2)).
#
#cnf(i_0_40, negated_conjecture, (p_3_1)).
#
#cnf(i_0_39, negated_conjecture, (p_2_5)).
#
#cnf(i_0_38, negated_conjecture, (p_2_4)).
#
#cnf(i_0_37, negated_conjecture, (p_2_3)).
#
#cnf(i_0_36, negated_conjecture, (p_2_2)).
#
#cnf(i_0_35, negated_conjecture, (p_2_1)).
#
#cnf(i_0_34, negated_conjecture, (p_1_5)).
#
#cnf(i_0_33, negated_conjecture, (p_1_4)).
#
#cnf(i_0_29, negated_conjecture, (~p_1_3|~p_1_2|~p_1_1)).
#
#cnf(i_0_32, negated_conjecture, (p_1_3)).
#
#cnf(i_0_31, negated_conjecture, (p_1_2)).
#
#cnf(i_0_30, negated_conjecture, (p_1_1)).
#
# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(clause_18, negated_conjecture, (p_4_1|~p_5_1|~p_5_2|~p_5_3|~p_5_4|~p_5_5), file('./p/SYN010-1.005.005.tptp', clause_18)).
cnf(clause_23, negated_conjecture, (p_5_1), file('./p/SYN010-1.005.005.tptp', clause_23)).
cnf(clause_24, negated_conjecture, (p_5_2), file('./p/SYN010-1.005.005.tptp', clause_24)).
cnf(clause_25, negated_conjecture, (p_5_3), file('./p/SYN010-1.005.005.tptp', clause_25)).
cnf(clause_26, negated_conjecture, (p_5_4), file('./p/SYN010-1.005.005.tptp', clause_26)).
cnf(clause_27, negated_conjecture, (p_5_5), file('./p/SYN010-1.005.005.tptp', clause_27)).
cnf(clause_16, negated_conjecture, (p_3_4|~p_4_1|~p_4_2|~p_4_3|~p_4_4|~p_4_5), file('./p/SYN010-1.005.005.tptp', clause_16)).
cnf(clause_19, negated_conjecture, (p_4_2|~p_5_1|~p_5_2|~p_5_3|~p_5_4|~p_5_5), file('./p/SYN010-1.005.005.tptp', clause_19)).
cnf(clause_17, negated_conjecture, (p_3_5|~p_4_1|~p_4_2|~p_4_3|~p_4_4|~p_4_5), file('./p/SYN010-1.005.005.tptp', clause_17)).
cnf(clause_15, negated_conjecture, (p_3_3|~p_4_1|~p_4_2|~p_4_3|~p_4_4|~p_4_5), file('./p/SYN010-1.005.005.tptp', clause_15)).
cnf(clause_14, negated_conjecture, (p_3_2|~p_4_1|~p_4_2|~p_4_3|~p_4_4|~p_4_5), file('./p/SYN010-1.005.005.tptp', clause_14)).
cnf(clause_13, negated_conjecture, (p_3_1|~p_4_1|~p_4_2|~p_4_3|~p_4_4|~p_4_5), file('./p/SYN010-1.005.005.tptp', clause_13)).
cnf(clause_20, negated_conjecture, (p_4_3|~p_5_1|~p_5_2|~p_5_3|~p_5_4|~p_5_5), file('./p/SYN010-1.005.005.tptp', clause_20)).
cnf(clause_21, negated_conjecture, (p_4_4|~p_5_1|~p_5_2|~p_5_3|~p_5_4|~p_5_5), file('./p/SYN010-1.005.005.tptp', clause_21)).
cnf(clause_22, negated_conjecture, (p_4_5|~p_5_1|~p_5_2|~p_5_3|~p_5_4|~p_5_5), file('./p/SYN010-1.005.005.tptp', clause_22)).
cnf(clause_12, negated_conjecture, (p_2_5|~p_3_1|~p_3_2|~p_3_3|~p_3_4|~p_3_5), file('./p/SYN010-1.005.005.tptp', clause_12)).
cnf(clause_11, negated_conjecture, (p_2_4|~p_3_1|~p_3_2|~p_3_3|~p_3_4|~p_3_5), file('./p/SYN010-1.005.005.tptp', clause_11)).
cnf(clause_10, negated_conjecture, (p_2_3|~p_3_1|~p_3_2|~p_3_3|~p_3_4|~p_3_5), file('./p/SYN010-1.005.005.tptp', clause_10)).
cnf(clause_9, negated_conjecture, (p_2_2|~p_3_1|~p_3_2|~p_3_3|~p_3_4|~p_3_5), file('./p/SYN010-1.005.005.tptp', clause_9)).
cnf(clause_8, negated_conjecture, (p_2_1|~p_3_1|~p_3_2|~p_3_3|~p_3_4|~p_3_5), file('./p/SYN010-1.005.005.tptp', clause_8)).
cnf(clause_2, negated_conjecture, (p_0|~p_1_1|~p_1_2|~p_1_3|~p_1_4|~p_1_5), file('./p/SYN010-1.005.005.tptp', clause_2)).
cnf(clause_1, negated_conjecture, (~p_0), file('./p/SYN010-1.005.005.tptp', clause_1)).
cnf(clause_7, negated_conjecture, (p_1_5|~p_2_1|~p_2_2|~p_2_3|~p_2_4|~p_2_5), file('./p/SYN010-1.005.005.tptp', clause_7)).
cnf(clause_6, negated_conjecture, (p_1_4|~p_2_1|~p_2_2|~p_2_3|~p_2_4|~p_2_5), file('./p/SYN010-1.005.005.tptp', clause_6)).
cnf(clause_5, negated_conjecture, (p_1_3|~p_2_1|~p_2_2|~p_2_3|~p_2_4|~p_2_5), file('./p/SYN010-1.005.005.tptp', clause_5)).
cnf(clause_4, negated_conjecture, (p_1_2|~p_2_1|~p_2_2|~p_2_3|~p_2_4|~p_2_5), file('./p/SYN010-1.005.005.tptp', clause_4)).
cnf(clause_3, negated_conjecture, (p_1_1|~p_2_1|~p_2_2|~p_2_3|~p_2_4|~p_2_5), file('./p/SYN010-1.005.005.tptp', clause_3)).
cnf(c_0_27, negated_conjecture, (p_4_1|~p_5_1|~p_5_2|~p_5_3|~p_5_4|~p_5_5), clause_18).
cnf(c_0_28, negated_conjecture, (p_5_1), clause_23).
cnf(c_0_29, negated_conjecture, (p_5_2), clause_24).
cnf(c_0_30, negated_conjecture, (p_5_3), clause_25).
cnf(c_0_31, negated_conjecture, (p_5_4), clause_26).
cnf(c_0_32, negated_conjecture, (p_5_5), clause_27).
cnf(c_0_33, negated_conjecture, (p_3_4|~p_4_1|~p_4_2|~p_4_3|~p_4_4|~p_4_5), clause_16).
cnf(c_0_34, negated_conjecture, (p_4_1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_27, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])])).
cnf(c_0_35, negated_conjecture, (p_4_2|~p_5_1|~p_5_2|~p_5_3|~p_5_4|~p_5_5), clause_19).
cnf(c_0_36, negated_conjecture, (p_3_5|~p_4_1|~p_4_2|~p_4_3|~p_4_4|~p_4_5), clause_17).
cnf(c_0_37, negated_conjecture, (p_3_3|~p_4_1|~p_4_2|~p_4_3|~p_4_4|~p_4_5), clause_15).
cnf(c_0_38, negated_conjecture, (p_3_2|~p_4_1|~p_4_2|~p_4_3|~p_4_4|~p_4_5), clause_14).
cnf(c_0_39, negated_conjecture, (p_3_1|~p_4_1|~p_4_2|~p_4_3|~p_4_4|~p_4_5), clause_13).
cnf(c_0_40, negated_conjecture, (p_3_4|~p_4_5|~p_4_4|~p_4_3|~p_4_2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_33, c_0_34])])).
cnf(c_0_41, negated_conjecture, (p_4_2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_35, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])])).
cnf(c_0_42, negated_conjecture, (p_4_3|~p_5_1|~p_5_2|~p_5_3|~p_5_4|~p_5_5), clause_20).
cnf(c_0_43, negated_conjecture, (p_3_5|~p_4_5|~p_4_4|~p_4_3|~p_4_2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_36, c_0_34])])).
cnf(c_0_44, negated_conjecture, (p_3_3|~p_4_5|~p_4_4|~p_4_3|~p_4_2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_37, c_0_34])])).
cnf(c_0_45, negated_conjecture, (p_3_2|~p_4_5|~p_4_4|~p_4_3|~p_4_2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_38, c_0_34])])).
cnf(c_0_46, negated_conjecture, (p_3_1|~p_4_5|~p_4_4|~p_4_3|~p_4_2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_39, c_0_34])])).
cnf(c_0_47, negated_conjecture, (p_3_4|~p_4_5|~p_4_4|~p_4_3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_40, c_0_41])])).
cnf(c_0_48, negated_conjecture, (p_4_3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_42, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])])).
cnf(c_0_49, negated_conjecture, (p_4_4|~p_5_1|~p_5_2|~p_5_3|~p_5_4|~p_5_5), clause_21).
cnf(c_0_50, negated_conjecture, (p_3_5|~p_4_5|~p_4_4|~p_4_3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_43, c_0_41])])).
cnf(c_0_51, negated_conjecture, (p_3_3|~p_4_5|~p_4_4|~p_4_3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_44, c_0_41])])).
cnf(c_0_52, negated_conjecture, (p_3_2|~p_4_5|~p_4_4|~p_4_3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_45, c_0_41])])).
cnf(c_0_53, negated_conjecture, (p_3_1|~p_4_5|~p_4_4|~p_4_3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_46, c_0_41])])).
cnf(c_0_54, negated_conjecture, (p_3_4|~p_4_5|~p_4_4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_47, c_0_48])])).
cnf(c_0_55, negated_conjecture, (p_4_4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_49, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])])).
cnf(c_0_56, negated_conjecture, (p_4_5|~p_5_1|~p_5_2|~p_5_3|~p_5_4|~p_5_5), clause_22).
cnf(c_0_57, negated_conjecture, (p_3_5|~p_4_5|~p_4_4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_50, c_0_48])])).
cnf(c_0_58, negated_conjecture, (p_3_3|~p_4_5|~p_4_4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_51, c_0_48])])).
cnf(c_0_59, negated_conjecture, (p_3_2|~p_4_5|~p_4_4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_52, c_0_48])])).
cnf(c_0_60, negated_conjecture, (p_3_1|~p_4_5|~p_4_4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_48])])).
cnf(c_0_61, negated_conjecture, (p_3_4|~p_4_5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_54, c_0_55])])).
cnf(c_0_62, negated_conjecture, (p_4_5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_56, c_0_28]), c_0_29]), c_0_30]), c_0_31]), c_0_32])])).
cnf(c_0_63, negated_conjecture, (p_3_5|~p_4_5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_57, c_0_55])])).
cnf(c_0_64, negated_conjecture, (p_3_3|~p_4_5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_58, c_0_55])])).
cnf(c_0_65, negated_conjecture, (p_3_2|~p_4_5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_59, c_0_55])])).
cnf(c_0_66, negated_conjecture, (p_3_1|~p_4_5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_55])])).
cnf(c_0_67, negated_conjecture, (p_2_5|~p_3_1|~p_3_2|~p_3_3|~p_3_4|~p_3_5), clause_12).
cnf(c_0_68, negated_conjecture, (p_3_4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_61, c_0_62])])).
cnf(c_0_69, negated_conjecture, (p_3_5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_63, c_0_62])])).
cnf(c_0_70, negated_conjecture, (p_3_3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_64, c_0_62])])).
cnf(c_0_71, negated_conjecture, (p_3_2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_65, c_0_62])])).
cnf(c_0_72, negated_conjecture, (p_3_1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_66, c_0_62])])).
cnf(c_0_73, negated_conjecture, (p_2_4|~p_3_1|~p_3_2|~p_3_3|~p_3_4|~p_3_5), clause_11).
cnf(c_0_74, negated_conjecture, (p_2_3|~p_3_1|~p_3_2|~p_3_3|~p_3_4|~p_3_5), clause_10).
cnf(c_0_75, negated_conjecture, (p_2_2|~p_3_1|~p_3_2|~p_3_3|~p_3_4|~p_3_5), clause_9).
cnf(c_0_76, negated_conjecture, (p_2_1|~p_3_1|~p_3_2|~p_3_3|~p_3_4|~p_3_5), clause_8).
cnf(c_0_77, negated_conjecture, (p_0|~p_1_1|~p_1_2|~p_1_3|~p_1_4|~p_1_5), clause_2).
cnf(c_0_78, negated_conjecture, (~p_0), clause_1).
cnf(c_0_79, negated_conjecture, (p_1_5|~p_2_1|~p_2_2|~p_2_3|~p_2_4|~p_2_5), clause_7).
cnf(c_0_80, negated_conjecture, (p_2_5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_67, c_0_68])]), c_0_69]), c_0_70]), c_0_71]), c_0_72])])).
cnf(c_0_81, negated_conjecture, (p_2_4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_68])]), c_0_69]), c_0_70]), c_0_71]), c_0_72])])).
cnf(c_0_82, negated_conjecture, (p_2_3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_74, c_0_68])]), c_0_69]), c_0_70]), c_0_71]), c_0_72])])).
cnf(c_0_83, negated_conjecture, (p_2_2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_75, c_0_68])]), c_0_69]), c_0_70]), c_0_71]), c_0_72])])).
cnf(c_0_84, negated_conjecture, (p_2_1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_76, c_0_68])]), c_0_69]), c_0_70]), c_0_71]), c_0_72])])).
cnf(c_0_85, negated_conjecture, (p_1_4|~p_2_1|~p_2_2|~p_2_3|~p_2_4|~p_2_5), clause_6).
cnf(c_0_86, negated_conjecture, (~p_1_5|~p_1_4|~p_1_3|~p_1_2|~p_1_1), inference(sr,[status(thm)],[c_0_77, c_0_78])).
cnf(c_0_87, negated_conjecture, (p_1_5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_80])]), c_0_81]), c_0_82]), c_0_83]), c_0_84])])).
cnf(c_0_88, negated_conjecture, (p_1_4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_85, c_0_80])]), c_0_81]), c_0_82]), c_0_83]), c_0_84])])).
cnf(c_0_89, negated_conjecture, (p_1_3|~p_2_1|~p_2_2|~p_2_3|~p_2_4|~p_2_5), clause_5).
cnf(c_0_90, negated_conjecture, (p_1_2|~p_2_1|~p_2_2|~p_2_3|~p_2_4|~p_2_5), clause_4).
cnf(c_0_91, negated_conjecture, (p_1_1|~p_2_1|~p_2_2|~p_2_3|~p_2_4|~p_2_5), clause_3).
cnf(c_0_92, negated_conjecture, (~p_1_3|~p_1_2|~p_1_1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_86, c_0_87])]), c_0_88])])).
cnf(c_0_93, negated_conjecture, (p_1_3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_80])]), c_0_81]), c_0_82]), c_0_83]), c_0_84])])).
cnf(c_0_94, negated_conjecture, (p_1_2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_90, c_0_80])]), c_0_81]), c_0_82]), c_0_83]), c_0_84])])).
cnf(c_0_95, negated_conjecture, (p_1_1), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_91, c_0_80])]), c_0_81]), c_0_82]), c_0_83]), c_0_84])])).
cnf(c_0_96, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_92, c_0_93])]), c_0_94]), c_0_95])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 97
# Proof object clause steps            : 70
# Proof object formula steps           : 27
# Proof object conjectures             : 97
# Proof object clause conjectures      : 70
# Proof object formula conjectures     : 27
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 27
# Proof object generating inferences   : 0
# Proof object simplifying inferences  : 160
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 27
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 27
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 27
# Processed clauses                    : 64
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 63
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 37
# Generated clauses                    : 0
# ...of the previous two non-trivial   : 37
# Contextual simplify-reflections      : 0
# Paramodulations                      : 0
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 26
#    Positive orientable unit clauses  : 25
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 37
# Clause-clause subsumption calls (NU) : 136
# Rec. Clause-clause subsumption calls : 40
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 66
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 9
# BW rewrite match successes           : 9
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1241
