# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_16, negated_conjecture, (t)).
#
#cnf(i_0_11, negated_conjecture, (l|~m)).
#
#cnf(i_0_12, negated_conjecture, (l|~q)).
#
#cnf(i_0_9, negated_conjecture, (~n)).
#
#cnf(i_0_10, negated_conjecture, (q|m)).
#
#cnf(i_0_14, negated_conjecture, (r|p)).
#
#cnf(i_0_13, negated_conjecture, (~p|~l)).
#
#cnf(i_0_15, negated_conjecture, (~r|~l)).
#
#cnf(i_0_17, negated_conjecture, (~p|~m)).
#
#cnf(i_0_18, negated_conjecture, (~p|~q)).
#
#cnf(i_0_19, negated_conjecture, (~r|~m)).
#
#cnf(i_0_20, negated_conjecture, (~r|~q)).
#
#cnf(i_0_21, negated_conjecture, (p|~m)).
#
#cnf(i_0_22, negated_conjecture, (p|~q)).
#
#cnf(i_0_23, negated_conjecture, (~m)).
#
#cnf(i_0_26, negated_conjecture, (~q)).
#
# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(clause_1, negated_conjecture, (~n|~t), file('./p/SYN011-1.tptp', clause_1)).
cnf(clause_8, negated_conjecture, (t), file('./p/SYN011-1.tptp', clause_8)).
cnf(clause_7, negated_conjecture, (~r|~l), file('./p/SYN011-1.tptp', clause_7)).
cnf(clause_3, negated_conjecture, (l|~m), file('./p/SYN011-1.tptp', clause_3)).
cnf(clause_6, negated_conjecture, (r|p|n), file('./p/SYN011-1.tptp', clause_6)).
cnf(clause_4, negated_conjecture, (l|~q), file('./p/SYN011-1.tptp', clause_4)).
cnf(clause_5, negated_conjecture, (~l|~p), file('./p/SYN011-1.tptp', clause_5)).
cnf(clause_2, negated_conjecture, (m|q|n), file('./p/SYN011-1.tptp', clause_2)).
cnf(c_0_8, negated_conjecture, (~n|~t), clause_1).
cnf(c_0_9, negated_conjecture, (t), clause_8).
cnf(c_0_10, negated_conjecture, (~r|~l), clause_7).
cnf(c_0_11, negated_conjecture, (l|~m), clause_3).
cnf(c_0_12, negated_conjecture, (r|p|n), clause_6).
cnf(c_0_13, negated_conjecture, (~n), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_9])])).
cnf(c_0_14, negated_conjecture, (l|~q), clause_4).
cnf(c_0_15, negated_conjecture, (~l|~p), clause_5).
cnf(c_0_16, negated_conjecture, (~r|~m), inference(pm,[status(thm)],[c_0_10, c_0_11])).
cnf(c_0_17, negated_conjecture, (r|p), inference(sr,[status(thm)],[c_0_12, c_0_13])).
cnf(c_0_18, negated_conjecture, (~r|~q), inference(pm,[status(thm)],[c_0_10, c_0_14])).
cnf(c_0_19, negated_conjecture, (m|q|n), clause_2).
cnf(c_0_20, negated_conjecture, (~p|~m), inference(pm,[status(thm)],[c_0_15, c_0_11])).
cnf(c_0_21, negated_conjecture, (p|~m), inference(pm,[status(thm)],[c_0_16, c_0_17])).
cnf(c_0_22, negated_conjecture, (~p|~q), inference(pm,[status(thm)],[c_0_15, c_0_14])).
cnf(c_0_23, negated_conjecture, (p|~q), inference(pm,[status(thm)],[c_0_18, c_0_17])).
cnf(c_0_24, negated_conjecture, (q|m), inference(sr,[status(thm)],[c_0_19, c_0_13])).
cnf(c_0_25, negated_conjecture, (~m), inference(pm,[status(thm)],[c_0_20, c_0_21])).
cnf(c_0_26, negated_conjecture, (~q), inference(pm,[status(thm)],[c_0_22, c_0_23])).
cnf(c_0_27, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_24, c_0_25]), c_0_26]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 28
# Proof object clause steps            : 20
# Proof object formula steps           : 8
# Proof object conjectures             : 28
# Proof object clause conjectures      : 20
# Proof object formula conjectures     : 8
# Proof object initial clauses used    : 8
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 6
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 8
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 8
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 8
# Processed clauses                    : 17
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 16
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 11
# ...of the previous two non-trivial   : 11
# Contextual simplify-reflections      : 0
# Paramodulations                      : 10
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 15
#    Positive orientable unit clauses  : 1
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 3
#    Non-unit-clauses                  : 11
# Current number of unprocessed clauses: 2
# ...number of literals in the above   : 4
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 8
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 158
