# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_6, negated_conjecture, (p)).
#
#cnf(i_0_10, negated_conjecture, (r|~q)).
#
#cnf(i_0_8, negated_conjecture, (s|~r)).
#
#cnf(i_0_9, negated_conjecture, (~s)).
#
#cnf(i_0_7, negated_conjecture, (q)).
#
#cnf(i_0_10, negated_conjecture, (r)).

# Proof found!
# SZS status Unsatisfiable
# SZS output start CNFRefutation
cnf(clause_2, negated_conjecture, (q|s|~p), file('./p/SYN047-1.tptp', clause_2)).
cnf(clause_1, negated_conjecture, (p), file('./p/SYN047-1.tptp', clause_1)).
cnf(clause_4, negated_conjecture, (~s), file('./p/SYN047-1.tptp', clause_4)).
cnf(clause_3, negated_conjecture, (s|~p|~r), file('./p/SYN047-1.tptp', clause_3)).
cnf(clause_5, negated_conjecture, (r|~q), file('./p/SYN047-1.tptp', clause_5)).
cnf(c_0_5, negated_conjecture, (q|s|~p), clause_2).
cnf(c_0_6, negated_conjecture, (p), clause_1).
cnf(c_0_7, negated_conjecture, (~s), clause_4).
cnf(c_0_8, negated_conjecture, (s|~p|~r), clause_3).
cnf(c_0_9, negated_conjecture, (r|~q), clause_5).
cnf(c_0_10, negated_conjecture, (q), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_6])]), c_0_7])).
cnf(c_0_11, negated_conjecture, (s|~r), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_8, c_0_6])])).
cnf(c_0_12, negated_conjecture, (r), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_9, c_0_10])])).
cnf(c_0_13, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_11, c_0_12])]), c_0_7]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 14
# Proof object clause steps            : 9
# Proof object formula steps           : 5
# Proof object conjectures             : 14
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 5
# Proof object initial clauses used    : 5
# Proof object initial formulas used   : 5
# Proof object generating inferences   : 0
# Proof object simplifying inferences  : 10
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 5
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 5
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 5
# Processed clauses                    : 6
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 6
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 2
# Generated clauses                    : 1
# ...of the previous two non-trivial   : 2
# Contextual simplify-reflections      : 0
# Paramodulations                      : 1
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 4
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 1
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 2
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 88
