# Initializing proof state
# Scanning for AC axioms
#
# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(thm119, conjecture, ~(((![X1]:((p(X1)|r(X1))&q(X1))&![X2]:?[X3]:(p(X2)|~(q(X2))|~(q(X3))|~(q(c))|~(q(d))))&(~(p(a))|~(p(b))))), file('./p/SYN722+1.tptp', thm119)).
fof(c_0_1, negated_conjecture, ~(~(((![X1]:((p(X1)|r(X1))&q(X1))&![X2]:?[X3]:(p(X2)|~(q(X2))|~(q(X3))|~(q(c))|~(q(d))))&(~(p(a))|~(p(b)))))), inference(assume_negation,[status(cth)],[thm119])).
fof(c_0_2, negated_conjecture, ![X4, X5]:((((p(X4)|r(X4))&q(X4))&(p(X5)|~q(X5)|~q(esk1_0)|~q(c)|~q(d)))&(~p(a)|~p(b))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_1])])])])])])).
cnf(c_0_3, negated_conjecture, (p(X1)|~q(X1)|~q(esk1_0)|~q(c)|~q(d)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_4, negated_conjecture, (q(X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_5, negated_conjecture, (~p(a)|~p(b)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_6, negated_conjecture, (p(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_3, c_0_4]), c_0_4]), c_0_4]), c_0_4])])).
cnf(c_0_7, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_5, c_0_6]), c_0_6])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 8
# Proof object clause steps            : 5
# Proof object formula steps           : 3
# Proof object conjectures             : 8
# Proof object clause conjectures      : 5
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 3
# Proof object initial formulas used   : 1
# Proof object generating inferences   : 0
# Proof object simplifying inferences  : 8
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 1
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 4
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 1
# Processed clauses                    : 1
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 0
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 0
# ...of the previous two non-trivial   : 0
# Contextual simplify-reflections      : 0
# Paramodulations                      : 0
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 0
#    Positive orientable unit clauses  : 0
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 0
#    Non-unit-clauses                  : 0
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 261
