# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_4, negated_conjecture, (p(esk1_0)|~c)).
#
#cnf(i_0_3, negated_conjecture, (p(esk2_0)|~c)).
#
#cnf(i_0_1, negated_conjecture, (c|~p(X2)|~p(X1))).
#
#cnf(i_0_2, negated_conjecture, (~c)).
#
#cnf(i_0_5, negated_conjecture, (p(esk1_0)|p(esk2_0))).
#
#cnf(i_0_10, negated_conjecture, (p(esk2_0)|~p(X1))).
#
#cnf(i_0_17, negated_conjecture, (p(esk2_0))).
#
#cnf(i_0_18, negated_conjecture, (~p(X1))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(prove_this, conjecture, (![X1]:(p(X1)=>c)<=>(?[X1]:p(X1)=>c)), file('./p/SYN937+1.tptp', prove_this)).
fof(c_0_1, negated_conjecture, ~((![X1]:(p(X1)=>c)<=>(?[X1]:p(X1)=>c))), inference(assume_negation,[status(cth)],[prove_this])).
fof(c_0_2, negated_conjecture, ![X4, X5]:((((p(esk2_0)|p(esk1_0))&(~c|p(esk1_0)))&((p(esk2_0)|~c)&(~c|~c)))&(~p(X4)|c|(~p(X5)|c))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
cnf(c_0_3, negated_conjecture, (c|c|~p(X1)|~p(X2)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_4, negated_conjecture, (~c|~c), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_5, negated_conjecture, (c|~p(X2)|~p(X1)), inference(cn,[status(thm)],[c_0_3])).
cnf(c_0_6, negated_conjecture, (p(esk2_0)|p(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_7, negated_conjecture, (~c), inference(cn,[status(thm)],[c_0_4])).
cnf(c_0_8, negated_conjecture, (p(esk2_0)|~p(X1)), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_5, c_0_6]), c_0_7])).
cnf(c_0_9, negated_conjecture, (p(esk2_0)), inference(pm,[status(thm)],[c_0_8, c_0_6])).
cnf(c_0_10, negated_conjecture, (~p(X1)), inference(sr,[status(thm)],[inference(pm,[status(thm)],[c_0_5, c_0_9]), c_0_7])).
cnf(c_0_11, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_9, c_0_10]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 12
# Proof object clause steps            : 9
# Proof object formula steps           : 3
# Proof object conjectures             : 12
# Proof object clause conjectures      : 9
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 3
# Proof object initial formulas used   : 1
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 5
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 1
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 5
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 5
# Processed clauses                    : 8
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 8
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 2
# Backward-rewritten                   : 1
# Generated clauses                    : 16
# ...of the previous two non-trivial   : 10
# Contextual simplify-reflections      : 0
# Paramodulations                      : 15
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 4
#    Positive orientable unit clauses  : 0
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 2
# Current number of unprocessed clauses: 7
# ...number of literals in the above   : 13
# Current number of archived formulas  : 0
# Current number of archived clauses   : 4
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 6
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 312
