# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_3, negated_conjecture, (p(f(X1),X1)|~r(X1))).
#
#cnf(i_0_4, negated_conjecture, (p(f(X1),X1)|r(esk1_0))).
#
#cnf(i_0_2, negated_conjecture, (q(f(esk1_0),esk1_0)|~p(X1,X2))).
#
#cnf(i_0_1, negated_conjecture, (~q(X1,X2)|~p(X1,X2))).
#
#cnf(i_0_6, negated_conjecture, (q(f(esk1_0),esk1_0)|~r(X1))).
#
#cnf(i_0_5, negated_conjecture, (q(f(esk1_0),esk1_0)|r(esk1_0))).
#
#cnf(i_0_9, negated_conjecture, (r(esk1_0)|~p(f(esk1_0),esk1_0))).
#
#cnf(i_0_8, negated_conjecture, (~p(f(esk1_0),esk1_0)|~r(X1))).
#
#cnf(i_0_10, negated_conjecture, (r(esk1_0))).
#
#cnf(i_0_13, negated_conjecture, (~r(X1))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(prove_this, conjecture, ![X1]:(![X2]:((r(X1)=>r(X2))=>p(f(X2),X2))=>?[X3, X2]:(p(X3,X2)&(q(f(X1),X1)=>q(X3,X2)))), file('./p/SYN980+1.tptp', prove_this)).
fof(c_0_1, negated_conjecture, ~(![X1]:(![X2]:((r(X1)=>r(X2))=>p(f(X2),X2))=>?[X3, X2]:(p(X3,X2)&(q(f(X1),X1)=>q(X3,X2))))), inference(assume_negation,[status(cth)],[prove_this])).
fof(c_0_2, negated_conjecture, ![X5, X6, X7]:(((r(esk1_0)|p(f(X5),X5))&(~r(X5)|p(f(X5),X5)))&((q(f(esk1_0),esk1_0)|~p(X6,X7))&(~q(X6,X7)|~p(X6,X7)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])])).
cnf(c_0_3, negated_conjecture, (q(f(esk1_0),esk1_0)|~p(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_4, negated_conjecture, (r(esk1_0)|p(f(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_5, negated_conjecture, (p(f(X1),X1)|~r(X1)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_6, negated_conjecture, (~q(X1,X2)|~p(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_7, negated_conjecture, (q(f(esk1_0),esk1_0)|r(esk1_0)), inference(pm,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_8, negated_conjecture, (q(f(esk1_0),esk1_0)|~r(X1)), inference(pm,[status(thm)],[c_0_3, c_0_5])).
cnf(c_0_9, negated_conjecture, (r(esk1_0)|~p(f(esk1_0),esk1_0)), inference(pm,[status(thm)],[c_0_6, c_0_7])).
cnf(c_0_10, negated_conjecture, (~p(f(esk1_0),esk1_0)|~r(X1)), inference(pm,[status(thm)],[c_0_6, c_0_8])).
cnf(c_0_11, negated_conjecture, (r(esk1_0)), inference(pm,[status(thm)],[c_0_9, c_0_4])).
cnf(c_0_12, negated_conjecture, (~r(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_10, c_0_5]), c_0_11])])).
cnf(c_0_13, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_11, c_0_12]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 14
# Proof object clause steps            : 11
# Proof object formula steps           : 3
# Proof object conjectures             : 14
# Proof object clause conjectures      : 11
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 4
# Proof object initial formulas used   : 1
# Proof object generating inferences   : 6
# Proof object simplifying inferences  : 3
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 1
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 4
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 4
# Processed clauses                    : 10
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 10
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 3
# Generated clauses                    : 10
# ...of the previous two non-trivial   : 8
# Contextual simplify-reflections      : 0
# Paramodulations                      : 9
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 6
#    Positive orientable unit clauses  : 0
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 5
# Current number of unprocessed clauses: 1
# ...number of literals in the above   : 2
# Current number of archived formulas  : 0
# Current number of archived clauses   : 4
# Clause-clause subsumption calls (NU) : 0
# Rec. Clause-clause subsumption calls : 0
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 3
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 331
