# Initializing proof state
# Scanning for AC axioms
#
#cnf(i_0_2, negated_conjecture, (r(esk1_0,esk2_0))).
#
#cnf(i_0_4, negated_conjecture, (r(X2,X1)|~r(X1,X2))).
#
#cnf(i_0_6, negated_conjecture, (r(esk2_0,esk1_0))).
#
#cnf(i_0_1, negated_conjecture, (~r(esk1_0,esk1_0))).
#
#cnf(i_0_3, negated_conjecture, (r(X1,X3)|~r(X2,X3)|~r(X1,X2))).
#
#cnf(i_0_14, negated_conjecture, (r(X1,esk2_0)|~r(X1,esk1_0))).
#
#cnf(i_0_16, negated_conjecture, (r(esk2_0,esk2_0))).
#
#cnf(i_0_7, negated_conjecture, (~r(X1,esk1_0)|~r(esk1_0,X1))).

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(prove_this, conjecture, ((![X1, X2]:(r(X1,X2)=>r(X2,X1))&![X1, X2, X3]:((r(X1,X2)&r(X2,X3))=>r(X1,X3)))=>![X1, X2]:(r(X1,X2)=>r(X1,X1))), file('./p/SYN919+1.tptp', prove_this)).
fof(c_0_1, negated_conjecture, ~(((![X1, X2]:(r(X1,X2)=>r(X2,X1))&![X1, X2, X3]:((r(X1,X2)&r(X2,X3))=>r(X1,X3)))=>![X1, X2]:(r(X1,X2)=>r(X1,X1)))), inference(assume_negation,[status(cth)],[prove_this])).
fof(c_0_2, negated_conjecture, ![X4, X5, X6, X7, X8]:(((~r(X4,X5)|r(X5,X4))&(~r(X6,X7)|~r(X7,X8)|r(X6,X8)))&(r(esk1_0,esk2_0)&~r(esk1_0,esk1_0))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_1])])])])).
cnf(c_0_3, negated_conjecture, (~r(esk1_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_4, negated_conjecture, (r(X1,X3)|~r(X1,X2)|~r(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_5, negated_conjecture, (r(X2,X1)|~r(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_6, negated_conjecture, (r(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_2])).
cnf(c_0_7, negated_conjecture, (~r(X1,esk1_0)|~r(esk1_0,X1)), inference(pm,[status(thm)],[c_0_3, c_0_4])).
cnf(c_0_8, negated_conjecture, (r(esk2_0,esk1_0)), inference(pm,[status(thm)],[c_0_5, c_0_6])).
cnf(c_0_9, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(pm,[status(thm)],[c_0_7, c_0_6]), c_0_8])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 10
# Proof object clause steps            : 7
# Proof object formula steps           : 3
# Proof object conjectures             : 10
# Proof object clause conjectures      : 7
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 4
# Proof object initial formulas used   : 1
# Proof object generating inferences   : 3
# Proof object simplifying inferences  : 2
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 1
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 4
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 4
# Processed clauses                    : 8
# ...of these trivial                  : 0
# ...subsumed                          : 0
# ...remaining for further processing  : 8
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 0
# Generated clauses                    : 22
# ...of the previous two non-trivial   : 15
# Contextual simplify-reflections      : 0
# Paramodulations                      : 22
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 8
#    Positive orientable unit clauses  : 3
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 1
#    Non-unit-clauses                  : 4
# Current number of unprocessed clauses: 9
# ...number of literals in the above   : 26
# Current number of archived formulas  : 0
# Current number of archived clauses   : 0
# Clause-clause subsumption calls (NU) : 2
# Rec. Clause-clause subsumption calls : 2
# Non-unit clause-clause subsumptions  : 0
# Unit Clause-clause subsumption calls : 0
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 0
# BW rewrite match successes           : 0
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 406
