% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for ALG012-1
% SZS output start Proof for ALG012-1
fof(f5826,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f140,f146,f734,f873,f910,f914,f943,f1022,f1655,f1744,f1973,f3486,f5471,f5535,f5556,f5746,f5773,f5822])).
fof(f5822,plain,(
  spl0_10 | spl0_7 | ~spl0_15 | ~spl0_18),
  inference(avatar_split_clause,[],[f2313,f912,f732,f380,f444])).
fof(f444,plain,(
  spl0_10 <=> ! [X3] : ~c(X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f380,plain,(
  spl0_7 <=> ! [X3,X0] : (~c(X0) | ~d1(X3) | ~d1(f(X3,X0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).
fof(f732,plain,(
  spl0_15 <=> ! [X5,X4] : (d1(f(X4,X5)) | ~c(X4) | ~c(X5))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_15])])).
fof(f912,plain,(
  spl0_18 <=> ! [X9,X8,X6] : (~c(X6) | ~d1(f(X8,X9)) | ~d1(f(X9,X6)) | ~d1(X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_18])])).
fof(f2313,plain,(
  ( ! [X28,X26,X27] : (~c(X26) | ~c(X27) | ~d1(f(X28,X26)) | ~d1(X28)) ) | (~spl0_15 | ~spl0_18)),
  inference(duplicate_literal_removal,[],[f2292])).
fof(f2292,plain,(
  ( ! [X28,X26,X27] : (~c(X26) | ~c(X27) | ~d1(f(X28,X26)) | ~c(X27) | ~d1(X28)) ) | (~spl0_15 | ~spl0_18)),
  inference(resolution,[],[f733,f913])).
fof(f913,plain,(
  ( ! [X6,X8,X9] : (~d1(f(X9,X6)) | ~d1(f(X8,X9)) | ~c(X6) | ~d1(X8)) ) | ~spl0_18),
  inference(avatar_component_clause,[],[f912])).
fof(f733,plain,(
  ( ! [X4,X5] : (d1(f(X4,X5)) | ~c(X4) | ~c(X5)) ) | ~spl0_15),
  inference(avatar_component_clause,[],[f732])).
fof(f5773,plain,(
  spl0_10 | spl0_10 | ~spl0_15 | ~spl0_103),
  inference(avatar_split_clause,[],[f5772,f5744,f732,f444,f444])).
fof(f5744,plain,(
  spl0_103 <=> ! [X25,X27] : (~c(X25) | ~d1(f(X27,X25)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_103])])).
fof(f5772,plain,(
  ( ! [X0,X1] : (~c(X0) | ~c(X1)) ) | (~spl0_15 | ~spl0_103)),
  inference(duplicate_literal_removal,[],[f5756])).
fof(f5756,plain,(
  ( ! [X0,X1] : (~c(X0) | ~c(X1) | ~c(X0)) ) | (~spl0_15 | ~spl0_103)),
  inference(resolution,[],[f5745,f733])).
fof(f5745,plain,(
  ( ! [X27,X25] : (~d1(f(X27,X25)) | ~c(X25)) ) | ~spl0_103),
  inference(avatar_component_clause,[],[f5744])).
fof(f5746,plain,(
  spl0_20 | spl0_103 | ~spl0_18 | ~spl0_94),
  inference(avatar_split_clause,[],[f5742,f5057,f912,f5744,f941])).
fof(f941,plain,(
  spl0_20 <=> ! [X3] : ~d1(X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_20])])).
fof(f5057,plain,(
  spl0_94 <=> ! [X209,X210,X212] : (~d2(f(X209,X210)) | ~c(f(X209,f(X210,X212))) | ~c(X212))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_94])])).
fof(f5742,plain,(
  ( ! [X26,X27,X25] : (~c(X25) | ~d1(X26) | ~d1(f(X27,X25))) ) | (~spl0_18 | ~spl0_94)),
  inference(subsumption_resolution,[],[f5741,f48])).
fof(f48,plain,(
  ( ! [X10,X11,X9] : (~d1(f(X11,X9)) | ~c(f(X10,X11)) | ~d1(X10) | ~c(X9)) )),
  inference(resolution,[],[f29,f10])).
fof(f10,axiom,(
  ( ! [X0,X1] : (c(f(X0,X1)) | ~d1(X0) | ~d1(X1)) )),
  file('./p/ALG012-1.tptp',unknown)).
fof(f29,plain,(
  ( ! [X2,X0,X1] : (~c(f(X0,f(X1,X2))) | ~c(X2) | ~c(f(X0,X1))) )),
  inference(superposition,[],[f25,f1])).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (f(X0,f(X1,X2)) = f(f(X0,X1),X2)) )),
  file('./p/ALG012-1.tptp',unknown)).
fof(f25,plain,(
  ( ! [X2,X3] : (~c(f(X2,X3)) | ~c(X3) | ~c(X2)) )),
  inference(subsumption_resolution,[],[f23,f3])).
fof(f3,axiom,(
  ( ! [X0] : (~d1(X0) | ~c(X0)) )),
  file('./p/ALG012-1.tptp',unknown)).
fof(f23,plain,(
  ( ! [X2,X3] : (~c(X2) | d1(f(X2,X3)) | ~c(X3) | ~c(f(X2,X3))) )),
  inference(resolution,[],[f9,f4])).
fof(f4,axiom,(
  ( ! [X0] : (~d2(X0) | ~c(X0)) )),
  file('./p/ALG012-1.tptp',unknown)).
fof(f9,axiom,(
  ( ! [X0,X1] : (d2(f(X0,X1)) | ~c(X0) | d1(f(X0,X1)) | ~c(X1)) )),
  file('./p/ALG012-1.tptp',unknown)).
fof(f5741,plain,(
  ( ! [X26,X27,X25] : (~c(X25) | ~d1(X26) | ~d1(f(X27,X25)) | c(f(X26,X27))) ) | (~spl0_18 | ~spl0_94)),
  inference(subsumption_resolution,[],[f5730,f913])).
fof(f5730,plain,(
  ( ! [X26,X27,X25] : (~c(X25) | ~d1(X26) | ~d1(f(X27,X25)) | d1(f(X26,X27)) | c(f(X26,X27))) ) | ~spl0_94),
  inference(resolution,[],[f5589,f2])).
fof(f2,axiom,(
  ( ! [X0] : (d2(X0) | d1(X0) | c(X0)) )),
  file('./p/ALG012-1.tptp',unknown)).
fof(f5589,plain,(
  ( ! [X39,X41,X40] : (~d2(f(X39,X40)) | ~c(X41) | ~d1(X39) | ~d1(f(X40,X41))) ) | ~spl0_94),
  inference(resolution,[],[f5058,f10])).
fof(f5058,plain,(
  ( ! [X212,X210,X209] : (~c(f(X209,f(X210,X212))) | ~d2(f(X209,X210)) | ~c(X212)) ) | ~spl0_94),
  inference(avatar_component_clause,[],[f5057])).
fof(f5556,plain,(
  spl0_94 | ~spl0_36),
  inference(avatar_split_clause,[],[f5554,f2021,f5057])).
fof(f2021,plain,(
  spl0_36 <=> ! [X7,X6] : (~c(f(X6,X7)) | ~d2(X6) | ~c(X7))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_36])])).
fof(f5554,plain,(
  ( ! [X2,X0,X1] : (~c(f(X0,f(X1,X2))) | ~d2(f(X0,X1)) | ~c(X2)) ) | ~spl0_36),
  inference(superposition,[],[f2022,f1])).
fof(f2022,plain,(
  ( ! [X6,X7] : (~c(f(X6,X7)) | ~d2(X6) | ~c(X7)) ) | ~spl0_36),
  inference(avatar_component_clause,[],[f2021])).
fof(f5535,plain,(
  spl0_10 | spl0_20 | ~spl0_7 | ~spl0_11 | ~spl0_102),
  inference(avatar_split_clause,[],[f5534,f5469,f447,f380,f941,f444])).
fof(f447,plain,(
  spl0_11 <=> ! [X2,X4] : (~c(X2) | ~d1(X4) | ~c(f(X4,X2)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_11])])).
fof(f5469,plain,(
  spl0_102 <=> ! [X7,X6] : (~c(X6) | ~d1(X7) | ~d2(f(X7,X6)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_102])])).
fof(f5534,plain,(
  ( ! [X21,X20] : (~d1(X20) | ~c(X21)) ) | (~spl0_7 | ~spl0_11 | ~spl0_102)),
  inference(subsumption_resolution,[],[f5533,f448])).
fof(f448,plain,(
  ( ! [X4,X2] : (~c(f(X4,X2)) | ~d1(X4) | ~c(X2)) ) | ~spl0_11),
  inference(avatar_component_clause,[],[f447])).
fof(f5533,plain,(
  ( ! [X21,X20] : (~d1(X20) | ~c(X21) | c(f(X20,X21))) ) | (~spl0_7 | ~spl0_102)),
  inference(subsumption_resolution,[],[f5528,f381])).
fof(f381,plain,(
  ( ! [X0,X3] : (~d1(f(X3,X0)) | ~d1(X3) | ~c(X0)) ) | ~spl0_7),
  inference(avatar_component_clause,[],[f380])).
fof(f5528,plain,(
  ( ! [X21,X20] : (~d1(X20) | ~c(X21) | d1(f(X20,X21)) | c(f(X20,X21))) ) | ~spl0_102),
  inference(resolution,[],[f5470,f2])).
fof(f5470,plain,(
  ( ! [X6,X7] : (~d2(f(X7,X6)) | ~d1(X7) | ~c(X6)) ) | ~spl0_102),
  inference(avatar_component_clause,[],[f5469])).
fof(f5471,plain,(
  spl0_102 | spl0_36 | ~spl0_15),
  inference(avatar_split_clause,[],[f5456,f732,f2021,f5469])).
fof(f5456,plain,(
  ( ! [X6,X4,X7,X5] : (~c(f(X4,X5)) | ~c(X6) | ~d1(X7) | ~c(X5) | ~d2(f(X7,X6)) | ~d2(X4)) ) | ~spl0_15),
  inference(resolution,[],[f2297,f18])).
fof(f18,plain,(
  ( ! [X2,X0,X1] : (c(f(X0,f(X1,X2))) | ~d2(f(X0,X1)) | ~d2(X2)) )),
  inference(superposition,[],[f11,f1])).
fof(f11,axiom,(
  ( ! [X0,X1] : (c(f(X0,X1)) | ~d2(X0) | ~d2(X1)) )),
  file('./p/ALG012-1.tptp',unknown)).
fof(f2297,plain,(
  ( ! [X45,X43,X44,X42] : (~c(f(X45,f(X42,X43))) | ~c(f(X43,X44)) | ~c(X42) | ~d1(X45) | ~c(X44)) ) | ~spl0_15),
  inference(resolution,[],[f733,f78])).
fof(f78,plain,(
  ( ! [X2,X0,X3,X1] : (~d1(f(X0,f(X1,X2))) | ~c(f(X3,f(X0,X1))) | ~d1(X3) | ~c(X2)) )),
  inference(superposition,[],[f48,f1])).
fof(f3486,plain,(
  spl0_10 | spl0_2 | spl0_2 | ~spl0_1 | ~spl0_17 | ~spl0_30),
  inference(avatar_split_clause,[],[f3485,f1742,f871,f135,f138,f138,f444])).
fof(f138,plain,(
  spl0_2 <=> ! [X0] : ~d2(X0)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f135,plain,(
  spl0_1 <=> ! [X1,X2] : (~c(X1) | ~d2(X2) | c(f(X1,X2)) | d1(f(X1,X2)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f871,plain,(
  spl0_17 <=> ! [X3,X0,X2] : (~c(f(X2,X3)) | ~d2(X0) | ~d2(X3) | d1(f(X2,f(X3,X0))) | ~c(X2))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_17])])).
fof(f1742,plain,(
  spl0_30 <=> ! [X1,X0] : (~c(X0) | ~d2(X1) | ~d1(f(X0,X1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_30])])).
fof(f3485,plain,(
  ( ! [X6,X8,X7] : (~d2(X6) | ~d2(X8) | ~c(X7)) ) | (~spl0_1 | ~spl0_17 | ~spl0_30)),
  inference(subsumption_resolution,[],[f3482,f2195])).
fof(f2195,plain,(
  ( ! [X4,X5] : (c(f(X5,X4)) | ~c(X5) | ~d2(X4)) ) | (~spl0_1 | ~spl0_30)),
  inference(duplicate_literal_removal,[],[f2184])).
fof(f2184,plain,(
  ( ! [X4,X5] : (~d2(X4) | ~c(X5) | ~d2(X4) | c(f(X5,X4)) | ~c(X5)) ) | (~spl0_1 | ~spl0_30)),
  inference(resolution,[],[f1743,f136])).
fof(f136,plain,(
  ( ! [X2,X1] : (d1(f(X1,X2)) | ~d2(X2) | c(f(X1,X2)) | ~c(X1)) ) | ~spl0_1),
  inference(avatar_component_clause,[],[f135])).
fof(f1743,plain,(
  ( ! [X0,X1] : (~d1(f(X0,X1)) | ~d2(X1) | ~c(X0)) ) | ~spl0_30),
  inference(avatar_component_clause,[],[f1742])).
fof(f3482,plain,(
  ( ! [X6,X8,X7] : (~d2(X6) | ~c(f(X7,X8)) | ~d2(X8) | ~c(X7)) ) | (~spl0_17 | ~spl0_30)),
  inference(duplicate_literal_removal,[],[f3470])).
fof(f3470,plain,(
  ( ! [X6,X8,X7] : (~d2(X6) | ~c(f(X7,X8)) | ~d2(X6) | ~d2(X8) | ~c(f(X7,X8)) | ~c(X7)) ) | (~spl0_17 | ~spl0_30)),
  inference(resolution,[],[f2190,f872])).
fof(f872,plain,(
  ( ! [X2,X0,X3] : (d1(f(X2,f(X3,X0))) | ~d2(X0) | ~d2(X3) | ~c(f(X2,X3)) | ~c(X2)) ) | ~spl0_17),
  inference(avatar_component_clause,[],[f871])).
fof(f2190,plain,(
  ( ! [X2,X0,X1] : (~d1(f(X0,f(X1,X2))) | ~d2(X2) | ~c(f(X0,X1))) ) | ~spl0_30),
  inference(superposition,[],[f1743,f1])).
fof(f1973,plain,(
  ~spl0_20),
  inference(avatar_contradiction_clause,[],[f1963])).
fof(f1963,plain,(
  $false | ~spl0_20),
  inference(resolution,[],[f942,f7])).
fof(f7,axiom,(
  d1(a2)),
  file('./p/ALG012-1.tptp',unknown)).
fof(f942,plain,(
  ( ! [X3] : (~d1(X3)) ) | ~spl0_20),
  inference(avatar_component_clause,[],[f941])).
fof(f1744,plain,(
  spl0_20 | spl0_30 | ~spl0_12 | ~spl0_14),
  inference(avatar_split_clause,[],[f1740,f729,f451,f1742,f941])).
fof(f451,plain,(
  spl0_12 <=> ! [X7,X6] : (~c(X6) | ~c(f(X6,X7)) | ~d1(X7))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_12])])).
fof(f729,plain,(
  spl0_14 <=> ! [X7,X6] : (~d1(X6) | ~d2(X7) | d1(f(X7,X6)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_14])])).
fof(f1740,plain,(
  ( ! [X2,X0,X1] : (~c(X0) | ~d1(f(X0,X1)) | ~d1(X2) | ~d2(X1)) ) | (~spl0_12 | ~spl0_14)),
  inference(duplicate_literal_removal,[],[f1730])).
fof(f1730,plain,(
  ( ! [X2,X0,X1] : (~c(X0) | ~d1(f(X0,X1)) | ~d1(X2) | ~d2(X1) | ~d1(X2)) ) | (~spl0_12 | ~spl0_14)),
  inference(resolution,[],[f1675,f730])).
fof(f730,plain,(
  ( ! [X6,X7] : (d1(f(X7,X6)) | ~d2(X7) | ~d1(X6)) ) | ~spl0_14),
  inference(avatar_component_clause,[],[f729])).
fof(f1675,plain,(
  ( ! [X6,X4,X5] : (~d1(f(X5,X6)) | ~c(X4) | ~d1(f(X4,X5)) | ~d1(X6)) ) | ~spl0_12),
  inference(resolution,[],[f452,f19])).
fof(f19,plain,(
  ( ! [X4,X5,X3] : (c(f(X3,f(X4,X5))) | ~d1(f(X3,X4)) | ~d1(X5)) )),
  inference(superposition,[],[f10,f1])).
fof(f452,plain,(
  ( ! [X6,X7] : (~c(f(X6,X7)) | ~c(X6) | ~d1(X7)) ) | ~spl0_12),
  inference(avatar_component_clause,[],[f451])).
fof(f1655,plain,(
  spl0_12 | spl0_12 | ~spl0_25),
  inference(avatar_split_clause,[],[f1647,f1020,f451,f451])).
fof(f1020,plain,(
  spl0_25 <=> ! [X9,X8,X10] : (~c(X8) | ~d1(X9) | ~c(X10) | ~c(f(X8,X9)) | d1(f(X8,f(X9,X10))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_25])])).
fof(f1647,plain,(
  ( ! [X21,X19,X22,X20] : (~d1(X19) | ~c(X20) | ~c(f(X21,X19)) | ~c(X21) | ~d1(X22) | ~c(f(X20,X22))) ) | ~spl0_25),
  inference(duplicate_literal_removal,[],[f1622])).
fof(f1622,plain,(
  ( ! [X21,X19,X22,X20] : (~d1(X19) | ~c(X20) | ~c(f(X21,X19)) | ~c(X21) | ~d1(X22) | ~c(f(X20,X22)) | ~c(f(X21,X19))) ) | ~spl0_25),
  inference(resolution,[],[f1021,f58])).
fof(f58,plain,(
  ( ! [X2,X0,X3,X1] : (~d1(f(X0,f(X1,X2))) | ~d1(X3) | ~c(f(X2,X3)) | ~c(f(X0,X1))) )),
  inference(superposition,[],[f34,f1])).
fof(f34,plain,(
  ( ! [X2,X0,X1] : (~d1(f(X0,X1)) | ~d1(X2) | ~c(f(X1,X2)) | ~c(X0)) )),
  inference(resolution,[],[f19,f25])).
fof(f1021,plain,(
  ( ! [X10,X8,X9] : (d1(f(X8,f(X9,X10))) | ~d1(X9) | ~c(X10) | ~c(f(X8,X9)) | ~c(X8)) ) | ~spl0_25),
  inference(avatar_component_clause,[],[f1020])).
fof(f1022,plain,(
  spl0_2 | spl0_25 | ~spl0_1),
  inference(avatar_split_clause,[],[f1018,f135,f1020,f138])).
fof(f1018,plain,(
  ( ! [X10,X8,X11,X9] : (~c(X8) | d1(f(X8,f(X9,X10))) | ~c(f(X8,X9)) | ~c(X10) | ~d1(X9) | ~d2(X11)) ) | ~spl0_1),
  inference(subsumption_resolution,[],[f1017,f153])).
fof(f153,plain,(
  ( ! [X2,X0,X3,X1] : (~c(f(X2,X3)) | ~c(f(X1,X0)) | ~d2(X0) | d1(f(X2,f(X3,X1))) | ~c(X1)) )),
  inference(duplicate_literal_removal,[],[f147])).
fof(f147,plain,(
  ( ! [X2,X0,X3,X1] : (~d2(X0) | ~c(f(X1,X0)) | ~c(f(X2,X3)) | d1(f(X2,f(X3,X1))) | ~c(f(X2,X3)) | ~c(X1)) )),
  inference(resolution,[],[f55,f26])).
fof(f26,plain,(
  ( ! [X2,X0,X1] : (d2(f(X0,f(X1,X2))) | d1(f(X0,f(X1,X2))) | ~c(f(X0,X1)) | ~c(X2)) )),
  inference(forward_demodulation,[],[f24,f1])).
fof(f24,plain,(
  ( ! [X2,X0,X1] : (d2(f(X0,f(X1,X2))) | ~c(f(X0,X1)) | d1(f(f(X0,X1),X2)) | ~c(X2)) )),
  inference(superposition,[],[f9,f1])).
fof(f55,plain,(
  ( ! [X2,X0,X3,X1] : (~d2(f(X0,f(X1,X2))) | ~d2(X3) | ~c(f(X2,X3)) | ~c(f(X0,X1))) )),
  inference(superposition,[],[f30,f1])).
fof(f30,plain,(
  ( ! [X2,X0,X1] : (~d2(f(X0,X1)) | ~d2(X2) | ~c(f(X1,X2)) | ~c(X0)) )),
  inference(resolution,[],[f18,f25])).
fof(f1017,plain,(
  ( ! [X10,X8,X11,X9] : (~c(X8) | d1(f(X8,f(X9,X10))) | ~c(f(X8,X9)) | ~c(X10) | ~d1(X9) | ~d2(X11) | c(f(X10,X11))) ) | ~spl0_1),
  inference(duplicate_literal_removal,[],[f1012])).
fof(f1012,plain,(
  ( ! [X10,X8,X11,X9] : (~c(X8) | d1(f(X8,f(X9,X10))) | ~c(f(X8,X9)) | ~c(X10) | ~d1(X9) | ~d2(X11) | ~d2(X11) | c(f(X10,X11)) | ~c(X10)) ) | ~spl0_1),
  inference(resolution,[],[f276,f136])).
fof(f276,plain,(
  ( ! [X37,X35,X36,X34] : (~d1(f(X37,X34)) | ~c(X35) | d1(f(X35,f(X36,X37))) | ~c(f(X35,X36)) | ~c(X37) | ~d1(X36) | ~d2(X34)) )),
  inference(resolution,[],[f57,f10])).
fof(f57,plain,(
  ( ! [X6,X4,X5,X3] : (~c(f(X4,f(X5,X3))) | ~d2(X3) | ~c(X6) | d1(f(X6,f(X4,X5))) | ~c(f(X6,X4)) | ~c(X5)) )),
  inference(forward_demodulation,[],[f53,f1])).
fof(f53,plain,(
  ( ! [X6,X4,X5,X3] : (~d2(X3) | ~c(f(f(X4,X5),X3)) | ~c(X6) | d1(f(X6,f(X4,X5))) | ~c(f(X6,X4)) | ~c(X5)) )),
  inference(resolution,[],[f30,f26])).
fof(f943,plain,(
  spl0_20 | spl0_20 | ~spl0_10),
  inference(avatar_split_clause,[],[f931,f444,f941,f941])).
fof(f931,plain,(
  ( ! [X2,X3] : (~d1(X2) | ~d1(X3)) ) | ~spl0_10),
  inference(resolution,[],[f445,f10])).
fof(f445,plain,(
  ( ! [X3] : (~c(X3)) ) | ~spl0_10),
  inference(avatar_component_clause,[],[f444])).
fof(f914,plain,(
  spl0_10 | spl0_18 | ~spl0_15),
  inference(avatar_split_clause,[],[f906,f732,f912,f444])).
fof(f906,plain,(
  ( ! [X6,X8,X7,X9] : (~c(X6) | ~c(X7) | ~d1(f(X8,X9)) | ~d1(X8) | ~d1(f(X9,X6))) ) | ~spl0_15),
  inference(duplicate_literal_removal,[],[f889])).
fof(f889,plain,(
  ( ! [X6,X8,X7,X9] : (~c(X6) | ~c(X7) | ~d1(f(X8,X9)) | ~c(X7) | ~d1(X8) | ~d1(f(X9,X6))) ) | ~spl0_15),
  inference(resolution,[],[f733,f372])).
fof(f372,plain,(
  ( ! [X37,X35,X36,X34] : (~d1(f(X37,X34)) | ~d1(f(X35,X36)) | ~c(X34) | ~d1(X35) | ~d1(f(X36,X37))) )),
  inference(resolution,[],[f103,f10])).
fof(f103,plain,(
  ( ! [X10,X8,X11,X9] : (~c(f(X9,f(X10,X11))) | ~c(X8) | ~d1(f(X9,X10)) | ~d1(f(X11,X8))) )),
  inference(resolution,[],[f49,f19])).
fof(f49,plain,(
  ( ! [X2,X0,X3,X1] : (~c(f(X3,f(X0,f(X1,X2)))) | ~c(X2) | ~c(f(X3,f(X0,X1)))) )),
  inference(superposition,[],[f29,f1])).
fof(f910,plain,(
  spl0_10 | spl0_11 | ~spl0_15),
  inference(avatar_split_clause,[],[f907,f732,f447,f444])).
fof(f907,plain,(
  ( ! [X4,X5,X3] : (~c(X3) | ~c(X4) | ~c(f(X5,X3)) | ~d1(X5)) ) | ~spl0_15),
  inference(duplicate_literal_removal,[],[f888])).
fof(f888,plain,(
  ( ! [X4,X5,X3] : (~c(X3) | ~c(X4) | ~c(f(X5,X3)) | ~d1(X5) | ~c(X4)) ) | ~spl0_15),
  inference(resolution,[],[f733,f48])).
fof(f873,plain,(
  spl0_2 | spl0_17),
  inference(avatar_split_clause,[],[f869,f871,f138])).
fof(f869,plain,(
  ( ! [X2,X0,X3,X1] : (~c(f(X2,X3)) | ~c(X2) | d1(f(X2,f(X3,X0))) | ~d2(X1) | ~d2(X3) | ~d2(X0)) )),
  inference(subsumption_resolution,[],[f860,f11])).
fof(f860,plain,(
  ( ! [X2,X0,X3,X1] : (~c(f(X0,X1)) | ~c(f(X2,X3)) | ~c(X2) | d1(f(X2,f(X3,X0))) | ~d2(X1) | ~d2(X3) | ~d2(X0)) )),
  inference(resolution,[],[f149,f11])).
fof(f149,plain,(
  ( ! [X12,X10,X11,X9] : (~c(f(X12,X10)) | ~c(f(X10,X9)) | ~c(f(X11,X12)) | ~c(X11) | d1(f(X11,f(X12,X10))) | ~d2(X9)) )),
  inference(resolution,[],[f55,f9])).
fof(f734,plain,(
  spl0_14 | spl0_15 | ~spl0_1),
  inference(avatar_split_clause,[],[f727,f135,f732,f729])).
fof(f727,plain,(
  ( ! [X6,X4,X7,X5] : (d1(f(X4,X5)) | ~c(X5) | ~c(X4) | ~d1(X6) | d1(f(X7,X6)) | ~d2(X7)) ) | ~spl0_1),
  inference(subsumption_resolution,[],[f726,f56])).
fof(f56,plain,(
  ( ! [X2,X0,X1] : (~c(f(X1,X0)) | ~d2(X0) | ~c(X2) | d1(f(X2,X1)) | ~c(X1)) )),
  inference(duplicate_literal_removal,[],[f52])).
fof(f52,plain,(
  ( ! [X2,X0,X1] : (~d2(X0) | ~c(f(X1,X0)) | ~c(X2) | ~c(X2) | d1(f(X2,X1)) | ~c(X1)) )),
  inference(resolution,[],[f30,f9])).
fof(f726,plain,(
  ( ! [X6,X4,X7,X5] : (d1(f(X4,X5)) | ~c(X5) | ~c(X4) | ~d1(X6) | d1(f(X7,X6)) | ~d2(X7) | c(f(X5,X7))) ) | ~spl0_1),
  inference(duplicate_literal_removal,[],[f721])).
fof(f721,plain,(
  ( ! [X6,X4,X7,X5] : (d1(f(X4,X5)) | ~c(X5) | ~c(X4) | ~d1(X6) | d1(f(X7,X6)) | ~d2(X7) | c(f(X5,X7)) | ~c(X5)) ) | ~spl0_1),
  inference(resolution,[],[f458,f136])).
fof(f458,plain,(
  ( ! [X24,X23,X25,X22] : (~d1(f(X23,X24)) | d1(f(X22,X23)) | ~c(X23) | ~c(X22) | ~d1(X25) | d1(f(X24,X25))) )),
  inference(subsumption_resolution,[],[f429,f34])).
fof(f429,plain,(
  ( ! [X24,X23,X25,X22] : (~c(X22) | d1(f(X22,X23)) | ~c(X23) | ~d1(f(X23,X24)) | ~d1(X25) | d1(f(X24,X25)) | c(f(X24,X25))) )),
  inference(resolution,[],[f95,f2])).
fof(f95,plain,(
  ( ! [X6,X8,X7,X9] : (~d2(f(X6,X7)) | ~c(X8) | d1(f(X8,X9)) | ~c(X9) | ~d1(f(X9,X6)) | ~d1(X7)) )),
  inference(resolution,[],[f56,f19])).
fof(f146,plain,(
  ~spl0_2),
  inference(avatar_contradiction_clause,[],[f141])).
fof(f141,plain,(
  $false | ~spl0_2),
  inference(resolution,[],[f139,f8])).
fof(f8,axiom,(
  d2(a3)),
  file('./p/ALG012-1.tptp',unknown)).
fof(f139,plain,(
  ( ! [X0] : (~d2(X0)) ) | ~spl0_2),
  inference(avatar_component_clause,[],[f138])).
fof(f140,plain,(
  spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f133,f138,f135])).
fof(f133,plain,(
  ( ! [X2,X0,X1] : (~d2(X0) | ~c(X1) | d1(f(X1,X2)) | c(f(X1,X2)) | ~d2(X2)) )),
  inference(duplicate_literal_removal,[],[f126])).
fof(f126,plain,(
  ( ! [X2,X0,X1] : (~d2(X0) | ~c(X1) | d1(f(X1,X2)) | c(f(X1,X2)) | ~d2(X2) | ~d2(X0)) )),
  inference(resolution,[],[f54,f11])).
fof(f54,plain,(
  ( ! [X8,X7,X9] : (~c(f(X8,X7)) | ~d2(X7) | ~c(X9) | d1(f(X9,X8)) | c(f(X9,X8))) )),
  inference(resolution,[],[f30,f2])).
% SZS output end Proof for ALG012-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 9083
% Time elapsed: 0.841 s
% ------------------------------
% ------------------------------
