% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for ALG190+1
% SZS output start Proof for ALG190+1
fof(f13084,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f631,f652,f673,f694,f715,f736,f757,f778,f799,f820,f841,f862,f883,f904,f925,f946,f967,f988,f1009,f1030,f1051,f1072,f1093,f1114,f1135,f1136,f1137,f1138,f1139,f1140,f1141,f1142,f1143,f1144,f1145,f1146,f1147,f1148,f1149,f1150,f1151,f1152,f1153,f1154,f1155,f1156,f1157,f1158,f1159,f1160,f1161,f1162,f1163,f1164,f1165,f1166,f1167,f1168,f1169,f1170,f1171,f1172,f1173,f1174,f1175,f1176,f1177,f1178,f1179,f1180,f1181,f1182,f1183,f1184,f1185,f1191,f1193,f1194,f1199,f1200,f1201,f1208,f1210,f1211,f1212,f1218,f1219,f1220,f1221,f1222,f1225,f1226,f1252,f1261,f1274,f1302,f1320,f1346,f1357,f1372,f1378,f1379,f1385,f1395,f1408,f1423,f1439,f1468,f1469,f1470,f1471,f1476,f1477,f1495,f1502,f1522,f1525,f1526,f1532,f1533,f1536,f1539,f1544,f1551,f1556,f1557,f1565,f1576,f1577,f1579,f1582,f1584,f1602,f1605,f1622,f1623,f1637,f1651,f1669,f1672,f1684,f1726,f1752,f1781,f1787,f1799,f1815,f1825,f1829,f1838,f1860,f1895,f1898,f1901,f1913,f1919,f1922,f1932,f1938,f1941,f1949,f1952,f1955,f1958,f1964,f1967,f1977,f1987,f1992,f1993,f1996,f1997,f2000,f2003,f2009,f2022,f2042,f2048,f2059,f2064,f2076,f2091,f2094,f2095,f2103,f2111,f2114,f2142,f2145,f2164,f2177,f2183,f2186,f2197,f2199,f2206,f2218,f2232,f2244,f2252,f2263,f2266,f2275,f2284,f2290,f2293,f2318,f2319,f2338,f2340,f2368,f2375,f2379,f2383,f2386,f2401,f2405,f2417,f2424,f2425,f2433,f2449,f2482,f2484,f2490,f2497,f2501,f2526,f2534,f2535,f2539,f2560,f2576,f2585,f2590,f2600,f2607,f2618,f2631,f2643,f2645,f2648,f2657,f2671,f2680,f2702,f2705,f2708,f2728,f2758,f2782,f2809,f2839,f2852,f2878,f2882,f2895,f2907,f2913,f2924,f2941,f2944,f2952,f2960,f2975,f2977,f2995,f3036,f3040,f3044,f3065,f3082,f3089,f3092,f3093,f3098,f3116,f3132,f3133,f3179,f3182,f3187,f3213,f3216,f3226,f3229,f3246,f3251,f3303,f3311,f3341,f3376,f3378,f3380,f3383,f3394,f3409,f3430,f3440,f3454,f3460,f3467,f3497,f3515,f3541,f3548,f3560,f3579,f3583,f3594,f3601,f3610,f3623,f3636,f3638,f3639,f3647,f3680,f3690,f3694,f3713,f3722,f3725,f3729,f3747,f3750,f3764,f3771,f3783,f3799,f3814,f3820,f3837,f3887,f3895,f3907,f3914,f3920,f3926,f3927,f3958,f3964,f3966,f3981,f3999,f4007,f4008,f4009,f4011,f4013,f4020,f4026,f4039,f4072,f4088,f4113,f4146,f4151,f4160,f4167,f4187,f4204,f4240,f4264,f4280,f4285,f4288,f4293,f4296,f4306,f4342,f4361,f4363,f4367,f4373,f4410,f4414,f4423,f4432,f4436,f4448,f4449,f4454,f4468,f4473,f4486,f4494,f4501,f4502,f4514,f4523,f4534,f4542,f4547,f4558,f4566,f4617,f4618,f4644,f4647,f4659,f4661,f4674,f4680,f4695,f4724,f4734,f4737,f4740,f4749,f4750,f4755,f4764,f4787,f4789,f4792,f4799,f4800,f4813,f4838,f4842,f4848,f4865,f4871,f4900,f4941,f4955,f4957,f4973,f4975,f4986,f4993,f4996,f4997,f5008,f5016,f5031,f5044,f5053,f5073,f5086,f5121,f5176,f5206,f5216,f5218,f5225,f5226,f5260,f5266,f5275,f5277,f5279,f5290,f5302,f5315,f5334,f5346,f5349,f5354,f5376,f5428,f5429,f5438,f5439,f5472,f5504,f5506,f5528,f5551,f5621,f5627,f5653,f5656,f5673,f5697,f5727,f5736,f5770,f5772,f5775,f5784,f5786,f5790,f5803,f5804,f5817,f5830,f5845,f5854,f5857,f5866,f5870,f5879,f5892,f5931,f5950,f5955,f5976,f5977,f5978,f5996,f5998,f6036,f6037,f6038,f6050,f6061,f6081,f6158,f6172,f6187,f6223,f6225,f6229,f6250,f6278,f6283,f6329,f6331,f6356,f6367,f6387,f6392,f6434,f6460,f6557,f6610,f6631,f6680,f6681,f6687,f6721,f6741,f6742,f6754,f6763,f6794,f6807,f6808,f6817,f6826,f6834,f6835,f6878,f6929,f6939,f6951,f6973,f7002,f7003,f7004,f7024,f7031,f7032,f7033,f7053,f7054,f7061,f7076,f7077,f7081,f7084,f7104,f7115,f7132,f7146,f7157,f7165,f7167,f7186,f7203,f7209,f7272,f7278,f7280,f7325,f7329,f7333,f7350,f7355,f7367,f7386,f7406,f7422,f7432,f7433,f7434,f7435,f7451,f7480,f7521,f7535,f7545,f7551,f7561,f7570,f7604,f7624,f7635,f7636,f7637,f7638,f7652,f7653,f7654,f7699,f7700,f7727,f7731,f7734,f7753,f7754,f7766,f7795,f7803,f7831,f7837,f7882,f7913,f7961,f7963,f7964,f7975,f7984,f7985,f7999,f8001,f8016,f8036,f8045,f8070,f8092,f8097,f8159,f8160,f8173,f8187,f8195,f8197,f8198,f8231,f8244,f8252,f8253,f8273,f8277,f8311,f8312,f8313,f8332,f8335,f8373,f8386,f8415,f8423,f8440,f8453,f8466,f8486,f8510,f8511,f8539,f8547,f8578,f8589,f8591,f8614,f8624,f8629,f8637,f8638,f8648,f8649,f8680,f8720,f8723,f8725,f8751,f8767,f8773,f8783,f8785,f8786,f8796,f8798,f8806,f8808,f8824,f8825,f8836,f8847,f8865,f8875,f8876,f8901,f8912,f8914,f8931,f8936,f8951,f8955,f8957,f8959,f8981,f8983,f9018,f9023,f9058,f9059,f9060,f9073,f9091,f9111,f9133,f9135,f9146,f9147,f9148,f9162,f9164,f9175,f9177,f9189,f9196,f9197,f9207,f9236,f9238,f9248,f9268,f9292,f9293,f9317,f9318,f9349,f9350,f9351,f9352,f9367,f9368,f9378,f9379,f9389,f9390,f9391,f9392,f9414,f9415,f9426,f9441,f9450,f9451,f9463,f9476,f9477,f9487,f9501,f9511,f9521,f9555,f9566,f9625,f9627,f9628,f9639,f9640,f9657,f9658,f9667,f9670,f9683,f9684,f9686,f9688,f9729,f9741,f9751,f9761,f9763,f9779,f9781,f9783,f9790,f9802,f9812,f9814,f9825,f9846,f9848,f9860,f9870,f9894,f9895,f9896,f9907,f9910,f9933,f9936,f9967,f9968,f9985,f9995,f9997,f9999,f10011,f10014,f10028,f10050,f10065,f10069,f10071,f10073,f10081,f10103,f10117,f10140,f10161,f10173,f10174,f10175,f10196,f10197,f10198,f10199,f10207,f10222,f10225,f10239,f10305,f10306,f10316,f10352,f10375,f10387,f10414,f10415,f10420,f10431,f10444,f10449,f10463,f10469,f10479,f10480,f10482,f10492,f10507,f10508,f10570,f10571,f10580,f10619,f10621,f10627,f10650,f10658,f10660,f10664,f10697,f10698,f10724,f10737,f10738,f10739,f10750,f10751,f10763,f10765,f10766,f10774,f10800,f10818,f10826,f10827,f10837,f10838,f10860,f10861,f10867,f10876,f10887,f10924,f10927,f10943,f10946,f10949,f10967,f10979,f10980,f10987,f10988,f10989,f11001,f11002,f11018,f11023,f11034,f11035,f11062,f11073,f11074,f11086,f11087,f11088,f11117,f11124,f11133,f11158,f11161,f11209,f11236,f11238,f11249,f11258,f11266,f11267,f11280,f11281,f11282,f11290,f11291,f11303,f11304,f11305,f11330,f11340,f11341,f11342,f11351,f11368,f11385,f11402,f11407,f11434,f11436,f11445,f11453,f11465,f11475,f11489,f11501,f11502,f11513,f11514,f11515,f11521,f11522,f11523,f11530,f11531,f11541,f11542,f11561,f11580,f11582,f11584,f11586,f11615,f11621,f11623,f11631,f11676,f11677,f11688,f11699,f11713,f11714,f11729,f11730,f11731,f11733,f11740,f11749,f11758,f11759,f11772,f11773,f11785,f11811,f11812,f11821,f11822,f11849,f11864,f11880,f11894,f11936,f11937,f11938,f11939,f11940,f11956,f11957,f11958,f11971,f11972,f11982,f11999,f12000,f12001,f12002,f12020,f12033,f12041,f12049,f12068,f12084,f12121,f12130,f12137,f12145,f12146,f12163,f12164,f12177,f12283,f12284,f12285,f12299,f12300,f12309,f12317,f12318,f12328,f12329,f12346,f12347,f12365,f12372,f12381,f12389,f12390,f12400,f12401,f12402,f12403,f12420,f12421,f12432,f12439,f12440,f12448,f12467,f12471,f12498,f12511,f12534,f12535,f12536,f12549,f12550,f12560,f12572,f12581,f12594,f12595,f12596,f12611,f12612,f12613,f12629,f12639,f12640,f12653,f12654,f12667,f12668,f12689,f12690,f12696,f12697,f12698,f12731,f12745,f12746,f12798,f12817,f12824,f12849,f12876,f12877,f12890,f12891,f12892,f12893,f12903,f12913,f12928,f12929,f12930,f12941,f12942,f12943,f12947,f12958,f12959,f12970,f12974,f12990,f13009,f13010,f13025,f13031,f13032,f13033,f13054,f13065,f13076,f13080])).
fof(f13080,plain,(
  ~spl4_10 | ~spl4_15),
  inference(avatar_split_clause,[],[f13077,f670,f649])).
fof(f649,plain,(
  spl4_10 <=> e4 = op(e4,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_10])])).
fof(f670,plain,(
  spl4_15 <=> e4 = op(e4,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_15])])).
fof(f13077,plain,(
  e4 != op(e4,e3) | ~spl4_15),
  inference(backward_demodulation,[],[f672,f453])).
fof(f453,plain,(
  op(e4,e2) != op(e4,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f4,axiom,(
  op(e4,e3) != op(e4,e4) & op(e4,e2) != op(e4,e4) & op(e4,e2) != op(e4,e3) & op(e4,e1) != op(e4,e4) & op(e4,e1) != op(e4,e3) & op(e4,e1) != op(e4,e2) & op(e4,e0) != op(e4,e4) & op(e4,e0) != op(e4,e3) & op(e4,e0) != op(e4,e2) & op(e4,e0) != op(e4,e1) & op(e3,e3) != op(e3,e4) & op(e3,e2) != op(e3,e4) & op(e3,e2) != op(e3,e3) & op(e3,e1) != op(e3,e4) & op(e3,e1) != op(e3,e3) & op(e3,e1) != op(e3,e2) & op(e3,e0) != op(e3,e4) & op(e3,e0) != op(e3,e3) & op(e3,e0) != op(e3,e2) & op(e3,e0) != op(e3,e1) & op(e2,e3) != op(e2,e4) & op(e2,e2) != op(e2,e4) & op(e2,e2) != op(e2,e3) & op(e2,e1) != op(e2,e4) & op(e2,e1) != op(e2,e3) & op(e2,e1) != op(e2,e2) & op(e2,e0) != op(e2,e4) & op(e2,e0) != op(e2,e3) & op(e2,e0) != op(e2,e2) & op(e2,e0) != op(e2,e1) & op(e1,e3) != op(e1,e4) & op(e1,e2) != op(e1,e4) & op(e1,e2) != op(e1,e3) & op(e1,e1) != op(e1,e4) & op(e1,e1) != op(e1,e3) & op(e1,e1) != op(e1,e2) & op(e1,e0) != op(e1,e4) & op(e1,e0) != op(e1,e3) & op(e1,e0) != op(e1,e2) & op(e1,e0) != op(e1,e1) & op(e0,e3) != op(e0,e4) & op(e0,e2) != op(e0,e4) & op(e0,e2) != op(e0,e3) & op(e0,e1) != op(e0,e4) & op(e0,e1) != op(e0,e3) & op(e0,e1) != op(e0,e2) & op(e0,e0) != op(e0,e4) & op(e0,e0) != op(e0,e3) & op(e0,e0) != op(e0,e2) & op(e0,e0) != op(e0,e1) & op(e3,e4) != op(e4,e4) & op(e2,e4) != op(e4,e4) & op(e2,e4) != op(e3,e4) & op(e1,e4) != op(e4,e4) & op(e1,e4) != op(e3,e4) & op(e1,e4) != op(e2,e4) & op(e0,e4) != op(e4,e4) & op(e0,e4) != op(e3,e4) & op(e0,e4) != op(e2,e4) & op(e0,e4) != op(e1,e4) & op(e3,e3) != op(e4,e3) & op(e2,e3) != op(e4,e3) & op(e2,e3) != op(e3,e3) & op(e1,e3) != op(e4,e3) & op(e1,e3) != op(e3,e3) & op(e1,e3) != op(e2,e3) & op(e0,e3) != op(e4,e3) & op(e0,e3) != op(e3,e3) & op(e0,e3) != op(e2,e3) & op(e0,e3) != op(e1,e3) & op(e3,e2) != op(e4,e2) & op(e2,e2) != op(e4,e2) & op(e2,e2) != op(e3,e2) & op(e1,e2) != op(e4,e2) & op(e1,e2) != op(e3,e2) & op(e1,e2) != op(e2,e2) & op(e0,e2) != op(e4,e2) & op(e0,e2) != op(e3,e2) & op(e0,e2) != op(e2,e2) & op(e0,e2) != op(e1,e2) & op(e3,e1) != op(e4,e1) & op(e2,e1) != op(e4,e1) & op(e2,e1) != op(e3,e1) & op(e1,e1) != op(e4,e1) & op(e1,e1) != op(e3,e1) & op(e1,e1) != op(e2,e1) & op(e0,e1) != op(e4,e1) & op(e0,e1) != op(e3,e1) & op(e0,e1) != op(e2,e1) & op(e0,e1) != op(e1,e1) & op(e3,e0) != op(e4,e0) & op(e2,e0) != op(e4,e0) & op(e2,e0) != op(e3,e0) & op(e1,e0) != op(e4,e0) & op(e1,e0) != op(e3,e0) & op(e1,e0) != op(e2,e0) & op(e0,e0) != op(e4,e0) & op(e0,e0) != op(e3,e0) & op(e0,e0) != op(e2,e0) & op(e0,e0) != op(e1,e0)),
  file('./p/ALG190+1.tptp',unknown)).
fof(f672,plain,(
  e4 = op(e4,e2) | ~spl4_15),
  inference(avatar_component_clause,[],[f670])).
fof(f13076,plain,(
  spl4_15 | ~spl4_18 | ~spl4_80),
  inference(avatar_split_clause,[],[f13075,f943,f683,f670])).
fof(f683,plain,(
  spl4_18 <=> e2 = op(e4,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_18])])).
fof(f943,plain,(
  spl4_80 <=> e4 = op(e1,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_80])])).
fof(f13075,plain,(
  e4 = op(e4,e2) | (~spl4_18 | ~spl4_80)),
  inference(backward_demodulation,[],[f685,f12972])).
fof(f12972,plain,(
  e4 = op(e4,op(e4,e1)) | ~spl4_80),
  inference(backward_demodulation,[],[f945,f352])).
fof(f352,plain,(
  e4 = op(op(e1,e4),op(e4,e1))),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  e4 = op(op(e4,e4),op(e4,e4)) & e4 = op(op(e3,e4),op(e4,e3)) & e4 = op(op(e2,e4),op(e4,e2)) & e4 = op(op(e1,e4),op(e4,e1)) & e4 = op(op(e0,e4),op(e4,e0)) & e3 = op(op(e4,e3),op(e3,e4)) & e3 = op(op(e3,e3),op(e3,e3)) & e3 = op(op(e2,e3),op(e3,e2)) & e3 = op(op(e1,e3),op(e3,e1)) & e3 = op(op(e0,e3),op(e3,e0)) & e2 = op(op(e4,e2),op(e2,e4)) & e2 = op(op(e3,e2),op(e2,e3)) & e2 = op(op(e2,e2),op(e2,e2)) & e2 = op(op(e1,e2),op(e2,e1)) & e2 = op(op(e0,e2),op(e2,e0)) & e1 = op(op(e4,e1),op(e1,e4)) & e1 = op(op(e3,e1),op(e1,e3)) & e1 = op(op(e2,e1),op(e1,e2)) & e1 = op(op(e1,e1),op(e1,e1)) & e1 = op(op(e0,e1),op(e1,e0)) & e0 = op(op(e4,e0),op(e0,e4)) & e0 = op(op(e3,e0),op(e0,e3)) & e0 = op(op(e2,e0),op(e0,e2)) & e0 = op(op(e1,e0),op(e0,e1)) & e0 = op(op(e0,e0),op(e0,e0))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f945,plain,(
  e4 = op(e1,e4) | ~spl4_80),
  inference(avatar_component_clause,[],[f943])).
fof(f685,plain,(
  e2 = op(e4,e1) | ~spl4_18),
  inference(avatar_component_clause,[],[f683])).
fof(f13065,plain,(
  ~spl4_17 | ~spl4_22),
  inference(avatar_split_clause,[],[f13055,f700,f679])).
fof(f679,plain,(
  spl4_17 <=> e1 = op(e4,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_17])])).
fof(f700,plain,(
  spl4_22 <=> e1 = op(e4,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_22])])).
fof(f13055,plain,(
  e1 != op(e4,e1) | ~spl4_22),
  inference(backward_demodulation,[],[f702,f446])).
fof(f446,plain,(
  op(e4,e0) != op(e4,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f702,plain,(
  e1 = op(e4,e0) | ~spl4_22),
  inference(avatar_component_clause,[],[f700])).
fof(f13054,plain,(
  ~spl4_32 | ~spl4_45 | spl4_186),
  inference(avatar_contradiction_clause,[],[f13053])).
fof(f13053,plain,(
  $false | (~spl4_32 | ~spl4_45 | spl4_186)),
  inference(subsumption_resolution,[],[f13045,f798])).
fof(f798,plain,(
  e4 = op(e3,e1) | ~spl4_45),
  inference(avatar_component_clause,[],[f796])).
fof(f796,plain,(
  spl4_45 <=> e4 = op(e3,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_45])])).
fof(f13045,plain,(
  e4 != op(e3,e1) | (~spl4_32 | spl4_186)),
  inference(backward_demodulation,[],[f744,f1550])).
fof(f1550,plain,(
  e4 != op(e3,op(e3,e3)) | spl4_186),
  inference(avatar_component_clause,[],[f1548])).
fof(f1548,plain,(
  spl4_186 <=> e4 = op(e3,op(e3,e3))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_186])])).
fof(f744,plain,(
  e1 = op(e3,e3) | ~spl4_32),
  inference(avatar_component_clause,[],[f742])).
fof(f742,plain,(
  spl4_32 <=> e1 = op(e3,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_32])])).
fof(f13033,plain,(
  spl4_9 | ~spl4_39 | ~spl4_60),
  inference(avatar_split_clause,[],[f13030,f859,f771,f645])).
fof(f645,plain,(
  spl4_9 <=> e3 = op(e4,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_9])])).
fof(f771,plain,(
  spl4_39 <=> e3 = op(e3,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_39])])).
fof(f859,plain,(
  spl4_60 <=> e4 = op(e2,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_60])])).
fof(f13030,plain,(
  e3 = op(e4,e3) | (~spl4_39 | ~spl4_60)),
  inference(backward_demodulation,[],[f773,f13005])).
fof(f13005,plain,(
  e3 = op(e4,op(e3,e2)) | ~spl4_60),
  inference(backward_demodulation,[],[f861,f348])).
fof(f348,plain,(
  e3 = op(op(e2,e3),op(e3,e2))),
  inference(cnf_transformation,[],[f3])).
fof(f861,plain,(
  e4 = op(e2,e3) | ~spl4_60),
  inference(avatar_component_clause,[],[f859])).
fof(f773,plain,(
  e3 = op(e3,e2) | ~spl4_39),
  inference(avatar_component_clause,[],[f771])).
fof(f13032,plain,(
  ~spl4_34 | ~spl4_39),
  inference(avatar_split_clause,[],[f13027,f771,f750])).
fof(f750,plain,(
  spl4_34 <=> e3 = op(e3,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_34])])).
fof(f13027,plain,(
  e3 != op(e3,e3) | ~spl4_39),
  inference(backward_demodulation,[],[f773,f443])).
fof(f443,plain,(
  op(e3,e2) != op(e3,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f13031,plain,(
  ~spl4_14 | ~spl4_39),
  inference(avatar_split_clause,[],[f13026,f771,f666])).
fof(f666,plain,(
  spl4_14 <=> e3 = op(e4,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_14])])).
fof(f13026,plain,(
  e3 != op(e4,e2) | ~spl4_39),
  inference(backward_demodulation,[],[f773,f385])).
fof(f385,plain,(
  op(e3,e2) != op(e4,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f13025,plain,(
  ~spl4_36 | ~spl4_46),
  inference(avatar_split_clause,[],[f13017,f801,f759])).
fof(f759,plain,(
  spl4_36 <=> e0 = op(e3,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_36])])).
fof(f801,plain,(
  spl4_46 <=> e0 = op(e3,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_46])])).
fof(f13017,plain,(
  e0 != op(e3,e2) | ~spl4_46),
  inference(backward_demodulation,[],[f803,f437])).
fof(f437,plain,(
  op(e3,e0) != op(e3,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f803,plain,(
  e0 = op(e3,e0) | ~spl4_46),
  inference(avatar_component_clause,[],[f801])).
fof(f13010,plain,(
  ~spl4_55 | ~spl4_60),
  inference(avatar_split_clause,[],[f13008,f859,f838])).
fof(f838,plain,(
  spl4_55 <=> e4 = op(e2,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_55])])).
fof(f13008,plain,(
  e4 != op(e2,e4) | ~spl4_60),
  inference(backward_demodulation,[],[f861,f435])).
fof(f435,plain,(
  op(e2,e3) != op(e2,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f13009,plain,(
  ~spl4_10 | ~spl4_60),
  inference(avatar_split_clause,[],[f13007,f859,f649])).
fof(f13007,plain,(
  e4 != op(e4,e3) | ~spl4_60),
  inference(backward_demodulation,[],[f861,f394])).
fof(f394,plain,(
  op(e2,e3) != op(e4,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f12990,plain,(
  ~spl4_58 | ~spl4_63),
  inference(avatar_split_clause,[],[f12978,f872,f851])).
fof(f851,plain,(
  spl4_58 <=> e2 = op(e2,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_58])])).
fof(f872,plain,(
  spl4_63 <=> e2 = op(e2,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_63])])).
fof(f12978,plain,(
  e2 != op(e2,e3) | ~spl4_63),
  inference(backward_demodulation,[],[f874,f433])).
fof(f433,plain,(
  op(e2,e2) != op(e2,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f874,plain,(
  e2 = op(e2,e2) | ~spl4_63),
  inference(avatar_component_clause,[],[f872])).
fof(f12974,plain,(
  ~spl4_55 | ~spl4_80),
  inference(avatar_split_clause,[],[f12973,f943,f838])).
fof(f12973,plain,(
  e4 != op(e2,e4) | ~spl4_80),
  inference(backward_demodulation,[],[f945,f400])).
fof(f400,plain,(
  op(e1,e4) != op(e2,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f12970,plain,(
  spl4_22 | ~spl4_45 | ~spl4_81),
  inference(avatar_split_clause,[],[f12969,f948,f796,f700])).
fof(f948,plain,(
  spl4_81 <=> e0 = op(e1,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_81])])).
fof(f12969,plain,(
  e1 = op(e4,e0) | (~spl4_45 | ~spl4_81)),
  inference(backward_demodulation,[],[f950,f12832])).
fof(f12832,plain,(
  e1 = op(e4,op(e1,e3)) | ~spl4_45),
  inference(forward_demodulation,[],[f339,f798])).
fof(f339,plain,(
  e1 = op(op(e3,e1),op(e1,e3))),
  inference(cnf_transformation,[],[f3])).
fof(f950,plain,(
  e0 = op(e1,e3) | ~spl4_81),
  inference(avatar_component_clause,[],[f948])).
fof(f12959,plain,(
  ~spl4_82 | ~spl4_87),
  inference(avatar_split_clause,[],[f12951,f973,f952])).
fof(f952,plain,(
  spl4_82 <=> e1 = op(e1,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_82])])).
fof(f973,plain,(
  spl4_87 <=> e1 = op(e1,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_87])])).
fof(f12951,plain,(
  e1 != op(e1,e3) | ~spl4_87),
  inference(backward_demodulation,[],[f975,f423])).
fof(f423,plain,(
  op(e1,e2) != op(e1,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f975,plain,(
  e1 = op(e1,e2) | ~spl4_87),
  inference(avatar_component_clause,[],[f973])).
fof(f12958,plain,(
  ~spl4_62 | ~spl4_87),
  inference(avatar_split_clause,[],[f12948,f973,f868])).
fof(f868,plain,(
  spl4_62 <=> e1 = op(e2,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_62])])).
fof(f12948,plain,(
  e1 != op(e2,e2) | ~spl4_87),
  inference(backward_demodulation,[],[f975,f380])).
fof(f380,plain,(
  op(e1,e2) != op(e2,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f12947,plain,(
  ~spl4_94 | ~spl4_108 | ~spl4_148 | spl4_155),
  inference(avatar_contradiction_clause,[],[f12946])).
fof(f12946,plain,(
  $false | (~spl4_94 | ~spl4_108 | ~spl4_148 | spl4_155)),
  inference(subsumption_resolution,[],[f12940,f1063])).
fof(f1063,plain,(
  e2 = op(e0,e3) | ~spl4_108),
  inference(avatar_component_clause,[],[f1061])).
fof(f1061,plain,(
  spl4_108 <=> e2 = op(e0,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_108])])).
fof(f12940,plain,(
  e2 != op(e0,e3) | (~spl4_94 | ~spl4_148 | spl4_155)),
  inference(backward_demodulation,[],[f1004,f12861])).
fof(f12861,plain,(
  e2 != op(e0,op(e1,e1)) | (~spl4_148 | spl4_155)),
  inference(backward_demodulation,[],[f1310,f1345])).
fof(f1345,plain,(
  e2 != op(op(e1,op(e1,e1)),op(e1,e1)) | spl4_155),
  inference(avatar_component_clause,[],[f1343])).
fof(f1343,plain,(
  spl4_155 <=> e2 = op(op(e1,op(e1,e1)),op(e1,e1))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_155])])).
fof(f1310,plain,(
  e0 = op(e1,op(e1,e1)) | ~spl4_148),
  inference(avatar_component_clause,[],[f1309])).
fof(f1309,plain,(
  spl4_148 <=> e0 = op(e1,op(e1,e1))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_148])])).
fof(f1004,plain,(
  e3 = op(e1,e1) | ~spl4_94),
  inference(avatar_component_clause,[],[f1002])).
fof(f1002,plain,(
  spl4_94 <=> e3 = op(e1,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_94])])).
fof(f12943,plain,(
  spl4_81 | ~spl4_94 | ~spl4_148),
  inference(avatar_split_clause,[],[f12937,f1309,f1002,f948])).
fof(f12937,plain,(
  e0 = op(e1,e3) | (~spl4_94 | ~spl4_148)),
  inference(backward_demodulation,[],[f1004,f1310])).
fof(f12942,plain,(
  ~spl4_79 | ~spl4_94),
  inference(avatar_split_clause,[],[f12935,f1002,f939])).
fof(f939,plain,(
  spl4_79 <=> e3 = op(e1,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_79])])).
fof(f12935,plain,(
  e3 != op(e1,e4) | ~spl4_94),
  inference(backward_demodulation,[],[f1004,f422])).
fof(f422,plain,(
  op(e1,e1) != op(e1,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f12941,plain,(
  spl4_32 | ~spl4_94),
  inference(avatar_split_clause,[],[f12931,f1002,f742])).
fof(f12931,plain,(
  e1 = op(e3,e3) | ~spl4_94),
  inference(backward_demodulation,[],[f1004,f337])).
fof(f337,plain,(
  e1 = op(op(e1,e1),op(e1,e1))),
  inference(cnf_transformation,[],[f3])).
fof(f12930,plain,(
  spl4_87 | ~spl4_98 | ~spl4_117),
  inference(avatar_split_clause,[],[f12927,f1099,f1019,f973])).
fof(f1019,plain,(
  spl4_98 <=> e2 = op(e1,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_98])])).
fof(f1099,plain,(
  spl4_117 <=> e1 = op(e0,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_117])])).
fof(f12927,plain,(
  e1 = op(e1,e2) | (~spl4_98 | ~spl4_117)),
  inference(backward_demodulation,[],[f1021,f12867])).
fof(f12867,plain,(
  e1 = op(e1,op(e1,e0)) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f336])).
fof(f336,plain,(
  e1 = op(op(e0,e1),op(e1,e0))),
  inference(cnf_transformation,[],[f3])).
fof(f1101,plain,(
  e1 = op(e0,e1) | ~spl4_117),
  inference(avatar_component_clause,[],[f1099])).
fof(f1021,plain,(
  e2 = op(e1,e0) | ~spl4_98),
  inference(avatar_component_clause,[],[f1019])).
fof(f12929,plain,(
  ~spl4_93 | ~spl4_98),
  inference(avatar_split_clause,[],[f12920,f1019,f998])).
fof(f998,plain,(
  spl4_93 <=> e2 = op(e1,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_93])])).
fof(f12920,plain,(
  e2 != op(e1,e1) | ~spl4_98),
  inference(backward_demodulation,[],[f1021,f416])).
fof(f416,plain,(
  op(e1,e0) != op(e1,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f12928,plain,(
  ~spl4_23 | ~spl4_98),
  inference(avatar_split_clause,[],[f12919,f1019,f704])).
fof(f704,plain,(
  spl4_23 <=> e2 = op(e4,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_23])])).
fof(f12919,plain,(
  e2 != op(e4,e0) | ~spl4_98),
  inference(backward_demodulation,[],[f1021,f362])).
fof(f362,plain,(
  op(e1,e0) != op(e4,e0)),
  inference(cnf_transformation,[],[f4])).
fof(f12913,plain,(
  ~spl4_79 | ~spl4_104),
  inference(avatar_split_clause,[],[f12906,f1044,f939])).
fof(f1044,plain,(
  spl4_104 <=> e3 = op(e0,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_104])])).
fof(f12906,plain,(
  e3 != op(e1,e4) | ~spl4_104),
  inference(backward_demodulation,[],[f1046,f396])).
fof(f396,plain,(
  op(e0,e4) != op(e1,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f1046,plain,(
  e3 = op(e0,e4) | ~spl4_104),
  inference(avatar_component_clause,[],[f1044])).
fof(f12903,plain,(
  ~spl4_58 | ~spl4_108),
  inference(avatar_split_clause,[],[f12897,f1061,f851])).
fof(f12897,plain,(
  e2 != op(e2,e3) | ~spl4_108),
  inference(backward_demodulation,[],[f1063,f387])).
fof(f387,plain,(
  op(e0,e3) != op(e2,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f12893,plain,(
  spl4_46 | ~spl4_74 | ~spl4_111),
  inference(avatar_split_clause,[],[f12889,f1074,f918,f801])).
fof(f918,plain,(
  spl4_74 <=> e3 = op(e2,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_74])])).
fof(f1074,plain,(
  spl4_111 <=> e0 = op(e0,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_111])])).
fof(f12889,plain,(
  e0 = op(e3,e0) | (~spl4_74 | ~spl4_111)),
  inference(backward_demodulation,[],[f1076,f12840])).
fof(f12840,plain,(
  e0 = op(e3,op(e0,e2)) | ~spl4_74),
  inference(forward_demodulation,[],[f333,f920])).
fof(f920,plain,(
  e3 = op(e2,e0) | ~spl4_74),
  inference(avatar_component_clause,[],[f918])).
fof(f333,plain,(
  e0 = op(op(e2,e0),op(e0,e2))),
  inference(cnf_transformation,[],[f3])).
fof(f1076,plain,(
  e0 = op(e0,e2) | ~spl4_111),
  inference(avatar_component_clause,[],[f1074])).
fof(f12892,plain,(
  spl4_108 | ~spl4_74 | ~spl4_111),
  inference(avatar_split_clause,[],[f12888,f1074,f918,f1061])).
fof(f12888,plain,(
  e2 = op(e0,e3) | (~spl4_74 | ~spl4_111)),
  inference(backward_demodulation,[],[f1076,f12839])).
fof(f12839,plain,(
  e2 = op(op(e0,e2),e3) | ~spl4_74),
  inference(forward_demodulation,[],[f341,f920])).
fof(f341,plain,(
  e2 = op(op(e0,e2),op(e2,e0))),
  inference(cnf_transformation,[],[f3])).
fof(f12891,plain,(
  ~spl4_106 | ~spl4_111),
  inference(avatar_split_clause,[],[f12882,f1074,f1053])).
fof(f1053,plain,(
  spl4_106 <=> e0 = op(e0,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_106])])).
fof(f12882,plain,(
  e0 != op(e0,e3) | ~spl4_111),
  inference(backward_demodulation,[],[f1076,f413])).
fof(f413,plain,(
  op(e0,e2) != op(e0,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f12890,plain,(
  ~spl4_36 | ~spl4_111),
  inference(avatar_split_clause,[],[f12880,f1074,f759])).
fof(f12880,plain,(
  e0 != op(e3,e2) | ~spl4_111),
  inference(backward_demodulation,[],[f1076,f378])).
fof(f378,plain,(
  op(e0,e2) != op(e3,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f12877,plain,(
  ~spl4_102 | ~spl4_117),
  inference(avatar_split_clause,[],[f12872,f1099,f1036])).
fof(f1036,plain,(
  spl4_102 <=> e1 = op(e0,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_102])])).
fof(f12872,plain,(
  e1 != op(e0,e4) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f412])).
fof(f412,plain,(
  op(e0,e1) != op(e0,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f12876,plain,(
  ~spl4_17 | ~spl4_117),
  inference(avatar_split_clause,[],[f12869,f1099,f679])).
fof(f12869,plain,(
  e1 != op(e4,e1) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f369])).
fof(f369,plain,(
  op(e0,e1) != op(e4,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f12849,plain,(
  spl4_104 | ~spl4_125 | ~spl4_185),
  inference(avatar_split_clause,[],[f12848,f1529,f1132,f1044])).
fof(f1132,plain,(
  spl4_125 <=> op(e0,e0) = e4),
  introduced(avatar_definition,[new_symbols(naming,[spl4_125])])).
fof(f1529,plain,(
  spl4_185 <=> e3 = op(e0,op(e0,e0))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_185])])).
fof(f12848,plain,(
  e3 = op(e0,e4) | (~spl4_125 | ~spl4_185)),
  inference(forward_demodulation,[],[f1530,f1134])).
fof(f1134,plain,(
  op(e0,e0) = e4 | ~spl4_125),
  inference(avatar_component_clause,[],[f1132])).
fof(f1530,plain,(
  e3 = op(e0,op(e0,e0)) | ~spl4_185),
  inference(avatar_component_clause,[],[f1529])).
fof(f12824,plain,(
  ~spl4_24 | ~spl4_1 | spl4_182),
  inference(avatar_split_clause,[],[f12762,f1508,f612,f708])).
fof(f708,plain,(
  spl4_24 <=> e3 = op(e4,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_24])])).
fof(f612,plain,(
  spl4_1 <=> e0 = op(e4,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_1])])).
fof(f1508,plain,(
  spl4_182 <=> e3 = op(e4,op(e4,e4))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_182])])).
fof(f12762,plain,(
  e3 != op(e4,e0) | (~spl4_1 | spl4_182)),
  inference(backward_demodulation,[],[f614,f1510])).
fof(f1510,plain,(
  e3 != op(e4,op(e4,e4)) | spl4_182),
  inference(avatar_component_clause,[],[f1508])).
fof(f614,plain,(
  e0 = op(e4,e4) | ~spl4_1),
  inference(avatar_component_clause,[],[f612])).
fof(f12817,plain,(
  ~spl4_14 | ~spl4_90 | ~spl4_93 | spl4_149),
  inference(avatar_contradiction_clause,[],[f12816])).
fof(f12816,plain,(
  $false | (~spl4_14 | ~spl4_90 | ~spl4_93 | spl4_149)),
  inference(subsumption_resolution,[],[f12815,f668])).
fof(f668,plain,(
  e3 = op(e4,e2) | ~spl4_14),
  inference(avatar_component_clause,[],[f666])).
fof(f12815,plain,(
  e3 != op(e4,e2) | (~spl4_90 | ~spl4_93 | spl4_149)),
  inference(forward_demodulation,[],[f12814,f987])).
fof(f987,plain,(
  e4 = op(e1,e2) | ~spl4_90),
  inference(avatar_component_clause,[],[f985])).
fof(f985,plain,(
  spl4_90 <=> e4 = op(e1,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_90])])).
fof(f12814,plain,(
  e3 != op(op(e1,e2),e2) | (~spl4_93 | spl4_149)),
  inference(forward_demodulation,[],[f1315,f1000])).
fof(f1000,plain,(
  e2 = op(e1,e1) | ~spl4_93),
  inference(avatar_component_clause,[],[f998])).
fof(f1315,plain,(
  e3 != op(op(e1,op(e1,e1)),op(e1,e1)) | spl4_149),
  inference(avatar_component_clause,[],[f1313])).
fof(f1313,plain,(
  spl4_149 <=> e3 = op(op(e1,op(e1,e1)),op(e1,e1))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_149])])).
fof(f12798,plain,(
  ~spl4_62 | ~spl4_66 | ~spl4_119 | spl4_139),
  inference(avatar_contradiction_clause,[],[f12797])).
fof(f12797,plain,(
  $false | (~spl4_62 | ~spl4_66 | ~spl4_119 | spl4_139)),
  inference(subsumption_resolution,[],[f12796,f1109])).
fof(f1109,plain,(
  e3 = op(e0,e1) | ~spl4_119),
  inference(avatar_component_clause,[],[f1107])).
fof(f1107,plain,(
  spl4_119 <=> e3 = op(e0,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_119])])).
fof(f12796,plain,(
  e3 != op(e0,e1) | (~spl4_62 | ~spl4_66 | spl4_139)),
  inference(forward_demodulation,[],[f12795,f887])).
fof(f887,plain,(
  e0 = op(e2,e1) | ~spl4_66),
  inference(avatar_component_clause,[],[f885])).
fof(f885,plain,(
  spl4_66 <=> e0 = op(e2,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_66])])).
fof(f12795,plain,(
  e3 != op(op(e2,e1),e1) | (~spl4_62 | spl4_139)),
  inference(forward_demodulation,[],[f1269,f870])).
fof(f870,plain,(
  e1 = op(e2,e2) | ~spl4_62),
  inference(avatar_component_clause,[],[f868])).
fof(f1269,plain,(
  e3 != op(op(e2,op(e2,e2)),op(e2,e2)) | spl4_139),
  inference(avatar_component_clause,[],[f1267])).
fof(f1267,plain,(
  spl4_139 <=> e3 = op(op(e2,op(e2,e2)),op(e2,e2))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_139])])).
fof(f12746,plain,(
  spl4_10 | ~spl4_14 | ~spl4_55),
  inference(avatar_split_clause,[],[f12744,f838,f666,f649])).
fof(f12744,plain,(
  e4 = op(e4,e3) | (~spl4_14 | ~spl4_55)),
  inference(backward_demodulation,[],[f668,f12678])).
fof(f12678,plain,(
  e4 = op(e4,op(e4,e2)) | ~spl4_55),
  inference(backward_demodulation,[],[f840,f353])).
fof(f353,plain,(
  e4 = op(op(e2,e4),op(e4,e2))),
  inference(cnf_transformation,[],[f3])).
fof(f840,plain,(
  e4 = op(e2,e4) | ~spl4_55),
  inference(avatar_component_clause,[],[f838])).
fof(f12745,plain,(
  ~spl4_4 | ~spl4_14),
  inference(avatar_split_clause,[],[f12741,f666,f624])).
fof(f624,plain,(
  spl4_4 <=> e3 = op(e4,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_4])])).
fof(f12741,plain,(
  e3 != op(e4,e4) | ~spl4_14),
  inference(backward_demodulation,[],[f668,f454])).
fof(f454,plain,(
  op(e4,e2) != op(e4,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f12731,plain,(
  ~spl4_13 | ~spl4_23),
  inference(avatar_split_clause,[],[f12724,f704,f662])).
fof(f662,plain,(
  spl4_13 <=> e2 = op(e4,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_13])])).
fof(f12724,plain,(
  e2 != op(e4,e2) | ~spl4_23),
  inference(backward_demodulation,[],[f706,f447])).
fof(f447,plain,(
  op(e4,e0) != op(e4,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f706,plain,(
  e2 = op(e4,e0) | ~spl4_23),
  inference(avatar_component_clause,[],[f704])).
fof(f12698,plain,(
  spl4_17 | ~spl4_45 | ~spl4_82),
  inference(avatar_split_clause,[],[f12694,f952,f796,f679])).
fof(f12694,plain,(
  e1 = op(e4,e1) | (~spl4_45 | ~spl4_82)),
  inference(backward_demodulation,[],[f798,f12615])).
fof(f12615,plain,(
  e1 = op(op(e3,e1),e1) | ~spl4_82),
  inference(backward_demodulation,[],[f954,f339])).
fof(f954,plain,(
  e1 = op(e1,e3) | ~spl4_82),
  inference(avatar_component_clause,[],[f952])).
fof(f12697,plain,(
  ~spl4_35 | ~spl4_45),
  inference(avatar_split_clause,[],[f12693,f796,f754])).
fof(f754,plain,(
  spl4_35 <=> e4 = op(e3,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_35])])).
fof(f12693,plain,(
  e4 != op(e3,e3) | ~spl4_45),
  inference(backward_demodulation,[],[f798,f441])).
fof(f441,plain,(
  op(e3,e1) != op(e3,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f12696,plain,(
  ~spl4_20 | ~spl4_45),
  inference(avatar_split_clause,[],[f12691,f796,f691])).
fof(f691,plain,(
  spl4_20 <=> e4 = op(e4,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_20])])).
fof(f12691,plain,(
  e4 != op(e4,e1) | ~spl4_45),
  inference(backward_demodulation,[],[f798,f375])).
fof(f375,plain,(
  op(e3,e1) != op(e4,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f12690,plain,(
  ~spl4_37 | ~spl4_47),
  inference(avatar_split_clause,[],[f12682,f805,f763])).
fof(f763,plain,(
  spl4_37 <=> e1 = op(e3,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_37])])).
fof(f805,plain,(
  spl4_47 <=> e1 = op(e3,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_47])])).
fof(f12682,plain,(
  e1 != op(e3,e2) | ~spl4_47),
  inference(backward_demodulation,[],[f807,f437])).
fof(f807,plain,(
  e1 = op(e3,e0) | ~spl4_47),
  inference(avatar_component_clause,[],[f805])).
fof(f12689,plain,(
  ~spl4_22 | ~spl4_47),
  inference(avatar_split_clause,[],[f12680,f805,f700])).
fof(f12680,plain,(
  e1 != op(e4,e0) | ~spl4_47),
  inference(backward_demodulation,[],[f807,f365])).
fof(f365,plain,(
  op(e3,e0) != op(e4,e0)),
  inference(cnf_transformation,[],[f4])).
fof(f12668,plain,(
  ~spl4_52 | ~spl4_62),
  inference(avatar_split_clause,[],[f12659,f868,f826])).
fof(f826,plain,(
  spl4_52 <=> e1 = op(e2,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_52])])).
fof(f12659,plain,(
  e1 != op(e2,e4) | ~spl4_62),
  inference(backward_demodulation,[],[f870,f434])).
fof(f434,plain,(
  op(e2,e2) != op(e2,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f12667,plain,(
  ~spl4_37 | ~spl4_62),
  inference(avatar_split_clause,[],[f12656,f868,f763])).
fof(f12656,plain,(
  e1 != op(e3,e2) | ~spl4_62),
  inference(backward_demodulation,[],[f870,f383])).
fof(f383,plain,(
  op(e2,e2) != op(e3,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f12654,plain,(
  spl4_23 | ~spl4_66 | ~spl4_90),
  inference(avatar_split_clause,[],[f12652,f985,f885,f704])).
fof(f12652,plain,(
  e2 = op(e4,e0) | (~spl4_66 | ~spl4_90)),
  inference(backward_demodulation,[],[f887,f12359])).
fof(f12359,plain,(
  e2 = op(e4,op(e2,e1)) | ~spl4_90),
  inference(backward_demodulation,[],[f987,f342])).
fof(f342,plain,(
  e2 = op(op(e1,e2),op(e2,e1))),
  inference(cnf_transformation,[],[f3])).
fof(f12653,plain,(
  ~spl4_61 | ~spl4_66),
  inference(avatar_split_clause,[],[f12644,f885,f864])).
fof(f864,plain,(
  spl4_61 <=> e0 = op(e2,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_61])])).
fof(f12644,plain,(
  e0 != op(e2,e2) | ~spl4_66),
  inference(backward_demodulation,[],[f887,f430])).
fof(f430,plain,(
  op(e2,e1) != op(e2,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f12640,plain,(
  spl4_36 | ~spl4_74 | ~spl4_113),
  inference(avatar_split_clause,[],[f12637,f1082,f918,f759])).
fof(f1082,plain,(
  spl4_113 <=> e2 = op(e0,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_113])])).
fof(f12637,plain,(
  e0 = op(e3,e2) | (~spl4_74 | ~spl4_113)),
  inference(backward_demodulation,[],[f920,f12551])).
fof(f12551,plain,(
  e0 = op(op(e2,e0),e2) | ~spl4_113),
  inference(backward_demodulation,[],[f1084,f333])).
fof(f1084,plain,(
  e2 = op(e0,e2) | ~spl4_113),
  inference(avatar_component_clause,[],[f1082])).
fof(f12639,plain,(
  ~spl4_59 | ~spl4_74),
  inference(avatar_split_clause,[],[f12634,f918,f855])).
fof(f855,plain,(
  spl4_59 <=> e3 = op(e2,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_59])])).
fof(f12634,plain,(
  e3 != op(e2,e3) | ~spl4_74),
  inference(backward_demodulation,[],[f920,f428])).
fof(f428,plain,(
  op(e2,e0) != op(e2,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f12629,plain,(
  ~spl4_4 | ~spl4_79),
  inference(avatar_split_clause,[],[f12627,f939,f624])).
fof(f12627,plain,(
  e3 != op(e4,e4) | ~spl4_79),
  inference(backward_demodulation,[],[f941,f402])).
fof(f402,plain,(
  op(e1,e4) != op(e4,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f941,plain,(
  e3 = op(e1,e4) | ~spl4_79),
  inference(avatar_component_clause,[],[f939])).
fof(f12613,plain,(
  ~spl4_83 | ~spl4_93),
  inference(avatar_split_clause,[],[f12601,f998,f956])).
fof(f956,plain,(
  spl4_83 <=> e2 = op(e1,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_83])])).
fof(f12601,plain,(
  e2 != op(e1,e3) | ~spl4_93),
  inference(backward_demodulation,[],[f1000,f421])).
fof(f421,plain,(
  op(e1,e1) != op(e1,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f12612,plain,(
  ~spl4_68 | ~spl4_93),
  inference(avatar_split_clause,[],[f12598,f998,f893])).
fof(f893,plain,(
  spl4_68 <=> e2 = op(e2,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_68])])).
fof(f12598,plain,(
  e2 != op(e2,e1) | ~spl4_93),
  inference(backward_demodulation,[],[f1000,f370])).
fof(f370,plain,(
  op(e1,e1) != op(e2,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f12611,plain,(
  spl4_62 | ~spl4_93),
  inference(avatar_split_clause,[],[f12597,f998,f868])).
fof(f12597,plain,(
  e1 = op(e2,e2) | ~spl4_93),
  inference(backward_demodulation,[],[f1000,f337])).
fof(f12596,plain,(
  spl4_47 | ~spl4_96 | ~spl4_119),
  inference(avatar_split_clause,[],[f12593,f1107,f1011,f805])).
fof(f1011,plain,(
  spl4_96 <=> e0 = op(e1,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_96])])).
fof(f12593,plain,(
  e1 = op(e3,e0) | (~spl4_96 | ~spl4_119)),
  inference(backward_demodulation,[],[f1013,f12540])).
fof(f12540,plain,(
  e1 = op(e3,op(e1,e0)) | ~spl4_119),
  inference(backward_demodulation,[],[f1109,f336])).
fof(f1013,plain,(
  e0 = op(e1,e0) | ~spl4_96),
  inference(avatar_component_clause,[],[f1011])).
fof(f12595,plain,(
  ~spl4_76 | ~spl4_96),
  inference(avatar_split_clause,[],[f12587,f1011,f927])).
fof(f927,plain,(
  spl4_76 <=> e0 = op(e1,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_76])])).
fof(f12587,plain,(
  e0 != op(e1,e4) | ~spl4_96),
  inference(backward_demodulation,[],[f1013,f419])).
fof(f419,plain,(
  op(e1,e0) != op(e1,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f12594,plain,(
  ~spl4_46 | ~spl4_96),
  inference(avatar_split_clause,[],[f12583,f1011,f801])).
fof(f12583,plain,(
  e0 != op(e3,e0) | ~spl4_96),
  inference(backward_demodulation,[],[f1013,f361])).
fof(f361,plain,(
  op(e1,e0) != op(e3,e0)),
  inference(cnf_transformation,[],[f4])).
fof(f12581,plain,(
  ~spl4_52 | ~spl4_102),
  inference(avatar_split_clause,[],[f12576,f1036,f826])).
fof(f12576,plain,(
  e1 != op(e2,e4) | ~spl4_102),
  inference(backward_demodulation,[],[f1038,f397])).
fof(f397,plain,(
  op(e0,e4) != op(e2,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f1038,plain,(
  e1 = op(e0,e4) | ~spl4_102),
  inference(avatar_component_clause,[],[f1036])).
fof(f12572,plain,(
  ~spl4_6 | ~spl4_106),
  inference(avatar_split_clause,[],[f12566,f1053,f633])).
fof(f633,plain,(
  spl4_6 <=> e0 = op(e4,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_6])])).
fof(f12566,plain,(
  e0 != op(e4,e3) | ~spl4_106),
  inference(backward_demodulation,[],[f1055,f389])).
fof(f389,plain,(
  op(e0,e3) != op(e4,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f1055,plain,(
  e0 = op(e0,e3) | ~spl4_106),
  inference(avatar_component_clause,[],[f1053])).
fof(f12560,plain,(
  ~spl4_13 | ~spl4_113),
  inference(avatar_split_clause,[],[f12555,f1082,f662])).
fof(f12555,plain,(
  e2 != op(e4,e2) | ~spl4_113),
  inference(backward_demodulation,[],[f1084,f379])).
fof(f379,plain,(
  op(e0,e2) != op(e4,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f12550,plain,(
  ~spl4_114 | ~spl4_119),
  inference(avatar_split_clause,[],[f12545,f1107,f1086])).
fof(f1086,plain,(
  spl4_114 <=> e3 = op(e0,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_114])])).
fof(f12545,plain,(
  e3 != op(e0,e2) | ~spl4_119),
  inference(backward_demodulation,[],[f1109,f410])).
fof(f410,plain,(
  op(e0,e1) != op(e0,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f12549,plain,(
  ~spl4_44 | ~spl4_119),
  inference(avatar_split_clause,[],[f12543,f1107,f792])).
fof(f792,plain,(
  spl4_44 <=> e3 = op(e3,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_44])])).
fof(f12543,plain,(
  e3 != op(e3,e1) | ~spl4_119),
  inference(backward_demodulation,[],[f1109,f368])).
fof(f368,plain,(
  op(e0,e1) != op(e3,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f12536,plain,(
  ~spl4_105 | ~spl4_125),
  inference(avatar_split_clause,[],[f12530,f1132,f1048])).
fof(f1048,plain,(
  spl4_105 <=> e4 = op(e0,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_105])])).
fof(f12530,plain,(
  e4 != op(e0,e4) | ~spl4_125),
  inference(backward_demodulation,[],[f1134,f409])).
fof(f409,plain,(
  op(e0,e0) != op(e0,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f12535,plain,(
  ~spl4_75 | ~spl4_125),
  inference(avatar_split_clause,[],[f12524,f1132,f922])).
fof(f922,plain,(
  spl4_75 <=> e4 = op(e2,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_75])])).
fof(f12524,plain,(
  e4 != op(e2,e0) | ~spl4_125),
  inference(backward_demodulation,[],[f1134,f357])).
fof(f357,plain,(
  op(e0,e0) != op(e2,e0)),
  inference(cnf_transformation,[],[f4])).
fof(f12534,plain,(
  spl4_1 | ~spl4_125),
  inference(avatar_split_clause,[],[f12522,f1132,f612])).
fof(f12522,plain,(
  e0 = op(e4,e4) | ~spl4_125),
  inference(backward_demodulation,[],[f1134,f331])).
fof(f331,plain,(
  e0 = op(op(e0,e0),op(e0,e0))),
  inference(cnf_transformation,[],[f3])).
fof(f12511,plain,(
  ~spl4_4 | spl4_153 | ~spl4_188),
  inference(avatar_split_clause,[],[f12510,f1559,f1331,f624])).
fof(f1331,plain,(
  spl4_153 <=> e3 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_153])])).
fof(f1559,plain,(
  spl4_188 <=> e4 = op(e1,op(e1,e1))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_188])])).
fof(f12510,plain,(
  e3 != op(e4,e4) | (spl4_153 | ~spl4_188)),
  inference(forward_demodulation,[],[f1333,f1560])).
fof(f1560,plain,(
  e4 = op(e1,op(e1,e1)) | ~spl4_188),
  inference(avatar_component_clause,[],[f1559])).
fof(f1333,plain,(
  e3 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | spl4_153),
  inference(avatar_component_clause,[],[f1331])).
fof(f12498,plain,(
  ~spl4_1 | spl4_180 | ~spl4_188),
  inference(avatar_split_clause,[],[f12497,f1559,f1492,f612])).
fof(f1492,plain,(
  spl4_180 <=> e0 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_180])])).
fof(f12497,plain,(
  e0 != op(e4,e4) | (spl4_180 | ~spl4_188)),
  inference(forward_demodulation,[],[f1494,f1560])).
fof(f1494,plain,(
  e0 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | spl4_180),
  inference(avatar_component_clause,[],[f1492])).
fof(f12471,plain,(
  ~spl4_15 | ~spl4_20),
  inference(avatar_split_clause,[],[f12468,f691,f670])).
fof(f12468,plain,(
  e4 != op(e4,e2) | ~spl4_20),
  inference(backward_demodulation,[],[f693,f450])).
fof(f450,plain,(
  op(e4,e1) != op(e4,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f693,plain,(
  e4 = op(e4,e1) | ~spl4_20),
  inference(avatar_component_clause,[],[f691])).
fof(f12467,plain,(
  spl4_20 | ~spl4_22 | ~spl4_105),
  inference(avatar_split_clause,[],[f12465,f1048,f700,f691])).
fof(f12465,plain,(
  e4 = op(e4,e1) | (~spl4_22 | ~spl4_105)),
  inference(backward_demodulation,[],[f702,f11984])).
fof(f11984,plain,(
  e4 = op(e4,op(e4,e0)) | ~spl4_105),
  inference(backward_demodulation,[],[f1050,f351])).
fof(f351,plain,(
  e4 = op(op(e0,e4),op(e4,e0))),
  inference(cnf_transformation,[],[f3])).
fof(f1050,plain,(
  e4 = op(e0,e4) | ~spl4_105),
  inference(avatar_component_clause,[],[f1048])).
fof(f12448,plain,(
  ~spl4_27 | ~spl4_37),
  inference(avatar_split_clause,[],[f12442,f763,f721])).
fof(f721,plain,(
  spl4_27 <=> e1 = op(e3,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_27])])).
fof(f12442,plain,(
  e1 != op(e3,e4) | ~spl4_37),
  inference(backward_demodulation,[],[f765,f444])).
fof(f444,plain,(
  op(e3,e2) != op(e3,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f765,plain,(
  e1 = op(e3,e2) | ~spl4_37),
  inference(avatar_component_clause,[],[f763])).
fof(f12440,plain,(
  spl4_37 | ~spl4_44 | ~spl4_83),
  inference(avatar_split_clause,[],[f12437,f956,f792,f763])).
fof(f12437,plain,(
  e1 = op(e3,e2) | (~spl4_44 | ~spl4_83)),
  inference(backward_demodulation,[],[f794,f12366])).
fof(f12366,plain,(
  e1 = op(op(e3,e1),e2) | ~spl4_83),
  inference(backward_demodulation,[],[f958,f339])).
fof(f958,plain,(
  e2 = op(e1,e3) | ~spl4_83),
  inference(avatar_component_clause,[],[f956])).
fof(f794,plain,(
  e3 = op(e3,e1) | ~spl4_44),
  inference(avatar_component_clause,[],[f792])).
fof(f12439,plain,(
  ~spl4_39 | ~spl4_44),
  inference(avatar_split_clause,[],[f12434,f792,f771])).
fof(f12434,plain,(
  e3 != op(e3,e2) | ~spl4_44),
  inference(backward_demodulation,[],[f794,f440])).
fof(f440,plain,(
  op(e3,e1) != op(e3,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f12432,plain,(
  ~spl4_27 | ~spl4_52),
  inference(avatar_split_clause,[],[f12428,f826,f721])).
fof(f12428,plain,(
  e1 != op(e3,e4) | ~spl4_52),
  inference(backward_demodulation,[],[f828,f403])).
fof(f403,plain,(
  op(e2,e4) != op(e3,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f828,plain,(
  e1 = op(e2,e4) | ~spl4_52),
  inference(avatar_component_clause,[],[f826])).
fof(f12421,plain,(
  ~spl4_22 | ~spl4_61 | spl4_154 | ~spl4_187),
  inference(avatar_split_clause,[],[f12419,f1553,f1337,f864,f700])).
fof(f1337,plain,(
  spl4_154 <=> e1 = op(op(e2,op(e2,e2)),op(e2,e2))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_154])])).
fof(f1553,plain,(
  spl4_187 <=> e4 = op(e2,op(e2,e2))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_187])])).
fof(f12419,plain,(
  e1 != op(e4,e0) | (~spl4_61 | spl4_154 | ~spl4_187)),
  inference(backward_demodulation,[],[f866,f12267])).
fof(f12267,plain,(
  e1 != op(e4,op(e2,e2)) | (spl4_154 | ~spl4_187)),
  inference(forward_demodulation,[],[f1339,f1554])).
fof(f1554,plain,(
  e4 = op(e2,op(e2,e2)) | ~spl4_187),
  inference(avatar_component_clause,[],[f1553])).
fof(f1339,plain,(
  e1 != op(op(e2,op(e2,e2)),op(e2,e2)) | spl4_154),
  inference(avatar_component_clause,[],[f1337])).
fof(f866,plain,(
  e0 = op(e2,e2) | ~spl4_61),
  inference(avatar_component_clause,[],[f864])).
fof(f12420,plain,(
  ~spl4_51 | ~spl4_61),
  inference(avatar_split_clause,[],[f12408,f864,f822])).
fof(f822,plain,(
  spl4_51 <=> e0 = op(e2,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_51])])).
fof(f12408,plain,(
  e0 != op(e2,e4) | ~spl4_61),
  inference(backward_demodulation,[],[f866,f434])).
fof(f12403,plain,(
  spl4_13 | ~spl4_68 | ~spl4_90),
  inference(avatar_split_clause,[],[f12399,f985,f893,f662])).
fof(f12399,plain,(
  e2 = op(e4,e2) | (~spl4_68 | ~spl4_90)),
  inference(backward_demodulation,[],[f895,f12359])).
fof(f895,plain,(
  e2 = op(e2,e1) | ~spl4_68),
  inference(avatar_component_clause,[],[f893])).
fof(f12402,plain,(
  spl4_52 | ~spl4_68 | ~spl4_90),
  inference(avatar_split_clause,[],[f12398,f985,f893,f826])).
fof(f12398,plain,(
  e1 = op(e2,e4) | (~spl4_68 | ~spl4_90)),
  inference(backward_demodulation,[],[f895,f12358])).
fof(f12358,plain,(
  e1 = op(op(e2,e1),e4) | ~spl4_90),
  inference(backward_demodulation,[],[f987,f338])).
fof(f338,plain,(
  e1 = op(op(e2,e1),op(e1,e2))),
  inference(cnf_transformation,[],[f3])).
fof(f12401,plain,(
  ~spl4_63 | ~spl4_68),
  inference(avatar_split_clause,[],[f12393,f893,f872])).
fof(f12393,plain,(
  e2 != op(e2,e2) | ~spl4_68),
  inference(backward_demodulation,[],[f895,f430])).
fof(f12400,plain,(
  ~spl4_43 | ~spl4_68),
  inference(avatar_split_clause,[],[f12391,f893,f788])).
fof(f788,plain,(
  spl4_43 <=> e2 = op(e3,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_43])])).
fof(f12391,plain,(
  e2 != op(e3,e1) | ~spl4_68),
  inference(backward_demodulation,[],[f895,f373])).
fof(f373,plain,(
  op(e2,e1) != op(e3,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f12390,plain,(
  spl4_28 | ~spl4_75 | ~spl4_114),
  inference(avatar_split_clause,[],[f12388,f1086,f922,f725])).
fof(f725,plain,(
  spl4_28 <=> e2 = op(e3,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_28])])).
fof(f12388,plain,(
  e2 = op(e3,e4) | (~spl4_75 | ~spl4_114)),
  inference(backward_demodulation,[],[f924,f12302])).
fof(f12302,plain,(
  e2 = op(e3,op(e2,e0)) | ~spl4_114),
  inference(backward_demodulation,[],[f1088,f341])).
fof(f1088,plain,(
  e3 = op(e0,e2) | ~spl4_114),
  inference(avatar_component_clause,[],[f1086])).
fof(f924,plain,(
  e4 = op(e2,e0) | ~spl4_75),
  inference(avatar_component_clause,[],[f922])).
fof(f12389,plain,(
  ~spl4_70 | ~spl4_75),
  inference(avatar_split_clause,[],[f12383,f922,f901])).
fof(f901,plain,(
  spl4_70 <=> e4 = op(e2,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_70])])).
fof(f12383,plain,(
  e4 != op(e2,e1) | ~spl4_75),
  inference(backward_demodulation,[],[f924,f426])).
fof(f426,plain,(
  op(e2,e0) != op(e2,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f12381,plain,(
  ~spl4_51 | ~spl4_76),
  inference(avatar_split_clause,[],[f12375,f927,f822])).
fof(f12375,plain,(
  e0 != op(e2,e4) | ~spl4_76),
  inference(backward_demodulation,[],[f929,f400])).
fof(f929,plain,(
  e0 = op(e1,e4) | ~spl4_76),
  inference(avatar_component_clause,[],[f927])).
fof(f12372,plain,(
  ~spl4_78 | ~spl4_83),
  inference(avatar_split_clause,[],[f12369,f956,f935])).
fof(f935,plain,(
  spl4_78 <=> e2 = op(e1,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_78])])).
fof(f12369,plain,(
  e2 != op(e1,e4) | ~spl4_83),
  inference(backward_demodulation,[],[f958,f425])).
fof(f425,plain,(
  op(e1,e3) != op(e1,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f12365,plain,(
  ~spl4_15 | ~spl4_90),
  inference(avatar_split_clause,[],[f12362,f985,f670])).
fof(f12362,plain,(
  e4 != op(e4,e2) | ~spl4_90),
  inference(backward_demodulation,[],[f987,f382])).
fof(f382,plain,(
  op(e1,e2) != op(e4,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f12347,plain,(
  ~spl4_87 | ~spl4_92),
  inference(avatar_split_clause,[],[f12334,f994,f973])).
fof(f994,plain,(
  spl4_92 <=> e1 = op(e1,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_92])])).
fof(f12334,plain,(
  e1 != op(e1,e2) | ~spl4_92),
  inference(backward_demodulation,[],[f996,f420])).
fof(f420,plain,(
  op(e1,e1) != op(e1,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f996,plain,(
  e1 = op(e1,e1) | ~spl4_92),
  inference(avatar_component_clause,[],[f994])).
fof(f12346,plain,(
  ~spl4_17 | ~spl4_92),
  inference(avatar_split_clause,[],[f12333,f994,f679])).
fof(f12333,plain,(
  e1 != op(e4,e1) | ~spl4_92),
  inference(backward_demodulation,[],[f996,f372])).
fof(f372,plain,(
  op(e1,e1) != op(e4,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f12329,plain,(
  ~spl4_84 | ~spl4_99),
  inference(avatar_split_clause,[],[f12323,f1023,f960])).
fof(f960,plain,(
  spl4_84 <=> e3 = op(e1,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_84])])).
fof(f1023,plain,(
  spl4_99 <=> e3 = op(e1,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_99])])).
fof(f12323,plain,(
  e3 != op(e1,e3) | ~spl4_99),
  inference(backward_demodulation,[],[f1025,f418])).
fof(f418,plain,(
  op(e1,e0) != op(e1,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f1025,plain,(
  e3 = op(e1,e0) | ~spl4_99),
  inference(avatar_component_clause,[],[f1023])).
fof(f12328,plain,(
  ~spl4_74 | ~spl4_99),
  inference(avatar_split_clause,[],[f12319,f1023,f918])).
fof(f12319,plain,(
  e3 != op(e2,e0) | ~spl4_99),
  inference(backward_demodulation,[],[f1025,f360])).
fof(f360,plain,(
  op(e1,e0) != op(e2,e0)),
  inference(cnf_transformation,[],[f4])).
fof(f12318,plain,(
  spl4_99 | ~spl4_46 | ~spl4_107),
  inference(avatar_split_clause,[],[f12315,f1057,f801,f1023])).
fof(f1057,plain,(
  spl4_107 <=> e1 = op(e0,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_107])])).
fof(f12315,plain,(
  e3 = op(e1,e0) | (~spl4_46 | ~spl4_107)),
  inference(backward_demodulation,[],[f1059,f12233])).
fof(f12233,plain,(
  e3 = op(op(e0,e3),e0) | ~spl4_46),
  inference(forward_demodulation,[],[f346,f803])).
fof(f346,plain,(
  e3 = op(op(e0,e3),op(e3,e0))),
  inference(cnf_transformation,[],[f3])).
fof(f1059,plain,(
  e1 = op(e0,e3) | ~spl4_107),
  inference(avatar_component_clause,[],[f1057])).
fof(f12317,plain,(
  ~spl4_57 | ~spl4_107),
  inference(avatar_split_clause,[],[f12311,f1057,f847])).
fof(f847,plain,(
  spl4_57 <=> e1 = op(e2,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_57])])).
fof(f12311,plain,(
  e1 != op(e2,e3) | ~spl4_107),
  inference(backward_demodulation,[],[f1059,f387])).
fof(f12309,plain,(
  ~spl4_39 | ~spl4_114),
  inference(avatar_split_clause,[],[f12305,f1086,f771])).
fof(f12305,plain,(
  e3 != op(e3,e2) | ~spl4_114),
  inference(backward_demodulation,[],[f1088,f378])).
fof(f12300,plain,(
  ~spl4_111 | ~spl4_116),
  inference(avatar_split_clause,[],[f12293,f1095,f1074])).
fof(f1095,plain,(
  spl4_116 <=> e0 = op(e0,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_116])])).
fof(f12293,plain,(
  e0 != op(e0,e2) | ~spl4_116),
  inference(backward_demodulation,[],[f1097,f410])).
fof(f1097,plain,(
  e0 = op(e0,e1) | ~spl4_116),
  inference(avatar_component_clause,[],[f1095])).
fof(f12299,plain,(
  ~spl4_91 | ~spl4_116),
  inference(avatar_split_clause,[],[f12289,f1095,f990])).
fof(f990,plain,(
  spl4_91 <=> e0 = op(e1,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_91])])).
fof(f12289,plain,(
  e0 != op(e1,e1) | ~spl4_116),
  inference(backward_demodulation,[],[f1097,f366])).
fof(f366,plain,(
  op(e0,e1) != op(e1,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f12285,plain,(
  ~spl4_108 | ~spl4_123),
  inference(avatar_split_clause,[],[f12276,f1124,f1061])).
fof(f1124,plain,(
  spl4_123 <=> op(e0,e0) = e2),
  introduced(avatar_definition,[new_symbols(naming,[spl4_123])])).
fof(f12276,plain,(
  e2 != op(e0,e3) | ~spl4_123),
  inference(backward_demodulation,[],[f1126,f408])).
fof(f408,plain,(
  op(e0,e0) != op(e0,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f1126,plain,(
  op(e0,e0) = e2 | ~spl4_123),
  inference(avatar_component_clause,[],[f1124])).
fof(f12284,plain,(
  ~spl4_23 | ~spl4_123),
  inference(avatar_split_clause,[],[f12273,f1124,f704])).
fof(f12273,plain,(
  e2 != op(e4,e0) | ~spl4_123),
  inference(backward_demodulation,[],[f1126,f359])).
fof(f359,plain,(
  op(e0,e0) != op(e4,e0)),
  inference(cnf_transformation,[],[f4])).
fof(f12283,plain,(
  spl4_61 | ~spl4_123),
  inference(avatar_split_clause,[],[f12270,f1124,f864])).
fof(f12270,plain,(
  e0 = op(e2,e2) | ~spl4_123),
  inference(backward_demodulation,[],[f1126,f331])).
fof(f12177,plain,(
  ~spl4_5 | ~spl4_15),
  inference(avatar_split_clause,[],[f12174,f670,f628])).
fof(f628,plain,(
  spl4_5 <=> e4 = op(e4,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_5])])).
fof(f12174,plain,(
  e4 != op(e4,e4) | ~spl4_15),
  inference(backward_demodulation,[],[f672,f454])).
fof(f12164,plain,(
  spl4_15 | ~spl4_23 | ~spl4_105),
  inference(avatar_split_clause,[],[f12162,f1048,f704,f670])).
fof(f12162,plain,(
  e4 = op(e4,e2) | (~spl4_23 | ~spl4_105)),
  inference(backward_demodulation,[],[f706,f11984])).
fof(f12163,plain,(
  ~spl4_8 | ~spl4_23),
  inference(avatar_split_clause,[],[f12157,f704,f641])).
fof(f641,plain,(
  spl4_8 <=> e2 = op(e4,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_8])])).
fof(f12157,plain,(
  e2 != op(e4,e3) | ~spl4_23),
  inference(backward_demodulation,[],[f706,f448])).
fof(f448,plain,(
  op(e4,e0) != op(e4,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f12146,plain,(
  spl4_27 | ~spl4_35 | ~spl4_159),
  inference(avatar_split_clause,[],[f12142,f1369,f754,f721])).
fof(f1369,plain,(
  spl4_159 <=> e1 = op(e3,op(e3,e3))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_159])])).
fof(f12142,plain,(
  e1 = op(e3,e4) | (~spl4_35 | ~spl4_159)),
  inference(backward_demodulation,[],[f756,f1370])).
fof(f1370,plain,(
  e1 = op(e3,op(e3,e3)) | ~spl4_159),
  inference(avatar_component_clause,[],[f1369])).
fof(f756,plain,(
  e4 = op(e3,e3) | ~spl4_35),
  inference(avatar_component_clause,[],[f754])).
fof(f12145,plain,(
  spl4_4 | ~spl4_35),
  inference(avatar_split_clause,[],[f12138,f754,f624])).
fof(f12138,plain,(
  e3 = op(e4,e4) | ~spl4_35),
  inference(backward_demodulation,[],[f756,f349])).
fof(f349,plain,(
  e3 = op(op(e3,e3),op(e3,e3))),
  inference(cnf_transformation,[],[f3])).
fof(f12137,plain,(
  ~spl4_29 | ~spl4_39),
  inference(avatar_split_clause,[],[f12133,f771,f729])).
fof(f729,plain,(
  spl4_29 <=> e3 = op(e3,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_29])])).
fof(f12133,plain,(
  e3 != op(e3,e4) | ~spl4_39),
  inference(backward_demodulation,[],[f773,f444])).
fof(f12130,plain,(
  spl4_39 | ~spl4_43 | ~spl4_84),
  inference(avatar_split_clause,[],[f12129,f960,f788,f771])).
fof(f12129,plain,(
  e3 = op(e3,e2) | (~spl4_43 | ~spl4_84)),
  inference(backward_demodulation,[],[f790,f12035])).
fof(f12035,plain,(
  e3 = op(e3,op(e3,e1)) | ~spl4_84),
  inference(backward_demodulation,[],[f962,f347])).
fof(f347,plain,(
  e3 = op(op(e1,e3),op(e3,e1))),
  inference(cnf_transformation,[],[f3])).
fof(f962,plain,(
  e3 = op(e1,e3) | ~spl4_84),
  inference(avatar_component_clause,[],[f960])).
fof(f790,plain,(
  e2 = op(e3,e1) | ~spl4_43),
  inference(avatar_component_clause,[],[f788])).
fof(f12121,plain,(
  ~spl4_31 | ~spl4_46),
  inference(avatar_split_clause,[],[f12113,f801,f738])).
fof(f738,plain,(
  spl4_31 <=> e0 = op(e3,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_31])])).
fof(f12113,plain,(
  e0 != op(e3,e3) | ~spl4_46),
  inference(backward_demodulation,[],[f803,f438])).
fof(f438,plain,(
  op(e3,e0) != op(e3,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f12084,plain,(
  ~spl4_53 | ~spl4_63),
  inference(avatar_split_clause,[],[f12073,f872,f830])).
fof(f830,plain,(
  spl4_53 <=> e2 = op(e2,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_53])])).
fof(f12073,plain,(
  e2 != op(e2,e4) | ~spl4_63),
  inference(backward_demodulation,[],[f874,f434])).
fof(f12068,plain,(
  spl4_17 | ~spl4_70 | ~spl4_87),
  inference(avatar_split_clause,[],[f12066,f973,f901,f679])).
fof(f12066,plain,(
  e1 = op(e4,e1) | (~spl4_70 | ~spl4_87)),
  inference(backward_demodulation,[],[f903,f12022])).
fof(f12022,plain,(
  e1 = op(op(e2,e1),e1) | ~spl4_87),
  inference(backward_demodulation,[],[f975,f338])).
fof(f903,plain,(
  e4 = op(e2,e1) | ~spl4_70),
  inference(avatar_component_clause,[],[f901])).
fof(f12049,plain,(
  ~spl4_53 | ~spl4_78),
  inference(avatar_split_clause,[],[f12044,f935,f830])).
fof(f12044,plain,(
  e2 != op(e2,e4) | ~spl4_78),
  inference(backward_demodulation,[],[f937,f400])).
fof(f937,plain,(
  e2 = op(e1,e4) | ~spl4_78),
  inference(avatar_component_clause,[],[f935])).
fof(f12041,plain,(
  ~spl4_59 | ~spl4_84),
  inference(avatar_split_clause,[],[f12036,f960,f855])).
fof(f12036,plain,(
  e3 != op(e2,e3) | ~spl4_84),
  inference(backward_demodulation,[],[f962,f390])).
fof(f390,plain,(
  op(e1,e3) != op(e2,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f12033,plain,(
  ~spl4_77 | ~spl4_87),
  inference(avatar_split_clause,[],[f12028,f973,f931])).
fof(f931,plain,(
  spl4_77 <=> e1 = op(e1,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_77])])).
fof(f12028,plain,(
  e1 != op(e1,e4) | ~spl4_87),
  inference(backward_demodulation,[],[f975,f424])).
fof(f424,plain,(
  op(e1,e2) != op(e1,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f12020,plain,(
  ~spl4_66 | ~spl4_91),
  inference(avatar_split_clause,[],[f12004,f990,f885])).
fof(f12004,plain,(
  e0 != op(e2,e1) | ~spl4_91),
  inference(backward_demodulation,[],[f992,f370])).
fof(f992,plain,(
  e0 = op(e1,e1) | ~spl4_91),
  inference(avatar_component_clause,[],[f990])).
fof(f12002,plain,(
  spl4_27 | ~spl4_100 | ~spl4_119),
  inference(avatar_split_clause,[],[f11998,f1107,f1027,f721])).
fof(f1027,plain,(
  spl4_100 <=> e4 = op(e1,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_100])])).
fof(f11998,plain,(
  e1 = op(e3,e4) | (~spl4_100 | ~spl4_119)),
  inference(backward_demodulation,[],[f1029,f11942])).
fof(f11942,plain,(
  e1 = op(e3,op(e1,e0)) | ~spl4_119),
  inference(backward_demodulation,[],[f1109,f336])).
fof(f1029,plain,(
  e4 = op(e1,e0) | ~spl4_100),
  inference(avatar_component_clause,[],[f1027])).
fof(f12001,plain,(
  spl4_6 | ~spl4_100 | ~spl4_119),
  inference(avatar_split_clause,[],[f11997,f1107,f1027,f633])).
fof(f11997,plain,(
  e0 = op(e4,e3) | (~spl4_100 | ~spl4_119)),
  inference(backward_demodulation,[],[f1029,f11941])).
fof(f11941,plain,(
  e0 = op(op(e1,e0),e3) | ~spl4_119),
  inference(backward_demodulation,[],[f1109,f332])).
fof(f332,plain,(
  e0 = op(op(e1,e0),op(e0,e1))),
  inference(cnf_transformation,[],[f3])).
fof(f12000,plain,(
  ~spl4_85 | ~spl4_100),
  inference(avatar_split_clause,[],[f11995,f1027,f964])).
fof(f964,plain,(
  spl4_85 <=> e4 = op(e1,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_85])])).
fof(f11995,plain,(
  e4 != op(e1,e3) | ~spl4_100),
  inference(backward_demodulation,[],[f1029,f418])).
fof(f11999,plain,(
  ~spl4_75 | ~spl4_100),
  inference(avatar_split_clause,[],[f11990,f1027,f922])).
fof(f11990,plain,(
  e4 != op(e2,e0) | ~spl4_100),
  inference(backward_demodulation,[],[f1029,f360])).
fof(f11982,plain,(
  ~spl4_8 | ~spl4_108),
  inference(avatar_split_clause,[],[f11978,f1061,f641])).
fof(f11978,plain,(
  e2 != op(e4,e3) | ~spl4_108),
  inference(backward_demodulation,[],[f1063,f389])).
fof(f11972,plain,(
  ~spl4_101 | ~spl4_111),
  inference(avatar_split_clause,[],[f11966,f1074,f1032])).
fof(f1032,plain,(
  spl4_101 <=> e0 = op(e0,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_101])])).
fof(f11966,plain,(
  e0 != op(e0,e4) | ~spl4_111),
  inference(backward_demodulation,[],[f1076,f414])).
fof(f414,plain,(
  op(e0,e2) != op(e0,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f11971,plain,(
  ~spl4_11 | ~spl4_111),
  inference(avatar_split_clause,[],[f11964,f1074,f654])).
fof(f654,plain,(
  spl4_11 <=> e0 = op(e4,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_11])])).
fof(f11964,plain,(
  e0 != op(e4,e2) | ~spl4_111),
  inference(backward_demodulation,[],[f1076,f379])).
fof(f11958,plain,(
  ~spl4_43 | ~spl4_119 | ~spl4_122 | spl4_171),
  inference(avatar_split_clause,[],[f11954,f1431,f1120,f1107,f788])).
fof(f1120,plain,(
  spl4_122 <=> op(e0,e0) = e1),
  introduced(avatar_definition,[new_symbols(naming,[spl4_122])])).
fof(f1431,plain,(
  spl4_171 <=> e2 = op(op(e0,op(e0,e0)),op(e0,e0))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_171])])).
fof(f11954,plain,(
  e2 != op(e3,e1) | (~spl4_119 | ~spl4_122 | spl4_171)),
  inference(backward_demodulation,[],[f1109,f11931])).
fof(f11931,plain,(
  e2 != op(op(e0,e1),e1) | (~spl4_122 | spl4_171)),
  inference(backward_demodulation,[],[f1122,f1433])).
fof(f1433,plain,(
  e2 != op(op(e0,op(e0,e0)),op(e0,e0)) | spl4_171),
  inference(avatar_component_clause,[],[f1431])).
fof(f1122,plain,(
  op(e0,e0) = e1 | ~spl4_122),
  inference(avatar_component_clause,[],[f1120])).
fof(f11957,plain,(
  ~spl4_35 | ~spl4_119 | ~spl4_122 | spl4_166),
  inference(avatar_split_clause,[],[f11952,f1405,f1120,f1107,f754])).
fof(f1405,plain,(
  spl4_166 <=> e4 = op(op(e0,op(e0,e0)),op(e0,op(e0,e0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_166])])).
fof(f11952,plain,(
  e4 != op(e3,e3) | (~spl4_119 | ~spl4_122 | spl4_166)),
  inference(backward_demodulation,[],[f1109,f11929])).
fof(f11929,plain,(
  e4 != op(op(e0,e1),op(e0,e1)) | (~spl4_122 | spl4_166)),
  inference(backward_demodulation,[],[f1122,f1407])).
fof(f1407,plain,(
  e4 != op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) | spl4_166),
  inference(avatar_component_clause,[],[f1405])).
fof(f11956,plain,(
  ~spl4_19 | ~spl4_119),
  inference(avatar_split_clause,[],[f11946,f1107,f687])).
fof(f687,plain,(
  spl4_19 <=> e3 = op(e4,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_19])])).
fof(f11946,plain,(
  e3 != op(e4,e1) | ~spl4_119),
  inference(backward_demodulation,[],[f1109,f369])).
fof(f11940,plain,(
  ~spl4_120 | ~spl4_122 | spl4_189),
  inference(avatar_split_clause,[],[f11935,f1569,f1120,f1111])).
fof(f1111,plain,(
  spl4_120 <=> e4 = op(e0,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_120])])).
fof(f1569,plain,(
  spl4_189 <=> e4 = op(e0,op(e0,e0))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_189])])).
fof(f11935,plain,(
  e4 != op(e0,e1) | (~spl4_122 | spl4_189)),
  inference(backward_demodulation,[],[f1122,f1571])).
fof(f1571,plain,(
  e4 != op(e0,op(e0,e0)) | spl4_189),
  inference(avatar_component_clause,[],[f1569])).
fof(f11939,plain,(
  ~spl4_119 | ~spl4_122 | spl4_185),
  inference(avatar_split_clause,[],[f11934,f1529,f1120,f1107])).
fof(f11934,plain,(
  e3 != op(e0,e1) | (~spl4_122 | spl4_185)),
  inference(backward_demodulation,[],[f1122,f1531])).
fof(f1531,plain,(
  e3 != op(e0,op(e0,e0)) | spl4_185),
  inference(avatar_component_clause,[],[f1529])).
fof(f11938,plain,(
  ~spl4_107 | ~spl4_122),
  inference(avatar_split_clause,[],[f11922,f1120,f1057])).
fof(f11922,plain,(
  e1 != op(e0,e3) | ~spl4_122),
  inference(backward_demodulation,[],[f1122,f408])).
fof(f11937,plain,(
  ~spl4_22 | ~spl4_122),
  inference(avatar_split_clause,[],[f11919,f1120,f700])).
fof(f11919,plain,(
  e1 != op(e4,e0) | ~spl4_122),
  inference(backward_demodulation,[],[f1122,f359])).
fof(f11936,plain,(
  spl4_91 | ~spl4_122),
  inference(avatar_split_clause,[],[f11915,f1120,f990])).
fof(f11915,plain,(
  e0 = op(e1,e1) | ~spl4_122),
  inference(backward_demodulation,[],[f1122,f331])).
fof(f11894,plain,(
  ~spl4_1 | ~spl4_5 | spl4_130),
  inference(avatar_split_clause,[],[f11887,f1228,f628,f612])).
fof(f1228,plain,(
  spl4_130 <=> e0 = op(e4,op(e4,e4))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_130])])).
fof(f11887,plain,(
  e0 != op(e4,e4) | (~spl4_5 | spl4_130)),
  inference(backward_demodulation,[],[f630,f1230])).
fof(f1230,plain,(
  e0 != op(e4,op(e4,e4)) | spl4_130),
  inference(avatar_component_clause,[],[f1228])).
fof(f630,plain,(
  e4 = op(e4,e4) | ~spl4_5),
  inference(avatar_component_clause,[],[f628])).
fof(f11880,plain,(
  ~spl4_1 | ~spl4_11),
  inference(avatar_split_clause,[],[f11873,f654,f612])).
fof(f11873,plain,(
  e0 != op(e4,e4) | ~spl4_11),
  inference(backward_demodulation,[],[f656,f454])).
fof(f656,plain,(
  e0 = op(e4,e2) | ~spl4_11),
  inference(avatar_component_clause,[],[f654])).
fof(f11864,plain,(
  ~spl4_12 | ~spl4_22),
  inference(avatar_split_clause,[],[f11856,f700,f658])).
fof(f658,plain,(
  spl4_12 <=> e1 = op(e4,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_12])])).
fof(f11856,plain,(
  e1 != op(e4,e2) | ~spl4_22),
  inference(backward_demodulation,[],[f702,f447])).
fof(f11849,plain,(
  ~spl4_31 | ~spl4_48 | spl4_175),
  inference(avatar_contradiction_clause,[],[f11848])).
fof(f11848,plain,(
  $false | (~spl4_31 | ~spl4_48 | spl4_175)),
  inference(subsumption_resolution,[],[f11836,f811])).
fof(f811,plain,(
  e2 = op(e3,e0) | ~spl4_48),
  inference(avatar_component_clause,[],[f809])).
fof(f809,plain,(
  spl4_48 <=> e2 = op(e3,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_48])])).
fof(f11836,plain,(
  e2 != op(e3,e0) | (~spl4_31 | spl4_175)),
  inference(backward_demodulation,[],[f740,f1459])).
fof(f1459,plain,(
  e2 != op(e3,op(e3,e3)) | spl4_175),
  inference(avatar_component_clause,[],[f1457])).
fof(f1457,plain,(
  spl4_175 <=> e2 = op(e3,op(e3,e3))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_175])])).
fof(f740,plain,(
  e0 = op(e3,e3) | ~spl4_31),
  inference(avatar_component_clause,[],[f738])).
fof(f11822,plain,(
  spl4_19 | ~spl4_42 | ~spl4_85),
  inference(avatar_split_clause,[],[f11820,f964,f784,f687])).
fof(f784,plain,(
  spl4_42 <=> e1 = op(e3,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_42])])).
fof(f11820,plain,(
  e3 = op(e4,e1) | (~spl4_42 | ~spl4_85)),
  inference(backward_demodulation,[],[f786,f11625])).
fof(f11625,plain,(
  e3 = op(e4,op(e3,e1)) | ~spl4_85),
  inference(forward_demodulation,[],[f347,f966])).
fof(f966,plain,(
  e4 = op(e1,e3) | ~spl4_85),
  inference(avatar_component_clause,[],[f964])).
fof(f786,plain,(
  e1 = op(e3,e1) | ~spl4_42),
  inference(avatar_component_clause,[],[f784])).
fof(f11821,plain,(
  ~spl4_27 | ~spl4_42),
  inference(avatar_split_clause,[],[f11815,f784,f721])).
fof(f11815,plain,(
  e1 != op(e3,e4) | ~spl4_42),
  inference(backward_demodulation,[],[f786,f442])).
fof(f442,plain,(
  op(e3,e1) != op(e3,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f11812,plain,(
  ~spl4_33 | ~spl4_48),
  inference(avatar_split_clause,[],[f11805,f809,f746])).
fof(f746,plain,(
  spl4_33 <=> e2 = op(e3,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_33])])).
fof(f11805,plain,(
  e2 != op(e3,e3) | ~spl4_48),
  inference(backward_demodulation,[],[f811,f438])).
fof(f11811,plain,(
  ~spl4_23 | ~spl4_48),
  inference(avatar_split_clause,[],[f11803,f809,f704])).
fof(f11803,plain,(
  e2 != op(e4,e0) | ~spl4_48),
  inference(backward_demodulation,[],[f811,f365])).
fof(f11785,plain,(
  ~spl4_12 | ~spl4_62),
  inference(avatar_split_clause,[],[f11775,f868,f658])).
fof(f11775,plain,(
  e1 != op(e4,e2) | ~spl4_62),
  inference(backward_demodulation,[],[f870,f384])).
fof(f384,plain,(
  op(e2,e2) != op(e4,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f11773,plain,(
  spl4_48 | ~spl4_66 | ~spl4_89),
  inference(avatar_split_clause,[],[f11771,f981,f885,f809])).
fof(f981,plain,(
  spl4_89 <=> e3 = op(e1,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_89])])).
fof(f11771,plain,(
  e2 = op(e3,e0) | (~spl4_66 | ~spl4_89)),
  inference(backward_demodulation,[],[f887,f11735])).
fof(f11735,plain,(
  e2 = op(e3,op(e2,e1)) | ~spl4_89),
  inference(backward_demodulation,[],[f983,f342])).
fof(f983,plain,(
  e3 = op(e1,e2) | ~spl4_89),
  inference(avatar_component_clause,[],[f981])).
fof(f11772,plain,(
  ~spl4_56 | ~spl4_66),
  inference(avatar_split_clause,[],[f11764,f885,f843])).
fof(f843,plain,(
  spl4_56 <=> e0 = op(e2,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_56])])).
fof(f11764,plain,(
  e0 != op(e2,e3) | ~spl4_66),
  inference(backward_demodulation,[],[f887,f431])).
fof(f431,plain,(
  op(e2,e1) != op(e2,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f11759,plain,(
  spl4_11 | ~spl4_75 | ~spl4_113),
  inference(avatar_split_clause,[],[f11756,f1082,f922,f654])).
fof(f11756,plain,(
  e0 = op(e4,e2) | (~spl4_75 | ~spl4_113)),
  inference(backward_demodulation,[],[f924,f11455])).
fof(f11455,plain,(
  e0 = op(op(e2,e0),e2) | ~spl4_113),
  inference(backward_demodulation,[],[f1084,f333])).
fof(f11758,plain,(
  ~spl4_55 | ~spl4_75),
  inference(avatar_split_clause,[],[f11755,f922,f838])).
fof(f11755,plain,(
  e4 != op(e2,e4) | ~spl4_75),
  inference(backward_demodulation,[],[f924,f429])).
fof(f429,plain,(
  op(e2,e0) != op(e2,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f11749,plain,(
  ~spl4_27 | ~spl4_77),
  inference(avatar_split_clause,[],[f11744,f931,f721])).
fof(f11744,plain,(
  e1 != op(e3,e4) | ~spl4_77),
  inference(backward_demodulation,[],[f933,f401])).
fof(f401,plain,(
  op(e1,e4) != op(e3,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f933,plain,(
  e1 = op(e1,e4) | ~spl4_77),
  inference(avatar_component_clause,[],[f931])).
fof(f11740,plain,(
  ~spl4_64 | ~spl4_89),
  inference(avatar_split_clause,[],[f11736,f981,f876])).
fof(f876,plain,(
  spl4_64 <=> e3 = op(e2,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_64])])).
fof(f11736,plain,(
  e3 != op(e2,e2) | ~spl4_89),
  inference(backward_demodulation,[],[f983,f380])).
fof(f11733,plain,(
  ~spl4_40 | ~spl4_93 | spl4_152 | ~spl4_184),
  inference(avatar_contradiction_clause,[],[f11732])).
fof(f11732,plain,(
  $false | (~spl4_40 | ~spl4_93 | spl4_152 | ~spl4_184)),
  inference(subsumption_resolution,[],[f11725,f777])).
fof(f777,plain,(
  e4 = op(e3,e2) | ~spl4_40),
  inference(avatar_component_clause,[],[f775])).
fof(f775,plain,(
  spl4_40 <=> e4 = op(e3,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_40])])).
fof(f11725,plain,(
  e4 != op(e3,e2) | (~spl4_93 | spl4_152 | ~spl4_184)),
  inference(backward_demodulation,[],[f1000,f11649])).
fof(f11649,plain,(
  e4 != op(e3,op(e1,e1)) | (spl4_152 | ~spl4_184)),
  inference(backward_demodulation,[],[f1520,f1329])).
fof(f1329,plain,(
  e4 != op(op(e1,op(e1,e1)),op(e1,e1)) | spl4_152),
  inference(avatar_component_clause,[],[f1327])).
fof(f1327,plain,(
  spl4_152 <=> e4 = op(op(e1,op(e1,e1)),op(e1,e1))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_152])])).
fof(f1520,plain,(
  e3 = op(e1,op(e1,e1)) | ~spl4_184),
  inference(avatar_component_clause,[],[f1519])).
fof(f1519,plain,(
  spl4_184 <=> e3 = op(e1,op(e1,e1))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_184])])).
fof(f11731,plain,(
  spl4_89 | ~spl4_93 | ~spl4_184),
  inference(avatar_split_clause,[],[f11724,f1519,f998,f981])).
fof(f11724,plain,(
  e3 = op(e1,e2) | (~spl4_93 | ~spl4_184)),
  inference(backward_demodulation,[],[f1000,f1520])).
fof(f11730,plain,(
  ~spl4_86 | ~spl4_93 | spl4_148),
  inference(avatar_split_clause,[],[f11723,f1309,f998,f969])).
fof(f969,plain,(
  spl4_86 <=> e0 = op(e1,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_86])])).
fof(f11723,plain,(
  e0 != op(e1,e2) | (~spl4_93 | spl4_148)),
  inference(backward_demodulation,[],[f1000,f1311])).
fof(f1311,plain,(
  e0 != op(e1,op(e1,e1)) | spl4_148),
  inference(avatar_component_clause,[],[f1309])).
fof(f11729,plain,(
  ~spl4_78 | ~spl4_93),
  inference(avatar_split_clause,[],[f11720,f998,f935])).
fof(f11720,plain,(
  e2 != op(e1,e4) | ~spl4_93),
  inference(backward_demodulation,[],[f1000,f422])).
fof(f11714,plain,(
  spl4_22 | ~spl4_96 | ~spl4_120),
  inference(avatar_split_clause,[],[f11711,f1111,f1011,f700])).
fof(f11711,plain,(
  e1 = op(e4,e0) | (~spl4_96 | ~spl4_120)),
  inference(backward_demodulation,[],[f1013,f11682])).
fof(f11682,plain,(
  e1 = op(e4,op(e1,e0)) | ~spl4_120),
  inference(backward_demodulation,[],[f1113,f336])).
fof(f1113,plain,(
  e4 = op(e0,e1) | ~spl4_120),
  inference(avatar_component_clause,[],[f1111])).
fof(f11713,plain,(
  ~spl4_86 | ~spl4_96),
  inference(avatar_split_clause,[],[f11704,f1011,f969])).
fof(f11704,plain,(
  e0 != op(e1,e2) | ~spl4_96),
  inference(backward_demodulation,[],[f1013,f417])).
fof(f417,plain,(
  op(e1,e0) != op(e1,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f11699,plain,(
  ~spl4_1 | ~spl4_101),
  inference(avatar_split_clause,[],[f11694,f1032,f612])).
fof(f11694,plain,(
  e0 != op(e4,e4) | ~spl4_101),
  inference(backward_demodulation,[],[f1034,f399])).
fof(f399,plain,(
  op(e0,e4) != op(e4,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f1034,plain,(
  e0 = op(e0,e4) | ~spl4_101),
  inference(avatar_component_clause,[],[f1032])).
fof(f11688,plain,(
  ~spl4_20 | ~spl4_120),
  inference(avatar_split_clause,[],[f11686,f1111,f691])).
fof(f11686,plain,(
  e4 != op(e4,e1) | ~spl4_120),
  inference(backward_demodulation,[],[f1113,f369])).
fof(f11677,plain,(
  ~spl4_99 | ~spl4_124),
  inference(avatar_split_clause,[],[f11665,f1128,f1023])).
fof(f1128,plain,(
  spl4_124 <=> op(e0,e0) = e3),
  introduced(avatar_definition,[new_symbols(naming,[spl4_124])])).
fof(f11665,plain,(
  e3 != op(e1,e0) | ~spl4_124),
  inference(backward_demodulation,[],[f1130,f356])).
fof(f356,plain,(
  op(e0,e0) != op(e1,e0)),
  inference(cnf_transformation,[],[f4])).
fof(f1130,plain,(
  op(e0,e0) = e3 | ~spl4_124),
  inference(avatar_component_clause,[],[f1128])).
fof(f11676,plain,(
  spl4_31 | ~spl4_124),
  inference(avatar_split_clause,[],[f11664,f1128,f738])).
fof(f11664,plain,(
  e0 = op(e3,e3) | ~spl4_124),
  inference(backward_demodulation,[],[f1130,f331])).
fof(f11631,plain,(
  ~spl4_125 | ~spl4_138 | spl4_140),
  inference(avatar_split_clause,[],[f11630,f1271,f1263,f1132])).
fof(f1263,plain,(
  spl4_138 <=> e0 = op(e2,op(e2,e2))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_138])])).
fof(f1271,plain,(
  spl4_140 <=> e4 = op(op(e2,op(e2,e2)),op(e2,op(e2,e2)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_140])])).
fof(f11630,plain,(
  op(e0,e0) != e4 | (~spl4_138 | spl4_140)),
  inference(forward_demodulation,[],[f1273,f1264])).
fof(f1264,plain,(
  e0 = op(e2,op(e2,e2)) | ~spl4_138),
  inference(avatar_component_clause,[],[f1263])).
fof(f1273,plain,(
  e4 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | spl4_140),
  inference(avatar_component_clause,[],[f1271])).
fof(f11623,plain,(
  ~spl4_35 | ~spl4_40),
  inference(avatar_split_clause,[],[f11622,f775,f754])).
fof(f11622,plain,(
  e4 != op(e3,e3) | ~spl4_40),
  inference(forward_demodulation,[],[f443,f777])).
fof(f11621,plain,(
  ~spl4_35 | ~spl4_85),
  inference(avatar_split_clause,[],[f11620,f964,f754])).
fof(f11620,plain,(
  e4 != op(e3,e3) | ~spl4_85),
  inference(forward_demodulation,[],[f391,f966])).
fof(f391,plain,(
  op(e1,e3) != op(e3,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f11615,plain,(
  ~spl4_7 | ~spl4_12),
  inference(avatar_split_clause,[],[f11609,f658,f637])).
fof(f637,plain,(
  spl4_7 <=> e1 = op(e4,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_7])])).
fof(f11609,plain,(
  e1 != op(e4,e3) | ~spl4_12),
  inference(backward_demodulation,[],[f660,f453])).
fof(f660,plain,(
  e1 = op(e4,e2) | ~spl4_12),
  inference(avatar_component_clause,[],[f658])).
fof(f11586,plain,(
  spl4_23 | ~spl4_40 | ~spl4_56),
  inference(avatar_split_clause,[],[f11578,f843,f775,f704])).
fof(f11578,plain,(
  e2 = op(e4,e0) | (~spl4_40 | ~spl4_56)),
  inference(backward_demodulation,[],[f777,f11543])).
fof(f11543,plain,(
  e2 = op(op(e3,e2),e0) | ~spl4_56),
  inference(backward_demodulation,[],[f845,f344])).
fof(f344,plain,(
  e2 = op(op(e3,e2),op(e2,e3))),
  inference(cnf_transformation,[],[f3])).
fof(f845,plain,(
  e0 = op(e2,e3) | ~spl4_56),
  inference(avatar_component_clause,[],[f843])).
fof(f11584,plain,(
  ~spl4_12 | ~spl4_33 | ~spl4_40 | spl4_146),
  inference(avatar_split_clause,[],[f11576,f1299,f775,f746,f658])).
fof(f1299,plain,(
  spl4_146 <=> e1 = op(op(e3,op(e3,e3)),op(e3,e3))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_146])])).
fof(f11576,plain,(
  e1 != op(e4,e2) | (~spl4_33 | ~spl4_40 | spl4_146)),
  inference(backward_demodulation,[],[f777,f11422])).
fof(f11422,plain,(
  e1 != op(op(e3,e2),e2) | (~spl4_33 | spl4_146)),
  inference(forward_demodulation,[],[f1301,f748])).
fof(f748,plain,(
  e2 = op(e3,e3) | ~spl4_33),
  inference(avatar_component_clause,[],[f746])).
fof(f1301,plain,(
  e1 != op(op(e3,op(e3,e3)),op(e3,e3)) | spl4_146),
  inference(avatar_component_clause,[],[f1299])).
fof(f11582,plain,(
  ~spl4_1 | ~spl4_33 | ~spl4_40 | spl4_167),
  inference(avatar_contradiction_clause,[],[f11581])).
fof(f11581,plain,(
  $false | (~spl4_1 | ~spl4_33 | ~spl4_40 | spl4_167)),
  inference(subsumption_resolution,[],[f11574,f614])).
fof(f11574,plain,(
  e0 != op(e4,e4) | (~spl4_33 | ~spl4_40 | spl4_167)),
  inference(backward_demodulation,[],[f777,f11419])).
fof(f11419,plain,(
  e0 != op(op(e3,e2),op(e3,e2)) | (~spl4_33 | spl4_167)),
  inference(forward_demodulation,[],[f1412,f748])).
fof(f1412,plain,(
  e0 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | spl4_167),
  inference(avatar_component_clause,[],[f1410])).
fof(f1410,plain,(
  spl4_167 <=> e0 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_167])])).
fof(f11580,plain,(
  ~spl4_30 | ~spl4_40),
  inference(avatar_split_clause,[],[f11573,f775,f733])).
fof(f733,plain,(
  spl4_30 <=> e4 = op(e3,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_30])])).
fof(f11573,plain,(
  e4 != op(e3,e4) | ~spl4_40),
  inference(backward_demodulation,[],[f777,f444])).
fof(f11561,plain,(
  ~spl4_30 | ~spl4_55),
  inference(avatar_split_clause,[],[f11560,f838,f733])).
fof(f11560,plain,(
  e4 != op(e3,e4) | ~spl4_55),
  inference(backward_demodulation,[],[f840,f403])).
fof(f11542,plain,(
  ~spl4_52 | ~spl4_72),
  inference(avatar_split_clause,[],[f11535,f910,f826])).
fof(f910,plain,(
  spl4_72 <=> e1 = op(e2,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_72])])).
fof(f11535,plain,(
  e1 != op(e2,e4) | ~spl4_72),
  inference(backward_demodulation,[],[f912,f429])).
fof(f912,plain,(
  e1 = op(e2,e0) | ~spl4_72),
  inference(avatar_component_clause,[],[f910])).
fof(f11541,plain,(
  ~spl4_47 | ~spl4_72),
  inference(avatar_split_clause,[],[f11532,f910,f805])).
fof(f11532,plain,(
  e1 != op(e3,e0) | ~spl4_72),
  inference(backward_demodulation,[],[f912,f363])).
fof(f363,plain,(
  op(e2,e0) != op(e3,e0)),
  inference(cnf_transformation,[],[f4])).
fof(f11531,plain,(
  spl4_12 | ~spl4_20 | ~spl4_78),
  inference(avatar_split_clause,[],[f11529,f935,f691,f658])).
fof(f11529,plain,(
  e1 = op(e4,e2) | (~spl4_20 | ~spl4_78)),
  inference(backward_demodulation,[],[f937,f11166])).
fof(f11166,plain,(
  e1 = op(e4,op(e1,e4)) | ~spl4_20),
  inference(forward_demodulation,[],[f340,f693])).
fof(f340,plain,(
  e1 = op(op(e4,e1),op(e1,e4))),
  inference(cnf_transformation,[],[f3])).
fof(f11530,plain,(
  spl4_55 | ~spl4_20 | ~spl4_78),
  inference(avatar_split_clause,[],[f11528,f935,f691,f838])).
fof(f11528,plain,(
  e4 = op(e2,e4) | (~spl4_20 | ~spl4_78)),
  inference(backward_demodulation,[],[f937,f11165])).
fof(f11165,plain,(
  e4 = op(op(e1,e4),e4) | ~spl4_20),
  inference(forward_demodulation,[],[f352,f693])).
fof(f11523,plain,(
  spl4_27 | ~spl4_44 | ~spl4_85),
  inference(avatar_split_clause,[],[f11520,f964,f792,f721])).
fof(f11520,plain,(
  e1 = op(e3,e4) | (~spl4_44 | ~spl4_85)),
  inference(backward_demodulation,[],[f966,f11168])).
fof(f11168,plain,(
  e1 = op(e3,op(e1,e3)) | ~spl4_44),
  inference(forward_demodulation,[],[f339,f794])).
fof(f11522,plain,(
  spl4_9 | ~spl4_44 | ~spl4_85),
  inference(avatar_split_clause,[],[f11519,f964,f792,f645])).
fof(f11519,plain,(
  e3 = op(e4,e3) | (~spl4_44 | ~spl4_85)),
  inference(backward_demodulation,[],[f966,f11167])).
fof(f11167,plain,(
  e3 = op(op(e1,e3),e3) | ~spl4_44),
  inference(forward_demodulation,[],[f347,f794])).
fof(f11521,plain,(
  ~spl4_60 | ~spl4_85),
  inference(avatar_split_clause,[],[f11516,f964,f859])).
fof(f11516,plain,(
  e4 != op(e2,e3) | ~spl4_85),
  inference(backward_demodulation,[],[f966,f390])).
fof(f11515,plain,(
  spl4_72 | ~spl4_68 | ~spl4_86),
  inference(avatar_split_clause,[],[f11512,f969,f893,f910])).
fof(f11512,plain,(
  e1 = op(e2,e0) | (~spl4_68 | ~spl4_86)),
  inference(backward_demodulation,[],[f971,f11170])).
fof(f11170,plain,(
  e1 = op(e2,op(e1,e2)) | ~spl4_68),
  inference(forward_demodulation,[],[f338,f895])).
fof(f971,plain,(
  e0 = op(e1,e2) | ~spl4_86),
  inference(avatar_component_clause,[],[f969])).
fof(f11514,plain,(
  ~spl4_81 | ~spl4_86),
  inference(avatar_split_clause,[],[f11505,f969,f948])).
fof(f11505,plain,(
  e0 != op(e1,e3) | ~spl4_86),
  inference(backward_demodulation,[],[f971,f423])).
fof(f11513,plain,(
  ~spl4_36 | ~spl4_86),
  inference(avatar_split_clause,[],[f11503,f969,f759])).
fof(f11503,plain,(
  e0 != op(e3,e2) | ~spl4_86),
  inference(backward_demodulation,[],[f971,f381])).
fof(f381,plain,(
  op(e1,e2) != op(e3,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f11502,plain,(
  spl4_46 | ~spl4_99 | ~spl4_116),
  inference(avatar_split_clause,[],[f11498,f1095,f1023,f801])).
fof(f11498,plain,(
  e0 = op(e3,e0) | (~spl4_99 | ~spl4_116)),
  inference(backward_demodulation,[],[f1025,f10955])).
fof(f10955,plain,(
  e0 = op(op(e1,e0),e0) | ~spl4_116),
  inference(backward_demodulation,[],[f1097,f332])).
fof(f11501,plain,(
  ~spl4_79 | ~spl4_99),
  inference(avatar_split_clause,[],[f11496,f1023,f939])).
fof(f11496,plain,(
  e3 != op(e1,e4) | ~spl4_99),
  inference(backward_demodulation,[],[f1025,f419])).
fof(f11489,plain,(
  ~spl4_33 | ~spl4_104 | ~spl4_125 | spl4_173),
  inference(avatar_contradiction_clause,[],[f11488])).
fof(f11488,plain,(
  $false | (~spl4_33 | ~spl4_104 | ~spl4_125 | spl4_173)),
  inference(subsumption_resolution,[],[f11484,f748])).
fof(f11484,plain,(
  e2 != op(e3,e3) | (~spl4_104 | ~spl4_125 | spl4_173)),
  inference(backward_demodulation,[],[f1046,f11427])).
fof(f11427,plain,(
  e2 != op(op(e0,e4),op(e0,e4)) | (~spl4_125 | spl4_173)),
  inference(forward_demodulation,[],[f1443,f1134])).
fof(f1443,plain,(
  e2 != op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) | spl4_173),
  inference(avatar_component_clause,[],[f1441])).
fof(f1441,plain,(
  spl4_173 <=> e2 = op(op(e0,op(e0,e0)),op(e0,op(e0,e0)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_173])])).
fof(f11475,plain,(
  ~spl4_7 | ~spl4_107),
  inference(avatar_split_clause,[],[f11470,f1057,f637])).
fof(f11470,plain,(
  e1 != op(e4,e3) | ~spl4_107),
  inference(backward_demodulation,[],[f1059,f389])).
fof(f11465,plain,(
  ~spl4_103 | ~spl4_113),
  inference(avatar_split_clause,[],[f11461,f1082,f1040])).
fof(f1040,plain,(
  spl4_103 <=> e2 = op(e0,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_103])])).
fof(f11461,plain,(
  e2 != op(e0,e4) | ~spl4_113),
  inference(backward_demodulation,[],[f1084,f414])).
fof(f11453,plain,(
  ~spl4_36 | ~spl4_33 | spl4_133),
  inference(avatar_split_clause,[],[f11452,f1241,f746,f759])).
fof(f1241,plain,(
  spl4_133 <=> e0 = op(e3,op(e3,e3))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_133])])).
fof(f11452,plain,(
  e0 != op(e3,e2) | (~spl4_33 | spl4_133)),
  inference(forward_demodulation,[],[f1243,f748])).
fof(f1243,plain,(
  e0 != op(e3,op(e3,e3)) | spl4_133),
  inference(avatar_component_clause,[],[f1241])).
fof(f11445,plain,(
  ~spl4_103 | ~spl4_125 | spl4_177),
  inference(avatar_split_clause,[],[f11444,f1473,f1132,f1040])).
fof(f1473,plain,(
  spl4_177 <=> e2 = op(e0,op(e0,e0))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_177])])).
fof(f11444,plain,(
  e2 != op(e0,e4) | (~spl4_125 | spl4_177)),
  inference(forward_demodulation,[],[f1475,f1134])).
fof(f1475,plain,(
  e2 != op(e0,op(e0,e0)) | spl4_177),
  inference(avatar_component_clause,[],[f1473])).
fof(f11436,plain,(
  spl4_40 | ~spl4_33 | ~spl4_186),
  inference(avatar_split_clause,[],[f11435,f1548,f746,f775])).
fof(f11435,plain,(
  e4 = op(e3,e2) | (~spl4_33 | ~spl4_186)),
  inference(forward_demodulation,[],[f1549,f748])).
fof(f1549,plain,(
  e4 = op(e3,op(e3,e3)) | ~spl4_186),
  inference(avatar_component_clause,[],[f1548])).
fof(f11434,plain,(
  ~spl4_60 | ~spl4_64 | spl4_187),
  inference(avatar_split_clause,[],[f11433,f1553,f876,f859])).
fof(f11433,plain,(
  e4 != op(e2,e3) | (~spl4_64 | spl4_187)),
  inference(forward_demodulation,[],[f1555,f878])).
fof(f878,plain,(
  e3 = op(e2,e2) | ~spl4_64),
  inference(avatar_component_clause,[],[f876])).
fof(f1555,plain,(
  e4 != op(e2,op(e2,e2)) | spl4_187),
  inference(avatar_component_clause,[],[f1553])).
fof(f11407,plain,(
  ~spl4_115 | ~spl4_125),
  inference(avatar_split_clause,[],[f11406,f1132,f1090])).
fof(f1090,plain,(
  spl4_115 <=> e4 = op(e0,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_115])])).
fof(f11406,plain,(
  e4 != op(e0,e2) | ~spl4_125),
  inference(forward_demodulation,[],[f407,f1134])).
fof(f407,plain,(
  op(e0,e0) != op(e0,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f11402,plain,(
  ~spl4_100 | ~spl4_125),
  inference(avatar_split_clause,[],[f11401,f1132,f1027])).
fof(f11401,plain,(
  e4 != op(e1,e0) | ~spl4_125),
  inference(forward_demodulation,[],[f356,f1134])).
fof(f11385,plain,(
  ~spl4_1 | spl4_21 | ~spl4_130),
  inference(avatar_contradiction_clause,[],[f11384])).
fof(f11384,plain,(
  $false | (~spl4_1 | spl4_21 | ~spl4_130)),
  inference(subsumption_resolution,[],[f11380,f697])).
fof(f697,plain,(
  e0 != op(e4,e0) | spl4_21),
  inference(avatar_component_clause,[],[f696])).
fof(f696,plain,(
  spl4_21 <=> e0 = op(e4,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_21])])).
fof(f11380,plain,(
  e0 = op(e4,e0) | (~spl4_1 | ~spl4_130)),
  inference(backward_demodulation,[],[f614,f1229])).
fof(f1229,plain,(
  e0 = op(e4,op(e4,e4)) | ~spl4_130),
  inference(avatar_component_clause,[],[f1228])).
fof(f11368,plain,(
  ~spl4_3 | ~spl4_13),
  inference(avatar_split_clause,[],[f11363,f662,f620])).
fof(f620,plain,(
  spl4_3 <=> e2 = op(e4,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_3])])).
fof(f11363,plain,(
  e2 != op(e4,e4) | ~spl4_13),
  inference(backward_demodulation,[],[f664,f454])).
fof(f664,plain,(
  e2 = op(e4,e2) | ~spl4_13),
  inference(avatar_component_clause,[],[f662])).
fof(f11351,plain,(
  ~spl4_33 | ~spl4_112 | ~spl4_133 | spl4_146),
  inference(avatar_contradiction_clause,[],[f11350])).
fof(f11350,plain,(
  $false | (~spl4_33 | ~spl4_112 | ~spl4_133 | spl4_146)),
  inference(subsumption_resolution,[],[f11349,f1080])).
fof(f1080,plain,(
  e1 = op(e0,e2) | ~spl4_112),
  inference(avatar_component_clause,[],[f1078])).
fof(f1078,plain,(
  spl4_112 <=> e1 = op(e0,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_112])])).
fof(f11349,plain,(
  e1 != op(e0,e2) | (~spl4_33 | ~spl4_133 | spl4_146)),
  inference(backward_demodulation,[],[f748,f11216])).
fof(f11216,plain,(
  e1 != op(e0,op(e3,e3)) | (~spl4_133 | spl4_146)),
  inference(backward_demodulation,[],[f1242,f1301])).
fof(f1242,plain,(
  e0 = op(e3,op(e3,e3)) | ~spl4_133),
  inference(avatar_component_clause,[],[f1241])).
fof(f11342,plain,(
  spl4_24 | ~spl4_36 | ~spl4_60),
  inference(avatar_split_clause,[],[f11339,f859,f759,f708])).
fof(f11339,plain,(
  e3 = op(e4,e0) | (~spl4_36 | ~spl4_60)),
  inference(backward_demodulation,[],[f761,f11308])).
fof(f11308,plain,(
  e3 = op(e4,op(e3,e2)) | ~spl4_60),
  inference(backward_demodulation,[],[f861,f348])).
fof(f761,plain,(
  e0 = op(e3,e2) | ~spl4_36),
  inference(avatar_component_clause,[],[f759])).
fof(f11341,plain,(
  ~spl4_31 | ~spl4_36),
  inference(avatar_split_clause,[],[f11332,f759,f738])).
fof(f11332,plain,(
  e0 != op(e3,e3) | ~spl4_36),
  inference(backward_demodulation,[],[f761,f443])).
fof(f11340,plain,(
  ~spl4_11 | ~spl4_36),
  inference(avatar_split_clause,[],[f11331,f759,f654])).
fof(f11331,plain,(
  e0 != op(e4,e2) | ~spl4_36),
  inference(backward_demodulation,[],[f761,f385])).
fof(f11330,plain,(
  ~spl4_27 | ~spl4_47),
  inference(avatar_split_clause,[],[f11323,f805,f721])).
fof(f11323,plain,(
  e1 != op(e3,e4) | ~spl4_47),
  inference(backward_demodulation,[],[f807,f439])).
fof(f439,plain,(
  op(e3,e0) != op(e3,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f11305,plain,(
  ~spl4_57 | ~spl4_64 | spl4_160),
  inference(avatar_split_clause,[],[f11299,f1375,f876,f847])).
fof(f1375,plain,(
  spl4_160 <=> e1 = op(e2,op(e2,e2))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_160])])).
fof(f11299,plain,(
  e1 != op(e2,e3) | (~spl4_64 | spl4_160)),
  inference(backward_demodulation,[],[f878,f1377])).
fof(f1377,plain,(
  e1 != op(e2,op(e2,e2)) | spl4_160),
  inference(avatar_component_clause,[],[f1375])).
fof(f11304,plain,(
  ~spl4_54 | ~spl4_64),
  inference(avatar_split_clause,[],[f11296,f876,f834])).
fof(f834,plain,(
  spl4_54 <=> e3 = op(e2,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_54])])).
fof(f11296,plain,(
  e3 != op(e2,e4) | ~spl4_64),
  inference(backward_demodulation,[],[f878,f434])).
fof(f11303,plain,(
  spl4_33 | ~spl4_64),
  inference(avatar_split_clause,[],[f11292,f876,f746])).
fof(f11292,plain,(
  e2 = op(e3,e3) | ~spl4_64),
  inference(backward_demodulation,[],[f878,f343])).
fof(f343,plain,(
  e2 = op(op(e2,e2),op(e2,e2))),
  inference(cnf_transformation,[],[f3])).
fof(f11291,plain,(
  spl4_7 | ~spl4_20 | ~spl4_79),
  inference(avatar_split_clause,[],[f11288,f939,f691,f637])).
fof(f11288,plain,(
  e1 = op(e4,e3) | (~spl4_20 | ~spl4_79)),
  inference(backward_demodulation,[],[f941,f11166])).
fof(f11290,plain,(
  spl4_30 | ~spl4_20 | ~spl4_79),
  inference(avatar_split_clause,[],[f11287,f939,f691,f733])).
fof(f11287,plain,(
  e4 = op(e3,e4) | (~spl4_20 | ~spl4_79)),
  inference(backward_demodulation,[],[f941,f11165])).
fof(f11282,plain,(
  spl4_47 | ~spl4_44 | ~spl4_81),
  inference(avatar_split_clause,[],[f11279,f948,f792,f805])).
fof(f11279,plain,(
  e1 = op(e3,e0) | (~spl4_44 | ~spl4_81)),
  inference(backward_demodulation,[],[f950,f11168])).
fof(f11281,plain,(
  ~spl4_76 | ~spl4_81),
  inference(avatar_split_clause,[],[f11273,f948,f927])).
fof(f11273,plain,(
  e0 != op(e1,e4) | ~spl4_81),
  inference(backward_demodulation,[],[f950,f425])).
fof(f11280,plain,(
  ~spl4_31 | ~spl4_81),
  inference(avatar_split_clause,[],[f11271,f948,f738])).
fof(f11271,plain,(
  e0 != op(e3,e3) | ~spl4_81),
  inference(backward_demodulation,[],[f950,f391])).
fof(f11267,plain,(
  ~spl4_85 | ~spl4_90),
  inference(avatar_split_clause,[],[f11262,f985,f964])).
fof(f11262,plain,(
  e4 != op(e1,e3) | ~spl4_90),
  inference(backward_demodulation,[],[f987,f423])).
fof(f11266,plain,(
  ~spl4_65 | ~spl4_90),
  inference(avatar_split_clause,[],[f11259,f985,f880])).
fof(f880,plain,(
  spl4_65 <=> e4 = op(e2,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_65])])).
fof(f11259,plain,(
  e4 != op(e2,e2) | ~spl4_90),
  inference(backward_demodulation,[],[f987,f380])).
fof(f11258,plain,(
  ~spl4_3 | ~spl4_103),
  inference(avatar_split_clause,[],[f11255,f1040,f620])).
fof(f11255,plain,(
  e2 != op(e4,e4) | ~spl4_103),
  inference(backward_demodulation,[],[f1042,f399])).
fof(f1042,plain,(
  e2 = op(e0,e4) | ~spl4_103),
  inference(avatar_component_clause,[],[f1040])).
fof(f11249,plain,(
  ~spl4_9 | ~spl4_109),
  inference(avatar_split_clause,[],[f11246,f1065,f645])).
fof(f1065,plain,(
  spl4_109 <=> e3 = op(e0,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_109])])).
fof(f11246,plain,(
  e3 != op(e4,e3) | ~spl4_109),
  inference(backward_demodulation,[],[f1067,f389])).
fof(f1067,plain,(
  e3 = op(e0,e3) | ~spl4_109),
  inference(avatar_component_clause,[],[f1065])).
fof(f11238,plain,(
  spl4_103 | ~spl4_125 | ~spl4_177),
  inference(avatar_split_clause,[],[f11230,f1473,f1132,f1040])).
fof(f11230,plain,(
  e2 = op(e0,e4) | (~spl4_125 | ~spl4_177)),
  inference(backward_demodulation,[],[f1134,f1474])).
fof(f1474,plain,(
  e2 = op(e0,op(e0,e0)) | ~spl4_177),
  inference(avatar_component_clause,[],[f1473])).
fof(f11236,plain,(
  ~spl4_50 | ~spl4_125),
  inference(avatar_split_clause,[],[f11225,f1132,f817])).
fof(f817,plain,(
  spl4_50 <=> e4 = op(e3,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_50])])).
fof(f11225,plain,(
  e4 != op(e3,e0) | ~spl4_125),
  inference(backward_demodulation,[],[f1134,f358])).
fof(f358,plain,(
  op(e0,e0) != op(e3,e0)),
  inference(cnf_transformation,[],[f4])).
fof(f11209,plain,(
  ~spl4_65 | spl4_166 | ~spl4_177),
  inference(avatar_split_clause,[],[f11208,f1473,f1405,f880])).
fof(f11208,plain,(
  e4 != op(e2,e2) | (spl4_166 | ~spl4_177)),
  inference(forward_demodulation,[],[f1407,f1474])).
fof(f11161,plain,(
  ~spl4_34 | ~spl4_44),
  inference(avatar_split_clause,[],[f11160,f792,f750])).
fof(f11160,plain,(
  e3 != op(e3,e3) | ~spl4_44),
  inference(forward_demodulation,[],[f441,f794])).
fof(f11158,plain,(
  ~spl4_10 | ~spl4_20),
  inference(avatar_split_clause,[],[f11157,f691,f649])).
fof(f11157,plain,(
  e4 != op(e4,e3) | ~spl4_20),
  inference(forward_demodulation,[],[f451,f693])).
fof(f451,plain,(
  op(e4,e1) != op(e4,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f11133,plain,(
  ~spl4_2 | ~spl4_22),
  inference(avatar_split_clause,[],[f11127,f700,f616])).
fof(f616,plain,(
  spl4_2 <=> e1 = op(e4,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_2])])).
fof(f11127,plain,(
  e1 != op(e4,e4) | ~spl4_22),
  inference(backward_demodulation,[],[f702,f449])).
fof(f449,plain,(
  op(e4,e0) != op(e4,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f11124,plain,(
  ~spl4_2 | ~spl4_27),
  inference(avatar_split_clause,[],[f11118,f721,f616])).
fof(f11118,plain,(
  e1 != op(e4,e4) | ~spl4_27),
  inference(backward_demodulation,[],[f723,f405])).
fof(f405,plain,(
  op(e3,e4) != op(e4,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f723,plain,(
  e1 = op(e3,e4) | ~spl4_27),
  inference(avatar_component_clause,[],[f721])).
fof(f11117,plain,(
  ~spl4_28 | ~spl4_38),
  inference(avatar_split_clause,[],[f11112,f767,f725])).
fof(f767,plain,(
  spl4_38 <=> e2 = op(e3,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_38])])).
fof(f11112,plain,(
  e2 != op(e3,e4) | ~spl4_38),
  inference(backward_demodulation,[],[f769,f444])).
fof(f769,plain,(
  e2 = op(e3,e2) | ~spl4_38),
  inference(avatar_component_clause,[],[f767])).
fof(f11088,plain,(
  ~spl4_60 | ~spl4_65),
  inference(avatar_split_clause,[],[f11078,f880,f859])).
fof(f11078,plain,(
  e4 != op(e2,e3) | ~spl4_65),
  inference(backward_demodulation,[],[f882,f433])).
fof(f882,plain,(
  e4 = op(e2,e2) | ~spl4_65),
  inference(avatar_component_clause,[],[f880])).
fof(f11087,plain,(
  ~spl4_15 | ~spl4_65),
  inference(avatar_split_clause,[],[f11077,f880,f670])).
fof(f11077,plain,(
  e4 != op(e4,e2) | ~spl4_65),
  inference(backward_demodulation,[],[f882,f384])).
fof(f11086,plain,(
  spl4_3 | ~spl4_65),
  inference(avatar_split_clause,[],[f11075,f880,f620])).
fof(f11075,plain,(
  e2 = op(e4,e4) | ~spl4_65),
  inference(backward_demodulation,[],[f882,f343])).
fof(f11074,plain,(
  spl4_38 | ~spl4_68 | ~spl4_89),
  inference(avatar_split_clause,[],[f11071,f981,f893,f767])).
fof(f11071,plain,(
  e2 = op(e3,e2) | (~spl4_68 | ~spl4_89)),
  inference(backward_demodulation,[],[f895,f11027])).
fof(f11027,plain,(
  e2 = op(e3,op(e2,e1)) | ~spl4_89),
  inference(backward_demodulation,[],[f983,f342])).
fof(f11073,plain,(
  spl4_57 | ~spl4_68 | ~spl4_89),
  inference(avatar_split_clause,[],[f11070,f981,f893,f847])).
fof(f11070,plain,(
  e1 = op(e2,e3) | (~spl4_68 | ~spl4_89)),
  inference(backward_demodulation,[],[f895,f11026])).
fof(f11026,plain,(
  e1 = op(op(e2,e1),e3) | ~spl4_89),
  inference(backward_demodulation,[],[f983,f338])).
fof(f11062,plain,(
  ~spl4_51 | ~spl4_71),
  inference(avatar_split_clause,[],[f11055,f906,f822])).
fof(f906,plain,(
  spl4_71 <=> e0 = op(e2,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_71])])).
fof(f11055,plain,(
  e0 != op(e2,e4) | ~spl4_71),
  inference(backward_demodulation,[],[f908,f429])).
fof(f908,plain,(
  e0 = op(e2,e0) | ~spl4_71),
  inference(avatar_component_clause,[],[f906])).
fof(f11035,plain,(
  ~spl4_79 | ~spl4_89),
  inference(avatar_split_clause,[],[f11032,f981,f939])).
fof(f11032,plain,(
  e3 != op(e1,e4) | ~spl4_89),
  inference(backward_demodulation,[],[f983,f424])).
fof(f11034,plain,(
  ~spl4_39 | ~spl4_89),
  inference(avatar_split_clause,[],[f11029,f981,f771])).
fof(f11029,plain,(
  e3 != op(e3,e2) | ~spl4_89),
  inference(backward_demodulation,[],[f983,f381])).
fof(f11023,plain,(
  ~spl4_92 | ~spl4_184),
  inference(avatar_contradiction_clause,[],[f11022])).
fof(f11022,plain,(
  $false | (~spl4_92 | ~spl4_184)),
  inference(subsumption_resolution,[],[f11021,f461])).
fof(f461,plain,(
  e1 != e3),
  inference(cnf_transformation,[],[f5])).
fof(f5,axiom,(
  e3 != e4 & e2 != e4 & e2 != e3 & e1 != e4 & e1 != e3 & e1 != e2 & e0 != e4 & e0 != e3 & e0 != e2 & e0 != e1),
  file('./p/ALG190+1.tptp',unknown)).
fof(f11021,plain,(
  e1 = e3 | (~spl4_92 | ~spl4_184)),
  inference(forward_demodulation,[],[f11015,f996])).
fof(f11015,plain,(
  e3 = op(e1,e1) | (~spl4_92 | ~spl4_184)),
  inference(backward_demodulation,[],[f996,f1520])).
fof(f11018,plain,(
  ~spl4_42 | ~spl4_92),
  inference(avatar_split_clause,[],[f11005,f994,f784])).
fof(f11005,plain,(
  e1 != op(e3,e1) | ~spl4_92),
  inference(backward_demodulation,[],[f996,f371])).
fof(f371,plain,(
  op(e1,e1) != op(e3,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f11002,plain,(
  spl4_71 | ~spl4_98 | ~spl4_116),
  inference(avatar_split_clause,[],[f10998,f1095,f1019,f906])).
fof(f10998,plain,(
  e0 = op(e2,e0) | (~spl4_98 | ~spl4_116)),
  inference(backward_demodulation,[],[f1021,f10955])).
fof(f11001,plain,(
  ~spl4_83 | ~spl4_98),
  inference(avatar_split_clause,[],[f10994,f1019,f956])).
fof(f10994,plain,(
  e2 != op(e1,e3) | ~spl4_98),
  inference(backward_demodulation,[],[f1021,f418])).
fof(f10989,plain,(
  spl4_11 | ~spl4_50 | ~spl4_108),
  inference(avatar_split_clause,[],[f10986,f1061,f817,f654])).
fof(f10986,plain,(
  e0 = op(e4,e2) | (~spl4_50 | ~spl4_108)),
  inference(backward_demodulation,[],[f1063,f10922])).
fof(f10922,plain,(
  e0 = op(e4,op(e0,e3)) | ~spl4_50),
  inference(forward_demodulation,[],[f334,f819])).
fof(f819,plain,(
  e4 = op(e3,e0) | ~spl4_50),
  inference(avatar_component_clause,[],[f817])).
fof(f334,plain,(
  e0 = op(op(e3,e0),op(e0,e3))),
  inference(cnf_transformation,[],[f3])).
fof(f10988,plain,(
  spl4_54 | ~spl4_50 | ~spl4_108),
  inference(avatar_split_clause,[],[f10985,f1061,f817,f834])).
fof(f10985,plain,(
  e3 = op(e2,e4) | (~spl4_50 | ~spl4_108)),
  inference(backward_demodulation,[],[f1063,f10921])).
fof(f10921,plain,(
  e3 = op(op(e0,e3),e4) | ~spl4_50),
  inference(forward_demodulation,[],[f346,f819])).
fof(f10987,plain,(
  ~spl4_83 | ~spl4_108),
  inference(avatar_split_clause,[],[f10981,f1061,f956])).
fof(f10981,plain,(
  e2 != op(e1,e3) | ~spl4_108),
  inference(backward_demodulation,[],[f1063,f386])).
fof(f386,plain,(
  op(e0,e3) != op(e1,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f10980,plain,(
  ~spl4_107 | ~spl4_112),
  inference(avatar_split_clause,[],[f10975,f1078,f1057])).
fof(f10975,plain,(
  e1 != op(e0,e3) | ~spl4_112),
  inference(backward_demodulation,[],[f1080,f413])).
fof(f10979,plain,(
  ~spl4_87 | ~spl4_112),
  inference(avatar_split_clause,[],[f10971,f1078,f973])).
fof(f10971,plain,(
  e1 != op(e1,e2) | ~spl4_112),
  inference(backward_demodulation,[],[f1080,f376])).
fof(f376,plain,(
  op(e0,e2) != op(e1,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f10967,plain,(
  ~spl4_16 | ~spl4_116),
  inference(avatar_split_clause,[],[f10960,f1095,f675])).
fof(f675,plain,(
  spl4_16 <=> e0 = op(e4,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_16])])).
fof(f10960,plain,(
  e0 != op(e4,e1) | ~spl4_116),
  inference(backward_demodulation,[],[f1097,f369])).
fof(f10949,plain,(
  ~spl4_22 | ~spl4_31 | ~spl4_50 | spl4_146),
  inference(avatar_split_clause,[],[f10948,f1299,f817,f738,f700])).
fof(f10948,plain,(
  e1 != op(e4,e0) | (~spl4_31 | ~spl4_50 | spl4_146)),
  inference(forward_demodulation,[],[f10947,f819])).
fof(f10947,plain,(
  e1 != op(op(e3,e0),e0) | (~spl4_31 | spl4_146)),
  inference(forward_demodulation,[],[f1301,f740])).
fof(f10946,plain,(
  ~spl4_2 | ~spl4_31 | ~spl4_50 | spl4_151),
  inference(avatar_split_clause,[],[f10945,f1322,f817,f738,f616])).
fof(f1322,plain,(
  spl4_151 <=> e1 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_151])])).
fof(f10945,plain,(
  e1 != op(e4,e4) | (~spl4_31 | ~spl4_50 | spl4_151)),
  inference(forward_demodulation,[],[f10944,f819])).
fof(f10944,plain,(
  e1 != op(op(e3,e0),op(e3,e0)) | (~spl4_31 | spl4_151)),
  inference(forward_demodulation,[],[f1324,f740])).
fof(f1324,plain,(
  e1 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | spl4_151),
  inference(avatar_component_clause,[],[f1322])).
fof(f10943,plain,(
  ~spl4_107 | ~spl4_124 | spl4_164),
  inference(avatar_split_clause,[],[f10942,f1397,f1128,f1057])).
fof(f1397,plain,(
  spl4_164 <=> e1 = op(e0,op(e0,e0))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_164])])).
fof(f10942,plain,(
  e1 != op(e0,e3) | (~spl4_124 | spl4_164)),
  inference(forward_demodulation,[],[f1399,f1130])).
fof(f1399,plain,(
  e1 != op(e0,op(e0,e0)) | spl4_164),
  inference(avatar_component_clause,[],[f1397])).
fof(f10927,plain,(
  spl4_108 | ~spl4_124 | ~spl4_177),
  inference(avatar_split_clause,[],[f10926,f1473,f1128,f1061])).
fof(f10926,plain,(
  e2 = op(e0,e3) | (~spl4_124 | ~spl4_177)),
  inference(forward_demodulation,[],[f1474,f1130])).
fof(f10924,plain,(
  ~spl4_119 | ~spl4_124),
  inference(avatar_split_clause,[],[f10923,f1128,f1107])).
fof(f10923,plain,(
  e3 != op(e0,e1) | ~spl4_124),
  inference(forward_demodulation,[],[f406,f1130])).
fof(f406,plain,(
  op(e0,e0) != op(e0,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f10887,plain,(
  ~spl4_11 | ~spl4_16),
  inference(avatar_split_clause,[],[f10878,f675,f654])).
fof(f10878,plain,(
  e0 != op(e4,e2) | ~spl4_16),
  inference(backward_demodulation,[],[f677,f450])).
fof(f677,plain,(
  e0 = op(e4,e1) | ~spl4_16),
  inference(avatar_component_clause,[],[f675])).
fof(f10876,plain,(
  ~spl4_3 | ~spl4_23),
  inference(avatar_split_clause,[],[f10871,f704,f620])).
fof(f10871,plain,(
  e2 != op(e4,e4) | ~spl4_23),
  inference(backward_demodulation,[],[f706,f449])).
fof(f10867,plain,(
  ~spl4_3 | ~spl4_28),
  inference(avatar_split_clause,[],[f10864,f725,f620])).
fof(f10864,plain,(
  e2 != op(e4,e4) | ~spl4_28),
  inference(backward_demodulation,[],[f727,f405])).
fof(f727,plain,(
  e2 = op(e3,e4) | ~spl4_28),
  inference(avatar_component_clause,[],[f725])).
fof(f10861,plain,(
  ~spl4_23 | ~spl4_31 | spl4_134 | ~spl4_186),
  inference(avatar_split_clause,[],[f10859,f1548,f1245,f738,f704])).
fof(f1245,plain,(
  spl4_134 <=> e2 = op(op(e3,op(e3,e3)),op(e3,e3))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_134])])).
fof(f10859,plain,(
  e2 != op(e4,e0) | (~spl4_31 | spl4_134 | ~spl4_186)),
  inference(backward_demodulation,[],[f740,f10646])).
fof(f10646,plain,(
  e2 != op(e4,op(e3,e3)) | (spl4_134 | ~spl4_186)),
  inference(backward_demodulation,[],[f1549,f1247])).
fof(f1247,plain,(
  e2 != op(op(e3,op(e3,e3)),op(e3,e3)) | spl4_134),
  inference(avatar_component_clause,[],[f1245])).
fof(f10860,plain,(
  ~spl4_26 | ~spl4_31),
  inference(avatar_split_clause,[],[f10850,f738,f717])).
fof(f717,plain,(
  spl4_26 <=> e0 = op(e3,e4)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_26])])).
fof(f10850,plain,(
  e0 != op(e3,e4) | ~spl4_31),
  inference(backward_demodulation,[],[f740,f445])).
fof(f445,plain,(
  op(e3,e3) != op(e3,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f10838,plain,(
  ~spl4_37 | ~spl4_42),
  inference(avatar_split_clause,[],[f10829,f784,f763])).
fof(f10829,plain,(
  e1 != op(e3,e2) | ~spl4_42),
  inference(backward_demodulation,[],[f786,f440])).
fof(f10837,plain,(
  ~spl4_17 | ~spl4_42),
  inference(avatar_split_clause,[],[f10828,f784,f679])).
fof(f10828,plain,(
  e1 != op(e4,e1) | ~spl4_42),
  inference(backward_demodulation,[],[f786,f375])).
fof(f10827,plain,(
  spl4_16 | ~spl4_50 | ~spl4_107),
  inference(avatar_split_clause,[],[f10824,f1057,f817,f675])).
fof(f10824,plain,(
  e0 = op(e4,e1) | (~spl4_50 | ~spl4_107)),
  inference(backward_demodulation,[],[f819,f10715])).
fof(f10715,plain,(
  e0 = op(op(e3,e0),e1) | ~spl4_107),
  inference(backward_demodulation,[],[f1059,f334])).
fof(f10826,plain,(
  ~spl4_45 | ~spl4_50),
  inference(avatar_split_clause,[],[f10820,f817,f796])).
fof(f10820,plain,(
  e4 != op(e3,e1) | ~spl4_50),
  inference(backward_demodulation,[],[f819,f436])).
fof(f436,plain,(
  op(e3,e0) != op(e3,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f10818,plain,(
  ~spl4_26 | ~spl4_51),
  inference(avatar_split_clause,[],[f10812,f822,f717])).
fof(f10812,plain,(
  e0 != op(e3,e4) | ~spl4_51),
  inference(backward_demodulation,[],[f824,f403])).
fof(f824,plain,(
  e0 = op(e2,e4) | ~spl4_51),
  inference(avatar_component_clause,[],[f822])).
fof(f10800,plain,(
  spl4_62 | ~spl4_63 | ~spl4_160),
  inference(avatar_contradiction_clause,[],[f10799])).
fof(f10799,plain,(
  $false | (spl4_62 | ~spl4_63 | ~spl4_160)),
  inference(subsumption_resolution,[],[f10794,f869])).
fof(f869,plain,(
  e1 != op(e2,e2) | spl4_62),
  inference(avatar_component_clause,[],[f868])).
fof(f10794,plain,(
  e1 = op(e2,e2) | (~spl4_63 | ~spl4_160)),
  inference(backward_demodulation,[],[f874,f1376])).
fof(f1376,plain,(
  e1 = op(e2,op(e2,e2)) | ~spl4_160),
  inference(avatar_component_clause,[],[f1375])).
fof(f10774,plain,(
  ~spl4_58 | ~spl4_83),
  inference(avatar_split_clause,[],[f10769,f956,f851])).
fof(f10769,plain,(
  e2 != op(e2,e3) | ~spl4_83),
  inference(backward_demodulation,[],[f958,f390])).
fof(f10766,plain,(
  spl4_42 | ~spl4_69 | ~spl4_87),
  inference(avatar_split_clause,[],[f10762,f973,f897,f784])).
fof(f897,plain,(
  spl4_69 <=> e3 = op(e2,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_69])])).
fof(f10762,plain,(
  e1 = op(e3,e1) | (~spl4_69 | ~spl4_87)),
  inference(backward_demodulation,[],[f975,f9856])).
fof(f9856,plain,(
  e1 = op(e3,op(e1,e2)) | ~spl4_69),
  inference(forward_demodulation,[],[f338,f899])).
fof(f899,plain,(
  e3 = op(e2,e1) | ~spl4_69),
  inference(avatar_component_clause,[],[f897])).
fof(f10765,plain,(
  spl4_83 | ~spl4_69 | ~spl4_87),
  inference(avatar_split_clause,[],[f10761,f973,f897,f956])).
fof(f10761,plain,(
  e2 = op(e1,e3) | (~spl4_69 | ~spl4_87)),
  inference(backward_demodulation,[],[f975,f9855])).
fof(f9855,plain,(
  e2 = op(op(e1,e2),e3) | ~spl4_69),
  inference(forward_demodulation,[],[f342,f899])).
fof(f10763,plain,(
  ~spl4_37 | ~spl4_87),
  inference(avatar_split_clause,[],[f10755,f973,f763])).
fof(f10755,plain,(
  e1 != op(e3,e2) | ~spl4_87),
  inference(backward_demodulation,[],[f975,f381])).
fof(f10751,plain,(
  ~spl4_45 | ~spl4_95),
  inference(avatar_split_clause,[],[f10741,f1006,f796])).
fof(f1006,plain,(
  spl4_95 <=> e4 = op(e1,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_95])])).
fof(f10741,plain,(
  e4 != op(e3,e1) | ~spl4_95),
  inference(backward_demodulation,[],[f1008,f371])).
fof(f1008,plain,(
  e4 = op(e1,e1) | ~spl4_95),
  inference(avatar_component_clause,[],[f1006])).
fof(f10750,plain,(
  spl4_2 | ~spl4_95),
  inference(avatar_split_clause,[],[f10740,f1006,f616])).
fof(f10740,plain,(
  e1 = op(e4,e4) | ~spl4_95),
  inference(backward_demodulation,[],[f1008,f337])).
fof(f10739,plain,(
  spl4_72 | ~spl4_96 | ~spl4_118),
  inference(avatar_split_clause,[],[f10736,f1103,f1011,f910])).
fof(f1103,plain,(
  spl4_118 <=> e2 = op(e0,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_118])])).
fof(f10736,plain,(
  e1 = op(e2,e0) | (~spl4_96 | ~spl4_118)),
  inference(backward_demodulation,[],[f1013,f10400])).
fof(f10400,plain,(
  e1 = op(e2,op(e1,e0)) | ~spl4_118),
  inference(backward_demodulation,[],[f1105,f336])).
fof(f1105,plain,(
  e2 = op(e0,e1) | ~spl4_118),
  inference(avatar_component_clause,[],[f1103])).
fof(f10738,plain,(
  ~spl4_91 | ~spl4_96),
  inference(avatar_split_clause,[],[f10728,f1011,f990])).
fof(f10728,plain,(
  e0 != op(e1,e1) | ~spl4_96),
  inference(backward_demodulation,[],[f1013,f416])).
fof(f10737,plain,(
  ~spl4_71 | ~spl4_96),
  inference(avatar_split_clause,[],[f10725,f1011,f906])).
fof(f10725,plain,(
  e0 != op(e2,e0) | ~spl4_96),
  inference(backward_demodulation,[],[f1013,f360])).
fof(f10724,plain,(
  ~spl4_82 | ~spl4_107),
  inference(avatar_split_clause,[],[f10717,f1057,f952])).
fof(f10717,plain,(
  e1 != op(e1,e3) | ~spl4_107),
  inference(backward_demodulation,[],[f1059,f386])).
fof(f10698,plain,(
  ~spl4_114 | ~spl4_124),
  inference(avatar_split_clause,[],[f10688,f1128,f1086])).
fof(f10688,plain,(
  e3 != op(e0,e2) | ~spl4_124),
  inference(backward_demodulation,[],[f1130,f407])).
fof(f10697,plain,(
  ~spl4_24 | ~spl4_124),
  inference(avatar_split_clause,[],[f10687,f1128,f708])).
fof(f10687,plain,(
  e3 != op(e4,e0) | ~spl4_124),
  inference(backward_demodulation,[],[f1130,f359])).
fof(f10664,plain,(
  ~spl4_164 | ~spl4_177),
  inference(avatar_contradiction_clause,[],[f10663])).
fof(f10663,plain,(
  $false | (~spl4_164 | ~spl4_177)),
  inference(subsumption_resolution,[],[f10662,f460])).
fof(f460,plain,(
  e1 != e2),
  inference(cnf_transformation,[],[f5])).
fof(f10662,plain,(
  e1 = e2 | (~spl4_164 | ~spl4_177)),
  inference(backward_demodulation,[],[f1398,f1474])).
fof(f1398,plain,(
  e1 = op(e0,op(e0,e0)) | ~spl4_164),
  inference(avatar_component_clause,[],[f1397])).
fof(f10660,plain,(
  ~spl4_91 | ~spl4_160 | spl4_172),
  inference(avatar_split_clause,[],[f10659,f1436,f1375,f990])).
fof(f1436,plain,(
  spl4_172 <=> e0 = op(op(e2,op(e2,e2)),op(e2,op(e2,e2)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_172])])).
fof(f10659,plain,(
  e0 != op(e1,e1) | (~spl4_160 | spl4_172)),
  inference(forward_demodulation,[],[f1438,f1376])).
fof(f1438,plain,(
  e0 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | spl4_172),
  inference(avatar_component_clause,[],[f1436])).
fof(f10658,plain,(
  ~spl4_31 | spl4_180 | ~spl4_184),
  inference(avatar_split_clause,[],[f10657,f1519,f1492,f738])).
fof(f10657,plain,(
  e0 != op(e3,e3) | (spl4_180 | ~spl4_184)),
  inference(forward_demodulation,[],[f1494,f1520])).
fof(f10650,plain,(
  ~spl4_3 | spl4_142 | ~spl4_186),
  inference(avatar_split_clause,[],[f10648,f1548,f1280,f620])).
fof(f1280,plain,(
  spl4_142 <=> e2 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_142])])).
fof(f10648,plain,(
  e2 != op(e4,e4) | (spl4_142 | ~spl4_186)),
  inference(backward_demodulation,[],[f1549,f1282])).
fof(f1282,plain,(
  e2 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | spl4_142),
  inference(avatar_component_clause,[],[f1280])).
fof(f10627,plain,(
  ~spl4_93 | ~spl4_118),
  inference(avatar_split_clause,[],[f10626,f1103,f998])).
fof(f10626,plain,(
  e2 != op(e1,e1) | ~spl4_118),
  inference(forward_demodulation,[],[f366,f1105])).
fof(f10621,plain,(
  ~spl4_55 | ~spl4_105),
  inference(avatar_split_clause,[],[f10620,f1048,f838])).
fof(f10620,plain,(
  e4 != op(e2,e4) | ~spl4_105),
  inference(forward_demodulation,[],[f397,f1050])).
fof(f10619,plain,(
  ~spl4_29 | ~spl4_79),
  inference(avatar_split_clause,[],[f10618,f939,f729])).
fof(f10618,plain,(
  e3 != op(e3,e4) | ~spl4_79),
  inference(forward_demodulation,[],[f401,f941])).
fof(f10580,plain,(
  ~spl4_7 | ~spl4_17),
  inference(avatar_split_clause,[],[f10573,f679,f637])).
fof(f10573,plain,(
  e1 != op(e4,e3) | ~spl4_17),
  inference(backward_demodulation,[],[f681,f451])).
fof(f681,plain,(
  e1 = op(e4,e1) | ~spl4_17),
  inference(avatar_component_clause,[],[f679])).
fof(f10571,plain,(
  spl4_10 | ~spl4_24 | ~spl4_105),
  inference(avatar_split_clause,[],[f10569,f1048,f708,f649])).
fof(f10569,plain,(
  e4 = op(e4,e3) | (~spl4_24 | ~spl4_105)),
  inference(backward_demodulation,[],[f710,f10446])).
fof(f10446,plain,(
  e4 = op(e4,op(e4,e0)) | ~spl4_105),
  inference(backward_demodulation,[],[f1050,f351])).
fof(f710,plain,(
  e3 = op(e4,e0) | ~spl4_24),
  inference(avatar_component_clause,[],[f708])).
fof(f10570,plain,(
  ~spl4_4 | ~spl4_24),
  inference(avatar_split_clause,[],[f10566,f708,f624])).
fof(f10566,plain,(
  e3 != op(e4,e4) | ~spl4_24),
  inference(backward_demodulation,[],[f710,f449])).
fof(f10508,plain,(
  ~spl4_56 | ~spl4_71),
  inference(avatar_split_clause,[],[f10500,f906,f843])).
fof(f10500,plain,(
  e0 != op(e2,e3) | ~spl4_71),
  inference(backward_demodulation,[],[f908,f428])).
fof(f10507,plain,(
  ~spl4_21 | ~spl4_71),
  inference(avatar_split_clause,[],[f10498,f906,f696])).
fof(f10498,plain,(
  e0 != op(e4,e0) | ~spl4_71),
  inference(backward_demodulation,[],[f908,f364])).
fof(f364,plain,(
  op(e2,e0) != op(e4,e0)),
  inference(cnf_transformation,[],[f4])).
fof(f10492,plain,(
  ~spl4_7 | ~spl4_82),
  inference(avatar_split_clause,[],[f10487,f952,f637])).
fof(f10487,plain,(
  e1 != op(e4,e3) | ~spl4_82),
  inference(backward_demodulation,[],[f954,f392])).
fof(f392,plain,(
  op(e1,e3) != op(e4,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f10482,plain,(
  spl4_37 | ~spl4_69 | ~spl4_88),
  inference(avatar_split_clause,[],[f10478,f977,f897,f763])).
fof(f977,plain,(
  spl4_88 <=> e2 = op(e1,e2)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_88])])).
fof(f10478,plain,(
  e1 = op(e3,e2) | (~spl4_69 | ~spl4_88)),
  inference(backward_demodulation,[],[f979,f9856])).
fof(f979,plain,(
  e2 = op(e1,e2) | ~spl4_88),
  inference(avatar_component_clause,[],[f977])).
fof(f10480,plain,(
  ~spl4_83 | ~spl4_88),
  inference(avatar_split_clause,[],[f10473,f977,f956])).
fof(f10473,plain,(
  e2 != op(e1,e3) | ~spl4_88),
  inference(backward_demodulation,[],[f979,f423])).
fof(f10479,plain,(
  ~spl4_63 | ~spl4_88),
  inference(avatar_split_clause,[],[f10470,f977,f872])).
fof(f10470,plain,(
  e2 != op(e2,e2) | ~spl4_88),
  inference(backward_demodulation,[],[f979,f380])).
fof(f10469,plain,(
  spl4_11 | ~spl4_100 | ~spl4_118),
  inference(avatar_split_clause,[],[f10460,f1103,f1027,f654])).
fof(f10460,plain,(
  e0 = op(e4,e2) | (~spl4_100 | ~spl4_118)),
  inference(backward_demodulation,[],[f1029,f10399])).
fof(f10399,plain,(
  e0 = op(op(e1,e0),e2) | ~spl4_118),
  inference(backward_demodulation,[],[f1105,f332])).
fof(f10463,plain,(
  ~spl4_80 | ~spl4_100),
  inference(avatar_split_clause,[],[f10454,f1027,f943])).
fof(f10454,plain,(
  e4 != op(e1,e4) | ~spl4_100),
  inference(backward_demodulation,[],[f1029,f419])).
fof(f10449,plain,(
  ~spl4_80 | ~spl4_105),
  inference(avatar_split_clause,[],[f10447,f1048,f943])).
fof(f10447,plain,(
  e4 != op(e1,e4) | ~spl4_105),
  inference(backward_demodulation,[],[f1050,f396])).
fof(f10444,plain,(
  spl4_71 | ~spl4_48 | ~spl4_106),
  inference(avatar_split_clause,[],[f10442,f1053,f809,f906])).
fof(f10442,plain,(
  e0 = op(e2,e0) | (~spl4_48 | ~spl4_106)),
  inference(backward_demodulation,[],[f1055,f10106])).
fof(f10106,plain,(
  e0 = op(e2,op(e0,e3)) | ~spl4_48),
  inference(forward_demodulation,[],[f334,f811])).
fof(f10431,plain,(
  ~spl4_109 | ~spl4_114),
  inference(avatar_split_clause,[],[f10427,f1086,f1065])).
fof(f10427,plain,(
  e3 != op(e0,e3) | ~spl4_114),
  inference(backward_demodulation,[],[f1088,f413])).
fof(f10420,plain,(
  ~spl4_69 | ~spl4_118 | ~spl4_122 | spl4_165),
  inference(avatar_contradiction_clause,[],[f10419])).
fof(f10419,plain,(
  $false | (~spl4_69 | ~spl4_118 | ~spl4_122 | spl4_165)),
  inference(subsumption_resolution,[],[f10413,f899])).
fof(f10413,plain,(
  e3 != op(e2,e1) | (~spl4_118 | ~spl4_122 | spl4_165)),
  inference(backward_demodulation,[],[f1105,f10391])).
fof(f10391,plain,(
  e3 != op(op(e0,e1),e1) | (~spl4_122 | spl4_165)),
  inference(forward_demodulation,[],[f1403,f1122])).
fof(f1403,plain,(
  e3 != op(op(e0,op(e0,e0)),op(e0,e0)) | spl4_165),
  inference(avatar_component_clause,[],[f1401])).
fof(f1401,plain,(
  spl4_165 <=> e3 = op(op(e0,op(e0,e0)),op(e0,e0))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_165])])).
fof(f10415,plain,(
  ~spl4_103 | ~spl4_118),
  inference(avatar_split_clause,[],[f10405,f1103,f1040])).
fof(f10405,plain,(
  e2 != op(e0,e4) | ~spl4_118),
  inference(backward_demodulation,[],[f1105,f412])).
fof(f10414,plain,(
  ~spl4_18 | ~spl4_118),
  inference(avatar_split_clause,[],[f10402,f1103,f683])).
fof(f10402,plain,(
  e2 != op(e4,e1) | ~spl4_118),
  inference(backward_demodulation,[],[f1105,f369])).
fof(f10387,plain,(
  ~spl4_99 | ~spl4_91 | spl4_184),
  inference(avatar_split_clause,[],[f10386,f1519,f990,f1023])).
fof(f10386,plain,(
  e3 != op(e1,e0) | (~spl4_91 | spl4_184)),
  inference(forward_demodulation,[],[f1521,f992])).
fof(f1521,plain,(
  e3 != op(e1,op(e1,e1)) | spl4_184),
  inference(avatar_component_clause,[],[f1519])).
fof(f10375,plain,(
  ~spl4_100 | ~spl4_91 | spl4_188),
  inference(avatar_split_clause,[],[f10374,f1559,f990,f1027])).
fof(f10374,plain,(
  e4 != op(e1,e0) | (~spl4_91 | spl4_188)),
  inference(forward_demodulation,[],[f1561,f992])).
fof(f1561,plain,(
  e4 != op(e1,op(e1,e1)) | spl4_188),
  inference(avatar_component_clause,[],[f1559])).
fof(f10352,plain,(
  ~spl4_4 | ~spl4_7 | spl4_158),
  inference(avatar_contradiction_clause,[],[f10351])).
fof(f10351,plain,(
  $false | (~spl4_4 | ~spl4_7 | spl4_158)),
  inference(subsumption_resolution,[],[f10342,f639])).
fof(f639,plain,(
  e1 = op(e4,e3) | ~spl4_7),
  inference(avatar_component_clause,[],[f637])).
fof(f10342,plain,(
  e1 != op(e4,e3) | (~spl4_4 | spl4_158)),
  inference(backward_demodulation,[],[f626,f1366])).
fof(f1366,plain,(
  e1 != op(e4,op(e4,e4)) | spl4_158),
  inference(avatar_component_clause,[],[f1364])).
fof(f1364,plain,(
  spl4_158 <=> e1 = op(e4,op(e4,e4))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_158])])).
fof(f626,plain,(
  e3 = op(e4,e4) | ~spl4_4),
  inference(avatar_component_clause,[],[f624])).
fof(f10316,plain,(
  ~spl4_1 | ~spl4_26),
  inference(avatar_split_clause,[],[f10311,f717,f612])).
fof(f10311,plain,(
  e0 != op(e4,e4) | ~spl4_26),
  inference(backward_demodulation,[],[f719,f405])).
fof(f719,plain,(
  e0 = op(e3,e4) | ~spl4_26),
  inference(avatar_component_clause,[],[f717])).
fof(f10306,plain,(
  spl4_26 | ~spl4_35 | ~spl4_133),
  inference(avatar_split_clause,[],[f10302,f1241,f754,f717])).
fof(f10302,plain,(
  e0 = op(e3,e4) | (~spl4_35 | ~spl4_133)),
  inference(backward_demodulation,[],[f756,f1242])).
fof(f10305,plain,(
  ~spl4_10 | ~spl4_35),
  inference(avatar_split_clause,[],[f10300,f754,f649])).
fof(f10300,plain,(
  e4 != op(e4,e3) | ~spl4_35),
  inference(backward_demodulation,[],[f756,f395])).
fof(f395,plain,(
  op(e3,e3) != op(e4,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f10239,plain,(
  ~spl4_65 | ~spl4_75),
  inference(avatar_split_clause,[],[f10234,f922,f880])).
fof(f10234,plain,(
  e4 != op(e2,e2) | ~spl4_75),
  inference(backward_demodulation,[],[f924,f427])).
fof(f427,plain,(
  op(e2,e0) != op(e2,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f10225,plain,(
  ~spl4_48 | ~spl4_91 | spl4_155 | ~spl4_184),
  inference(avatar_contradiction_clause,[],[f10224])).
fof(f10224,plain,(
  $false | (~spl4_48 | ~spl4_91 | spl4_155 | ~spl4_184)),
  inference(subsumption_resolution,[],[f10220,f811])).
fof(f10220,plain,(
  e2 != op(e3,e0) | (~spl4_91 | spl4_155 | ~spl4_184)),
  inference(backward_demodulation,[],[f992,f10120])).
fof(f10120,plain,(
  e2 != op(e3,op(e1,e1)) | (spl4_155 | ~spl4_184)),
  inference(backward_demodulation,[],[f1520,f1345])).
fof(f10222,plain,(
  ~spl4_81 | ~spl4_91),
  inference(avatar_split_clause,[],[f10211,f990,f948])).
fof(f10211,plain,(
  e0 != op(e1,e3) | ~spl4_91),
  inference(backward_demodulation,[],[f992,f421])).
fof(f10207,plain,(
  ~spl4_53 | ~spl4_103),
  inference(avatar_split_clause,[],[f10202,f1040,f830])).
fof(f10202,plain,(
  e2 != op(e2,e4) | ~spl4_103),
  inference(backward_demodulation,[],[f1042,f397])).
fof(f10199,plain,(
  spl4_56 | ~spl4_48 | ~spl4_109),
  inference(avatar_split_clause,[],[f10195,f1065,f809,f843])).
fof(f10195,plain,(
  e0 = op(e2,e3) | (~spl4_48 | ~spl4_109)),
  inference(backward_demodulation,[],[f1067,f10106])).
fof(f10198,plain,(
  spl4_39 | ~spl4_48 | ~spl4_109),
  inference(avatar_split_clause,[],[f10194,f1065,f809,f771])).
fof(f10194,plain,(
  e3 = op(e3,e2) | (~spl4_48 | ~spl4_109)),
  inference(backward_demodulation,[],[f1067,f10105])).
fof(f10105,plain,(
  e3 = op(op(e0,e3),e2) | ~spl4_48),
  inference(forward_demodulation,[],[f346,f811])).
fof(f10197,plain,(
  ~spl4_104 | ~spl4_109),
  inference(avatar_split_clause,[],[f10192,f1065,f1044])).
fof(f10192,plain,(
  e3 != op(e0,e4) | ~spl4_109),
  inference(backward_demodulation,[],[f1067,f415])).
fof(f415,plain,(
  op(e0,e3) != op(e0,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f10196,plain,(
  ~spl4_34 | ~spl4_109),
  inference(avatar_split_clause,[],[f10190,f1065,f750])).
fof(f10190,plain,(
  e3 != op(e3,e3) | ~spl4_109),
  inference(backward_demodulation,[],[f1067,f388])).
fof(f388,plain,(
  op(e0,e3) != op(e3,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f10175,plain,(
  spl4_7 | ~spl4_99 | ~spl4_120),
  inference(avatar_split_clause,[],[f10172,f1111,f1023,f637])).
fof(f10172,plain,(
  e1 = op(e4,e3) | (~spl4_99 | ~spl4_120)),
  inference(backward_demodulation,[],[f1113,f10094])).
fof(f10094,plain,(
  e1 = op(op(e0,e1),e3) | ~spl4_99),
  inference(forward_demodulation,[],[f336,f1025])).
fof(f10174,plain,(
  spl4_26 | ~spl4_99 | ~spl4_120),
  inference(avatar_split_clause,[],[f10171,f1111,f1023,f717])).
fof(f10171,plain,(
  e0 = op(e3,e4) | (~spl4_99 | ~spl4_120)),
  inference(backward_demodulation,[],[f1113,f10093])).
fof(f10093,plain,(
  e0 = op(e3,op(e0,e1)) | ~spl4_99),
  inference(forward_demodulation,[],[f332,f1025])).
fof(f10173,plain,(
  ~spl4_45 | ~spl4_120),
  inference(avatar_split_clause,[],[f10167,f1111,f796])).
fof(f10167,plain,(
  e4 != op(e3,e1) | ~spl4_120),
  inference(backward_demodulation,[],[f1113,f368])).
fof(f10161,plain,(
  ~spl4_112 | ~spl4_122),
  inference(avatar_split_clause,[],[f10147,f1120,f1078])).
fof(f10147,plain,(
  e1 != op(e0,e2) | ~spl4_122),
  inference(backward_demodulation,[],[f1122,f407])).
fof(f10140,plain,(
  ~spl4_125 | ~spl4_133 | spl4_135),
  inference(avatar_split_clause,[],[f10133,f1249,f1241,f1132])).
fof(f1249,plain,(
  spl4_135 <=> e4 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_135])])).
fof(f10133,plain,(
  op(e0,e0) != e4 | (~spl4_133 | spl4_135)),
  inference(backward_demodulation,[],[f1242,f1251])).
fof(f1251,plain,(
  e4 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | spl4_135),
  inference(avatar_component_clause,[],[f1249])).
fof(f10117,plain,(
  ~spl4_3 | spl4_173 | ~spl4_189),
  inference(avatar_split_clause,[],[f10114,f1569,f1441,f620])).
fof(f10114,plain,(
  e2 != op(e4,e4) | (spl4_173 | ~spl4_189)),
  inference(backward_demodulation,[],[f1570,f1443])).
fof(f1570,plain,(
  e4 = op(e0,op(e0,e0)) | ~spl4_189),
  inference(avatar_component_clause,[],[f1569])).
fof(f10103,plain,(
  ~spl4_98 | ~spl4_48),
  inference(avatar_split_clause,[],[f10102,f809,f1019])).
fof(f10102,plain,(
  e2 != op(e1,e0) | ~spl4_48),
  inference(forward_demodulation,[],[f361,f811])).
fof(f10081,plain,(
  spl4_52 | ~spl4_18 | ~spl4_80),
  inference(avatar_split_clause,[],[f10054,f943,f683,f826])).
fof(f10054,plain,(
  e1 = op(e2,e4) | (~spl4_18 | ~spl4_80)),
  inference(backward_demodulation,[],[f685,f9738])).
fof(f9738,plain,(
  e1 = op(op(e4,e1),e4) | ~spl4_80),
  inference(backward_demodulation,[],[f945,f340])).
fof(f10073,plain,(
  ~spl4_3 | ~spl4_18),
  inference(avatar_split_clause,[],[f10072,f683,f620])).
fof(f10072,plain,(
  e2 != op(e4,e4) | ~spl4_18),
  inference(forward_demodulation,[],[f452,f685])).
fof(f452,plain,(
  op(e4,e1) != op(e4,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f10071,plain,(
  ~spl4_5 | ~spl4_80),
  inference(avatar_split_clause,[],[f10070,f943,f628])).
fof(f10070,plain,(
  e4 != op(e4,e4) | ~spl4_80),
  inference(forward_demodulation,[],[f402,f945])).
fof(f10069,plain,(
  ~spl4_1 | ~spl4_3),
  inference(avatar_contradiction_clause,[],[f10068])).
fof(f10068,plain,(
  $false | (~spl4_1 | ~spl4_3)),
  inference(subsumption_resolution,[],[f10067,f457])).
fof(f457,plain,(
  e0 != e2),
  inference(cnf_transformation,[],[f5])).
fof(f10067,plain,(
  e0 = e2 | (~spl4_1 | ~spl4_3)),
  inference(backward_demodulation,[],[f614,f622])).
fof(f622,plain,(
  e2 = op(e4,e4) | ~spl4_3),
  inference(avatar_component_clause,[],[f620])).
fof(f10065,plain,(
  ~spl4_10 | ~spl4_18 | ~spl4_27),
  inference(avatar_contradiction_clause,[],[f10064])).
fof(f10064,plain,(
  $false | (~spl4_10 | ~spl4_18 | ~spl4_27)),
  inference(subsumption_resolution,[],[f10063,f463])).
fof(f463,plain,(
  e2 != e3),
  inference(cnf_transformation,[],[f5])).
fof(f10063,plain,(
  e2 = e3 | (~spl4_10 | ~spl4_18 | ~spl4_27)),
  inference(forward_demodulation,[],[f10061,f685])).
fof(f10061,plain,(
  e3 = op(e4,e1) | (~spl4_10 | ~spl4_27)),
  inference(backward_demodulation,[],[f651,f10038])).
fof(f10038,plain,(
  e3 = op(op(e4,e3),e1) | ~spl4_27),
  inference(backward_demodulation,[],[f723,f350])).
fof(f350,plain,(
  e3 = op(op(e4,e3),op(e3,e4))),
  inference(cnf_transformation,[],[f3])).
fof(f651,plain,(
  e4 = op(e4,e3) | ~spl4_10),
  inference(avatar_component_clause,[],[f649])).
fof(f10050,plain,(
  ~spl4_7 | ~spl4_22),
  inference(avatar_split_clause,[],[f10044,f700,f637])).
fof(f10044,plain,(
  e1 != op(e4,e3) | ~spl4_22),
  inference(backward_demodulation,[],[f702,f448])).
fof(f10028,plain,(
  ~spl4_29 | ~spl4_34),
  inference(avatar_split_clause,[],[f10017,f750,f729])).
fof(f10017,plain,(
  e3 != op(e3,e4) | ~spl4_34),
  inference(backward_demodulation,[],[f752,f445])).
fof(f752,plain,(
  e3 = op(e3,e3) | ~spl4_34),
  inference(avatar_component_clause,[],[f750])).
fof(f10014,plain,(
  ~spl4_36 | ~spl4_57 | ~spl4_98),
  inference(avatar_contradiction_clause,[],[f10013])).
fof(f10013,plain,(
  $false | (~spl4_36 | ~spl4_57 | ~spl4_98)),
  inference(subsumption_resolution,[],[f10012,f463])).
fof(f10012,plain,(
  e2 = e3 | (~spl4_36 | ~spl4_57 | ~spl4_98)),
  inference(forward_demodulation,[],[f10007,f1021])).
fof(f10007,plain,(
  e3 = op(e1,e0) | (~spl4_36 | ~spl4_57)),
  inference(backward_demodulation,[],[f761,f9960])).
fof(f9960,plain,(
  e3 = op(e1,op(e3,e2)) | ~spl4_57),
  inference(backward_demodulation,[],[f849,f348])).
fof(f849,plain,(
  e1 = op(e2,e3) | ~spl4_57),
  inference(avatar_component_clause,[],[f847])).
fof(f10011,plain,(
  ~spl4_36 | ~spl4_57 | ~spl4_116),
  inference(avatar_contradiction_clause,[],[f10010])).
fof(f10010,plain,(
  $false | (~spl4_36 | ~spl4_57 | ~spl4_116)),
  inference(subsumption_resolution,[],[f10009,f457])).
fof(f10009,plain,(
  e0 = e2 | (~spl4_36 | ~spl4_57 | ~spl4_116)),
  inference(forward_demodulation,[],[f10006,f1097])).
fof(f10006,plain,(
  e2 = op(e0,e1) | (~spl4_36 | ~spl4_57)),
  inference(backward_demodulation,[],[f761,f9959])).
fof(f9959,plain,(
  e2 = op(op(e3,e2),e1) | ~spl4_57),
  inference(backward_demodulation,[],[f849,f344])).
fof(f9999,plain,(
  ~spl4_48 | spl4_64 | ~spl4_108),
  inference(avatar_contradiction_clause,[],[f9998])).
fof(f9998,plain,(
  $false | (~spl4_48 | spl4_64 | ~spl4_108)),
  inference(subsumption_resolution,[],[f9994,f877])).
fof(f877,plain,(
  e3 != op(e2,e2) | spl4_64),
  inference(avatar_component_clause,[],[f876])).
fof(f9994,plain,(
  e3 = op(e2,e2) | (~spl4_48 | ~spl4_108)),
  inference(backward_demodulation,[],[f811,f9912])).
fof(f9912,plain,(
  e3 = op(e2,op(e3,e0)) | ~spl4_108),
  inference(backward_demodulation,[],[f1063,f346])).
fof(f9997,plain,(
  ~spl4_48 | spl4_61 | ~spl4_108),
  inference(avatar_contradiction_clause,[],[f9996])).
fof(f9996,plain,(
  $false | (~spl4_48 | spl4_61 | ~spl4_108)),
  inference(subsumption_resolution,[],[f9993,f865])).
fof(f865,plain,(
  e0 != op(e2,e2) | spl4_61),
  inference(avatar_component_clause,[],[f864])).
fof(f9993,plain,(
  e0 = op(e2,e2) | (~spl4_48 | ~spl4_108)),
  inference(backward_demodulation,[],[f811,f9911])).
fof(f9911,plain,(
  e0 = op(op(e3,e0),e2) | ~spl4_108),
  inference(backward_demodulation,[],[f1063,f334])).
fof(f9995,plain,(
  ~spl4_38 | ~spl4_48),
  inference(avatar_split_clause,[],[f9988,f809,f767])).
fof(f9988,plain,(
  e2 != op(e3,e2) | ~spl4_48),
  inference(backward_demodulation,[],[f811,f437])).
fof(f9985,plain,(
  ~spl4_14 | ~spl4_53 | ~spl4_57),
  inference(avatar_contradiction_clause,[],[f9984])).
fof(f9984,plain,(
  $false | (~spl4_14 | ~spl4_53 | ~spl4_57)),
  inference(subsumption_resolution,[],[f9983,f462])).
fof(f462,plain,(
  e1 != e4),
  inference(cnf_transformation,[],[f5])).
fof(f9983,plain,(
  e1 = e4 | (~spl4_14 | ~spl4_53 | ~spl4_57)),
  inference(forward_demodulation,[],[f9977,f849])).
fof(f9977,plain,(
  e4 = op(e2,e3) | (~spl4_14 | ~spl4_53)),
  inference(backward_demodulation,[],[f832,f9849])).
fof(f9849,plain,(
  e4 = op(op(e2,e4),e3) | ~spl4_14),
  inference(forward_demodulation,[],[f353,f668])).
fof(f832,plain,(
  e2 = op(e2,e4) | ~spl4_53),
  inference(avatar_component_clause,[],[f830])).
fof(f9968,plain,(
  ~spl4_52 | ~spl4_57),
  inference(avatar_split_clause,[],[f9963,f847,f826])).
fof(f9963,plain,(
  e1 != op(e2,e4) | ~spl4_57),
  inference(backward_demodulation,[],[f849,f435])).
fof(f9967,plain,(
  ~spl4_7 | ~spl4_57),
  inference(avatar_split_clause,[],[f9962,f847,f637])).
fof(f9962,plain,(
  e1 != op(e4,e3) | ~spl4_57),
  inference(backward_demodulation,[],[f849,f394])).
fof(f9936,plain,(
  ~spl4_98 | ~spl4_99),
  inference(avatar_contradiction_clause,[],[f9935])).
fof(f9935,plain,(
  $false | (~spl4_98 | ~spl4_99)),
  inference(subsumption_resolution,[],[f9934,f463])).
fof(f9934,plain,(
  e2 = e3 | (~spl4_98 | ~spl4_99)),
  inference(forward_demodulation,[],[f1025,f1021])).
fof(f9933,plain,(
  ~spl4_27 | ~spl4_104 | ~spl4_125 | spl4_179),
  inference(avatar_split_clause,[],[f9930,f1487,f1132,f1044,f721])).
fof(f1487,plain,(
  spl4_179 <=> e1 = op(op(e0,op(e0,e0)),op(e0,e0))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_179])])).
fof(f9930,plain,(
  e1 != op(e3,e4) | (~spl4_104 | ~spl4_125 | spl4_179)),
  inference(backward_demodulation,[],[f1046,f9889])).
fof(f9889,plain,(
  e1 != op(op(e0,e4),e4) | (~spl4_125 | spl4_179)),
  inference(backward_demodulation,[],[f1134,f1489])).
fof(f1489,plain,(
  e1 != op(op(e0,op(e0,e0)),op(e0,e0)) | spl4_179),
  inference(avatar_component_clause,[],[f1487])).
fof(f9910,plain,(
  ~spl4_112 | ~spl4_113),
  inference(avatar_contradiction_clause,[],[f9909])).
fof(f9909,plain,(
  $false | (~spl4_112 | ~spl4_113)),
  inference(subsumption_resolution,[],[f9908,f460])).
fof(f9908,plain,(
  e1 = e2 | (~spl4_112 | ~spl4_113)),
  inference(forward_demodulation,[],[f1084,f1080])).
fof(f9907,plain,(
  ~spl4_101 | ~spl4_116),
  inference(avatar_split_clause,[],[f9899,f1095,f1032])).
fof(f9899,plain,(
  e0 != op(e0,e4) | ~spl4_116),
  inference(backward_demodulation,[],[f1097,f412])).
fof(f9896,plain,(
  ~spl4_104 | ~spl4_125 | spl4_185),
  inference(avatar_split_clause,[],[f9891,f1529,f1132,f1044])).
fof(f9891,plain,(
  e3 != op(e0,e4) | (~spl4_125 | spl4_185)),
  inference(backward_demodulation,[],[f1134,f1531])).
fof(f9895,plain,(
  ~spl4_110 | ~spl4_125),
  inference(avatar_split_clause,[],[f9881,f1132,f1069])).
fof(f1069,plain,(
  spl4_110 <=> e4 = op(e0,e3)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_110])])).
fof(f9881,plain,(
  e4 != op(e0,e3) | ~spl4_125),
  inference(backward_demodulation,[],[f1134,f408])).
fof(f9894,plain,(
  ~spl4_25 | ~spl4_125),
  inference(avatar_split_clause,[],[f9879,f1132,f712])).
fof(f712,plain,(
  spl4_25 <=> e4 = op(e4,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_25])])).
fof(f9879,plain,(
  e4 != op(e4,e0) | ~spl4_125),
  inference(backward_demodulation,[],[f1134,f359])).
fof(f9870,plain,(
  ~spl4_52 | ~spl4_65 | spl4_160),
  inference(avatar_split_clause,[],[f9869,f1375,f880,f826])).
fof(f9869,plain,(
  e1 != op(e2,e4) | (~spl4_65 | spl4_160)),
  inference(forward_demodulation,[],[f1377,f882])).
fof(f9860,plain,(
  spl4_116 | ~spl4_71 | ~spl4_112),
  inference(avatar_split_clause,[],[f9749,f1078,f906,f1095])).
fof(f9749,plain,(
  e0 = op(e0,e1) | (~spl4_71 | ~spl4_112)),
  inference(backward_demodulation,[],[f908,f8973])).
fof(f8973,plain,(
  e0 = op(op(e2,e0),e1) | ~spl4_112),
  inference(forward_demodulation,[],[f333,f1080])).
fof(f9848,plain,(
  ~spl4_53 | ~spl4_3),
  inference(avatar_split_clause,[],[f9847,f620,f830])).
fof(f9847,plain,(
  e2 != op(e2,e4) | ~spl4_3),
  inference(forward_demodulation,[],[f404,f622])).
fof(f404,plain,(
  op(e2,e4) != op(e4,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f9846,plain,(
  ~spl4_54 | ~spl4_69),
  inference(avatar_split_clause,[],[f9845,f897,f834])).
fof(f9845,plain,(
  e3 != op(e2,e4) | ~spl4_69),
  inference(forward_demodulation,[],[f432,f899])).
fof(f432,plain,(
  op(e2,e1) != op(e2,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f9825,plain,(
  ~spl4_6 | ~spl4_16),
  inference(avatar_split_clause,[],[f9818,f675,f633])).
fof(f9818,plain,(
  e0 != op(e4,e3) | ~spl4_16),
  inference(backward_demodulation,[],[f677,f451])).
fof(f9814,plain,(
  ~spl4_31 | ~spl4_98 | spl4_134 | ~spl4_159),
  inference(avatar_contradiction_clause,[],[f9813])).
fof(f9813,plain,(
  $false | (~spl4_31 | ~spl4_98 | spl4_134 | ~spl4_159)),
  inference(subsumption_resolution,[],[f9811,f1021])).
fof(f9811,plain,(
  e2 != op(e1,e0) | (~spl4_31 | spl4_134 | ~spl4_159)),
  inference(backward_demodulation,[],[f740,f9598])).
fof(f9598,plain,(
  e2 != op(e1,op(e3,e3)) | (spl4_134 | ~spl4_159)),
  inference(backward_demodulation,[],[f1370,f1247])).
fof(f9812,plain,(
  ~spl4_6 | ~spl4_31),
  inference(avatar_split_clause,[],[f9804,f738,f633])).
fof(f9804,plain,(
  e0 != op(e4,e3) | ~spl4_31),
  inference(backward_demodulation,[],[f740,f395])).
fof(f9802,plain,(
  ~spl4_38 | ~spl4_58 | spl4_64),
  inference(avatar_contradiction_clause,[],[f9801])).
fof(f9801,plain,(
  $false | (~spl4_38 | ~spl4_58 | spl4_64)),
  inference(subsumption_resolution,[],[f9798,f877])).
fof(f9798,plain,(
  e3 = op(e2,e2) | (~spl4_38 | ~spl4_58)),
  inference(backward_demodulation,[],[f769,f9765])).
fof(f9765,plain,(
  e3 = op(e2,op(e3,e2)) | ~spl4_58),
  inference(backward_demodulation,[],[f853,f348])).
fof(f853,plain,(
  e2 = op(e2,e3) | ~spl4_58),
  inference(avatar_component_clause,[],[f851])).
fof(f9790,plain,(
  ~spl4_30 | ~spl4_45),
  inference(avatar_split_clause,[],[f9787,f796,f733])).
fof(f9787,plain,(
  e4 != op(e3,e4) | ~spl4_45),
  inference(backward_demodulation,[],[f798,f442])).
fof(f9783,plain,(
  spl4_19 | ~spl4_47 | ~spl4_110),
  inference(avatar_contradiction_clause,[],[f9782])).
fof(f9782,plain,(
  $false | (spl4_19 | ~spl4_47 | ~spl4_110)),
  inference(subsumption_resolution,[],[f9778,f688])).
fof(f688,plain,(
  e3 != op(e4,e1) | spl4_19),
  inference(avatar_component_clause,[],[f687])).
fof(f9778,plain,(
  e3 = op(e4,e1) | (~spl4_47 | ~spl4_110)),
  inference(backward_demodulation,[],[f807,f9642])).
fof(f9642,plain,(
  e3 = op(e4,op(e3,e0)) | ~spl4_110),
  inference(backward_demodulation,[],[f1071,f346])).
fof(f1071,plain,(
  e4 = op(e0,e3) | ~spl4_110),
  inference(avatar_component_clause,[],[f1069])).
fof(f9781,plain,(
  ~spl4_47 | spl4_76 | ~spl4_110),
  inference(avatar_contradiction_clause,[],[f9780])).
fof(f9780,plain,(
  $false | (~spl4_47 | spl4_76 | ~spl4_110)),
  inference(subsumption_resolution,[],[f9777,f928])).
fof(f928,plain,(
  e0 != op(e1,e4) | spl4_76),
  inference(avatar_component_clause,[],[f927])).
fof(f9777,plain,(
  e0 = op(e1,e4) | (~spl4_47 | ~spl4_110)),
  inference(backward_demodulation,[],[f807,f9641])).
fof(f9641,plain,(
  e0 = op(op(e3,e0),e4) | ~spl4_110),
  inference(backward_demodulation,[],[f1071,f334])).
fof(f9779,plain,(
  ~spl4_32 | ~spl4_47),
  inference(avatar_split_clause,[],[f9772,f805,f742])).
fof(f9772,plain,(
  e1 != op(e3,e3) | ~spl4_47),
  inference(backward_demodulation,[],[f807,f438])).
fof(f9763,plain,(
  ~spl4_69 | ~spl4_86 | spl4_108),
  inference(avatar_contradiction_clause,[],[f9762])).
fof(f9762,plain,(
  $false | (~spl4_69 | ~spl4_86 | spl4_108)),
  inference(subsumption_resolution,[],[f9759,f1062])).
fof(f1062,plain,(
  e2 != op(e0,e3) | spl4_108),
  inference(avatar_component_clause,[],[f1061])).
fof(f9759,plain,(
  e2 = op(e0,e3) | (~spl4_69 | ~spl4_86)),
  inference(backward_demodulation,[],[f899,f9720])).
fof(f9720,plain,(
  e2 = op(e0,op(e2,e1)) | ~spl4_86),
  inference(backward_demodulation,[],[f971,f342])).
fof(f9761,plain,(
  spl4_47 | ~spl4_69 | ~spl4_86),
  inference(avatar_split_clause,[],[f9758,f969,f897,f805])).
fof(f9758,plain,(
  e1 = op(e3,e0) | (~spl4_69 | ~spl4_86)),
  inference(backward_demodulation,[],[f899,f9719])).
fof(f9719,plain,(
  e1 = op(op(e2,e1),e0) | ~spl4_86),
  inference(backward_demodulation,[],[f971,f338])).
fof(f9751,plain,(
  ~spl4_66 | ~spl4_71),
  inference(avatar_split_clause,[],[f9743,f906,f885])).
fof(f9743,plain,(
  e0 != op(e2,e1) | ~spl4_71),
  inference(backward_demodulation,[],[f908,f426])).
fof(f9741,plain,(
  ~spl4_30 | ~spl4_80),
  inference(avatar_split_clause,[],[f9740,f943,f733])).
fof(f9740,plain,(
  e4 != op(e3,e4) | ~spl4_80),
  inference(backward_demodulation,[],[f945,f401])).
fof(f9729,plain,(
  ~spl4_76 | ~spl4_86),
  inference(avatar_split_clause,[],[f9723,f969,f927])).
fof(f9723,plain,(
  e0 != op(e1,e4) | ~spl4_86),
  inference(backward_demodulation,[],[f971,f424])).
fof(f9688,plain,(
  spl4_62 | ~spl4_98 | ~spl4_118),
  inference(avatar_contradiction_clause,[],[f9687])).
fof(f9687,plain,(
  $false | (spl4_62 | ~spl4_98 | ~spl4_118)),
  inference(subsumption_resolution,[],[f9682,f869])).
fof(f9682,plain,(
  e1 = op(e2,e2) | (~spl4_98 | ~spl4_118)),
  inference(backward_demodulation,[],[f1021,f9630])).
fof(f9630,plain,(
  e1 = op(e2,op(e1,e0)) | ~spl4_118),
  inference(backward_demodulation,[],[f1105,f336])).
fof(f9686,plain,(
  spl4_61 | ~spl4_98 | ~spl4_118),
  inference(avatar_contradiction_clause,[],[f9685])).
fof(f9685,plain,(
  $false | (spl4_61 | ~spl4_98 | ~spl4_118)),
  inference(subsumption_resolution,[],[f9681,f865])).
fof(f9681,plain,(
  e0 = op(e2,e2) | (~spl4_98 | ~spl4_118)),
  inference(backward_demodulation,[],[f1021,f9629])).
fof(f9629,plain,(
  e0 = op(op(e1,e0),e2) | ~spl4_118),
  inference(backward_demodulation,[],[f1105,f332])).
fof(f9684,plain,(
  ~spl4_88 | ~spl4_98),
  inference(avatar_split_clause,[],[f9676,f1019,f977])).
fof(f9676,plain,(
  e2 != op(e1,e2) | ~spl4_98),
  inference(backward_demodulation,[],[f1021,f417])).
fof(f9683,plain,(
  ~spl4_73 | ~spl4_98),
  inference(avatar_split_clause,[],[f9673,f1019,f914])).
fof(f914,plain,(
  spl4_73 <=> e2 = op(e2,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_73])])).
fof(f9673,plain,(
  e2 != op(e2,e0) | ~spl4_98),
  inference(backward_demodulation,[],[f1021,f360])).
fof(f9670,plain,(
  ~spl4_25 | ~spl4_101),
  inference(avatar_contradiction_clause,[],[f9669])).
fof(f9669,plain,(
  $false | (~spl4_25 | ~spl4_101)),
  inference(subsumption_resolution,[],[f9668,f459])).
fof(f459,plain,(
  e0 != e4),
  inference(cnf_transformation,[],[f5])).
fof(f9668,plain,(
  e0 = e4 | (~spl4_25 | ~spl4_101)),
  inference(forward_demodulation,[],[f9665,f1034])).
fof(f9665,plain,(
  e4 = op(e0,e4) | (~spl4_25 | ~spl4_101)),
  inference(backward_demodulation,[],[f1034,f9585])).
fof(f9585,plain,(
  e4 = op(op(e0,e4),e4) | ~spl4_25),
  inference(forward_demodulation,[],[f351,f714])).
fof(f714,plain,(
  e4 = op(e4,e0) | ~spl4_25),
  inference(avatar_component_clause,[],[f712])).
fof(f9667,plain,(
  ~spl4_76 | ~spl4_101),
  inference(avatar_split_clause,[],[f9659,f1032,f927])).
fof(f9659,plain,(
  e0 != op(e1,e4) | ~spl4_101),
  inference(backward_demodulation,[],[f1034,f396])).
fof(f9658,plain,(
  ~spl4_7 | ~spl4_110 | ~spl4_124 | spl4_179),
  inference(avatar_split_clause,[],[f9652,f1487,f1128,f1069,f637])).
fof(f9652,plain,(
  e1 != op(e4,e3) | (~spl4_110 | ~spl4_124 | spl4_179)),
  inference(backward_demodulation,[],[f1071,f9620])).
fof(f9620,plain,(
  e1 != op(op(e0,e3),e3) | (~spl4_124 | spl4_179)),
  inference(backward_demodulation,[],[f1130,f1489])).
fof(f9657,plain,(
  ~spl4_3 | ~spl4_110 | ~spl4_124 | spl4_173),
  inference(avatar_contradiction_clause,[],[f9656])).
fof(f9656,plain,(
  $false | (~spl4_3 | ~spl4_110 | ~spl4_124 | spl4_173)),
  inference(subsumption_resolution,[],[f9651,f622])).
fof(f9651,plain,(
  e2 != op(e4,e4) | (~spl4_110 | ~spl4_124 | spl4_173)),
  inference(backward_demodulation,[],[f1071,f9618])).
fof(f9618,plain,(
  e2 != op(op(e0,e3),op(e0,e3)) | (~spl4_124 | spl4_173)),
  inference(backward_demodulation,[],[f1130,f1443])).
fof(f9640,plain,(
  ~spl4_108 | ~spl4_118),
  inference(avatar_split_clause,[],[f9635,f1103,f1061])).
fof(f9635,plain,(
  e2 != op(e0,e3) | ~spl4_118),
  inference(backward_demodulation,[],[f1105,f411])).
fof(f411,plain,(
  op(e0,e1) != op(e0,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f9639,plain,(
  ~spl4_43 | ~spl4_118),
  inference(avatar_split_clause,[],[f9633,f1103,f788])).
fof(f9633,plain,(
  e2 != op(e3,e1) | ~spl4_118),
  inference(backward_demodulation,[],[f1105,f368])).
fof(f9628,plain,(
  ~spl4_110 | ~spl4_124 | spl4_189),
  inference(avatar_split_clause,[],[f9623,f1569,f1128,f1069])).
fof(f9623,plain,(
  e4 != op(e0,e3) | (~spl4_124 | spl4_189)),
  inference(backward_demodulation,[],[f1130,f1571])).
fof(f9627,plain,(
  ~spl4_108 | ~spl4_124 | spl4_177),
  inference(avatar_split_clause,[],[f9619,f1473,f1128,f1061])).
fof(f9619,plain,(
  e2 != op(e0,e3) | (~spl4_124 | spl4_177)),
  inference(backward_demodulation,[],[f1130,f1475])).
fof(f9625,plain,(
  ~spl4_49 | ~spl4_124),
  inference(avatar_split_clause,[],[f9609,f1128,f813])).
fof(f813,plain,(
  spl4_49 <=> e3 = op(e3,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_49])])).
fof(f9609,plain,(
  e3 != op(e3,e0) | ~spl4_124),
  inference(backward_demodulation,[],[f1130,f358])).
fof(f9566,plain,(
  spl4_43 | ~spl4_14 | ~spl4_52),
  inference(avatar_split_clause,[],[f9525,f826,f666,f788])).
fof(f9525,plain,(
  e2 = op(e3,e1) | (~spl4_14 | ~spl4_52)),
  inference(backward_demodulation,[],[f668,f9452])).
fof(f9452,plain,(
  e2 = op(op(e4,e2),e1) | ~spl4_52),
  inference(backward_demodulation,[],[f828,f345])).
fof(f345,plain,(
  e2 = op(op(e4,e2),op(e2,e4))),
  inference(cnf_transformation,[],[f3])).
fof(f9555,plain,(
  ~spl4_3 | ~spl4_14 | spl4_182),
  inference(avatar_contradiction_clause,[],[f9554])).
fof(f9554,plain,(
  $false | (~spl4_3 | ~spl4_14 | spl4_182)),
  inference(subsumption_resolution,[],[f9543,f668])).
fof(f9543,plain,(
  e3 != op(e4,e2) | (~spl4_3 | spl4_182)),
  inference(backward_demodulation,[],[f622,f1510])).
fof(f9521,plain,(
  ~spl4_2 | ~spl4_17),
  inference(avatar_split_clause,[],[f9515,f679,f616])).
fof(f9515,plain,(
  e1 != op(e4,e4) | ~spl4_17),
  inference(backward_demodulation,[],[f681,f452])).
fof(f9511,plain,(
  ~spl4_10 | ~spl4_25),
  inference(avatar_split_clause,[],[f9507,f712,f649])).
fof(f9507,plain,(
  e4 != op(e4,e3) | ~spl4_25),
  inference(backward_demodulation,[],[f714,f448])).
fof(f9501,plain,(
  ~spl4_32 | ~spl4_66 | spl4_162 | ~spl4_175),
  inference(avatar_contradiction_clause,[],[f9500])).
fof(f9500,plain,(
  $false | (~spl4_32 | ~spl4_66 | spl4_162 | ~spl4_175)),
  inference(subsumption_resolution,[],[f9498,f887])).
fof(f9498,plain,(
  e0 != op(e2,e1) | (~spl4_32 | spl4_162 | ~spl4_175)),
  inference(backward_demodulation,[],[f744,f9269])).
fof(f9269,plain,(
  e0 != op(e2,op(e3,e3)) | (spl4_162 | ~spl4_175)),
  inference(forward_demodulation,[],[f1389,f1458])).
fof(f1458,plain,(
  e2 = op(e3,op(e3,e3)) | ~spl4_175),
  inference(avatar_component_clause,[],[f1457])).
fof(f1389,plain,(
  e0 != op(op(e3,op(e3,e3)),op(e3,e3)) | spl4_162),
  inference(avatar_component_clause,[],[f1387])).
fof(f1387,plain,(
  spl4_162 <=> e0 = op(op(e3,op(e3,e3)),op(e3,e3))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_162])])).
fof(f9487,plain,(
  ~spl4_26 | ~spl4_36),
  inference(avatar_split_clause,[],[f9480,f759,f717])).
fof(f9480,plain,(
  e0 != op(e3,e4) | ~spl4_36),
  inference(backward_demodulation,[],[f761,f444])).
fof(f9477,plain,(
  spl4_36 | ~spl4_49 | ~spl4_108),
  inference(avatar_split_clause,[],[f9473,f1061,f813,f759])).
fof(f9473,plain,(
  e0 = op(e3,e2) | (~spl4_49 | ~spl4_108)),
  inference(backward_demodulation,[],[f815,f9319])).
fof(f9319,plain,(
  e0 = op(op(e3,e0),e2) | ~spl4_108),
  inference(backward_demodulation,[],[f1063,f334])).
fof(f815,plain,(
  e3 = op(e3,e0) | ~spl4_49),
  inference(avatar_component_clause,[],[f813])).
fof(f9476,plain,(
  ~spl4_34 | ~spl4_49),
  inference(avatar_split_clause,[],[f9470,f813,f750])).
fof(f9470,plain,(
  e3 != op(e3,e3) | ~spl4_49),
  inference(backward_demodulation,[],[f815,f438])).
fof(f9463,plain,(
  ~spl4_2 | ~spl4_52),
  inference(avatar_split_clause,[],[f9455,f826,f616])).
fof(f9455,plain,(
  e1 != op(e4,e4) | ~spl4_52),
  inference(backward_demodulation,[],[f828,f404])).
fof(f9451,plain,(
  ~spl4_54 | ~spl4_59),
  inference(avatar_split_clause,[],[f9448,f855,f834])).
fof(f9448,plain,(
  e3 != op(e2,e4) | ~spl4_59),
  inference(backward_demodulation,[],[f857,f435])).
fof(f857,plain,(
  e3 = op(e2,e3) | ~spl4_59),
  inference(avatar_component_clause,[],[f855])).
fof(f9450,plain,(
  ~spl4_34 | ~spl4_59),
  inference(avatar_split_clause,[],[f9446,f855,f750])).
fof(f9446,plain,(
  e3 != op(e3,e3) | ~spl4_59),
  inference(backward_demodulation,[],[f857,f393])).
fof(f393,plain,(
  op(e2,e3) != op(e3,e3)),
  inference(cnf_transformation,[],[f4])).
fof(f9441,plain,(
  ~spl4_40 | ~spl4_65),
  inference(avatar_split_clause,[],[f9429,f880,f775])).
fof(f9429,plain,(
  e4 != op(e3,e2) | ~spl4_65),
  inference(backward_demodulation,[],[f882,f383])).
fof(f9426,plain,(
  ~spl4_16 | ~spl4_66),
  inference(avatar_split_clause,[],[f9416,f885,f675])).
fof(f9416,plain,(
  e0 != op(e4,e1) | ~spl4_66),
  inference(backward_demodulation,[],[f887,f374])).
fof(f374,plain,(
  op(e2,e1) != op(e4,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f9415,plain,(
  spl4_66 | ~spl4_73 | ~spl4_112),
  inference(avatar_split_clause,[],[f9412,f1078,f914,f885])).
fof(f9412,plain,(
  e0 = op(e2,e1) | (~spl4_73 | ~spl4_112)),
  inference(backward_demodulation,[],[f916,f8973])).
fof(f916,plain,(
  e2 = op(e2,e0) | ~spl4_73),
  inference(avatar_component_clause,[],[f914])).
fof(f9414,plain,(
  ~spl4_58 | ~spl4_73),
  inference(avatar_split_clause,[],[f9408,f914,f851])).
fof(f9408,plain,(
  e2 != op(e2,e3) | ~spl4_73),
  inference(backward_demodulation,[],[f916,f428])).
fof(f9392,plain,(
  ~spl4_3 | ~spl4_85 | ~spl4_94 | spl4_157),
  inference(avatar_split_clause,[],[f9387,f1353,f1002,f964,f620])).
fof(f1353,plain,(
  spl4_157 <=> e2 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_157])])).
fof(f9387,plain,(
  e2 != op(e4,e4) | (~spl4_85 | ~spl4_94 | spl4_157)),
  inference(backward_demodulation,[],[f966,f9362])).
fof(f9362,plain,(
  e2 != op(op(e1,e3),op(e1,e3)) | (~spl4_94 | spl4_157)),
  inference(backward_demodulation,[],[f1004,f1355])).
fof(f1355,plain,(
  e2 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | spl4_157),
  inference(avatar_component_clause,[],[f1353])).
fof(f9391,plain,(
  spl4_52 | ~spl4_43 | ~spl4_85),
  inference(avatar_split_clause,[],[f9385,f964,f788,f826])).
fof(f9385,plain,(
  e1 = op(e2,e4) | (~spl4_43 | ~spl4_85)),
  inference(backward_demodulation,[],[f966,f9253])).
fof(f9253,plain,(
  e1 = op(e2,op(e1,e3)) | ~spl4_43),
  inference(forward_demodulation,[],[f339,f790])).
fof(f9390,plain,(
  spl4_14 | ~spl4_43 | ~spl4_85),
  inference(avatar_split_clause,[],[f9384,f964,f788,f666])).
fof(f9384,plain,(
  e3 = op(e4,e2) | (~spl4_43 | ~spl4_85)),
  inference(backward_demodulation,[],[f966,f9252])).
fof(f9252,plain,(
  e3 = op(op(e1,e3),e2) | ~spl4_43),
  inference(forward_demodulation,[],[f347,f790])).
fof(f9389,plain,(
  ~spl4_10 | ~spl4_85),
  inference(avatar_split_clause,[],[f9382,f964,f649])).
fof(f9382,plain,(
  e4 != op(e4,e3) | ~spl4_85),
  inference(backward_demodulation,[],[f966,f392])).
fof(f9379,plain,(
  ~spl4_78 | ~spl4_88),
  inference(avatar_split_clause,[],[f9375,f977,f935])).
fof(f9375,plain,(
  e2 != op(e1,e4) | ~spl4_88),
  inference(backward_demodulation,[],[f979,f424])).
fof(f9378,plain,(
  ~spl4_13 | ~spl4_88),
  inference(avatar_split_clause,[],[f9373,f977,f662])).
fof(f9373,plain,(
  e2 != op(e4,e2) | ~spl4_88),
  inference(backward_demodulation,[],[f979,f382])).
fof(f9368,plain,(
  ~spl4_85 | ~spl4_94 | spl4_188),
  inference(avatar_split_clause,[],[f9365,f1559,f1002,f964])).
fof(f9365,plain,(
  e4 != op(e1,e3) | (~spl4_94 | spl4_188)),
  inference(backward_demodulation,[],[f1004,f1561])).
fof(f9367,plain,(
  ~spl4_89 | ~spl4_94),
  inference(avatar_split_clause,[],[f9356,f1002,f981])).
fof(f9356,plain,(
  e3 != op(e1,e2) | ~spl4_94),
  inference(backward_demodulation,[],[f1004,f420])).
fof(f9352,plain,(
  spl4_17 | ~spl4_97 | ~spl4_120),
  inference(avatar_split_clause,[],[f9348,f1111,f1015,f679])).
fof(f1015,plain,(
  spl4_97 <=> e1 = op(e1,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_97])])).
fof(f9348,plain,(
  e1 = op(e4,e1) | (~spl4_97 | ~spl4_120)),
  inference(backward_demodulation,[],[f1017,f9311])).
fof(f9311,plain,(
  e1 = op(e4,op(e1,e0)) | ~spl4_120),
  inference(backward_demodulation,[],[f1113,f336])).
fof(f1017,plain,(
  e1 = op(e1,e0) | ~spl4_97),
  inference(avatar_component_clause,[],[f1015])).
fof(f9351,plain,(
  spl4_76 | ~spl4_97 | ~spl4_120),
  inference(avatar_split_clause,[],[f9347,f1111,f1015,f927])).
fof(f9347,plain,(
  e0 = op(e1,e4) | (~spl4_97 | ~spl4_120)),
  inference(backward_demodulation,[],[f1017,f9310])).
fof(f9310,plain,(
  e0 = op(op(e1,e0),e4) | ~spl4_120),
  inference(backward_demodulation,[],[f1113,f332])).
fof(f9350,plain,(
  ~spl4_82 | ~spl4_97),
  inference(avatar_split_clause,[],[f9342,f1015,f952])).
fof(f9342,plain,(
  e1 != op(e1,e3) | ~spl4_97),
  inference(backward_demodulation,[],[f1017,f418])).
fof(f9349,plain,(
  ~spl4_47 | ~spl4_97),
  inference(avatar_split_clause,[],[f9338,f1015,f805])).
fof(f9338,plain,(
  e1 != op(e3,e0) | ~spl4_97),
  inference(backward_demodulation,[],[f1017,f361])).
fof(f9318,plain,(
  ~spl4_105 | ~spl4_120),
  inference(avatar_split_clause,[],[f9316,f1111,f1048])).
fof(f9316,plain,(
  e4 != op(e0,e4) | ~spl4_120),
  inference(backward_demodulation,[],[f1113,f412])).
fof(f9317,plain,(
  ~spl4_95 | ~spl4_120),
  inference(avatar_split_clause,[],[f9312,f1111,f1006])).
fof(f9312,plain,(
  e4 != op(e1,e1) | ~spl4_120),
  inference(backward_demodulation,[],[f1113,f366])).
fof(f9293,plain,(
  ~spl4_106 | ~spl4_121),
  inference(avatar_split_clause,[],[f9276,f1116,f1053])).
fof(f1116,plain,(
  spl4_121 <=> e0 = op(e0,e0)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_121])])).
fof(f9276,plain,(
  e0 != op(e0,e3) | ~spl4_121),
  inference(backward_demodulation,[],[f1118,f408])).
fof(f1118,plain,(
  e0 = op(e0,e0) | ~spl4_121),
  inference(avatar_component_clause,[],[f1116])).
fof(f9292,plain,(
  ~spl4_96 | ~spl4_121),
  inference(avatar_split_clause,[],[f9271,f1116,f1011])).
fof(f9271,plain,(
  e0 != op(e1,e0) | ~spl4_121),
  inference(backward_demodulation,[],[f1118,f356])).
fof(f9268,plain,(
  ~spl4_61 | spl4_167 | ~spl4_175),
  inference(avatar_split_clause,[],[f9267,f1457,f1410,f864])).
fof(f9267,plain,(
  e0 != op(e2,e2) | (spl4_167 | ~spl4_175)),
  inference(forward_demodulation,[],[f1412,f1458])).
fof(f9248,plain,(
  ~spl4_28 | ~spl4_43),
  inference(avatar_split_clause,[],[f9247,f788,f725])).
fof(f9247,plain,(
  e2 != op(e3,e4) | ~spl4_43),
  inference(forward_demodulation,[],[f442,f790])).
fof(f9238,plain,(
  ~spl4_2 | ~spl4_16 | ~spl4_119 | spl4_136),
  inference(avatar_contradiction_clause,[],[f9237])).
fof(f9237,plain,(
  $false | (~spl4_2 | ~spl4_16 | ~spl4_119 | spl4_136)),
  inference(subsumption_resolution,[],[f9229,f1109])).
fof(f9229,plain,(
  e3 != op(e0,e1) | (~spl4_2 | ~spl4_16 | spl4_136)),
  inference(backward_demodulation,[],[f677,f9007])).
fof(f9007,plain,(
  e3 != op(op(e4,e1),e1) | (~spl4_2 | spl4_136)),
  inference(forward_demodulation,[],[f1256,f618])).
fof(f618,plain,(
  e1 = op(e4,e4) | ~spl4_2),
  inference(avatar_component_clause,[],[f616])).
fof(f1256,plain,(
  e3 != op(op(e4,op(e4,e4)),op(e4,e4)) | spl4_136),
  inference(avatar_component_clause,[],[f1254])).
fof(f1254,plain,(
  spl4_136 <=> e3 = op(op(e4,op(e4,e4)),op(e4,e4))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_136])])).
fof(f9236,plain,(
  ~spl4_2 | ~spl4_16 | ~spl4_123 | spl4_137),
  inference(avatar_contradiction_clause,[],[f9235])).
fof(f9235,plain,(
  $false | (~spl4_2 | ~spl4_16 | ~spl4_123 | spl4_137)),
  inference(subsumption_resolution,[],[f9228,f1126])).
fof(f9228,plain,(
  op(e0,e0) != e2 | (~spl4_2 | ~spl4_16 | spl4_137)),
  inference(backward_demodulation,[],[f677,f9006])).
fof(f9006,plain,(
  e2 != op(op(e4,e1),op(e4,e1)) | (~spl4_2 | spl4_137)),
  inference(forward_demodulation,[],[f1260,f618])).
fof(f1260,plain,(
  e2 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | spl4_137),
  inference(avatar_component_clause,[],[f1258])).
fof(f1258,plain,(
  spl4_137 <=> e2 = op(op(e4,op(e4,e4)),op(e4,op(e4,e4)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_137])])).
fof(f9207,plain,(
  ~spl4_9 | ~spl4_34),
  inference(avatar_split_clause,[],[f9199,f750,f645])).
fof(f9199,plain,(
  e3 != op(e4,e3) | ~spl4_34),
  inference(backward_demodulation,[],[f752,f395])).
fof(f9197,plain,(
  ~spl4_33 | ~spl4_43),
  inference(avatar_split_clause,[],[f9191,f788,f746])).
fof(f9191,plain,(
  e2 != op(e3,e3) | ~spl4_43),
  inference(backward_demodulation,[],[f790,f441])).
fof(f9196,plain,(
  ~spl4_18 | ~spl4_43),
  inference(avatar_split_clause,[],[f9190,f788,f683])).
fof(f9190,plain,(
  e2 != op(e4,e1) | ~spl4_43),
  inference(backward_demodulation,[],[f790,f375])).
fof(f9189,plain,(
  ~spl4_42 | ~spl4_47),
  inference(avatar_split_clause,[],[f9182,f805,f784])).
fof(f9182,plain,(
  e1 != op(e3,e1) | ~spl4_47),
  inference(backward_demodulation,[],[f807,f436])).
fof(f9177,plain,(
  spl4_13 | ~spl4_40 | ~spl4_58),
  inference(avatar_split_clause,[],[f9174,f851,f775,f662])).
fof(f9174,plain,(
  e2 = op(e4,e2) | (~spl4_40 | ~spl4_58)),
  inference(backward_demodulation,[],[f853,f8971])).
fof(f8971,plain,(
  e2 = op(e4,op(e2,e3)) | ~spl4_40),
  inference(forward_demodulation,[],[f344,f777])).
fof(f9175,plain,(
  ~spl4_33 | ~spl4_58),
  inference(avatar_split_clause,[],[f9168,f851,f746])).
fof(f9168,plain,(
  e2 != op(e3,e3) | ~spl4_58),
  inference(backward_demodulation,[],[f853,f393])).
fof(f9164,plain,(
  ~spl4_61 | spl4_72 | ~spl4_160),
  inference(avatar_contradiction_clause,[],[f9163])).
fof(f9163,plain,(
  $false | (~spl4_61 | spl4_72 | ~spl4_160)),
  inference(subsumption_resolution,[],[f9158,f911])).
fof(f911,plain,(
  e1 != op(e2,e0) | spl4_72),
  inference(avatar_component_clause,[],[f910])).
fof(f9158,plain,(
  e1 = op(e2,e0) | (~spl4_61 | ~spl4_160)),
  inference(backward_demodulation,[],[f866,f1376])).
fof(f9162,plain,(
  ~spl4_11 | ~spl4_61),
  inference(avatar_split_clause,[],[f9150,f864,f654])).
fof(f9150,plain,(
  e0 != op(e4,e2) | ~spl4_61),
  inference(backward_demodulation,[],[f866,f384])).
fof(f9148,plain,(
  spl4_43 | ~spl4_67 | ~spl4_89),
  inference(avatar_split_clause,[],[f9145,f981,f889,f788])).
fof(f889,plain,(
  spl4_67 <=> e1 = op(e2,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_67])])).
fof(f9145,plain,(
  e2 = op(e3,e1) | (~spl4_67 | ~spl4_89)),
  inference(backward_demodulation,[],[f891,f9094])).
fof(f9094,plain,(
  e2 = op(e3,op(e2,e1)) | ~spl4_89),
  inference(backward_demodulation,[],[f983,f342])).
fof(f891,plain,(
  e1 = op(e2,e1) | ~spl4_67),
  inference(avatar_component_clause,[],[f889])).
fof(f9147,plain,(
  ~spl4_57 | ~spl4_67),
  inference(avatar_split_clause,[],[f9139,f889,f847])).
fof(f9139,plain,(
  e1 != op(e2,e3) | ~spl4_67),
  inference(backward_demodulation,[],[f891,f431])).
fof(f9146,plain,(
  ~spl4_42 | ~spl4_67),
  inference(avatar_split_clause,[],[f9136,f889,f784])).
fof(f9136,plain,(
  e1 != op(e3,e1) | ~spl4_67),
  inference(backward_demodulation,[],[f891,f373])).
fof(f9135,plain,(
  spl4_16 | ~spl4_75 | ~spl4_112),
  inference(avatar_split_clause,[],[f9131,f1078,f922,f675])).
fof(f9131,plain,(
  e0 = op(e4,e1) | (~spl4_75 | ~spl4_112)),
  inference(backward_demodulation,[],[f924,f8973])).
fof(f9133,plain,(
  ~spl4_25 | ~spl4_75),
  inference(avatar_split_clause,[],[f9126,f922,f712])).
fof(f9126,plain,(
  e4 != op(e4,e0) | ~spl4_75),
  inference(backward_demodulation,[],[f924,f364])).
fof(f9111,plain,(
  ~spl4_57 | ~spl4_82),
  inference(avatar_split_clause,[],[f9104,f952,f847])).
fof(f9104,plain,(
  e1 != op(e2,e3) | ~spl4_82),
  inference(backward_demodulation,[],[f954,f390])).
fof(f9091,plain,(
  ~spl4_81 | ~spl4_96),
  inference(avatar_split_clause,[],[f9083,f1011,f948])).
fof(f9083,plain,(
  e0 != op(e1,e3) | ~spl4_96),
  inference(backward_demodulation,[],[f1013,f418])).
fof(f9073,plain,(
  ~spl4_81 | ~spl4_106),
  inference(avatar_split_clause,[],[f9064,f1053,f948])).
fof(f9064,plain,(
  e0 != op(e1,e3) | ~spl4_106),
  inference(backward_demodulation,[],[f1055,f386])).
fof(f9060,plain,(
  ~spl4_89 | ~spl4_123 | ~spl4_164 | spl4_165),
  inference(avatar_split_clause,[],[f9053,f1401,f1397,f1124,f981])).
fof(f9053,plain,(
  e3 != op(e1,e2) | (~spl4_123 | ~spl4_164 | spl4_165)),
  inference(backward_demodulation,[],[f1126,f9024])).
fof(f9024,plain,(
  e3 != op(e1,op(e0,e0)) | (~spl4_164 | spl4_165)),
  inference(backward_demodulation,[],[f1398,f1403])).
fof(f9059,plain,(
  ~spl4_103 | ~spl4_123),
  inference(avatar_split_clause,[],[f9049,f1124,f1040])).
fof(f9049,plain,(
  e2 != op(e0,e4) | ~spl4_123),
  inference(backward_demodulation,[],[f1126,f409])).
fof(f9058,plain,(
  ~spl4_73 | ~spl4_123),
  inference(avatar_split_clause,[],[f9045,f1124,f914])).
fof(f9045,plain,(
  e2 != op(e2,e0) | ~spl4_123),
  inference(backward_demodulation,[],[f1126,f357])).
fof(f9023,plain,(
  spl4_78 | ~spl4_95 | ~spl4_176),
  inference(avatar_split_clause,[],[f9022,f1463,f1006,f935])).
fof(f1463,plain,(
  spl4_176 <=> e2 = op(e1,op(e1,e1))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_176])])).
fof(f9022,plain,(
  e2 = op(e1,e4) | (~spl4_95 | ~spl4_176)),
  inference(forward_demodulation,[],[f1464,f1008])).
fof(f1464,plain,(
  e2 = op(e1,op(e1,e1)) | ~spl4_176),
  inference(avatar_component_clause,[],[f1463])).
fof(f9018,plain,(
  ~spl4_79 | ~spl4_95 | spl4_184),
  inference(avatar_split_clause,[],[f9017,f1519,f1006,f939])).
fof(f9017,plain,(
  e3 != op(e1,e4) | (~spl4_95 | spl4_184)),
  inference(forward_demodulation,[],[f1521,f1008])).
fof(f8983,plain,(
  ~spl4_104 | ~spl4_119),
  inference(avatar_split_clause,[],[f8982,f1107,f1044])).
fof(f8982,plain,(
  e3 != op(e0,e4) | ~spl4_119),
  inference(forward_demodulation,[],[f412,f1109])).
fof(f8981,plain,(
  ~spl4_100 | ~spl4_95),
  inference(avatar_split_clause,[],[f8980,f1006,f1027])).
fof(f8980,plain,(
  e4 != op(e1,e0) | ~spl4_95),
  inference(forward_demodulation,[],[f416,f1008])).
fof(f8959,plain,(
  ~spl4_16 | ~spl4_2 | spl4_130),
  inference(avatar_split_clause,[],[f8941,f1228,f616,f675])).
fof(f8941,plain,(
  e0 != op(e4,e1) | (~spl4_2 | spl4_130)),
  inference(backward_demodulation,[],[f618,f1230])).
fof(f8957,plain,(
  ~spl4_15 | ~spl4_40),
  inference(avatar_split_clause,[],[f8956,f775,f670])).
fof(f8956,plain,(
  e4 != op(e4,e2) | ~spl4_40),
  inference(forward_demodulation,[],[f385,f777])).
fof(f8955,plain,(
  ~spl4_12 | ~spl4_112),
  inference(avatar_split_clause,[],[f8954,f1078,f658])).
fof(f8954,plain,(
  e1 != op(e4,e2) | ~spl4_112),
  inference(forward_demodulation,[],[f379,f1080])).
fof(f8951,plain,(
  ~spl4_2 | ~spl4_66 | spl4_161 | ~spl4_174),
  inference(avatar_contradiction_clause,[],[f8950])).
fof(f8950,plain,(
  $false | (~spl4_2 | ~spl4_66 | spl4_161 | ~spl4_174)),
  inference(subsumption_resolution,[],[f8946,f887])).
fof(f8946,plain,(
  e0 != op(e2,e1) | (~spl4_2 | spl4_161 | ~spl4_174)),
  inference(backward_demodulation,[],[f618,f8460])).
fof(f8460,plain,(
  e0 != op(e2,op(e4,e4)) | (spl4_161 | ~spl4_174)),
  inference(backward_demodulation,[],[f1453,f1384])).
fof(f1384,plain,(
  e0 != op(op(e4,op(e4,e4)),op(e4,e4)) | spl4_161),
  inference(avatar_component_clause,[],[f1382])).
fof(f1382,plain,(
  spl4_161 <=> e0 = op(op(e4,op(e4,e4)),op(e4,e4))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_161])])).
fof(f1453,plain,(
  e2 = op(e4,op(e4,e4)) | ~spl4_174),
  inference(avatar_component_clause,[],[f1452])).
fof(f1452,plain,(
  spl4_174 <=> e2 = op(e4,op(e4,e4))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_174])])).
fof(f8936,plain,(
  ~spl4_4 | ~spl4_9),
  inference(avatar_split_clause,[],[f8932,f645,f624])).
fof(f8932,plain,(
  e3 != op(e4,e4) | ~spl4_9),
  inference(backward_demodulation,[],[f647,f455])).
fof(f455,plain,(
  op(e4,e3) != op(e4,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f647,plain,(
  e3 = op(e4,e3) | ~spl4_9),
  inference(avatar_component_clause,[],[f645])).
fof(f8931,plain,(
  ~spl4_8 | ~spl4_18),
  inference(avatar_split_clause,[],[f8925,f683,f641])).
fof(f8925,plain,(
  e2 != op(e4,e3) | ~spl4_18),
  inference(backward_demodulation,[],[f685,f451])).
fof(f8914,plain,(
  ~spl4_33 | spl4_37 | ~spl4_159),
  inference(avatar_contradiction_clause,[],[f8913])).
fof(f8913,plain,(
  $false | (~spl4_33 | spl4_37 | ~spl4_159)),
  inference(subsumption_resolution,[],[f8907,f764])).
fof(f764,plain,(
  e1 != op(e3,e2) | spl4_37),
  inference(avatar_component_clause,[],[f763])).
fof(f8907,plain,(
  e1 = op(e3,e2) | (~spl4_33 | ~spl4_159)),
  inference(backward_demodulation,[],[f748,f1370])).
fof(f8912,plain,(
  ~spl4_8 | ~spl4_33),
  inference(avatar_split_clause,[],[f8903,f746,f641])).
fof(f8903,plain,(
  e2 != op(e4,e3) | ~spl4_33),
  inference(backward_demodulation,[],[f748,f395])).
fof(f8901,plain,(
  spl4_18 | ~spl4_40 | ~spl4_57),
  inference(avatar_split_clause,[],[f8898,f847,f775,f683])).
fof(f8898,plain,(
  e2 = op(e4,e1) | (~spl4_40 | ~spl4_57)),
  inference(backward_demodulation,[],[f777,f8878])).
fof(f8878,plain,(
  e2 = op(op(e3,e2),e1) | ~spl4_57),
  inference(backward_demodulation,[],[f849,f344])).
fof(f8876,plain,(
  spl4_57 | ~spl4_64 | ~spl4_160),
  inference(avatar_split_clause,[],[f8872,f1375,f876,f847])).
fof(f8872,plain,(
  e1 = op(e2,e3) | (~spl4_64 | ~spl4_160)),
  inference(backward_demodulation,[],[f878,f1376])).
fof(f8875,plain,(
  ~spl4_56 | ~spl4_64 | spl4_138),
  inference(avatar_split_clause,[],[f8871,f1263,f876,f843])).
fof(f8871,plain,(
  e0 != op(e2,e3) | (~spl4_64 | spl4_138)),
  inference(backward_demodulation,[],[f878,f1265])).
fof(f1265,plain,(
  e0 != op(e2,op(e2,e2)) | spl4_138),
  inference(avatar_component_clause,[],[f1263])).
fof(f8865,plain,(
  ~spl4_41 | ~spl4_66),
  inference(avatar_split_clause,[],[f8855,f885,f780])).
fof(f780,plain,(
  spl4_41 <=> e0 = op(e3,e1)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_41])])).
fof(f8855,plain,(
  e0 != op(e3,e1) | ~spl4_66),
  inference(backward_demodulation,[],[f887,f373])).
fof(f8847,plain,(
  ~spl4_56 | ~spl4_81),
  inference(avatar_split_clause,[],[f8839,f948,f843])).
fof(f8839,plain,(
  e0 != op(e2,e3) | ~spl4_81),
  inference(backward_demodulation,[],[f950,f390])).
fof(f8836,plain,(
  ~spl4_38 | ~spl4_88),
  inference(avatar_split_clause,[],[f8831,f977,f767])).
fof(f8831,plain,(
  e2 != op(e3,e2) | ~spl4_88),
  inference(backward_demodulation,[],[f979,f381])).
fof(f8825,plain,(
  ~spl4_76 | ~spl4_95 | spl4_148),
  inference(avatar_split_clause,[],[f8816,f1309,f1006,f927])).
fof(f8816,plain,(
  e0 != op(e1,e4) | (~spl4_95 | spl4_148)),
  inference(backward_demodulation,[],[f1008,f1311])).
fof(f8824,plain,(
  ~spl4_90 | ~spl4_95),
  inference(avatar_split_clause,[],[f8813,f1006,f985])).
fof(f8813,plain,(
  e4 != op(e1,e2) | ~spl4_95),
  inference(backward_demodulation,[],[f1008,f420])).
fof(f8808,plain,(
  spl4_26 | ~spl4_49 | ~spl4_110),
  inference(avatar_split_clause,[],[f8805,f1069,f813,f717])).
fof(f8805,plain,(
  e0 = op(e3,e4) | (~spl4_49 | ~spl4_110)),
  inference(backward_demodulation,[],[f1071,f8741])).
fof(f8741,plain,(
  e0 = op(e3,op(e0,e3)) | ~spl4_49),
  inference(forward_demodulation,[],[f334,f815])).
fof(f8806,plain,(
  ~spl4_35 | ~spl4_110),
  inference(avatar_split_clause,[],[f8802,f1069,f754])).
fof(f8802,plain,(
  e4 != op(e3,e3) | ~spl4_110),
  inference(backward_demodulation,[],[f1071,f388])).
fof(f8798,plain,(
  spl4_88 | ~spl4_73 | ~spl4_112),
  inference(avatar_split_clause,[],[f8794,f1078,f914,f977])).
fof(f8794,plain,(
  e2 = op(e1,e2) | (~spl4_73 | ~spl4_112)),
  inference(backward_demodulation,[],[f1080,f8744])).
fof(f8744,plain,(
  e2 = op(op(e0,e2),e2) | ~spl4_73),
  inference(forward_demodulation,[],[f341,f916])).
fof(f8796,plain,(
  ~spl4_62 | ~spl4_112),
  inference(avatar_split_clause,[],[f8788,f1078,f868])).
fof(f8788,plain,(
  e1 != op(e2,e2) | ~spl4_112),
  inference(backward_demodulation,[],[f1080,f377])).
fof(f377,plain,(
  op(e0,e2) != op(e2,e2)),
  inference(cnf_transformation,[],[f4])).
fof(f8786,plain,(
  spl4_81 | ~spl4_97 | ~spl4_119),
  inference(avatar_split_clause,[],[f8782,f1107,f1015,f948])).
fof(f8782,plain,(
  e0 = op(e1,e3) | (~spl4_97 | ~spl4_119)),
  inference(backward_demodulation,[],[f1109,f8748])).
fof(f8748,plain,(
  e0 = op(e1,op(e0,e1)) | ~spl4_97),
  inference(forward_demodulation,[],[f332,f1017])).
fof(f8785,plain,(
  spl4_42 | ~spl4_97 | ~spl4_119),
  inference(avatar_split_clause,[],[f8781,f1107,f1015,f784])).
fof(f8781,plain,(
  e1 = op(e3,e1) | (~spl4_97 | ~spl4_119)),
  inference(backward_demodulation,[],[f1109,f8747])).
fof(f8747,plain,(
  e1 = op(op(e0,e1),e1) | ~spl4_97),
  inference(forward_demodulation,[],[f336,f1017])).
fof(f8783,plain,(
  ~spl4_69 | ~spl4_119),
  inference(avatar_split_clause,[],[f8775,f1107,f897])).
fof(f8775,plain,(
  e3 != op(e2,e1) | ~spl4_119),
  inference(backward_demodulation,[],[f1109,f367])).
fof(f367,plain,(
  op(e0,e1) != op(e2,e1)),
  inference(cnf_transformation,[],[f4])).
fof(f8773,plain,(
  spl4_64 | ~spl4_132 | ~spl4_174),
  inference(avatar_split_clause,[],[f8772,f1452,f1236,f876])).
fof(f1236,plain,(
  spl4_132 <=> e3 = op(op(e4,op(e4,e4)),op(e4,op(e4,e4)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_132])])).
fof(f8772,plain,(
  e3 = op(e2,e2) | (~spl4_132 | ~spl4_174)),
  inference(forward_demodulation,[],[f1237,f1453])).
fof(f1237,plain,(
  e3 = op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | ~spl4_132),
  inference(avatar_component_clause,[],[f1236])).
fof(f8767,plain,(
  ~spl4_95 | spl4_140 | ~spl4_160),
  inference(avatar_split_clause,[],[f8766,f1375,f1271,f1006])).
fof(f8766,plain,(
  e4 != op(e1,e1) | (spl4_140 | ~spl4_160)),
  inference(backward_demodulation,[],[f1376,f1273])).
fof(f8751,plain,(
  ~spl4_159 | ~spl4_186),
  inference(avatar_contradiction_clause,[],[f8750])).
fof(f8750,plain,(
  $false | (~spl4_159 | ~spl4_186)),
  inference(subsumption_resolution,[],[f8749,f462])).
fof(f8749,plain,(
  e1 = e4 | (~spl4_159 | ~spl4_186)),
  inference(forward_demodulation,[],[f1549,f1370])).
fof(f8725,plain,(
  ~spl4_63 | ~spl4_73),
  inference(avatar_split_clause,[],[f8724,f914,f872])).
fof(f8724,plain,(
  e2 != op(e2,e2) | ~spl4_73),
  inference(forward_demodulation,[],[f427,f916])).
fof(f8723,plain,(
  ~spl4_63 | spl4_137 | ~spl4_174),
  inference(avatar_split_clause,[],[f8459,f1452,f1258,f872])).
fof(f8459,plain,(
  e2 != op(e2,e2) | (spl4_137 | ~spl4_174)),
  inference(backward_demodulation,[],[f1453,f1260])).
fof(f8720,plain,(
  ~spl4_39 | ~spl4_49),
  inference(avatar_split_clause,[],[f8719,f813,f771])).
fof(f8719,plain,(
  e3 != op(e3,e2) | ~spl4_49),
  inference(forward_demodulation,[],[f437,f815])).
fof(f8680,plain,(
  ~spl4_5 | ~spl4_25),
  inference(avatar_split_clause,[],[f8677,f712,f628])).
fof(f8677,plain,(
  e4 != op(e4,e4) | ~spl4_25),
  inference(backward_demodulation,[],[f714,f449])).
fof(f8649,plain,(
  ~spl4_31 | ~spl4_41),
  inference(avatar_split_clause,[],[f8640,f780,f738])).
fof(f8640,plain,(
  e0 != op(e3,e3) | ~spl4_41),
  inference(backward_demodulation,[],[f782,f441])).
fof(f782,plain,(
  e0 = op(e3,e1) | ~spl4_41),
  inference(avatar_component_clause,[],[f780])).
fof(f8648,plain,(
  ~spl4_16 | ~spl4_41),
  inference(avatar_split_clause,[],[f8639,f780,f675])).
fof(f8639,plain,(
  e0 != op(e4,e1) | ~spl4_41),
  inference(backward_demodulation,[],[f782,f375])).
fof(f8638,plain,(
  spl4_41 | ~spl4_49 | ~spl4_107),
  inference(avatar_split_clause,[],[f8635,f1057,f813,f780])).
fof(f8635,plain,(
  e0 = op(e3,e1) | (~spl4_49 | ~spl4_107)),
  inference(backward_demodulation,[],[f815,f8548])).
fof(f8548,plain,(
  e0 = op(op(e3,e0),e1) | ~spl4_107),
  inference(backward_demodulation,[],[f1059,f334])).
fof(f8637,plain,(
  ~spl4_29 | ~spl4_49),
  inference(avatar_split_clause,[],[f8633,f813,f729])).
fof(f8633,plain,(
  e3 != op(e3,e4) | ~spl4_49),
  inference(backward_demodulation,[],[f815,f439])).
fof(f8629,plain,(
  ~spl4_5 | ~spl4_55),
  inference(avatar_split_clause,[],[f8628,f838,f628])).
fof(f8628,plain,(
  e4 != op(e4,e4) | ~spl4_55),
  inference(backward_demodulation,[],[f840,f404])).
fof(f8624,plain,(
  ~spl4_31 | ~spl4_56),
  inference(avatar_split_clause,[],[f8615,f843,f738])).
fof(f8615,plain,(
  e0 != op(e3,e3) | ~spl4_56),
  inference(backward_demodulation,[],[f845,f393])).
fof(f8614,plain,(
  spl4_56 | ~spl4_73 | ~spl4_114),
  inference(avatar_split_clause,[],[f8610,f1086,f914,f843])).
fof(f8610,plain,(
  e0 = op(e2,e3) | (~spl4_73 | ~spl4_114)),
  inference(backward_demodulation,[],[f916,f8540])).
fof(f8540,plain,(
  e0 = op(op(e2,e0),e3) | ~spl4_114),
  inference(backward_demodulation,[],[f1088,f333])).
fof(f8591,plain,(
  spl4_27 | ~spl4_69 | ~spl4_90),
  inference(avatar_split_clause,[],[f8586,f985,f897,f721])).
fof(f8586,plain,(
  e1 = op(e3,e4) | (~spl4_69 | ~spl4_90)),
  inference(backward_demodulation,[],[f987,f8446])).
fof(f8446,plain,(
  e1 = op(e3,op(e1,e2)) | ~spl4_69),
  inference(forward_demodulation,[],[f338,f899])).
fof(f8589,plain,(
  ~spl4_11 | ~spl4_90 | ~spl4_93 | spl4_178),
  inference(avatar_split_clause,[],[f8584,f1481,f998,f985,f654])).
fof(f1481,plain,(
  spl4_178 <=> e0 = op(op(e1,op(e1,e1)),op(e1,e1))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_178])])).
fof(f8584,plain,(
  e0 != op(e4,e2) | (~spl4_90 | ~spl4_93 | spl4_178)),
  inference(backward_demodulation,[],[f987,f8210])).
fof(f8210,plain,(
  e0 != op(op(e1,e2),e2) | (~spl4_93 | spl4_178)),
  inference(backward_demodulation,[],[f1000,f1483])).
fof(f1483,plain,(
  e0 != op(op(e1,op(e1,e1)),op(e1,e1)) | spl4_178),
  inference(avatar_component_clause,[],[f1481])).
fof(f8578,plain,(
  ~spl4_77 | ~spl4_97),
  inference(avatar_split_clause,[],[f8571,f1015,f931])).
fof(f8571,plain,(
  e1 != op(e1,e4) | ~spl4_97),
  inference(backward_demodulation,[],[f1017,f419])).
fof(f8547,plain,(
  ~spl4_14 | ~spl4_114),
  inference(avatar_split_clause,[],[f8543,f1086,f666])).
fof(f8543,plain,(
  e3 != op(e4,e2) | ~spl4_114),
  inference(backward_demodulation,[],[f1088,f379])).
fof(f8539,plain,(
  ~spl4_115 | ~spl4_120),
  inference(avatar_split_clause,[],[f8535,f1111,f1090])).
fof(f8535,plain,(
  e4 != op(e0,e2) | ~spl4_120),
  inference(backward_demodulation,[],[f1113,f410])).
fof(f8511,plain,(
  ~spl4_101 | ~spl4_121),
  inference(avatar_split_clause,[],[f8495,f1116,f1032])).
fof(f8495,plain,(
  e0 != op(e0,e4) | ~spl4_121),
  inference(backward_demodulation,[],[f1118,f409])).
fof(f8510,plain,(
  ~spl4_71 | ~spl4_121),
  inference(avatar_split_clause,[],[f8489,f1116,f906])).
fof(f8489,plain,(
  e0 != op(e2,e0) | ~spl4_121),
  inference(backward_demodulation,[],[f1118,f357])).
fof(f8486,plain,(
  ~spl4_35 | ~spl4_62 | ~spl4_69 | spl4_140),
  inference(avatar_split_clause,[],[f8485,f1271,f897,f868,f754])).
fof(f8485,plain,(
  e4 != op(e3,e3) | (~spl4_62 | ~spl4_69 | spl4_140)),
  inference(forward_demodulation,[],[f8484,f899])).
fof(f8484,plain,(
  e4 != op(op(e2,e1),op(e2,e1)) | (~spl4_62 | spl4_140)),
  inference(forward_demodulation,[],[f1273,f870])).
fof(f8466,plain,(
  ~spl4_31 | ~spl4_62 | ~spl4_69 | spl4_172),
  inference(avatar_split_clause,[],[f8465,f1436,f897,f868,f738])).
fof(f8465,plain,(
  e0 != op(e3,e3) | (~spl4_62 | ~spl4_69 | spl4_172)),
  inference(forward_demodulation,[],[f8464,f899])).
fof(f8464,plain,(
  e0 != op(op(e2,e1),op(e2,e1)) | (~spl4_62 | spl4_172)),
  inference(forward_demodulation,[],[f1438,f870])).
fof(f8453,plain,(
  spl4_90 | ~spl4_93 | ~spl4_188),
  inference(avatar_split_clause,[],[f8452,f1559,f998,f985])).
fof(f8452,plain,(
  e4 = op(e1,e2) | (~spl4_93 | ~spl4_188)),
  inference(forward_demodulation,[],[f1560,f1000])).
fof(f8440,plain,(
  ~spl4_88 | ~spl4_93),
  inference(avatar_split_clause,[],[f8439,f998,f977])).
fof(f8439,plain,(
  e2 != op(e1,e2) | ~spl4_93),
  inference(forward_demodulation,[],[f420,f1000])).
fof(f8423,plain,(
  ~spl4_45 | ~spl4_62 | ~spl4_69 | spl4_143),
  inference(avatar_contradiction_clause,[],[f8422])).
fof(f8422,plain,(
  $false | (~spl4_45 | ~spl4_62 | ~spl4_69 | spl4_143)),
  inference(subsumption_resolution,[],[f8421,f798])).
fof(f8421,plain,(
  e4 != op(e3,e1) | (~spl4_62 | ~spl4_69 | spl4_143)),
  inference(forward_demodulation,[],[f8420,f899])).
fof(f8420,plain,(
  e4 != op(op(e2,e1),e1) | (~spl4_62 | spl4_143)),
  inference(forward_demodulation,[],[f1287,f870])).
fof(f1287,plain,(
  e4 != op(op(e2,op(e2,e2)),op(e2,e2)) | spl4_143),
  inference(avatar_component_clause,[],[f1285])).
fof(f1285,plain,(
  spl4_143 <=> e4 = op(op(e2,op(e2,e2)),op(e2,e2))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_143])])).
fof(f8415,plain,(
  ~spl4_31 | ~spl4_47 | spl4_159),
  inference(avatar_contradiction_clause,[],[f8414])).
fof(f8414,plain,(
  $false | (~spl4_31 | ~spl4_47 | spl4_159)),
  inference(subsumption_resolution,[],[f8413,f807])).
fof(f8413,plain,(
  e1 != op(e3,e0) | (~spl4_31 | spl4_159)),
  inference(forward_demodulation,[],[f1371,f740])).
fof(f1371,plain,(
  e1 != op(e3,op(e3,e3)) | spl4_159),
  inference(avatar_component_clause,[],[f1369])).
fof(f8386,plain,(
  spl4_3 | ~spl4_5 | ~spl4_137),
  inference(avatar_contradiction_clause,[],[f8385])).
fof(f8385,plain,(
  $false | (spl4_3 | ~spl4_5 | ~spl4_137)),
  inference(subsumption_resolution,[],[f8384,f621])).
fof(f621,plain,(
  e2 != op(e4,e4) | spl4_3),
  inference(avatar_component_clause,[],[f620])).
fof(f8384,plain,(
  e2 = op(e4,e4) | (~spl4_5 | ~spl4_137)),
  inference(forward_demodulation,[],[f8378,f630])).
fof(f8378,plain,(
  e2 = op(op(e4,e4),op(e4,e4)) | (~spl4_5 | ~spl4_137)),
  inference(backward_demodulation,[],[f630,f1259])).
fof(f1259,plain,(
  e2 = op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | ~spl4_137),
  inference(avatar_component_clause,[],[f1258])).
fof(f8373,plain,(
  ~spl4_2 | ~spl4_7),
  inference(avatar_split_clause,[],[f8367,f637,f616])).
fof(f8367,plain,(
  e1 != op(e4,e4) | ~spl4_7),
  inference(backward_demodulation,[],[f639,f455])).
fof(f8335,plain,(
  ~spl4_31 | ~spl4_47 | ~spl4_93 | spl4_142),
  inference(avatar_contradiction_clause,[],[f8334])).
fof(f8334,plain,(
  $false | (~spl4_31 | ~spl4_47 | ~spl4_93 | spl4_142)),
  inference(subsumption_resolution,[],[f8333,f1000])).
fof(f8333,plain,(
  e2 != op(e1,e1) | (~spl4_31 | ~spl4_47 | spl4_142)),
  inference(forward_demodulation,[],[f8324,f807])).
fof(f8324,plain,(
  e2 != op(op(e3,e0),op(e3,e0)) | (~spl4_31 | spl4_142)),
  inference(backward_demodulation,[],[f740,f1282])).
fof(f8332,plain,(
  ~spl4_31 | ~spl4_47 | ~spl4_100 | spl4_141),
  inference(avatar_contradiction_clause,[],[f8331])).
fof(f8331,plain,(
  $false | (~spl4_31 | ~spl4_47 | ~spl4_100 | spl4_141)),
  inference(subsumption_resolution,[],[f8330,f1029])).
fof(f8330,plain,(
  e4 != op(e1,e0) | (~spl4_31 | ~spl4_47 | spl4_141)),
  inference(forward_demodulation,[],[f8323,f807])).
fof(f8323,plain,(
  e4 != op(op(e3,e0),e0) | (~spl4_31 | spl4_141)),
  inference(backward_demodulation,[],[f740,f1278])).
fof(f1278,plain,(
  e4 != op(op(e3,op(e3,e3)),op(e3,e3)) | spl4_141),
  inference(avatar_component_clause,[],[f1276])).
fof(f1276,plain,(
  spl4_141 <=> e4 = op(op(e3,op(e3,e3)),op(e3,e3))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_141])])).
fof(f8313,plain,(
  spl4_14 | ~spl4_38 | ~spl4_60),
  inference(avatar_split_clause,[],[f8310,f859,f767,f666])).
fof(f8310,plain,(
  e3 = op(e4,e2) | (~spl4_38 | ~spl4_60)),
  inference(backward_demodulation,[],[f769,f8113])).
fof(f8113,plain,(
  e3 = op(e4,op(e3,e2)) | ~spl4_60),
  inference(forward_demodulation,[],[f348,f861])).
fof(f8312,plain,(
  ~spl4_33 | ~spl4_38),
  inference(avatar_split_clause,[],[f8305,f767,f746])).
fof(f8305,plain,(
  e2 != op(e3,e3) | ~spl4_38),
  inference(backward_demodulation,[],[f769,f443])).
fof(f8311,plain,(
  ~spl4_13 | ~spl4_38),
  inference(avatar_split_clause,[],[f8304,f767,f662])).
fof(f8304,plain,(
  e2 != op(e4,e2) | ~spl4_38),
  inference(backward_demodulation,[],[f769,f385])).
fof(f8277,plain,(
  ~spl4_62 | ~spl4_69 | spl4_183),
  inference(avatar_contradiction_clause,[],[f8276])).
fof(f8276,plain,(
  $false | (~spl4_62 | ~spl4_69 | spl4_183)),
  inference(subsumption_resolution,[],[f8268,f899])).
fof(f8268,plain,(
  e3 != op(e2,e1) | (~spl4_62 | spl4_183)),
  inference(backward_demodulation,[],[f870,f1515])).
fof(f1515,plain,(
  e3 != op(e2,op(e2,e2)) | spl4_183),
  inference(avatar_component_clause,[],[f1513])).
fof(f1513,plain,(
  spl4_183 <=> e3 = op(e2,op(e2,e2))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_183])])).
fof(f8273,plain,(
  ~spl4_41 | ~spl4_62 | ~spl4_69 | spl4_170),
  inference(avatar_split_clause,[],[f8272,f1425,f897,f868,f780])).
fof(f1425,plain,(
  spl4_170 <=> e0 = op(op(e2,op(e2,e2)),op(e2,e2))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_170])])).
fof(f8272,plain,(
  e0 != op(e3,e1) | (~spl4_62 | ~spl4_69 | spl4_170)),
  inference(forward_demodulation,[],[f8266,f899])).
fof(f8266,plain,(
  e0 != op(op(e2,e1),e1) | (~spl4_62 | spl4_170)),
  inference(backward_demodulation,[],[f870,f1427])).
fof(f1427,plain,(
  e0 != op(op(e2,op(e2,e2)),op(e2,e2)) | spl4_170),
  inference(avatar_component_clause,[],[f1425])).
fof(f8253,plain,(
  ~spl4_64 | ~spl4_69),
  inference(avatar_split_clause,[],[f8247,f897,f876])).
fof(f8247,plain,(
  e3 != op(e2,e2) | ~spl4_69),
  inference(backward_demodulation,[],[f899,f430])).
fof(f8252,plain,(
  ~spl4_19 | ~spl4_69),
  inference(avatar_split_clause,[],[f8246,f897,f687])).
fof(f8246,plain,(
  e3 != op(e4,e1) | ~spl4_69),
  inference(backward_demodulation,[],[f899,f374])).
fof(f8244,plain,(
  spl4_23 | ~spl4_71 | ~spl4_115),
  inference(avatar_split_clause,[],[f8242,f1090,f906,f704])).
fof(f8242,plain,(
  e2 = op(e4,e0) | (~spl4_71 | ~spl4_115)),
  inference(backward_demodulation,[],[f908,f7294])).
fof(f7294,plain,(
  e2 = op(e4,op(e2,e0)) | ~spl4_115),
  inference(forward_demodulation,[],[f341,f1092])).
fof(f1092,plain,(
  e4 = op(e0,e2) | ~spl4_115),
  inference(avatar_component_clause,[],[f1090])).
fof(f8231,plain,(
  ~spl4_2 | ~spl4_77),
  inference(avatar_split_clause,[],[f8227,f931,f616])).
fof(f8227,plain,(
  e1 != op(e4,e4) | ~spl4_77),
  inference(backward_demodulation,[],[f933,f402])).
fof(f8198,plain,(
  spl4_77 | ~spl4_100 | ~spl4_117),
  inference(avatar_split_clause,[],[f8194,f1099,f1027,f931])).
fof(f8194,plain,(
  e1 = op(e1,e4) | (~spl4_100 | ~spl4_117)),
  inference(backward_demodulation,[],[f1029,f8149])).
fof(f8149,plain,(
  e1 = op(e1,op(e1,e0)) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f336])).
fof(f8197,plain,(
  spl4_16 | ~spl4_100 | ~spl4_117),
  inference(avatar_split_clause,[],[f8193,f1099,f1027,f675])).
fof(f8193,plain,(
  e0 = op(e4,e1) | (~spl4_100 | ~spl4_117)),
  inference(backward_demodulation,[],[f1029,f8148])).
fof(f8148,plain,(
  e0 = op(op(e1,e0),e1) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f332])).
fof(f8195,plain,(
  ~spl4_25 | ~spl4_100),
  inference(avatar_split_clause,[],[f8190,f1027,f712])).
fof(f8190,plain,(
  e4 != op(e4,e0) | ~spl4_100),
  inference(backward_demodulation,[],[f1029,f362])).
fof(f8187,plain,(
  ~spl4_51 | ~spl4_101),
  inference(avatar_split_clause,[],[f8180,f1032,f822])).
fof(f8180,plain,(
  e0 != op(e2,e4) | ~spl4_101),
  inference(backward_demodulation,[],[f1034,f397])).
fof(f8173,plain,(
  ~spl4_33 | ~spl4_108),
  inference(avatar_split_clause,[],[f8163,f1061,f746])).
fof(f8163,plain,(
  e2 != op(e3,e3) | ~spl4_108),
  inference(backward_demodulation,[],[f1063,f388])).
fof(f8160,plain,(
  ~spl4_107 | ~spl4_117),
  inference(avatar_split_clause,[],[f8154,f1099,f1057])).
fof(f8154,plain,(
  e1 != op(e0,e3) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f411])).
fof(f8159,plain,(
  ~spl4_67 | ~spl4_117),
  inference(avatar_split_clause,[],[f8151,f1099,f889])).
fof(f8151,plain,(
  e1 != op(e2,e1) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f367])).
fof(f8097,plain,(
  ~spl4_2 | ~spl4_19 | ~spl4_41 | spl4_161),
  inference(avatar_contradiction_clause,[],[f8096])).
fof(f8096,plain,(
  $false | (~spl4_2 | ~spl4_19 | ~spl4_41 | spl4_161)),
  inference(subsumption_resolution,[],[f8088,f782])).
fof(f8088,plain,(
  e0 != op(e3,e1) | (~spl4_2 | ~spl4_19 | spl4_161)),
  inference(backward_demodulation,[],[f689,f7876])).
fof(f7876,plain,(
  e0 != op(op(e4,e1),e1) | (~spl4_2 | spl4_161)),
  inference(forward_demodulation,[],[f1384,f618])).
fof(f689,plain,(
  e3 = op(e4,e1) | ~spl4_19),
  inference(avatar_component_clause,[],[f687])).
fof(f8092,plain,(
  ~spl4_14 | ~spl4_19),
  inference(avatar_split_clause,[],[f8083,f687,f666])).
fof(f8083,plain,(
  e3 != op(e4,e2) | ~spl4_19),
  inference(backward_demodulation,[],[f689,f450])).
fof(f8070,plain,(
  ~spl4_28 | ~spl4_33),
  inference(avatar_split_clause,[],[f8062,f746,f725])).
fof(f8062,plain,(
  e2 != op(e3,e4) | ~spl4_33),
  inference(backward_demodulation,[],[f748,f445])).
fof(f8045,plain,(
  spl4_19 | ~spl4_37 | ~spl4_60),
  inference(avatar_split_clause,[],[f8041,f859,f763,f687])).
fof(f8041,plain,(
  e3 = op(e4,e1) | (~spl4_37 | ~spl4_60)),
  inference(backward_demodulation,[],[f861,f7848])).
fof(f7848,plain,(
  e3 = op(op(e2,e3),e1) | ~spl4_37),
  inference(forward_demodulation,[],[f348,f765])).
fof(f8036,plain,(
  ~spl4_14 | ~spl4_64),
  inference(avatar_split_clause,[],[f8026,f876,f666])).
fof(f8026,plain,(
  e3 != op(e4,e2) | ~spl4_64),
  inference(backward_demodulation,[],[f878,f384])).
fof(f8016,plain,(
  spl4_13 | ~spl4_73 | ~spl4_115),
  inference(avatar_split_clause,[],[f8013,f1090,f914,f662])).
fof(f8013,plain,(
  e2 = op(e4,e2) | (~spl4_73 | ~spl4_115)),
  inference(backward_demodulation,[],[f916,f7294])).
fof(f8001,plain,(
  ~spl4_51 | ~spl4_78 | ~spl4_95 | spl4_178),
  inference(avatar_split_clause,[],[f7993,f1481,f1006,f935,f822])).
fof(f7993,plain,(
  e0 != op(e2,e4) | (~spl4_78 | ~spl4_95 | spl4_178)),
  inference(backward_demodulation,[],[f937,f7863])).
fof(f7863,plain,(
  e0 != op(op(e1,e4),e4) | (~spl4_95 | spl4_178)),
  inference(forward_demodulation,[],[f1483,f1008])).
fof(f7999,plain,(
  ~spl4_28 | ~spl4_78),
  inference(avatar_split_clause,[],[f7989,f935,f725])).
fof(f7989,plain,(
  e2 != op(e3,e4) | ~spl4_78),
  inference(backward_demodulation,[],[f937,f401])).
fof(f7985,plain,(
  spl4_49 | ~spl4_41 | ~spl4_84),
  inference(avatar_split_clause,[],[f7982,f960,f780,f813])).
fof(f7982,plain,(
  e3 = op(e3,e0) | (~spl4_41 | ~spl4_84)),
  inference(backward_demodulation,[],[f962,f7853])).
fof(f7853,plain,(
  e3 = op(op(e1,e3),e0) | ~spl4_41),
  inference(forward_demodulation,[],[f347,f782])).
fof(f7984,plain,(
  ~spl4_34 | ~spl4_84),
  inference(avatar_split_clause,[],[f7978,f960,f750])).
fof(f7978,plain,(
  e3 != op(e3,e3) | ~spl4_84),
  inference(backward_demodulation,[],[f962,f391])).
fof(f7975,plain,(
  ~spl4_61 | ~spl4_86),
  inference(avatar_split_clause,[],[f7967,f969,f864])).
fof(f7967,plain,(
  e0 != op(e2,e2) | ~spl4_86),
  inference(backward_demodulation,[],[f971,f380])).
fof(f7964,plain,(
  spl4_67 | ~spl4_97 | ~spl4_118),
  inference(avatar_split_clause,[],[f7960,f1103,f1015,f889])).
fof(f7960,plain,(
  e1 = op(e2,e1) | (~spl4_97 | ~spl4_118)),
  inference(backward_demodulation,[],[f1017,f7933])).
fof(f7933,plain,(
  e1 = op(e2,op(e1,e0)) | ~spl4_118),
  inference(backward_demodulation,[],[f1105,f336])).
fof(f7963,plain,(
  spl4_86 | ~spl4_97 | ~spl4_118),
  inference(avatar_split_clause,[],[f7959,f1103,f1015,f969])).
fof(f7959,plain,(
  e0 = op(e1,e2) | (~spl4_97 | ~spl4_118)),
  inference(backward_demodulation,[],[f1017,f7932])).
fof(f7932,plain,(
  e0 = op(op(e1,e0),e2) | ~spl4_118),
  inference(backward_demodulation,[],[f1105,f332])).
fof(f7961,plain,(
  ~spl4_72 | ~spl4_97),
  inference(avatar_split_clause,[],[f7950,f1015,f910])).
fof(f7950,plain,(
  e1 != op(e2,e0) | ~spl4_97),
  inference(backward_demodulation,[],[f1017,f360])).
fof(f7913,plain,(
  ~spl4_21 | ~spl4_121),
  inference(avatar_split_clause,[],[f7897,f1116,f696])).
fof(f7897,plain,(
  e0 != op(e4,e0) | ~spl4_121),
  inference(backward_demodulation,[],[f1118,f359])).
fof(f7882,plain,(
  ~spl4_95 | spl4_135 | ~spl4_159),
  inference(avatar_contradiction_clause,[],[f7881])).
fof(f7881,plain,(
  $false | (~spl4_95 | spl4_135 | ~spl4_159)),
  inference(subsumption_resolution,[],[f7879,f1008])).
fof(f7879,plain,(
  e4 != op(e1,e1) | (spl4_135 | ~spl4_159)),
  inference(backward_demodulation,[],[f1370,f1251])).
fof(f7837,plain,(
  ~spl4_18 | ~spl4_2 | spl4_174),
  inference(avatar_split_clause,[],[f7822,f1452,f616,f683])).
fof(f7822,plain,(
  e2 != op(e4,e1) | (~spl4_2 | spl4_174)),
  inference(backward_demodulation,[],[f618,f1454])).
fof(f1454,plain,(
  e2 != op(e4,op(e4,e4)) | spl4_174),
  inference(avatar_component_clause,[],[f1452])).
fof(f7831,plain,(
  ~spl4_2 | ~spl4_18 | ~spl4_61 | spl4_163),
  inference(avatar_contradiction_clause,[],[f7830])).
fof(f7830,plain,(
  $false | (~spl4_2 | ~spl4_18 | ~spl4_61 | spl4_163)),
  inference(subsumption_resolution,[],[f7829,f866])).
fof(f7829,plain,(
  e0 != op(e2,e2) | (~spl4_2 | ~spl4_18 | spl4_163)),
  inference(forward_demodulation,[],[f7821,f685])).
fof(f7821,plain,(
  e0 != op(op(e4,e1),op(e4,e1)) | (~spl4_2 | spl4_163)),
  inference(backward_demodulation,[],[f618,f1394])).
fof(f1394,plain,(
  e0 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | spl4_163),
  inference(avatar_component_clause,[],[f1392])).
fof(f1392,plain,(
  spl4_163 <=> e0 = op(op(e4,op(e4,e4)),op(e4,op(e4,e4)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_163])])).
fof(f7803,plain,(
  ~spl4_13 | ~spl4_18),
  inference(avatar_split_clause,[],[f7796,f683,f662])).
fof(f7796,plain,(
  e2 != op(e4,e2) | ~spl4_18),
  inference(backward_demodulation,[],[f685,f450])).
fof(f7795,plain,(
  ~spl4_16 | ~spl4_21),
  inference(avatar_split_clause,[],[f7785,f696,f675])).
fof(f7785,plain,(
  e0 != op(e4,e1) | ~spl4_21),
  inference(backward_demodulation,[],[f698,f446])).
fof(f698,plain,(
  e0 = op(e4,e0) | ~spl4_21),
  inference(avatar_component_clause,[],[f696])).
fof(f7766,plain,(
  ~spl4_36 | ~spl4_41),
  inference(avatar_split_clause,[],[f7757,f780,f759])).
fof(f7757,plain,(
  e0 != op(e3,e2) | ~spl4_41),
  inference(backward_demodulation,[],[f782,f440])).
fof(f7754,plain,(
  ~spl4_35 | ~spl4_50),
  inference(avatar_split_clause,[],[f7750,f817,f754])).
fof(f7750,plain,(
  e4 != op(e3,e3) | ~spl4_50),
  inference(backward_demodulation,[],[f819,f438])).
fof(f7753,plain,(
  ~spl4_25 | ~spl4_50),
  inference(avatar_split_clause,[],[f7747,f817,f712])).
fof(f7747,plain,(
  e4 != op(e4,e0) | ~spl4_50),
  inference(backward_demodulation,[],[f819,f365])).
fof(f7734,plain,(
  ~spl4_61 | ~spl4_72 | spl4_160),
  inference(avatar_contradiction_clause,[],[f7733])).
fof(f7733,plain,(
  $false | (~spl4_61 | ~spl4_72 | spl4_160)),
  inference(subsumption_resolution,[],[f7724,f912])).
fof(f7724,plain,(
  e1 != op(e2,e0) | (~spl4_61 | spl4_160)),
  inference(backward_demodulation,[],[f866,f1377])).
fof(f7731,plain,(
  ~spl4_61 | ~spl4_72 | ~spl4_99 | spl4_139),
  inference(avatar_contradiction_clause,[],[f7730])).
fof(f7730,plain,(
  $false | (~spl4_61 | ~spl4_72 | ~spl4_99 | spl4_139)),
  inference(subsumption_resolution,[],[f7729,f1025])).
fof(f7729,plain,(
  e3 != op(e1,e0) | (~spl4_61 | ~spl4_72 | spl4_139)),
  inference(forward_demodulation,[],[f7722,f912])).
fof(f7722,plain,(
  e3 != op(op(e2,e0),e0) | (~spl4_61 | spl4_139)),
  inference(backward_demodulation,[],[f866,f1269])).
fof(f7727,plain,(
  ~spl4_36 | ~spl4_61),
  inference(avatar_split_clause,[],[f7713,f864,f759])).
fof(f7713,plain,(
  e0 != op(e3,e2) | ~spl4_61),
  inference(backward_demodulation,[],[f866,f383])).
fof(f7700,plain,(
  spl4_18 | ~spl4_72 | ~spl4_115),
  inference(avatar_split_clause,[],[f7698,f1090,f910,f683])).
fof(f7698,plain,(
  e2 = op(e4,e1) | (~spl4_72 | ~spl4_115)),
  inference(backward_demodulation,[],[f912,f7294])).
fof(f7699,plain,(
  ~spl4_62 | ~spl4_72),
  inference(avatar_split_clause,[],[f7691,f910,f868])).
fof(f7691,plain,(
  e1 != op(e2,e2) | ~spl4_72),
  inference(backward_demodulation,[],[f912,f427])).
fof(f7654,plain,(
  spl4_76 | ~spl4_95 | ~spl4_148),
  inference(avatar_split_clause,[],[f7646,f1309,f1006,f927])).
fof(f7646,plain,(
  e0 = op(e1,e4) | (~spl4_95 | ~spl4_148)),
  inference(backward_demodulation,[],[f1008,f1310])).
fof(f7653,plain,(
  ~spl4_80 | ~spl4_95),
  inference(avatar_split_clause,[],[f7645,f1006,f943])).
fof(f7645,plain,(
  e4 != op(e1,e4) | ~spl4_95),
  inference(backward_demodulation,[],[f1008,f422])).
fof(f7652,plain,(
  ~spl4_70 | ~spl4_95),
  inference(avatar_split_clause,[],[f7640,f1006,f901])).
fof(f7640,plain,(
  e4 != op(e2,e1) | ~spl4_95),
  inference(backward_demodulation,[],[f1008,f370])).
fof(f7638,plain,(
  spl4_82 | ~spl4_99 | ~spl4_117),
  inference(avatar_split_clause,[],[f7634,f1099,f1023,f952])).
fof(f7634,plain,(
  e1 = op(e1,e3) | (~spl4_99 | ~spl4_117)),
  inference(backward_demodulation,[],[f1025,f7594])).
fof(f7594,plain,(
  e1 = op(e1,op(e1,e0)) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f336])).
fof(f7637,plain,(
  spl4_41 | ~spl4_99 | ~spl4_117),
  inference(avatar_split_clause,[],[f7633,f1099,f1023,f780])).
fof(f7633,plain,(
  e0 = op(e3,e1) | (~spl4_99 | ~spl4_117)),
  inference(backward_demodulation,[],[f1025,f7593])).
fof(f7593,plain,(
  e0 = op(op(e1,e0),e1) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f332])).
fof(f7636,plain,(
  ~spl4_89 | ~spl4_99),
  inference(avatar_split_clause,[],[f7629,f1023,f981])).
fof(f7629,plain,(
  e3 != op(e1,e2) | ~spl4_99),
  inference(backward_demodulation,[],[f1025,f417])).
fof(f7635,plain,(
  ~spl4_49 | ~spl4_99),
  inference(avatar_split_clause,[],[f7626,f1023,f813])).
fof(f7626,plain,(
  e3 != op(e3,e0) | ~spl4_99),
  inference(backward_demodulation,[],[f1025,f361])).
fof(f7624,plain,(
  ~spl4_4 | ~spl4_104),
  inference(avatar_split_clause,[],[f7622,f1044,f624])).
fof(f7622,plain,(
  e3 != op(e4,e4) | ~spl4_104),
  inference(backward_demodulation,[],[f1046,f399])).
fof(f7604,plain,(
  ~spl4_42 | ~spl4_117),
  inference(avatar_split_clause,[],[f7597,f1099,f784])).
fof(f7597,plain,(
  e1 != op(e3,e1) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f368])).
fof(f7570,plain,(
  ~spl4_148 | ~spl4_184),
  inference(avatar_contradiction_clause,[],[f7569])).
fof(f7569,plain,(
  $false | (~spl4_148 | ~spl4_184)),
  inference(subsumption_resolution,[],[f7568,f458])).
fof(f458,plain,(
  e0 != e3),
  inference(cnf_transformation,[],[f5])).
fof(f7568,plain,(
  e0 = e3 | (~spl4_148 | ~spl4_184)),
  inference(backward_demodulation,[],[f1310,f1520])).
fof(f7561,plain,(
  ~spl4_90 | ~spl4_115),
  inference(avatar_split_clause,[],[f7560,f1090,f985])).
fof(f7560,plain,(
  e4 != op(e1,e2) | ~spl4_115),
  inference(forward_demodulation,[],[f376,f1092])).
fof(f7551,plain,(
  ~spl4_65 | spl4_135 | ~spl4_175),
  inference(avatar_split_clause,[],[f7550,f1457,f1249,f880])).
fof(f7550,plain,(
  e4 != op(e2,e2) | (spl4_135 | ~spl4_175)),
  inference(forward_demodulation,[],[f1251,f1458])).
fof(f7545,plain,(
  ~spl4_4 | ~spl4_7 | ~spl4_81 | spl4_161),
  inference(avatar_contradiction_clause,[],[f7544])).
fof(f7544,plain,(
  $false | (~spl4_4 | ~spl4_7 | ~spl4_81 | spl4_161)),
  inference(subsumption_resolution,[],[f7540,f950])).
fof(f7540,plain,(
  e0 != op(e1,e3) | (~spl4_4 | ~spl4_7 | spl4_161)),
  inference(backward_demodulation,[],[f639,f7330])).
fof(f7330,plain,(
  e0 != op(op(e4,e3),e3) | (~spl4_4 | spl4_161)),
  inference(forward_demodulation,[],[f1384,f626])).
fof(f7535,plain,(
  ~spl4_8 | ~spl4_13),
  inference(avatar_split_clause,[],[f7530,f662,f641])).
fof(f7530,plain,(
  e2 != op(e4,e3) | ~spl4_13),
  inference(backward_demodulation,[],[f664,f453])).
fof(f7521,plain,(
  spl4_16 | ~spl4_25 | ~spl4_102),
  inference(avatar_split_clause,[],[f7518,f1036,f712,f675])).
fof(f7518,plain,(
  e0 = op(e4,e1) | (~spl4_25 | ~spl4_102)),
  inference(backward_demodulation,[],[f714,f7283])).
fof(f7283,plain,(
  e0 = op(op(e4,e0),e1) | ~spl4_102),
  inference(forward_demodulation,[],[f335,f1038])).
fof(f335,plain,(
  e0 = op(op(e4,e0),op(e0,e4))),
  inference(cnf_transformation,[],[f3])).
fof(f7480,plain,(
  ~spl4_44 | ~spl4_49),
  inference(avatar_split_clause,[],[f7474,f813,f792])).
fof(f7474,plain,(
  e3 != op(e3,e1) | ~spl4_49),
  inference(backward_demodulation,[],[f815,f436])).
fof(f7451,plain,(
  ~spl4_57 | ~spl4_62),
  inference(avatar_split_clause,[],[f7439,f868,f847])).
fof(f7439,plain,(
  e1 != op(e2,e3) | ~spl4_62),
  inference(backward_demodulation,[],[f870,f433])).
fof(f7435,plain,(
  spl4_28 | ~spl4_70 | ~spl4_89),
  inference(avatar_split_clause,[],[f7431,f981,f901,f725])).
fof(f7431,plain,(
  e2 = op(e3,e4) | (~spl4_70 | ~spl4_89)),
  inference(backward_demodulation,[],[f903,f6949])).
fof(f6949,plain,(
  e2 = op(e3,op(e2,e1)) | ~spl4_89),
  inference(forward_demodulation,[],[f342,f983])).
fof(f7434,plain,(
  spl4_7 | ~spl4_70 | ~spl4_89),
  inference(avatar_split_clause,[],[f7430,f981,f901,f637])).
fof(f7430,plain,(
  e1 = op(e4,e3) | (~spl4_70 | ~spl4_89)),
  inference(backward_demodulation,[],[f903,f6948])).
fof(f6948,plain,(
  e1 = op(op(e2,e1),e3) | ~spl4_89),
  inference(forward_demodulation,[],[f338,f983])).
fof(f7433,plain,(
  ~spl4_55 | ~spl4_70),
  inference(avatar_split_clause,[],[f7429,f901,f838])).
fof(f7429,plain,(
  e4 != op(e2,e4) | ~spl4_70),
  inference(backward_demodulation,[],[f903,f432])).
fof(f7432,plain,(
  ~spl4_20 | ~spl4_70),
  inference(avatar_split_clause,[],[f7426,f901,f691])).
fof(f7426,plain,(
  e4 != op(e4,e1) | ~spl4_70),
  inference(backward_demodulation,[],[f903,f374])).
fof(f7422,plain,(
  ~spl4_68 | ~spl4_73),
  inference(avatar_split_clause,[],[f7414,f914,f893])).
fof(f7414,plain,(
  e2 != op(e2,e1) | ~spl4_73),
  inference(backward_demodulation,[],[f916,f426])).
fof(f7406,plain,(
  ~spl4_36 | ~spl4_93 | spl4_178 | ~spl4_184),
  inference(avatar_split_clause,[],[f7401,f1519,f1481,f998,f759])).
fof(f7401,plain,(
  e0 != op(e3,e2) | (~spl4_93 | spl4_178 | ~spl4_184)),
  inference(backward_demodulation,[],[f1000,f7322])).
fof(f7322,plain,(
  e0 != op(e3,op(e1,e1)) | (spl4_178 | ~spl4_184)),
  inference(backward_demodulation,[],[f1520,f1483])).
fof(f7386,plain,(
  ~spl4_92 | ~spl4_97),
  inference(avatar_split_clause,[],[f7378,f1015,f994])).
fof(f7378,plain,(
  e1 != op(e1,e1) | ~spl4_97),
  inference(backward_demodulation,[],[f1017,f416])).
fof(f7367,plain,(
  ~spl4_109 | ~spl4_119),
  inference(avatar_split_clause,[],[f7364,f1107,f1065])).
fof(f7364,plain,(
  e3 != op(e0,e3) | ~spl4_119),
  inference(backward_demodulation,[],[f1109,f411])).
fof(f7355,plain,(
  ~spl4_121 | ~spl4_189),
  inference(avatar_contradiction_clause,[],[f7354])).
fof(f7354,plain,(
  $false | (~spl4_121 | ~spl4_189)),
  inference(subsumption_resolution,[],[f7353,f459])).
fof(f7353,plain,(
  e0 = e4 | (~spl4_121 | ~spl4_189)),
  inference(forward_demodulation,[],[f7347,f1118])).
fof(f7347,plain,(
  op(e0,e0) = e4 | (~spl4_121 | ~spl4_189)),
  inference(backward_demodulation,[],[f1118,f1570])).
fof(f7350,plain,(
  ~spl4_116 | ~spl4_121),
  inference(avatar_split_clause,[],[f7339,f1116,f1095])).
fof(f7339,plain,(
  e0 != op(e0,e1) | ~spl4_121),
  inference(backward_demodulation,[],[f1118,f406])).
fof(f7333,plain,(
  ~spl4_26 | ~spl4_35 | spl4_133),
  inference(avatar_split_clause,[],[f7332,f1241,f754,f717])).
fof(f7332,plain,(
  e0 != op(e3,e4) | (~spl4_35 | spl4_133)),
  inference(forward_demodulation,[],[f1243,f756])).
fof(f7329,plain,(
  ~spl4_8 | ~spl4_4 | spl4_174),
  inference(avatar_split_clause,[],[f7328,f1452,f624,f641])).
fof(f7328,plain,(
  e2 != op(e4,e3) | (~spl4_4 | spl4_174)),
  inference(forward_demodulation,[],[f1454,f626])).
fof(f7325,plain,(
  ~spl4_35 | spl4_150 | ~spl4_184),
  inference(avatar_contradiction_clause,[],[f7324])).
fof(f7324,plain,(
  $false | (~spl4_35 | spl4_150 | ~spl4_184)),
  inference(subsumption_resolution,[],[f7318,f756])).
fof(f7318,plain,(
  e4 != op(e3,e3) | (spl4_150 | ~spl4_184)),
  inference(backward_demodulation,[],[f1520,f1319])).
fof(f1319,plain,(
  e4 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | spl4_150),
  inference(avatar_component_clause,[],[f1317])).
fof(f1317,plain,(
  spl4_150 <=> e4 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_150])])).
fof(f7280,plain,(
  ~spl4_14 | ~spl4_89),
  inference(avatar_split_clause,[],[f7279,f981,f666])).
fof(f7279,plain,(
  e3 != op(e4,e2) | ~spl4_89),
  inference(forward_demodulation,[],[f382,f983])).
fof(f7278,plain,(
  ~spl4_15 | ~spl4_115),
  inference(avatar_split_clause,[],[f7277,f1090,f670])).
fof(f7277,plain,(
  e4 != op(e4,e2) | ~spl4_115),
  inference(forward_demodulation,[],[f379,f1092])).
fof(f7272,plain,(
  ~spl4_26 | ~spl4_35 | ~spl4_123 | spl4_142),
  inference(avatar_contradiction_clause,[],[f7271])).
fof(f7271,plain,(
  $false | (~spl4_26 | ~spl4_35 | ~spl4_123 | spl4_142)),
  inference(subsumption_resolution,[],[f7270,f1126])).
fof(f7270,plain,(
  op(e0,e0) != e2 | (~spl4_26 | ~spl4_35 | spl4_142)),
  inference(forward_demodulation,[],[f7269,f719])).
fof(f7269,plain,(
  e2 != op(op(e3,e4),op(e3,e4)) | (~spl4_35 | spl4_142)),
  inference(forward_demodulation,[],[f1282,f756])).
fof(f7209,plain,(
  ~spl4_4 | spl4_7 | ~spl4_158),
  inference(avatar_contradiction_clause,[],[f7208])).
fof(f7208,plain,(
  $false | (~spl4_4 | spl4_7 | ~spl4_158)),
  inference(subsumption_resolution,[],[f7207,f638])).
fof(f638,plain,(
  e1 != op(e4,e3) | spl4_7),
  inference(avatar_component_clause,[],[f637])).
fof(f7207,plain,(
  e1 = op(e4,e3) | (~spl4_4 | ~spl4_158)),
  inference(backward_demodulation,[],[f626,f1365])).
fof(f1365,plain,(
  e1 = op(e4,op(e4,e4)) | ~spl4_158),
  inference(avatar_component_clause,[],[f1364])).
fof(f7203,plain,(
  ~spl4_3 | ~spl4_8),
  inference(avatar_split_clause,[],[f7198,f641,f620])).
fof(f7198,plain,(
  e2 != op(e4,e4) | ~spl4_8),
  inference(backward_demodulation,[],[f643,f455])).
fof(f643,plain,(
  e2 = op(e4,e3) | ~spl4_8),
  inference(avatar_component_clause,[],[f641])).
fof(f7186,plain,(
  ~spl4_35 | ~spl4_102 | ~spl4_133 | spl4_146),
  inference(avatar_contradiction_clause,[],[f7185])).
fof(f7185,plain,(
  $false | (~spl4_35 | ~spl4_102 | ~spl4_133 | spl4_146)),
  inference(subsumption_resolution,[],[f7179,f1038])).
fof(f7179,plain,(
  e1 != op(e0,e4) | (~spl4_35 | ~spl4_133 | spl4_146)),
  inference(backward_demodulation,[],[f756,f6980])).
fof(f6980,plain,(
  e1 != op(e0,op(e3,e3)) | (~spl4_133 | spl4_146)),
  inference(backward_demodulation,[],[f1242,f1301])).
fof(f7167,plain,(
  ~spl4_29 | ~spl4_44),
  inference(avatar_split_clause,[],[f7161,f792,f729])).
fof(f7161,plain,(
  e3 != op(e3,e4) | ~spl4_44),
  inference(backward_demodulation,[],[f794,f442])).
fof(f7165,plain,(
  ~spl4_19 | ~spl4_44),
  inference(avatar_split_clause,[],[f7158,f792,f687])).
fof(f7158,plain,(
  e3 != op(e4,e1) | ~spl4_44),
  inference(backward_demodulation,[],[f794,f375])).
fof(f7157,plain,(
  spl4_44 | ~spl4_47 | ~spl4_109),
  inference(avatar_split_clause,[],[f7155,f1065,f805,f792])).
fof(f7155,plain,(
  e3 = op(e3,e1) | (~spl4_47 | ~spl4_109)),
  inference(backward_demodulation,[],[f807,f7035])).
fof(f7035,plain,(
  e3 = op(e3,op(e3,e0)) | ~spl4_109),
  inference(backward_demodulation,[],[f1067,f346])).
fof(f7146,plain,(
  spl4_20 | ~spl4_12 | ~spl4_55),
  inference(avatar_split_clause,[],[f7144,f838,f658,f691])).
fof(f7144,plain,(
  e4 = op(e4,e1) | (~spl4_12 | ~spl4_55)),
  inference(backward_demodulation,[],[f840,f4830])).
fof(f4830,plain,(
  e4 = op(op(e2,e4),e1) | ~spl4_12),
  inference(forward_demodulation,[],[f353,f660])).
fof(f7132,plain,(
  ~spl4_61 | spl4_71 | ~spl4_138),
  inference(avatar_contradiction_clause,[],[f7131])).
fof(f7131,plain,(
  $false | (~spl4_61 | spl4_71 | ~spl4_138)),
  inference(subsumption_resolution,[],[f7126,f907])).
fof(f907,plain,(
  e0 != op(e2,e0) | spl4_71),
  inference(avatar_component_clause,[],[f906])).
fof(f7126,plain,(
  e0 = op(e2,e0) | (~spl4_61 | ~spl4_138)),
  inference(backward_demodulation,[],[f866,f1264])).
fof(f7115,plain,(
  ~spl4_58 | ~spl4_68),
  inference(avatar_split_clause,[],[f7108,f893,f851])).
fof(f7108,plain,(
  e2 != op(e2,e3) | ~spl4_68),
  inference(backward_demodulation,[],[f895,f431])).
fof(f7104,plain,(
  ~spl4_3 | ~spl4_78),
  inference(avatar_split_clause,[],[f7101,f935,f620])).
fof(f7101,plain,(
  e2 != op(e4,e4) | ~spl4_78),
  inference(backward_demodulation,[],[f937,f402])).
fof(f7084,plain,(
  spl4_91 | ~spl4_92 | ~spl4_178),
  inference(avatar_split_clause,[],[f7083,f1481,f994,f990])).
fof(f7083,plain,(
  e0 = op(e1,e1) | (~spl4_92 | ~spl4_178)),
  inference(forward_demodulation,[],[f7074,f996])).
fof(f7074,plain,(
  e0 = op(op(e1,e1),e1) | (~spl4_92 | ~spl4_178)),
  inference(backward_demodulation,[],[f996,f1482])).
fof(f1482,plain,(
  e0 = op(op(e1,op(e1,e1)),op(e1,e1)) | ~spl4_178),
  inference(avatar_component_clause,[],[f1481])).
fof(f7081,plain,(
  ~spl4_92 | ~spl4_157),
  inference(avatar_contradiction_clause,[],[f7080])).
fof(f7080,plain,(
  $false | (~spl4_92 | ~spl4_157)),
  inference(subsumption_resolution,[],[f7079,f460])).
fof(f7079,plain,(
  e1 = e2 | (~spl4_92 | ~spl4_157)),
  inference(forward_demodulation,[],[f7078,f996])).
fof(f7078,plain,(
  e2 = op(e1,e1) | (~spl4_92 | ~spl4_157)),
  inference(forward_demodulation,[],[f7072,f996])).
fof(f7072,plain,(
  e2 = op(op(e1,e1),op(e1,e1)) | (~spl4_92 | ~spl4_157)),
  inference(backward_demodulation,[],[f996,f1354])).
fof(f1354,plain,(
  e2 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | ~spl4_157),
  inference(avatar_component_clause,[],[f1353])).
fof(f7077,plain,(
  ~spl4_82 | ~spl4_92),
  inference(avatar_split_clause,[],[f7067,f994,f952])).
fof(f7067,plain,(
  e1 != op(e1,e3) | ~spl4_92),
  inference(backward_demodulation,[],[f996,f421])).
fof(f7076,plain,(
  ~spl4_67 | ~spl4_92),
  inference(avatar_split_clause,[],[f7064,f994,f889])).
fof(f7064,plain,(
  e1 != op(e2,e1) | ~spl4_92),
  inference(backward_demodulation,[],[f996,f370])).
fof(f7061,plain,(
  ~spl4_50 | ~spl4_100),
  inference(avatar_split_clause,[],[f7055,f1027,f817])).
fof(f7055,plain,(
  e4 != op(e3,e0) | ~spl4_100),
  inference(backward_demodulation,[],[f1029,f361])).
fof(f7054,plain,(
  spl4_100 | ~spl4_21 | ~spl4_102),
  inference(avatar_split_clause,[],[f7051,f1036,f696,f1027])).
fof(f7051,plain,(
  e4 = op(e1,e0) | (~spl4_21 | ~spl4_102)),
  inference(backward_demodulation,[],[f1038,f6957])).
fof(f6957,plain,(
  e4 = op(op(e0,e4),e0) | ~spl4_21),
  inference(forward_demodulation,[],[f351,f698])).
fof(f7053,plain,(
  ~spl4_27 | ~spl4_102),
  inference(avatar_split_clause,[],[f7046,f1036,f721])).
fof(f7046,plain,(
  e1 != op(e3,e4) | ~spl4_102),
  inference(backward_demodulation,[],[f1038,f398])).
fof(f398,plain,(
  op(e0,e4) != op(e3,e4)),
  inference(cnf_transformation,[],[f4])).
fof(f7033,plain,(
  spl4_26 | ~spl4_74 | ~spl4_115),
  inference(avatar_split_clause,[],[f7030,f1090,f918,f717])).
fof(f7030,plain,(
  e0 = op(e3,e4) | (~spl4_74 | ~spl4_115)),
  inference(backward_demodulation,[],[f1092,f6961])).
fof(f6961,plain,(
  e0 = op(e3,op(e0,e2)) | ~spl4_74),
  inference(forward_demodulation,[],[f333,f920])).
fof(f7032,plain,(
  spl4_8 | ~spl4_74 | ~spl4_115),
  inference(avatar_split_clause,[],[f7029,f1090,f918,f641])).
fof(f7029,plain,(
  e2 = op(e4,e3) | (~spl4_74 | ~spl4_115)),
  inference(backward_demodulation,[],[f1092,f6960])).
fof(f6960,plain,(
  e2 = op(op(e0,e2),e3) | ~spl4_74),
  inference(forward_demodulation,[],[f341,f920])).
fof(f7031,plain,(
  ~spl4_65 | ~spl4_115),
  inference(avatar_split_clause,[],[f7025,f1090,f880])).
fof(f7025,plain,(
  e4 != op(e2,e2) | ~spl4_115),
  inference(backward_demodulation,[],[f1092,f377])).
fof(f7024,plain,(
  ~spl4_106 | ~spl4_116),
  inference(avatar_split_clause,[],[f7017,f1095,f1053])).
fof(f7017,plain,(
  e0 != op(e0,e3) | ~spl4_116),
  inference(backward_demodulation,[],[f1097,f411])).
fof(f7004,plain,(
  spl4_115 | ~spl4_123 | ~spl4_189),
  inference(avatar_split_clause,[],[f6996,f1569,f1124,f1090])).
fof(f6996,plain,(
  e4 = op(e0,e2) | (~spl4_123 | ~spl4_189)),
  inference(backward_demodulation,[],[f1126,f1570])).
fof(f7003,plain,(
  ~spl4_113 | ~spl4_123),
  inference(avatar_split_clause,[],[f6989,f1124,f1082])).
fof(f6989,plain,(
  e2 != op(e0,e2) | ~spl4_123),
  inference(backward_demodulation,[],[f1126,f407])).
fof(f7002,plain,(
  ~spl4_98 | ~spl4_123),
  inference(avatar_split_clause,[],[f6986,f1124,f1019])).
fof(f6986,plain,(
  e2 != op(e1,e0) | ~spl4_123),
  inference(backward_demodulation,[],[f1126,f356])).
fof(f6973,plain,(
  ~spl4_158 | ~spl4_174),
  inference(avatar_contradiction_clause,[],[f6972])).
fof(f6972,plain,(
  $false | (~spl4_158 | ~spl4_174)),
  inference(subsumption_resolution,[],[f6971,f460])).
fof(f6971,plain,(
  e1 = e2 | (~spl4_158 | ~spl4_174)),
  inference(forward_demodulation,[],[f1453,f1365])).
fof(f6951,plain,(
  ~spl4_69 | ~spl4_74),
  inference(avatar_split_clause,[],[f6950,f918,f897])).
fof(f6950,plain,(
  e3 != op(e2,e1) | ~spl4_74),
  inference(forward_demodulation,[],[f426,f920])).
fof(f6939,plain,(
  ~spl4_93 | spl4_137 | ~spl4_158),
  inference(avatar_split_clause,[],[f6938,f1364,f1258,f998])).
fof(f6938,plain,(
  e2 != op(e1,e1) | (spl4_137 | ~spl4_158)),
  inference(forward_demodulation,[],[f1260,f1365])).
fof(f6929,plain,(
  spl4_10 | ~spl4_19 | ~spl4_80),
  inference(avatar_split_clause,[],[f6927,f943,f687,f649])).
fof(f6927,plain,(
  e4 = op(e4,e3) | (~spl4_19 | ~spl4_80)),
  inference(backward_demodulation,[],[f689,f6512])).
fof(f6512,plain,(
  e4 = op(e4,op(e4,e1)) | ~spl4_80),
  inference(backward_demodulation,[],[f945,f352])).
fof(f6878,plain,(
  ~spl4_38 | ~spl4_43),
  inference(avatar_split_clause,[],[f6871,f788,f767])).
fof(f6871,plain,(
  e2 != op(e3,e2) | ~spl4_43),
  inference(backward_demodulation,[],[f790,f440])).
fof(f6835,plain,(
  ~spl4_54 | ~spl4_74),
  inference(avatar_split_clause,[],[f6830,f918,f834])).
fof(f6830,plain,(
  e3 != op(e2,e4) | ~spl4_74),
  inference(backward_demodulation,[],[f920,f429])).
fof(f6834,plain,(
  ~spl4_49 | ~spl4_74),
  inference(avatar_split_clause,[],[f6827,f918,f813])).
fof(f6827,plain,(
  e3 != op(e3,e0) | ~spl4_74),
  inference(backward_demodulation,[],[f920,f363])).
fof(f6826,plain,(
  ~spl4_32 | ~spl4_82),
  inference(avatar_split_clause,[],[f6821,f952,f742])).
fof(f6821,plain,(
  e1 != op(e3,e3) | ~spl4_82),
  inference(backward_demodulation,[],[f954,f391])).
fof(f6817,plain,(
  spl4_82 | ~spl4_67 | ~spl4_89),
  inference(avatar_split_clause,[],[f6815,f981,f889,f952])).
fof(f6815,plain,(
  e1 = op(e1,e3) | (~spl4_67 | ~spl4_89)),
  inference(backward_demodulation,[],[f983,f6686])).
fof(f6686,plain,(
  e1 = op(e1,op(e1,e2)) | ~spl4_67),
  inference(forward_demodulation,[],[f338,f891])).
fof(f6808,plain,(
  ~spl4_86 | ~spl4_91),
  inference(avatar_split_clause,[],[f6798,f990,f969])).
fof(f6798,plain,(
  e0 != op(e1,e2) | ~spl4_91),
  inference(backward_demodulation,[],[f992,f420])).
fof(f6807,plain,(
  ~spl4_16 | ~spl4_91),
  inference(avatar_split_clause,[],[f6797,f990,f675])).
fof(f6797,plain,(
  e0 != op(e4,e1) | ~spl4_91),
  inference(backward_demodulation,[],[f992,f372])).
fof(f6794,plain,(
  spl4_51 | ~spl4_98 | ~spl4_120),
  inference(avatar_split_clause,[],[f6790,f1111,f1019,f822])).
fof(f6790,plain,(
  e0 = op(e2,e4) | (~spl4_98 | ~spl4_120)),
  inference(backward_demodulation,[],[f1021,f6745])).
fof(f6745,plain,(
  e0 = op(op(e1,e0),e4) | ~spl4_120),
  inference(backward_demodulation,[],[f1113,f332])).
fof(f6763,plain,(
  ~spl4_38 | ~spl4_113),
  inference(avatar_split_clause,[],[f6758,f1082,f767])).
fof(f6758,plain,(
  e2 != op(e3,e2) | ~spl4_113),
  inference(backward_demodulation,[],[f1084,f378])).
fof(f6754,plain,(
  ~spl4_110 | ~spl4_120),
  inference(avatar_split_clause,[],[f6751,f1111,f1069])).
fof(f6751,plain,(
  e4 != op(e0,e3) | ~spl4_120),
  inference(backward_demodulation,[],[f1113,f411])).
fof(f6742,plain,(
  ~spl4_102 | ~spl4_122),
  inference(avatar_split_clause,[],[f6730,f1120,f1036])).
fof(f6730,plain,(
  e1 != op(e0,e4) | ~spl4_122),
  inference(backward_demodulation,[],[f1122,f409])).
fof(f6741,plain,(
  ~spl4_97 | ~spl4_122),
  inference(avatar_split_clause,[],[f6723,f1120,f1015])).
fof(f6723,plain,(
  e1 != op(e1,e0) | ~spl4_122),
  inference(backward_demodulation,[],[f1122,f356])).
fof(f6721,plain,(
  spl4_51 | ~spl4_65 | ~spl4_138),
  inference(avatar_split_clause,[],[f6720,f1263,f880,f822])).
fof(f6720,plain,(
  e0 = op(e2,e4) | (~spl4_65 | ~spl4_138)),
  inference(forward_demodulation,[],[f1264,f882])).
fof(f6687,plain,(
  ~spl4_91 | ~spl4_3 | ~spl4_12 | spl4_163),
  inference(avatar_split_clause,[],[f6611,f1392,f658,f620,f990])).
fof(f6611,plain,(
  e0 != op(e1,e1) | (~spl4_3 | ~spl4_12 | spl4_163)),
  inference(forward_demodulation,[],[f6600,f660])).
fof(f6600,plain,(
  e0 != op(op(e4,e2),op(e4,e2)) | (~spl4_3 | spl4_163)),
  inference(backward_demodulation,[],[f622,f1394])).
fof(f6681,plain,(
  ~spl4_89 | ~spl4_3 | ~spl4_12 | spl4_136),
  inference(avatar_split_clause,[],[f6606,f1254,f658,f620,f981])).
fof(f6606,plain,(
  e3 != op(e1,e2) | (~spl4_3 | ~spl4_12 | spl4_136)),
  inference(forward_demodulation,[],[f6597,f660])).
fof(f6597,plain,(
  e3 != op(op(e4,e2),e2) | (~spl4_3 | spl4_136)),
  inference(backward_demodulation,[],[f622,f1256])).
fof(f6680,plain,(
  ~spl4_86 | ~spl4_3 | ~spl4_12 | spl4_161),
  inference(avatar_split_clause,[],[f6657,f1382,f658,f620,f969])).
fof(f6657,plain,(
  e0 != op(e1,e2) | (~spl4_3 | ~spl4_12 | spl4_161)),
  inference(forward_demodulation,[],[f6656,f660])).
fof(f6656,plain,(
  e0 != op(op(e4,e2),e2) | (~spl4_3 | spl4_161)),
  inference(forward_demodulation,[],[f1384,f622])).
fof(f6631,plain,(
  ~spl4_16 | ~spl4_32 | ~spl4_45 | spl4_162),
  inference(avatar_contradiction_clause,[],[f6630])).
fof(f6630,plain,(
  $false | (~spl4_16 | ~spl4_32 | ~spl4_45 | spl4_162)),
  inference(subsumption_resolution,[],[f6629,f677])).
fof(f6629,plain,(
  e0 != op(e4,e1) | (~spl4_32 | ~spl4_45 | spl4_162)),
  inference(forward_demodulation,[],[f6628,f798])).
fof(f6628,plain,(
  e0 != op(op(e3,e1),e1) | (~spl4_32 | spl4_162)),
  inference(forward_demodulation,[],[f1389,f744])).
fof(f6610,plain,(
  ~spl4_3 | ~spl4_12 | spl4_158),
  inference(avatar_contradiction_clause,[],[f6609])).
fof(f6609,plain,(
  $false | (~spl4_3 | ~spl4_12 | spl4_158)),
  inference(subsumption_resolution,[],[f6599,f660])).
fof(f6599,plain,(
  e1 != op(e4,e2) | (~spl4_3 | spl4_158)),
  inference(backward_demodulation,[],[f622,f1366])).
fof(f6557,plain,(
  ~spl4_40 | ~spl4_45),
  inference(avatar_split_clause,[],[f6552,f796,f775])).
fof(f6552,plain,(
  e4 != op(e3,e2) | ~spl4_45),
  inference(backward_demodulation,[],[f798,f440])).
fof(f6460,plain,(
  spl4_83 | ~spl4_94 | ~spl4_176),
  inference(avatar_split_clause,[],[f6459,f1463,f1002,f956])).
fof(f6459,plain,(
  e2 = op(e1,e3) | (~spl4_94 | ~spl4_176)),
  inference(forward_demodulation,[],[f1464,f1004])).
fof(f6434,plain,(
  ~spl4_48 | ~spl4_73),
  inference(avatar_split_clause,[],[f6433,f914,f809])).
fof(f6433,plain,(
  e2 != op(e3,e0) | ~spl4_73),
  inference(forward_demodulation,[],[f363,f916])).
fof(f6392,plain,(
  ~spl4_5 | ~spl4_20),
  inference(avatar_split_clause,[],[f6389,f691,f628])).
fof(f6389,plain,(
  e4 != op(e4,e4) | ~spl4_20),
  inference(backward_demodulation,[],[f693,f452])).
fof(f6387,plain,(
  ~spl4_19 | ~spl4_24),
  inference(avatar_split_clause,[],[f6381,f708,f687])).
fof(f6381,plain,(
  e3 != op(e4,e1) | ~spl4_24),
  inference(backward_demodulation,[],[f710,f446])).
fof(f6367,plain,(
  spl4_24 | ~spl4_41 | ~spl4_85),
  inference(avatar_split_clause,[],[f6365,f964,f780,f708])).
fof(f6365,plain,(
  e3 = op(e4,e0) | (~spl4_41 | ~spl4_85)),
  inference(backward_demodulation,[],[f782,f6017])).
fof(f6017,plain,(
  e3 = op(e4,op(e3,e1)) | ~spl4_85),
  inference(backward_demodulation,[],[f966,f347])).
fof(f6356,plain,(
  ~spl4_48 | ~spl4_106 | ~spl4_113),
  inference(avatar_contradiction_clause,[],[f6355])).
fof(f6355,plain,(
  $false | (~spl4_48 | ~spl4_106 | ~spl4_113)),
  inference(subsumption_resolution,[],[f6354,f463])).
fof(f6354,plain,(
  e2 = e3 | (~spl4_48 | ~spl4_106 | ~spl4_113)),
  inference(forward_demodulation,[],[f6351,f1084])).
fof(f6351,plain,(
  e3 = op(e0,e2) | (~spl4_48 | ~spl4_106)),
  inference(backward_demodulation,[],[f811,f6231])).
fof(f6231,plain,(
  e3 = op(e0,op(e3,e0)) | ~spl4_106),
  inference(backward_demodulation,[],[f1055,f346])).
fof(f6331,plain,(
  spl4_52 | ~spl4_65 | ~spl4_160),
  inference(avatar_contradiction_clause,[],[f6330])).
fof(f6330,plain,(
  $false | (spl4_52 | ~spl4_65 | ~spl4_160)),
  inference(subsumption_resolution,[],[f6322,f827])).
fof(f827,plain,(
  e1 != op(e2,e4) | spl4_52),
  inference(avatar_component_clause,[],[f826])).
fof(f6322,plain,(
  e1 = op(e2,e4) | (~spl4_65 | ~spl4_160)),
  inference(backward_demodulation,[],[f882,f1376])).
fof(f6329,plain,(
  ~spl4_55 | ~spl4_65),
  inference(avatar_split_clause,[],[f6321,f880,f838])).
fof(f6321,plain,(
  e4 != op(e2,e4) | ~spl4_65),
  inference(backward_demodulation,[],[f882,f434])).
fof(f6283,plain,(
  ~spl4_6 | ~spl4_85 | ~spl4_94 | spl4_178),
  inference(avatar_split_clause,[],[f6282,f1481,f1002,f964,f633])).
fof(f6282,plain,(
  e0 != op(e4,e3) | (~spl4_85 | ~spl4_94 | spl4_178)),
  inference(forward_demodulation,[],[f6275,f966])).
fof(f6275,plain,(
  e0 != op(op(e1,e3),e3) | (~spl4_94 | spl4_178)),
  inference(backward_demodulation,[],[f1004,f1483])).
fof(f6278,plain,(
  ~spl4_19 | ~spl4_94),
  inference(avatar_split_clause,[],[f6267,f1002,f687])).
fof(f6267,plain,(
  e3 != op(e4,e1) | ~spl4_94),
  inference(backward_demodulation,[],[f1004,f372])).
fof(f6250,plain,(
  ~spl4_77 | ~spl4_102),
  inference(avatar_split_clause,[],[f6244,f1036,f931])).
fof(f6244,plain,(
  e1 != op(e1,e4) | ~spl4_102),
  inference(backward_demodulation,[],[f1038,f396])).
fof(f6229,plain,(
  ~spl4_113 | ~spl4_114),
  inference(avatar_contradiction_clause,[],[f6228])).
fof(f6228,plain,(
  $false | (~spl4_113 | ~spl4_114)),
  inference(subsumption_resolution,[],[f6227,f463])).
fof(f6227,plain,(
  e2 = e3 | (~spl4_113 | ~spl4_114)),
  inference(backward_demodulation,[],[f1084,f1088])).
fof(f6225,plain,(
  spl4_48 | ~spl4_71 | ~spl4_114),
  inference(avatar_split_clause,[],[f6221,f1086,f906,f809])).
fof(f6221,plain,(
  e2 = op(e3,e0) | (~spl4_71 | ~spl4_114)),
  inference(backward_demodulation,[],[f1088,f6173])).
fof(f6173,plain,(
  e2 = op(op(e0,e2),e0) | ~spl4_71),
  inference(forward_demodulation,[],[f341,f908])).
fof(f6223,plain,(
  ~spl4_64 | ~spl4_114),
  inference(avatar_split_clause,[],[f6216,f1086,f876])).
fof(f6216,plain,(
  e3 != op(e2,e2) | ~spl4_114),
  inference(backward_demodulation,[],[f1088,f377])).
fof(f6187,plain,(
  ~spl4_33 | spl4_137 | ~spl4_182),
  inference(avatar_split_clause,[],[f6185,f1508,f1258,f746])).
fof(f6185,plain,(
  e2 != op(e3,e3) | (spl4_137 | ~spl4_182)),
  inference(backward_demodulation,[],[f1509,f1260])).
fof(f1509,plain,(
  e3 = op(e4,op(e4,e4)) | ~spl4_182),
  inference(avatar_component_clause,[],[f1508])).
fof(f6172,plain,(
  ~spl4_113 | ~spl4_118),
  inference(avatar_split_clause,[],[f5720,f1103,f1082])).
fof(f5720,plain,(
  e2 != op(e0,e2) | ~spl4_118),
  inference(backward_demodulation,[],[f1105,f410])).
fof(f6158,plain,(
  ~spl4_57 | ~spl4_64 | ~spl4_85 | spl4_143),
  inference(avatar_contradiction_clause,[],[f6157])).
fof(f6157,plain,(
  $false | (~spl4_57 | ~spl4_64 | ~spl4_85 | spl4_143)),
  inference(subsumption_resolution,[],[f6156,f966])).
fof(f6156,plain,(
  e4 != op(e1,e3) | (~spl4_57 | ~spl4_64 | spl4_143)),
  inference(forward_demodulation,[],[f6155,f849])).
fof(f6155,plain,(
  e4 != op(op(e2,e3),e3) | (~spl4_64 | spl4_143)),
  inference(forward_demodulation,[],[f1287,f878])).
fof(f6081,plain,(
  ~spl4_32 | ~spl4_42),
  inference(avatar_split_clause,[],[f6075,f784,f742])).
fof(f6075,plain,(
  e1 != op(e3,e3) | ~spl4_42),
  inference(backward_demodulation,[],[f786,f441])).
fof(f6061,plain,(
  ~spl4_32 | ~spl4_57),
  inference(avatar_split_clause,[],[f6055,f847,f742])).
fof(f6055,plain,(
  e1 != op(e3,e3) | ~spl4_57),
  inference(backward_demodulation,[],[f849,f393])).
fof(f6050,plain,(
  ~spl4_39 | ~spl4_64),
  inference(avatar_split_clause,[],[f6040,f876,f771])).
fof(f6040,plain,(
  e3 != op(e3,e2) | ~spl4_64),
  inference(backward_demodulation,[],[f878,f383])).
fof(f6038,plain,(
  spl4_53 | ~spl4_70 | ~spl4_88),
  inference(avatar_split_clause,[],[f6035,f977,f901,f830])).
fof(f6035,plain,(
  e2 = op(e2,e4) | (~spl4_70 | ~spl4_88)),
  inference(backward_demodulation,[],[f903,f6006])).
fof(f6006,plain,(
  e2 = op(e2,op(e2,e1)) | ~spl4_88),
  inference(backward_demodulation,[],[f979,f342])).
fof(f6037,plain,(
  ~spl4_60 | ~spl4_70),
  inference(avatar_split_clause,[],[f6032,f901,f859])).
fof(f6032,plain,(
  e4 != op(e2,e3) | ~spl4_70),
  inference(backward_demodulation,[],[f903,f431])).
fof(f6036,plain,(
  ~spl4_45 | ~spl4_70),
  inference(avatar_split_clause,[],[f6029,f901,f796])).
fof(f6029,plain,(
  e4 != op(e3,e1) | ~spl4_70),
  inference(backward_demodulation,[],[f903,f373])).
fof(f5998,plain,(
  ~spl4_91 | spl4_98 | ~spl4_176),
  inference(avatar_contradiction_clause,[],[f5997])).
fof(f5997,plain,(
  $false | (~spl4_91 | spl4_98 | ~spl4_176)),
  inference(subsumption_resolution,[],[f5991,f1020])).
fof(f1020,plain,(
  e2 != op(e1,e0) | spl4_98),
  inference(avatar_component_clause,[],[f1019])).
fof(f5991,plain,(
  e2 = op(e1,e0) | (~spl4_91 | ~spl4_176)),
  inference(backward_demodulation,[],[f992,f1464])).
fof(f5996,plain,(
  ~spl4_76 | ~spl4_91),
  inference(avatar_split_clause,[],[f5985,f990,f927])).
fof(f5985,plain,(
  e0 != op(e1,e4) | ~spl4_91),
  inference(backward_demodulation,[],[f992,f422])).
fof(f5978,plain,(
  spl4_57 | ~spl4_99 | ~spl4_118),
  inference(avatar_split_clause,[],[f5975,f1103,f1023,f847])).
fof(f5975,plain,(
  e1 = op(e2,e3) | (~spl4_99 | ~spl4_118)),
  inference(backward_demodulation,[],[f1025,f5912])).
fof(f5912,plain,(
  e1 = op(e2,op(e1,e0)) | ~spl4_118),
  inference(forward_demodulation,[],[f336,f1105])).
fof(f5977,plain,(
  spl4_36 | ~spl4_99 | ~spl4_118),
  inference(avatar_split_clause,[],[f5974,f1103,f1023,f759])).
fof(f5974,plain,(
  e0 = op(e3,e2) | (~spl4_99 | ~spl4_118)),
  inference(backward_demodulation,[],[f1025,f5911])).
fof(f5911,plain,(
  e0 = op(op(e1,e0),e2) | ~spl4_118),
  inference(forward_demodulation,[],[f332,f1105])).
fof(f5976,plain,(
  ~spl4_94 | ~spl4_99),
  inference(avatar_split_clause,[],[f5969,f1023,f1002])).
fof(f5969,plain,(
  e3 != op(e1,e1) | ~spl4_99),
  inference(backward_demodulation,[],[f1025,f416])).
fof(f5955,plain,(
  ~spl4_118 | ~spl4_122 | spl4_177),
  inference(avatar_contradiction_clause,[],[f5954])).
fof(f5954,plain,(
  $false | (~spl4_118 | ~spl4_122 | spl4_177)),
  inference(subsumption_resolution,[],[f5946,f1105])).
fof(f5946,plain,(
  e2 != op(e0,e1) | (~spl4_122 | spl4_177)),
  inference(backward_demodulation,[],[f1122,f1475])).
fof(f5950,plain,(
  ~spl4_47 | ~spl4_122),
  inference(avatar_split_clause,[],[f5938,f1120,f805])).
fof(f5938,plain,(
  e1 != op(e3,e0) | ~spl4_122),
  inference(backward_demodulation,[],[f1122,f358])).
fof(f5931,plain,(
  ~spl4_64 | spl4_153 | ~spl4_176),
  inference(avatar_split_clause,[],[f5930,f1463,f1331,f876])).
fof(f5930,plain,(
  e3 != op(e2,e2) | (spl4_153 | ~spl4_176)),
  inference(forward_demodulation,[],[f1333,f1464])).
fof(f5892,plain,(
  ~spl4_18 | ~spl4_23),
  inference(avatar_split_clause,[],[f5891,f704,f683])).
fof(f5891,plain,(
  e2 != op(e4,e1) | ~spl4_23),
  inference(forward_demodulation,[],[f446,f706])).
fof(f5879,plain,(
  ~spl4_11 | ~spl4_12),
  inference(avatar_contradiction_clause,[],[f5878])).
fof(f5878,plain,(
  $false | (~spl4_11 | ~spl4_12)),
  inference(subsumption_resolution,[],[f5877,f456])).
fof(f456,plain,(
  e0 != e1),
  inference(cnf_transformation,[],[f5])).
fof(f5877,plain,(
  e0 = e1 | (~spl4_11 | ~spl4_12)),
  inference(backward_demodulation,[],[f656,f660])).
fof(f5870,plain,(
  ~spl4_27 | ~spl4_29),
  inference(avatar_contradiction_clause,[],[f5869])).
fof(f5869,plain,(
  $false | (~spl4_27 | ~spl4_29)),
  inference(subsumption_resolution,[],[f5868,f461])).
fof(f5868,plain,(
  e1 = e3 | (~spl4_27 | ~spl4_29)),
  inference(backward_demodulation,[],[f723,f731])).
fof(f731,plain,(
  e3 = op(e3,e4) | ~spl4_29),
  inference(avatar_component_clause,[],[f729])).
fof(f5866,plain,(
  ~spl4_31 | ~spl4_32),
  inference(avatar_contradiction_clause,[],[f5865])).
fof(f5865,plain,(
  $false | (~spl4_31 | ~spl4_32)),
  inference(subsumption_resolution,[],[f5864,f456])).
fof(f5864,plain,(
  e0 = e1 | (~spl4_31 | ~spl4_32)),
  inference(backward_demodulation,[],[f740,f744])).
fof(f5857,plain,(
  ~spl4_43 | ~spl4_45),
  inference(avatar_contradiction_clause,[],[f5856])).
fof(f5856,plain,(
  $false | (~spl4_43 | ~spl4_45)),
  inference(subsumption_resolution,[],[f5855,f464])).
fof(f464,plain,(
  e2 != e4),
  inference(cnf_transformation,[],[f5])).
fof(f5855,plain,(
  e2 = e4 | (~spl4_43 | ~spl4_45)),
  inference(forward_demodulation,[],[f798,f790])).
fof(f5854,plain,(
  ~spl4_47 | spl4_91 | ~spl4_107),
  inference(avatar_contradiction_clause,[],[f5853])).
fof(f5853,plain,(
  $false | (~spl4_47 | spl4_91 | ~spl4_107)),
  inference(subsumption_resolution,[],[f5851,f991])).
fof(f991,plain,(
  e0 != op(e1,e1) | spl4_91),
  inference(avatar_component_clause,[],[f990])).
fof(f5851,plain,(
  e0 = op(e1,e1) | (~spl4_47 | ~spl4_107)),
  inference(backward_demodulation,[],[f807,f5737])).
fof(f5737,plain,(
  e0 = op(op(e3,e0),e1) | ~spl4_107),
  inference(backward_demodulation,[],[f1059,f334])).
fof(f5845,plain,(
  ~spl4_12 | ~spl4_43 | ~spl4_54),
  inference(avatar_contradiction_clause,[],[f5844])).
fof(f5844,plain,(
  $false | (~spl4_12 | ~spl4_43 | ~spl4_54)),
  inference(subsumption_resolution,[],[f5843,f464])).
fof(f5843,plain,(
  e2 = e4 | (~spl4_12 | ~spl4_43 | ~spl4_54)),
  inference(forward_demodulation,[],[f5841,f790])).
fof(f5841,plain,(
  e4 = op(e3,e1) | (~spl4_12 | ~spl4_54)),
  inference(backward_demodulation,[],[f836,f4830])).
fof(f836,plain,(
  e3 = op(e2,e4) | ~spl4_54),
  inference(avatar_component_clause,[],[f834])).
fof(f5830,plain,(
  ~spl4_61 | ~spl4_63 | spl4_138),
  inference(avatar_split_clause,[],[f5824,f1263,f872,f864])).
fof(f5824,plain,(
  e0 != op(e2,e2) | (~spl4_63 | spl4_138)),
  inference(backward_demodulation,[],[f874,f1265])).
fof(f5817,plain,(
  ~spl4_67 | ~spl4_87 | spl4_93),
  inference(avatar_contradiction_clause,[],[f5816])).
fof(f5816,plain,(
  $false | (~spl4_67 | ~spl4_87 | spl4_93)),
  inference(subsumption_resolution,[],[f5813,f999])).
fof(f999,plain,(
  e2 != op(e1,e1) | spl4_93),
  inference(avatar_component_clause,[],[f998])).
fof(f5813,plain,(
  e2 = op(e1,e1) | (~spl4_67 | ~spl4_87)),
  inference(backward_demodulation,[],[f891,f5746])).
fof(f5746,plain,(
  e2 = op(e1,op(e2,e1)) | ~spl4_87),
  inference(backward_demodulation,[],[f975,f342])).
fof(f5804,plain,(
  ~spl4_61 | ~spl4_71),
  inference(avatar_split_clause,[],[f5794,f906,f864])).
fof(f5794,plain,(
  e0 != op(e2,e2) | ~spl4_71),
  inference(backward_demodulation,[],[f908,f427])).
fof(f5803,plain,(
  ~spl4_46 | ~spl4_71),
  inference(avatar_split_clause,[],[f5791,f906,f801])).
fof(f5791,plain,(
  e0 != op(e3,e0) | ~spl4_71),
  inference(backward_demodulation,[],[f908,f363])).
fof(f5790,plain,(
  ~spl4_76 | ~spl4_78),
  inference(avatar_contradiction_clause,[],[f5789])).
fof(f5789,plain,(
  $false | (~spl4_76 | ~spl4_78)),
  inference(subsumption_resolution,[],[f5788,f457])).
fof(f5788,plain,(
  e0 = e2 | (~spl4_76 | ~spl4_78)),
  inference(backward_demodulation,[],[f929,f937])).
fof(f5786,plain,(
  ~spl4_16 | ~spl4_78 | spl4_112),
  inference(avatar_contradiction_clause,[],[f5785])).
fof(f5785,plain,(
  $false | (~spl4_16 | ~spl4_78 | spl4_112)),
  inference(subsumption_resolution,[],[f5781,f1079])).
fof(f1079,plain,(
  e1 != op(e0,e2) | spl4_112),
  inference(avatar_component_clause,[],[f1078])).
fof(f5781,plain,(
  e1 = op(e0,e2) | (~spl4_16 | ~spl4_78)),
  inference(backward_demodulation,[],[f937,f5649])).
fof(f5649,plain,(
  e1 = op(e0,op(e1,e4)) | ~spl4_16),
  inference(forward_demodulation,[],[f340,f677])).
fof(f5784,plain,(
  ~spl4_16 | spl4_75 | ~spl4_78),
  inference(avatar_contradiction_clause,[],[f5783])).
fof(f5783,plain,(
  $false | (~spl4_16 | spl4_75 | ~spl4_78)),
  inference(subsumption_resolution,[],[f5780,f923])).
fof(f923,plain,(
  e4 != op(e2,e0) | spl4_75),
  inference(avatar_component_clause,[],[f922])).
fof(f5780,plain,(
  e4 = op(e2,e0) | (~spl4_16 | ~spl4_78)),
  inference(backward_demodulation,[],[f937,f5648])).
fof(f5648,plain,(
  e4 = op(op(e1,e4),e0) | ~spl4_16),
  inference(forward_demodulation,[],[f352,f677])).
fof(f5775,plain,(
  ~spl4_61 | ~spl4_83 | ~spl4_94 | spl4_180),
  inference(avatar_split_clause,[],[f5765,f1492,f1002,f956,f864])).
fof(f5765,plain,(
  e0 != op(e2,e2) | (~spl4_83 | ~spl4_94 | spl4_180)),
  inference(backward_demodulation,[],[f958,f5682])).
fof(f5682,plain,(
  e0 != op(op(e1,e3),op(e1,e3)) | (~spl4_94 | spl4_180)),
  inference(forward_demodulation,[],[f1494,f1004])).
fof(f5772,plain,(
  ~spl4_43 | spl4_62 | ~spl4_83),
  inference(avatar_contradiction_clause,[],[f5771])).
fof(f5771,plain,(
  $false | (~spl4_43 | spl4_62 | ~spl4_83)),
  inference(subsumption_resolution,[],[f5761,f869])).
fof(f5761,plain,(
  e1 = op(e2,e2) | (~spl4_43 | ~spl4_83)),
  inference(backward_demodulation,[],[f958,f5652])).
fof(f5652,plain,(
  e1 = op(e2,op(e1,e3)) | ~spl4_43),
  inference(forward_demodulation,[],[f339,f790])).
fof(f5770,plain,(
  ~spl4_43 | spl4_64 | ~spl4_83),
  inference(avatar_contradiction_clause,[],[f5769])).
fof(f5769,plain,(
  $false | (~spl4_43 | spl4_64 | ~spl4_83)),
  inference(subsumption_resolution,[],[f5760,f877])).
fof(f5760,plain,(
  e3 = op(e2,e2) | (~spl4_43 | ~spl4_83)),
  inference(backward_demodulation,[],[f958,f5651])).
fof(f5651,plain,(
  e3 = op(op(e1,e3),e2) | ~spl4_43),
  inference(forward_demodulation,[],[f347,f790])).
fof(f5736,plain,(
  ~spl4_110 | ~spl4_115),
  inference(avatar_split_clause,[],[f5734,f1090,f1069])).
fof(f5734,plain,(
  e4 != op(e0,e3) | ~spl4_115),
  inference(backward_demodulation,[],[f1092,f413])).
fof(f5727,plain,(
  spl4_52 | ~spl4_100 | ~spl4_118),
  inference(avatar_contradiction_clause,[],[f5726])).
fof(f5726,plain,(
  $false | (spl4_52 | ~spl4_100 | ~spl4_118)),
  inference(subsumption_resolution,[],[f5724,f827])).
fof(f5724,plain,(
  e1 = op(e2,e4) | (~spl4_100 | ~spl4_118)),
  inference(backward_demodulation,[],[f1105,f5663])).
fof(f5663,plain,(
  e1 = op(op(e0,e1),e4) | ~spl4_100),
  inference(forward_demodulation,[],[f336,f1029])).
fof(f5697,plain,(
  ~spl4_81 | ~spl4_94 | spl4_148),
  inference(avatar_split_clause,[],[f5696,f1309,f1002,f948])).
fof(f5696,plain,(
  e0 != op(e1,e3) | (~spl4_94 | spl4_148)),
  inference(forward_demodulation,[],[f1311,f1004])).
fof(f5673,plain,(
  ~spl4_70 | ~spl4_32 | ~spl4_43 | spl4_141),
  inference(avatar_split_clause,[],[f5672,f1276,f788,f742,f901])).
fof(f5672,plain,(
  e4 != op(e2,e1) | (~spl4_32 | ~spl4_43 | spl4_141)),
  inference(forward_demodulation,[],[f5671,f790])).
fof(f5671,plain,(
  e4 != op(op(e3,e1),e1) | (~spl4_32 | spl4_141)),
  inference(forward_demodulation,[],[f1278,f744])).
fof(f5656,plain,(
  ~spl4_107 | ~spl4_32),
  inference(avatar_split_clause,[],[f5655,f742,f1057])).
fof(f5655,plain,(
  e1 != op(e0,e3) | ~spl4_32),
  inference(forward_demodulation,[],[f388,f744])).
fof(f5653,plain,(
  ~spl4_87 | ~spl4_12),
  inference(avatar_split_clause,[],[f4872,f658,f973])).
fof(f4872,plain,(
  e1 != op(e1,e2) | ~spl4_12),
  inference(forward_demodulation,[],[f382,f660])).
fof(f5627,plain,(
  ~spl4_48 | ~spl4_43),
  inference(avatar_split_clause,[],[f5626,f788,f809])).
fof(f5626,plain,(
  e2 != op(e3,e0) | ~spl4_43),
  inference(forward_demodulation,[],[f436,f790])).
fof(f5621,plain,(
  ~spl4_27 | ~spl4_32),
  inference(avatar_split_clause,[],[f5563,f742,f721])).
fof(f5563,plain,(
  e1 != op(e3,e4) | ~spl4_32),
  inference(backward_demodulation,[],[f744,f445])).
fof(f5551,plain,(
  ~spl4_28 | ~spl4_53),
  inference(avatar_split_clause,[],[f5545,f830,f725])).
fof(f5545,plain,(
  e2 != op(e3,e4) | ~spl4_53),
  inference(backward_demodulation,[],[f832,f403])).
fof(f5528,plain,(
  ~spl4_56 | ~spl4_61),
  inference(avatar_split_clause,[],[f5516,f864,f843])).
fof(f5516,plain,(
  e0 != op(e2,e3) | ~spl4_61),
  inference(backward_demodulation,[],[f866,f433])).
fof(f5506,plain,(
  spl4_43 | ~spl4_72 | ~spl4_114),
  inference(avatar_split_clause,[],[f5503,f1086,f910,f788])).
fof(f5503,plain,(
  e2 = op(e3,e1) | (~spl4_72 | ~spl4_114)),
  inference(backward_demodulation,[],[f912,f5422])).
fof(f5422,plain,(
  e2 = op(e3,op(e2,e0)) | ~spl4_114),
  inference(backward_demodulation,[],[f1088,f341])).
fof(f5504,plain,(
  ~spl4_57 | ~spl4_72),
  inference(avatar_split_clause,[],[f5497,f910,f847])).
fof(f5497,plain,(
  e1 != op(e2,e3) | ~spl4_72),
  inference(backward_demodulation,[],[f912,f428])).
fof(f5472,plain,(
  ~spl4_69 | ~spl4_94),
  inference(avatar_split_clause,[],[f5460,f1002,f897])).
fof(f5460,plain,(
  e3 != op(e2,e1) | ~spl4_94),
  inference(backward_demodulation,[],[f1004,f370])).
fof(f5439,plain,(
  spl4_101 | ~spl4_46 | ~spl4_110),
  inference(avatar_split_clause,[],[f5435,f1069,f801,f1032])).
fof(f5435,plain,(
  e0 = op(e0,e4) | (~spl4_46 | ~spl4_110)),
  inference(backward_demodulation,[],[f1071,f5109])).
fof(f5109,plain,(
  e0 = op(e0,op(e0,e3)) | ~spl4_46),
  inference(forward_demodulation,[],[f334,f803])).
fof(f5438,plain,(
  ~spl4_10 | ~spl4_110),
  inference(avatar_split_clause,[],[f5433,f1069,f649])).
fof(f5433,plain,(
  e4 != op(e4,e3) | ~spl4_110),
  inference(backward_demodulation,[],[f1071,f389])).
fof(f5429,plain,(
  ~spl4_104 | ~spl4_114),
  inference(avatar_split_clause,[],[f5426,f1086,f1044])).
fof(f5426,plain,(
  e3 != op(e0,e4) | ~spl4_114),
  inference(backward_demodulation,[],[f1088,f414])).
fof(f5428,plain,(
  ~spl4_89 | ~spl4_114),
  inference(avatar_split_clause,[],[f5423,f1086,f981])).
fof(f5423,plain,(
  e3 != op(e1,e2) | ~spl4_114),
  inference(backward_demodulation,[],[f1088,f376])).
fof(f5376,plain,(
  ~spl4_159 | ~spl4_175),
  inference(avatar_contradiction_clause,[],[f5375])).
fof(f5375,plain,(
  $false | (~spl4_159 | ~spl4_175)),
  inference(subsumption_resolution,[],[f5374,f460])).
fof(f5374,plain,(
  e1 = e2 | (~spl4_159 | ~spl4_175)),
  inference(forward_demodulation,[],[f1458,f1370])).
fof(f5354,plain,(
  ~spl4_10 | ~spl4_12 | ~spl4_28),
  inference(avatar_contradiction_clause,[],[f5353])).
fof(f5353,plain,(
  $false | (~spl4_10 | ~spl4_12 | ~spl4_28)),
  inference(subsumption_resolution,[],[f5352,f461])).
fof(f5352,plain,(
  e1 = e3 | (~spl4_10 | ~spl4_12 | ~spl4_28)),
  inference(forward_demodulation,[],[f5350,f660])).
fof(f5350,plain,(
  e3 = op(e4,e2) | (~spl4_10 | ~spl4_28)),
  inference(backward_demodulation,[],[f651,f5335])).
fof(f5335,plain,(
  e3 = op(op(e4,e3),e2) | ~spl4_28),
  inference(backward_demodulation,[],[f727,f350])).
fof(f5349,plain,(
  ~spl4_18 | ~spl4_77 | ~spl4_89),
  inference(avatar_contradiction_clause,[],[f5348])).
fof(f5348,plain,(
  $false | (~spl4_18 | ~spl4_77 | ~spl4_89)),
  inference(subsumption_resolution,[],[f5347,f465])).
fof(f465,plain,(
  e3 != e4),
  inference(cnf_transformation,[],[f5])).
fof(f5347,plain,(
  e3 = e4 | (~spl4_18 | ~spl4_77 | ~spl4_89)),
  inference(forward_demodulation,[],[f5343,f983])).
fof(f5343,plain,(
  e4 = op(e1,e2) | (~spl4_18 | ~spl4_77)),
  inference(backward_demodulation,[],[f685,f4999])).
fof(f4999,plain,(
  e4 = op(e1,op(e4,e1)) | ~spl4_77),
  inference(backward_demodulation,[],[f933,f352])).
fof(f5346,plain,(
  ~spl4_18 | spl4_67 | ~spl4_77),
  inference(avatar_contradiction_clause,[],[f5345])).
fof(f5345,plain,(
  $false | (~spl4_18 | spl4_67 | ~spl4_77)),
  inference(subsumption_resolution,[],[f5342,f890])).
fof(f890,plain,(
  e1 != op(e2,e1) | spl4_67),
  inference(avatar_component_clause,[],[f889])).
fof(f5342,plain,(
  e1 = op(e2,e1) | (~spl4_18 | ~spl4_77)),
  inference(backward_demodulation,[],[f685,f4998])).
fof(f4998,plain,(
  e1 = op(op(e4,e1),e1) | ~spl4_77),
  inference(backward_demodulation,[],[f933,f340])).
fof(f5334,plain,(
  spl4_32 | ~spl4_34 | ~spl4_159),
  inference(avatar_contradiction_clause,[],[f5333])).
fof(f5333,plain,(
  $false | (spl4_32 | ~spl4_34 | ~spl4_159)),
  inference(subsumption_resolution,[],[f5329,f743])).
fof(f743,plain,(
  e1 != op(e3,e3) | spl4_32),
  inference(avatar_component_clause,[],[f742])).
fof(f5329,plain,(
  e1 = op(e3,e3) | (~spl4_34 | ~spl4_159)),
  inference(backward_demodulation,[],[f752,f1370])).
fof(f5315,plain,(
  ~spl4_12 | ~spl4_52 | ~spl4_91),
  inference(avatar_contradiction_clause,[],[f5314])).
fof(f5314,plain,(
  $false | (~spl4_12 | ~spl4_52 | ~spl4_91)),
  inference(subsumption_resolution,[],[f5313,f457])).
fof(f5313,plain,(
  e0 = e2 | (~spl4_12 | ~spl4_52 | ~spl4_91)),
  inference(forward_demodulation,[],[f5308,f992])).
fof(f5308,plain,(
  e2 = op(e1,e1) | (~spl4_12 | ~spl4_52)),
  inference(backward_demodulation,[],[f828,f4831])).
fof(f4831,plain,(
  e2 = op(e1,op(e2,e4)) | ~spl4_12),
  inference(forward_demodulation,[],[f345,f660])).
fof(f5302,plain,(
  ~spl4_56 | ~spl4_57),
  inference(avatar_contradiction_clause,[],[f5301])).
fof(f5301,plain,(
  $false | (~spl4_56 | ~spl4_57)),
  inference(subsumption_resolution,[],[f5300,f456])).
fof(f5300,plain,(
  e0 = e1 | (~spl4_56 | ~spl4_57)),
  inference(forward_demodulation,[],[f849,f845])).
fof(f5290,plain,(
  ~spl4_77 | ~spl4_80),
  inference(avatar_contradiction_clause,[],[f5289])).
fof(f5289,plain,(
  $false | (~spl4_77 | ~spl4_80)),
  inference(subsumption_resolution,[],[f5288,f462])).
fof(f5288,plain,(
  e1 = e4 | (~spl4_77 | ~spl4_80)),
  inference(forward_demodulation,[],[f945,f933])).
fof(f5279,plain,(
  spl4_33 | ~spl4_69 | ~spl4_89),
  inference(avatar_contradiction_clause,[],[f5278])).
fof(f5278,plain,(
  $false | (spl4_33 | ~spl4_69 | ~spl4_89)),
  inference(subsumption_resolution,[],[f5273,f747])).
fof(f747,plain,(
  e2 != op(e3,e3) | spl4_33),
  inference(avatar_component_clause,[],[f746])).
fof(f5273,plain,(
  e2 = op(e3,e3) | (~spl4_69 | ~spl4_89)),
  inference(backward_demodulation,[],[f983,f5128])).
fof(f5128,plain,(
  e2 = op(op(e1,e2),e3) | ~spl4_69),
  inference(forward_demodulation,[],[f342,f899])).
fof(f5277,plain,(
  spl4_32 | ~spl4_69 | ~spl4_89),
  inference(avatar_contradiction_clause,[],[f5276])).
fof(f5276,plain,(
  $false | (spl4_32 | ~spl4_69 | ~spl4_89)),
  inference(subsumption_resolution,[],[f5272,f743])).
fof(f5272,plain,(
  e1 = op(e3,e3) | (~spl4_69 | ~spl4_89)),
  inference(backward_demodulation,[],[f983,f5127])).
fof(f5127,plain,(
  e1 = op(e3,op(e1,e2)) | ~spl4_69),
  inference(forward_demodulation,[],[f338,f899])).
fof(f5275,plain,(
  ~spl4_84 | ~spl4_89),
  inference(avatar_split_clause,[],[f5270,f981,f960])).
fof(f5270,plain,(
  e3 != op(e1,e3) | ~spl4_89),
  inference(backward_demodulation,[],[f983,f423])).
fof(f5266,plain,(
  ~spl4_91 | ~spl4_98 | spl4_176),
  inference(avatar_contradiction_clause,[],[f5265])).
fof(f5265,plain,(
  $false | (~spl4_91 | ~spl4_98 | spl4_176)),
  inference(subsumption_resolution,[],[f5252,f1021])).
fof(f5252,plain,(
  e2 != op(e1,e0) | (~spl4_91 | spl4_176)),
  inference(backward_demodulation,[],[f992,f1465])).
fof(f1465,plain,(
  e2 != op(e1,op(e1,e1)) | spl4_176),
  inference(avatar_component_clause,[],[f1463])).
fof(f5260,plain,(
  ~spl4_75 | ~spl4_91 | ~spl4_98 | spl4_152),
  inference(avatar_split_clause,[],[f5259,f1327,f1019,f990,f922])).
fof(f5259,plain,(
  e4 != op(e2,e0) | (~spl4_91 | ~spl4_98 | spl4_152)),
  inference(forward_demodulation,[],[f5249,f1021])).
fof(f5249,plain,(
  e4 != op(op(e1,e0),e0) | (~spl4_91 | spl4_152)),
  inference(backward_demodulation,[],[f992,f1329])).
fof(f5226,plain,(
  spl4_35 | ~spl4_24 | ~spl4_104),
  inference(avatar_split_clause,[],[f5223,f1044,f708,f754])).
fof(f5223,plain,(
  e4 = op(e3,e3) | (~spl4_24 | ~spl4_104)),
  inference(backward_demodulation,[],[f1046,f5147])).
fof(f5147,plain,(
  e4 = op(op(e0,e4),e3) | ~spl4_24),
  inference(forward_demodulation,[],[f351,f710])).
fof(f5225,plain,(
  ~spl4_24 | spl4_31 | ~spl4_104),
  inference(avatar_contradiction_clause,[],[f5224])).
fof(f5224,plain,(
  $false | (~spl4_24 | spl4_31 | ~spl4_104)),
  inference(subsumption_resolution,[],[f5222,f739])).
fof(f739,plain,(
  e0 != op(e3,e3) | spl4_31),
  inference(avatar_component_clause,[],[f738])).
fof(f5222,plain,(
  e0 = op(e3,e3) | (~spl4_24 | ~spl4_104)),
  inference(backward_demodulation,[],[f1046,f5146])).
fof(f5146,plain,(
  e0 = op(e3,op(e0,e4)) | ~spl4_24),
  inference(forward_demodulation,[],[f335,f710])).
fof(f5218,plain,(
  ~spl4_46 | spl4_74 | ~spl4_108),
  inference(avatar_contradiction_clause,[],[f5217])).
fof(f5217,plain,(
  $false | (~spl4_46 | spl4_74 | ~spl4_108)),
  inference(subsumption_resolution,[],[f5214,f919])).
fof(f919,plain,(
  e3 != op(e2,e0) | spl4_74),
  inference(avatar_component_clause,[],[f918])).
fof(f5214,plain,(
  e3 = op(e2,e0) | (~spl4_46 | ~spl4_108)),
  inference(backward_demodulation,[],[f1063,f5110])).
fof(f5110,plain,(
  e3 = op(op(e0,e3),e0) | ~spl4_46),
  inference(forward_demodulation,[],[f346,f803])).
fof(f5216,plain,(
  ~spl4_103 | ~spl4_108),
  inference(avatar_split_clause,[],[f5210,f1061,f1040])).
fof(f5210,plain,(
  e2 != op(e0,e4) | ~spl4_108),
  inference(backward_demodulation,[],[f1063,f415])).
fof(f5206,plain,(
  ~spl4_86 | ~spl4_111),
  inference(avatar_split_clause,[],[f5198,f1074,f969])).
fof(f5198,plain,(
  e0 != op(e1,e2) | ~spl4_111),
  inference(backward_demodulation,[],[f1076,f376])).
fof(f5176,plain,(
  ~spl4_72 | ~spl4_122),
  inference(avatar_split_clause,[],[f5159,f1120,f910])).
fof(f5159,plain,(
  e1 != op(e2,e0) | ~spl4_122),
  inference(backward_demodulation,[],[f1122,f357])).
fof(f5121,plain,(
  ~spl4_52 | ~spl4_77),
  inference(avatar_split_clause,[],[f5000,f931,f826])).
fof(f5000,plain,(
  e1 != op(e2,e4) | ~spl4_77),
  inference(backward_demodulation,[],[f933,f400])).
fof(f5086,plain,(
  spl4_125 | ~spl4_1),
  inference(avatar_split_clause,[],[f5085,f612,f1132])).
fof(f5085,plain,(
  op(e0,e0) = e4 | ~spl4_1),
  inference(forward_demodulation,[],[f355,f614])).
fof(f355,plain,(
  e4 = op(op(e4,e4),op(e4,e4))),
  inference(cnf_transformation,[],[f3])).
fof(f5073,plain,(
  ~spl4_1 | ~spl4_4),
  inference(avatar_contradiction_clause,[],[f5072])).
fof(f5072,plain,(
  $false | (~spl4_1 | ~spl4_4)),
  inference(subsumption_resolution,[],[f5071,f458])).
fof(f5071,plain,(
  e0 = e3 | (~spl4_1 | ~spl4_4)),
  inference(backward_demodulation,[],[f614,f626])).
fof(f5053,plain,(
  ~spl4_20 | ~spl4_77),
  inference(avatar_contradiction_clause,[],[f5052])).
fof(f5052,plain,(
  $false | (~spl4_20 | ~spl4_77)),
  inference(subsumption_resolution,[],[f5051,f462])).
fof(f5051,plain,(
  e1 = e4 | (~spl4_20 | ~spl4_77)),
  inference(forward_demodulation,[],[f5047,f933])).
fof(f5047,plain,(
  e4 = op(e1,e4) | (~spl4_20 | ~spl4_77)),
  inference(backward_demodulation,[],[f693,f4999])).
fof(f5044,plain,(
  ~spl4_24 | spl4_36 | ~spl4_103),
  inference(avatar_contradiction_clause,[],[f5043])).
fof(f5043,plain,(
  $false | (~spl4_24 | spl4_36 | ~spl4_103)),
  inference(subsumption_resolution,[],[f5037,f760])).
fof(f760,plain,(
  e0 != op(e3,e2) | spl4_36),
  inference(avatar_component_clause,[],[f759])).
fof(f5037,plain,(
  e0 = op(e3,e2) | (~spl4_24 | ~spl4_103)),
  inference(backward_demodulation,[],[f710,f4809])).
fof(f4809,plain,(
  e0 = op(op(e4,e0),e2) | ~spl4_103),
  inference(forward_demodulation,[],[f335,f1042])).
fof(f5031,plain,(
  ~spl4_46 | ~spl4_48),
  inference(avatar_contradiction_clause,[],[f5030])).
fof(f5030,plain,(
  $false | (~spl4_46 | ~spl4_48)),
  inference(subsumption_resolution,[],[f5029,f457])).
fof(f5029,plain,(
  e0 = e2 | (~spl4_46 | ~spl4_48)),
  inference(backward_demodulation,[],[f803,f811])).
fof(f5016,plain,(
  ~spl4_51 | ~spl4_56),
  inference(avatar_split_clause,[],[f5009,f843,f822])).
fof(f5009,plain,(
  e0 != op(e2,e4) | ~spl4_56),
  inference(backward_demodulation,[],[f845,f435])).
fof(f5008,plain,(
  ~spl4_69 | ~spl4_70),
  inference(avatar_contradiction_clause,[],[f5007])).
fof(f5007,plain,(
  $false | (~spl4_69 | ~spl4_70)),
  inference(subsumption_resolution,[],[f5006,f465])).
fof(f5006,plain,(
  e3 = e4 | (~spl4_69 | ~spl4_70)),
  inference(backward_demodulation,[],[f899,f903])).
fof(f4997,plain,(
  spl4_82 | ~spl4_42 | ~spl4_84),
  inference(avatar_split_clause,[],[f4991,f960,f784,f952])).
fof(f4991,plain,(
  e1 = op(e1,e3) | (~spl4_42 | ~spl4_84)),
  inference(backward_demodulation,[],[f962,f4863])).
fof(f4863,plain,(
  e1 = op(e1,op(e1,e3)) | ~spl4_42),
  inference(forward_demodulation,[],[f339,f786])).
fof(f4996,plain,(
  ~spl4_42 | ~spl4_84),
  inference(avatar_contradiction_clause,[],[f4995])).
fof(f4995,plain,(
  $false | (~spl4_42 | ~spl4_84)),
  inference(subsumption_resolution,[],[f4994,f461])).
fof(f4994,plain,(
  e1 = e3 | (~spl4_42 | ~spl4_84)),
  inference(forward_demodulation,[],[f4990,f786])).
fof(f4990,plain,(
  e3 = op(e3,e1) | (~spl4_42 | ~spl4_84)),
  inference(backward_demodulation,[],[f962,f4862])).
fof(f4862,plain,(
  e3 = op(op(e1,e3),e1) | ~spl4_42),
  inference(forward_demodulation,[],[f347,f786])).
fof(f4993,plain,(
  ~spl4_79 | ~spl4_84),
  inference(avatar_split_clause,[],[f4988,f960,f939])).
fof(f4988,plain,(
  e3 != op(e1,e4) | ~spl4_84),
  inference(backward_demodulation,[],[f962,f425])).
fof(f4986,plain,(
  spl4_22 | ~spl4_70 | ~spl4_86),
  inference(avatar_contradiction_clause,[],[f4985])).
fof(f4985,plain,(
  $false | (spl4_22 | ~spl4_70 | ~spl4_86)),
  inference(subsumption_resolution,[],[f4984,f701])).
fof(f701,plain,(
  e1 != op(e4,e0) | spl4_22),
  inference(avatar_component_clause,[],[f700])).
fof(f4984,plain,(
  e1 = op(e4,e0) | (~spl4_70 | ~spl4_86)),
  inference(backward_demodulation,[],[f971,f4875])).
fof(f4875,plain,(
  e1 = op(e4,op(e1,e2)) | ~spl4_70),
  inference(forward_demodulation,[],[f338,f903])).
fof(f4975,plain,(
  ~spl4_93 | ~spl4_114 | ~spl4_148 | spl4_149),
  inference(avatar_contradiction_clause,[],[f4974])).
fof(f4974,plain,(
  $false | (~spl4_93 | ~spl4_114 | ~spl4_148 | spl4_149)),
  inference(subsumption_resolution,[],[f4966,f1088])).
fof(f4966,plain,(
  e3 != op(e0,e2) | (~spl4_93 | ~spl4_148 | spl4_149)),
  inference(backward_demodulation,[],[f1000,f4561])).
fof(f4561,plain,(
  e3 != op(e0,op(e1,e1)) | (~spl4_148 | spl4_149)),
  inference(forward_demodulation,[],[f1315,f1310])).
fof(f4973,plain,(
  spl4_86 | ~spl4_93 | ~spl4_148),
  inference(avatar_split_clause,[],[f4965,f1309,f998,f969])).
fof(f4965,plain,(
  e0 = op(e1,e2) | (~spl4_93 | ~spl4_148)),
  inference(backward_demodulation,[],[f1000,f1310])).
fof(f4957,plain,(
  ~spl4_48 | spl4_66 | ~spl4_107),
  inference(avatar_contradiction_clause,[],[f4956])).
fof(f4956,plain,(
  $false | (~spl4_48 | spl4_66 | ~spl4_107)),
  inference(subsumption_resolution,[],[f4952,f886])).
fof(f886,plain,(
  e0 != op(e2,e1) | spl4_66),
  inference(avatar_component_clause,[],[f885])).
fof(f4952,plain,(
  e0 = op(e2,e1) | (~spl4_48 | ~spl4_107)),
  inference(backward_demodulation,[],[f1059,f4883])).
fof(f4883,plain,(
  e0 = op(e2,op(e0,e3)) | ~spl4_48),
  inference(forward_demodulation,[],[f334,f811])).
fof(f4955,plain,(
  ~spl4_48 | spl4_89 | ~spl4_107),
  inference(avatar_contradiction_clause,[],[f4954])).
fof(f4954,plain,(
  $false | (~spl4_48 | spl4_89 | ~spl4_107)),
  inference(subsumption_resolution,[],[f4951,f982])).
fof(f982,plain,(
  e3 != op(e1,e2) | spl4_89),
  inference(avatar_component_clause,[],[f981])).
fof(f4951,plain,(
  e3 = op(e1,e2) | (~spl4_48 | ~spl4_107)),
  inference(backward_demodulation,[],[f1059,f4881])).
fof(f4881,plain,(
  e3 = op(op(e0,e3),e2) | ~spl4_48),
  inference(forward_demodulation,[],[f346,f811])).
fof(f4941,plain,(
  ~spl4_100 | spl4_102 | ~spl4_116),
  inference(avatar_contradiction_clause,[],[f4940])).
fof(f4940,plain,(
  $false | (~spl4_100 | spl4_102 | ~spl4_116)),
  inference(subsumption_resolution,[],[f4936,f1037])).
fof(f1037,plain,(
  e1 != op(e0,e4) | spl4_102),
  inference(avatar_component_clause,[],[f1036])).
fof(f4936,plain,(
  e1 = op(e0,e4) | (~spl4_100 | ~spl4_116)),
  inference(backward_demodulation,[],[f1097,f4890])).
fof(f4890,plain,(
  e1 = op(op(e0,e1),e4) | ~spl4_100),
  inference(forward_demodulation,[],[f336,f1029])).
fof(f4900,plain,(
  ~spl4_125 | ~spl4_148 | spl4_150),
  inference(avatar_split_clause,[],[f4567,f1317,f1309,f1132])).
fof(f4567,plain,(
  op(e0,e0) != e4 | (~spl4_148 | spl4_150)),
  inference(forward_demodulation,[],[f1319,f1310])).
fof(f4871,plain,(
  ~spl4_90 | ~spl4_100),
  inference(avatar_split_clause,[],[f4870,f1027,f985])).
fof(f4870,plain,(
  e4 != op(e1,e2) | ~spl4_100),
  inference(forward_demodulation,[],[f417,f1029])).
fof(f4865,plain,(
  spl4_90 | ~spl4_8 | ~spl4_27),
  inference(avatar_split_clause,[],[f4770,f721,f641,f985])).
fof(f4770,plain,(
  e4 = op(e1,e2) | (~spl4_8 | ~spl4_27)),
  inference(backward_demodulation,[],[f723,f4320])).
fof(f4320,plain,(
  e4 = op(op(e3,e4),e2) | ~spl4_8),
  inference(forward_demodulation,[],[f354,f643])).
fof(f354,plain,(
  e4 = op(op(e3,e4),op(e4,e3))),
  inference(cnf_transformation,[],[f3])).
fof(f4848,plain,(
  ~spl4_72 | spl4_76 | ~spl4_115),
  inference(avatar_contradiction_clause,[],[f4847])).
fof(f4847,plain,(
  $false | (~spl4_72 | spl4_76 | ~spl4_115)),
  inference(subsumption_resolution,[],[f4846,f928])).
fof(f4846,plain,(
  e0 = op(e1,e4) | (~spl4_72 | ~spl4_115)),
  inference(forward_demodulation,[],[f4599,f912])).
fof(f4599,plain,(
  e0 = op(op(e2,e0),e4) | ~spl4_115),
  inference(backward_demodulation,[],[f1092,f333])).
fof(f4842,plain,(
  spl4_69 | ~spl4_8 | ~spl4_27),
  inference(avatar_split_clause,[],[f4771,f721,f641,f897])).
fof(f4771,plain,(
  e3 = op(e2,e1) | (~spl4_8 | ~spl4_27)),
  inference(backward_demodulation,[],[f723,f4321])).
fof(f4321,plain,(
  e3 = op(e2,op(e3,e4)) | ~spl4_8),
  inference(forward_demodulation,[],[f350,f643])).
fof(f4838,plain,(
  ~spl4_60 | ~spl4_35),
  inference(avatar_split_clause,[],[f4837,f754,f859])).
fof(f4837,plain,(
  e4 != op(e2,e3) | ~spl4_35),
  inference(forward_demodulation,[],[f393,f756])).
fof(f4813,plain,(
  ~spl4_22 | ~spl4_72),
  inference(avatar_split_clause,[],[f4812,f910,f700])).
fof(f4812,plain,(
  e1 != op(e4,e0) | ~spl4_72),
  inference(forward_demodulation,[],[f364,f912])).
fof(f4800,plain,(
  spl4_4 | ~spl4_5 | ~spl4_182),
  inference(avatar_split_clause,[],[f4797,f1508,f628,f624])).
fof(f4797,plain,(
  e3 = op(e4,e4) | (~spl4_5 | ~spl4_182)),
  inference(backward_demodulation,[],[f630,f1509])).
fof(f4799,plain,(
  ~spl4_2 | ~spl4_5 | spl4_158),
  inference(avatar_split_clause,[],[f4795,f1364,f628,f616])).
fof(f4795,plain,(
  e1 != op(e4,e4) | (~spl4_5 | spl4_158)),
  inference(backward_demodulation,[],[f630,f1366])).
fof(f4792,plain,(
  ~spl4_12 | ~spl4_51 | ~spl4_119),
  inference(avatar_contradiction_clause,[],[f4791])).
fof(f4791,plain,(
  $false | (~spl4_12 | ~spl4_51 | ~spl4_119)),
  inference(subsumption_resolution,[],[f4790,f465])).
fof(f4790,plain,(
  e3 = e4 | (~spl4_12 | ~spl4_51 | ~spl4_119)),
  inference(forward_demodulation,[],[f4786,f1109])).
fof(f4786,plain,(
  e4 = op(e0,e1) | (~spl4_12 | ~spl4_51)),
  inference(backward_demodulation,[],[f660,f4705])).
fof(f4705,plain,(
  e4 = op(e0,op(e4,e2)) | ~spl4_51),
  inference(backward_demodulation,[],[f824,f353])).
fof(f4789,plain,(
  ~spl4_12 | ~spl4_51 | spl4_98),
  inference(avatar_contradiction_clause,[],[f4788])).
fof(f4788,plain,(
  $false | (~spl4_12 | ~spl4_51 | spl4_98)),
  inference(subsumption_resolution,[],[f4785,f1020])).
fof(f4785,plain,(
  e2 = op(e1,e0) | (~spl4_12 | ~spl4_51)),
  inference(backward_demodulation,[],[f660,f4704])).
fof(f4704,plain,(
  e2 = op(op(e4,e2),e0) | ~spl4_51),
  inference(backward_demodulation,[],[f824,f345])).
fof(f4787,plain,(
  ~spl4_2 | ~spl4_12),
  inference(avatar_split_clause,[],[f4781,f658,f616])).
fof(f4781,plain,(
  e1 != op(e4,e4) | ~spl4_12),
  inference(backward_demodulation,[],[f660,f454])).
fof(f4764,plain,(
  ~spl4_30 | ~spl4_35),
  inference(avatar_split_clause,[],[f4757,f754,f733])).
fof(f4757,plain,(
  e4 != op(e3,e4) | ~spl4_35),
  inference(backward_demodulation,[],[f756,f445])).
fof(f4755,plain,(
  ~spl4_36 | ~spl4_39),
  inference(avatar_contradiction_clause,[],[f4754])).
fof(f4754,plain,(
  $false | (~spl4_36 | ~spl4_39)),
  inference(subsumption_resolution,[],[f4753,f458])).
fof(f4753,plain,(
  e0 = e3 | (~spl4_36 | ~spl4_39)),
  inference(backward_demodulation,[],[f761,f773])).
fof(f4750,plain,(
  spl4_34 | ~spl4_39 | ~spl4_59),
  inference(avatar_split_clause,[],[f4746,f855,f771,f750])).
fof(f4746,plain,(
  e3 = op(e3,e3) | (~spl4_39 | ~spl4_59)),
  inference(backward_demodulation,[],[f773,f4266])).
fof(f4266,plain,(
  e3 = op(e3,op(e3,e2)) | ~spl4_59),
  inference(backward_demodulation,[],[f857,f348])).
fof(f4749,plain,(
  spl4_33 | ~spl4_39 | ~spl4_59),
  inference(avatar_contradiction_clause,[],[f4748])).
fof(f4748,plain,(
  $false | (spl4_33 | ~spl4_39 | ~spl4_59)),
  inference(subsumption_resolution,[],[f4745,f747])).
fof(f4745,plain,(
  e2 = op(e3,e3) | (~spl4_39 | ~spl4_59)),
  inference(backward_demodulation,[],[f773,f4265])).
fof(f4265,plain,(
  e2 = op(op(e3,e2),e3) | ~spl4_59),
  inference(backward_demodulation,[],[f857,f344])).
fof(f4740,plain,(
  ~spl4_42 | ~spl4_43),
  inference(avatar_contradiction_clause,[],[f4739])).
fof(f4739,plain,(
  $false | (~spl4_42 | ~spl4_43)),
  inference(subsumption_resolution,[],[f4738,f460])).
fof(f4738,plain,(
  e1 = e2 | (~spl4_42 | ~spl4_43)),
  inference(backward_demodulation,[],[f786,f790])).
fof(f4737,plain,(
  ~spl4_43 | ~spl4_82 | ~spl4_88),
  inference(avatar_contradiction_clause,[],[f4736])).
fof(f4736,plain,(
  $false | (~spl4_43 | ~spl4_82 | ~spl4_88)),
  inference(subsumption_resolution,[],[f4735,f463])).
fof(f4735,plain,(
  e2 = e3 | (~spl4_43 | ~spl4_82 | ~spl4_88)),
  inference(forward_demodulation,[],[f4731,f979])).
fof(f4731,plain,(
  e3 = op(e1,e2) | (~spl4_43 | ~spl4_82)),
  inference(backward_demodulation,[],[f790,f4538])).
fof(f4538,plain,(
  e3 = op(e1,op(e3,e1)) | ~spl4_82),
  inference(forward_demodulation,[],[f347,f954])).
fof(f4734,plain,(
  ~spl4_43 | spl4_67 | ~spl4_82),
  inference(avatar_contradiction_clause,[],[f4733])).
fof(f4733,plain,(
  $false | (~spl4_43 | spl4_67 | ~spl4_82)),
  inference(subsumption_resolution,[],[f4730,f890])).
fof(f4730,plain,(
  e1 = op(e2,e1) | (~spl4_43 | ~spl4_82)),
  inference(backward_demodulation,[],[f790,f4537])).
fof(f4537,plain,(
  e1 = op(op(e3,e1),e1) | ~spl4_82),
  inference(forward_demodulation,[],[f339,f954])).
fof(f4724,plain,(
  ~spl4_48 | ~spl4_106 | spl4_114),
  inference(avatar_contradiction_clause,[],[f4723])).
fof(f4723,plain,(
  $false | (~spl4_48 | ~spl4_106 | spl4_114)),
  inference(subsumption_resolution,[],[f4718,f1087])).
fof(f1087,plain,(
  e3 != op(e0,e2) | spl4_114),
  inference(avatar_component_clause,[],[f1086])).
fof(f4718,plain,(
  e3 = op(e0,e2) | (~spl4_48 | ~spl4_106)),
  inference(backward_demodulation,[],[f811,f4014])).
fof(f4014,plain,(
  e3 = op(e0,op(e3,e0)) | ~spl4_106),
  inference(forward_demodulation,[],[f346,f1055])).
fof(f4695,plain,(
  ~spl4_38 | ~spl4_63),
  inference(avatar_split_clause,[],[f4683,f872,f767])).
fof(f4683,plain,(
  e2 != op(e3,e2) | ~spl4_63),
  inference(backward_demodulation,[],[f874,f383])).
fof(f4680,plain,(
  spl4_12 | ~spl4_70 | ~spl4_88),
  inference(avatar_split_clause,[],[f4678,f977,f901,f658])).
fof(f4678,plain,(
  e1 = op(e4,e2) | (~spl4_70 | ~spl4_88)),
  inference(backward_demodulation,[],[f903,f4216])).
fof(f4216,plain,(
  e1 = op(op(e2,e1),e2) | ~spl4_88),
  inference(backward_demodulation,[],[f979,f338])).
fof(f4674,plain,(
  ~spl4_72 | ~spl4_74),
  inference(avatar_contradiction_clause,[],[f4673])).
fof(f4673,plain,(
  $false | (~spl4_72 | ~spl4_74)),
  inference(subsumption_resolution,[],[f4672,f461])).
fof(f4672,plain,(
  e1 = e3 | (~spl4_72 | ~spl4_74)),
  inference(backward_demodulation,[],[f912,f920])).
fof(f4661,plain,(
  ~spl4_19 | spl4_32 | ~spl4_79),
  inference(avatar_contradiction_clause,[],[f4660])).
fof(f4660,plain,(
  $false | (~spl4_19 | spl4_32 | ~spl4_79)),
  inference(subsumption_resolution,[],[f4657,f743])).
fof(f4657,plain,(
  e1 = op(e3,e3) | (~spl4_19 | ~spl4_79)),
  inference(backward_demodulation,[],[f941,f4548])).
fof(f4548,plain,(
  e1 = op(e3,op(e1,e4)) | ~spl4_19),
  inference(forward_demodulation,[],[f340,f689])).
fof(f4659,plain,(
  spl4_35 | ~spl4_19 | ~spl4_79),
  inference(avatar_split_clause,[],[f4656,f939,f687,f754])).
fof(f4656,plain,(
  e4 = op(e3,e3) | (~spl4_19 | ~spl4_79)),
  inference(backward_demodulation,[],[f941,f4545])).
fof(f4545,plain,(
  e4 = op(op(e1,e4),e3) | ~spl4_19),
  inference(forward_demodulation,[],[f352,f689])).
fof(f4647,plain,(
  ~spl4_91 | spl4_96 | ~spl4_148),
  inference(avatar_contradiction_clause,[],[f4646])).
fof(f4646,plain,(
  $false | (~spl4_91 | spl4_96 | ~spl4_148)),
  inference(subsumption_resolution,[],[f4639,f1012])).
fof(f1012,plain,(
  e0 != op(e1,e0) | spl4_96),
  inference(avatar_component_clause,[],[f1011])).
fof(f4639,plain,(
  e0 = op(e1,e0) | (~spl4_91 | ~spl4_148)),
  inference(backward_demodulation,[],[f992,f1310])).
fof(f4644,plain,(
  ~spl4_41 | ~spl4_91),
  inference(avatar_split_clause,[],[f4633,f990,f780])).
fof(f4633,plain,(
  e0 != op(e3,e1) | ~spl4_91),
  inference(backward_demodulation,[],[f992,f371])).
fof(f4618,plain,(
  spl4_75 | ~spl4_21 | ~spl4_103),
  inference(avatar_split_clause,[],[f4614,f1040,f696,f922])).
fof(f4614,plain,(
  e4 = op(e2,e0) | (~spl4_21 | ~spl4_103)),
  inference(backward_demodulation,[],[f1042,f4555])).
fof(f4555,plain,(
  e4 = op(op(e0,e4),e0) | ~spl4_21),
  inference(forward_demodulation,[],[f351,f698])).
fof(f4617,plain,(
  ~spl4_21 | ~spl4_103 | spl4_111),
  inference(avatar_contradiction_clause,[],[f4616])).
fof(f4616,plain,(
  $false | (~spl4_21 | ~spl4_103 | spl4_111)),
  inference(subsumption_resolution,[],[f4613,f1075])).
fof(f1075,plain,(
  e0 != op(e0,e2) | spl4_111),
  inference(avatar_component_clause,[],[f1074])).
fof(f4613,plain,(
  e0 = op(e0,e2) | (~spl4_21 | ~spl4_103)),
  inference(backward_demodulation,[],[f1042,f4554])).
fof(f4554,plain,(
  e0 = op(e0,op(e0,e4)) | ~spl4_21),
  inference(forward_demodulation,[],[f335,f698])).
fof(f4566,plain,(
  ~spl4_123 | ~spl4_148 | spl4_157),
  inference(avatar_split_clause,[],[f4565,f1353,f1309,f1124])).
fof(f4565,plain,(
  op(e0,e0) != e2 | (~spl4_148 | spl4_157)),
  inference(forward_demodulation,[],[f1355,f1310])).
fof(f4558,plain,(
  ~spl4_113 | ~spl4_88),
  inference(avatar_split_clause,[],[f4557,f977,f1082])).
fof(f4557,plain,(
  e2 != op(e0,e2) | ~spl4_88),
  inference(forward_demodulation,[],[f376,f979])).
fof(f4547,plain,(
  ~spl4_77 | ~spl4_82),
  inference(avatar_split_clause,[],[f4546,f952,f931])).
fof(f4546,plain,(
  e1 != op(e1,e4) | ~spl4_82),
  inference(forward_demodulation,[],[f425,f954])).
fof(f4542,plain,(
  ~spl4_64 | ~spl4_59),
  inference(avatar_split_clause,[],[f4541,f855,f876])).
fof(f4541,plain,(
  e3 != op(e2,e2) | ~spl4_59),
  inference(forward_demodulation,[],[f433,f857])).
fof(f4534,plain,(
  ~spl4_4 | ~spl4_19),
  inference(avatar_split_clause,[],[f4507,f687,f624])).
fof(f4507,plain,(
  e3 != op(e4,e4) | ~spl4_19),
  inference(backward_demodulation,[],[f689,f452])).
fof(f4523,plain,(
  ~spl4_15 | ~spl4_53),
  inference(avatar_contradiction_clause,[],[f4522])).
fof(f4522,plain,(
  $false | (~spl4_15 | ~spl4_53)),
  inference(subsumption_resolution,[],[f4521,f464])).
fof(f4521,plain,(
  e2 = e4 | (~spl4_15 | ~spl4_53)),
  inference(forward_demodulation,[],[f4517,f832])).
fof(f4517,plain,(
  e4 = op(e2,e4) | (~spl4_15 | ~spl4_53)),
  inference(backward_demodulation,[],[f672,f4456])).
fof(f4456,plain,(
  e4 = op(e2,op(e4,e2)) | ~spl4_53),
  inference(backward_demodulation,[],[f832,f353])).
fof(f4514,plain,(
  ~spl4_19 | ~spl4_76 | ~spl4_106),
  inference(avatar_contradiction_clause,[],[f4513])).
fof(f4513,plain,(
  $false | (~spl4_19 | ~spl4_76 | ~spl4_106)),
  inference(subsumption_resolution,[],[f4512,f459])).
fof(f4512,plain,(
  e0 = e4 | (~spl4_19 | ~spl4_76 | ~spl4_106)),
  inference(forward_demodulation,[],[f4510,f1055])).
fof(f4510,plain,(
  e4 = op(e0,e3) | (~spl4_19 | ~spl4_76)),
  inference(backward_demodulation,[],[f689,f4313])).
fof(f4313,plain,(
  e4 = op(e0,op(e4,e1)) | ~spl4_76),
  inference(forward_demodulation,[],[f352,f929])).
fof(f4502,plain,(
  spl4_15 | ~spl4_8 | ~spl4_30),
  inference(avatar_split_clause,[],[f4499,f733,f641,f670])).
fof(f4499,plain,(
  e4 = op(e4,e2) | (~spl4_8 | ~spl4_30)),
  inference(backward_demodulation,[],[f735,f4320])).
fof(f735,plain,(
  e4 = op(e3,e4) | ~spl4_30),
  inference(avatar_component_clause,[],[f733])).
fof(f4501,plain,(
  ~spl4_5 | ~spl4_30),
  inference(avatar_split_clause,[],[f4498,f733,f628])).
fof(f4498,plain,(
  e4 != op(e4,e4) | ~spl4_30),
  inference(backward_demodulation,[],[f735,f405])).
fof(f4494,plain,(
  spl4_31 | ~spl4_34 | ~spl4_133),
  inference(avatar_contradiction_clause,[],[f4493])).
fof(f4493,plain,(
  $false | (spl4_31 | ~spl4_34 | ~spl4_133)),
  inference(subsumption_resolution,[],[f4490,f739])).
fof(f4490,plain,(
  e0 = op(e3,e3) | (~spl4_34 | ~spl4_133)),
  inference(backward_demodulation,[],[f752,f1242])).
fof(f4486,plain,(
  ~spl4_38 | ~spl4_59),
  inference(avatar_contradiction_clause,[],[f4485])).
fof(f4485,plain,(
  $false | (~spl4_38 | ~spl4_59)),
  inference(subsumption_resolution,[],[f4484,f463])).
fof(f4484,plain,(
  e2 = e3 | (~spl4_38 | ~spl4_59)),
  inference(forward_demodulation,[],[f4480,f769])).
fof(f4480,plain,(
  e3 = op(e3,e2) | (~spl4_38 | ~spl4_59)),
  inference(backward_demodulation,[],[f769,f4266])).
fof(f4473,plain,(
  ~spl4_47 | ~spl4_50),
  inference(avatar_contradiction_clause,[],[f4472])).
fof(f4472,plain,(
  $false | (~spl4_47 | ~spl4_50)),
  inference(subsumption_resolution,[],[f4471,f462])).
fof(f4471,plain,(
  e1 = e4 | (~spl4_47 | ~spl4_50)),
  inference(backward_demodulation,[],[f807,f819])).
fof(f4468,plain,(
  ~spl4_30 | ~spl4_50),
  inference(avatar_split_clause,[],[f4464,f817,f733])).
fof(f4464,plain,(
  e4 != op(e3,e4) | ~spl4_50),
  inference(backward_demodulation,[],[f819,f439])).
fof(f4454,plain,(
  ~spl4_67 | ~spl4_88),
  inference(avatar_contradiction_clause,[],[f4453])).
fof(f4453,plain,(
  $false | (~spl4_67 | ~spl4_88)),
  inference(subsumption_resolution,[],[f4452,f460])).
fof(f4452,plain,(
  e1 = e2 | (~spl4_67 | ~spl4_88)),
  inference(forward_demodulation,[],[f4447,f891])).
fof(f4447,plain,(
  e2 = op(e2,e1) | (~spl4_67 | ~spl4_88)),
  inference(backward_demodulation,[],[f891,f4217])).
fof(f4217,plain,(
  e2 = op(e2,op(e2,e1)) | ~spl4_88),
  inference(backward_demodulation,[],[f979,f342])).
fof(f4449,plain,(
  ~spl4_52 | ~spl4_67),
  inference(avatar_split_clause,[],[f4442,f889,f826])).
fof(f4442,plain,(
  e1 != op(e2,e4) | ~spl4_67),
  inference(backward_demodulation,[],[f891,f432])).
fof(f4448,plain,(
  ~spl4_17 | ~spl4_67),
  inference(avatar_split_clause,[],[f4441,f889,f679])).
fof(f4441,plain,(
  e1 != op(e4,e1) | ~spl4_67),
  inference(backward_demodulation,[],[f891,f374])).
fof(f4436,plain,(
  ~spl4_75 | ~spl4_76 | ~spl4_112),
  inference(avatar_contradiction_clause,[],[f4435])).
fof(f4435,plain,(
  $false | (~spl4_75 | ~spl4_76 | ~spl4_112)),
  inference(subsumption_resolution,[],[f4434,f457])).
fof(f4434,plain,(
  e0 = e2 | (~spl4_75 | ~spl4_76 | ~spl4_112)),
  inference(forward_demodulation,[],[f4427,f929])).
fof(f4427,plain,(
  e2 = op(e1,e4) | (~spl4_75 | ~spl4_112)),
  inference(backward_demodulation,[],[f924,f3861])).
fof(f3861,plain,(
  e2 = op(e1,op(e2,e0)) | ~spl4_112),
  inference(forward_demodulation,[],[f341,f1080])).
fof(f4432,plain,(
  ~spl4_50 | ~spl4_75),
  inference(avatar_split_clause,[],[f4424,f922,f817])).
fof(f4424,plain,(
  e4 != op(e3,e0) | ~spl4_75),
  inference(backward_demodulation,[],[f924,f363])).
fof(f4423,plain,(
  ~spl4_41 | ~spl4_82 | spl4_117),
  inference(avatar_contradiction_clause,[],[f4422])).
fof(f4422,plain,(
  $false | (~spl4_41 | ~spl4_82 | spl4_117)),
  inference(subsumption_resolution,[],[f4420,f1100])).
fof(f1100,plain,(
  e1 != op(e0,e1) | spl4_117),
  inference(avatar_component_clause,[],[f1099])).
fof(f4420,plain,(
  e1 = op(e0,e1) | (~spl4_41 | ~spl4_82)),
  inference(backward_demodulation,[],[f954,f4169])).
fof(f4169,plain,(
  e1 = op(e0,op(e1,e3)) | ~spl4_41),
  inference(forward_demodulation,[],[f339,f782])).
fof(f4414,plain,(
  ~spl4_88 | ~spl4_89),
  inference(avatar_contradiction_clause,[],[f4413])).
fof(f4413,plain,(
  $false | (~spl4_88 | ~spl4_89)),
  inference(subsumption_resolution,[],[f4412,f463])).
fof(f4412,plain,(
  e2 = e3 | (~spl4_88 | ~spl4_89)),
  inference(forward_demodulation,[],[f983,f979])).
fof(f4410,plain,(
  ~spl4_85 | ~spl4_95),
  inference(avatar_split_clause,[],[f4402,f1006,f964])).
fof(f4402,plain,(
  e4 != op(e1,e3) | ~spl4_95),
  inference(backward_demodulation,[],[f1008,f421])).
fof(f4373,plain,(
  ~spl4_148 | ~spl4_176),
  inference(avatar_contradiction_clause,[],[f4372])).
fof(f4372,plain,(
  $false | (~spl4_148 | ~spl4_176)),
  inference(subsumption_resolution,[],[f4371,f457])).
fof(f4371,plain,(
  e0 = e2 | (~spl4_148 | ~spl4_176)),
  inference(backward_demodulation,[],[f1310,f1464])).
fof(f4367,plain,(
  ~spl4_34 | spl4_132 | ~spl4_182),
  inference(avatar_split_clause,[],[f4365,f1508,f1236,f750])).
fof(f4365,plain,(
  e3 != op(e3,e3) | (spl4_132 | ~spl4_182)),
  inference(backward_demodulation,[],[f1509,f1238])).
fof(f1238,plain,(
  e3 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | spl4_132),
  inference(avatar_component_clause,[],[f1236])).
fof(f4363,plain,(
  spl4_75 | ~spl4_61 | ~spl4_187),
  inference(avatar_split_clause,[],[f4362,f1553,f864,f922])).
fof(f4362,plain,(
  e4 = op(e2,e0) | (~spl4_61 | ~spl4_187)),
  inference(forward_demodulation,[],[f1554,f866])).
fof(f4361,plain,(
  ~spl4_123 | ~spl4_118),
  inference(avatar_split_clause,[],[f4360,f1103,f1124])).
fof(f4360,plain,(
  op(e0,e0) != e2 | ~spl4_118),
  inference(forward_demodulation,[],[f406,f1105])).
fof(f4342,plain,(
  ~spl4_95 | ~spl4_164 | spl4_166),
  inference(avatar_split_clause,[],[f4185,f1405,f1397,f1006])).
fof(f4185,plain,(
  e4 != op(e1,e1) | (~spl4_164 | spl4_166)),
  inference(forward_demodulation,[],[f1407,f1398])).
fof(f4306,plain,(
  ~spl4_8 | ~spl4_28 | ~spl4_61),
  inference(avatar_contradiction_clause,[],[f4305])).
fof(f4305,plain,(
  $false | (~spl4_8 | ~spl4_28 | ~spl4_61)),
  inference(subsumption_resolution,[],[f4304,f459])).
fof(f4304,plain,(
  e0 = e4 | (~spl4_8 | ~spl4_28 | ~spl4_61)),
  inference(forward_demodulation,[],[f4300,f866])).
fof(f4300,plain,(
  e4 = op(e2,e2) | (~spl4_8 | ~spl4_28)),
  inference(backward_demodulation,[],[f643,f3477])).
fof(f3477,plain,(
  e4 = op(e2,op(e4,e3)) | ~spl4_28),
  inference(forward_demodulation,[],[f354,f727])).
fof(f4296,plain,(
  ~spl4_14 | ~spl4_41 | ~spl4_52),
  inference(avatar_contradiction_clause,[],[f4295])).
fof(f4295,plain,(
  $false | (~spl4_14 | ~spl4_41 | ~spl4_52)),
  inference(subsumption_resolution,[],[f4294,f457])).
fof(f4294,plain,(
  e0 = e2 | (~spl4_14 | ~spl4_41 | ~spl4_52)),
  inference(forward_demodulation,[],[f4291,f782])).
fof(f4291,plain,(
  e2 = op(e3,e1) | (~spl4_14 | ~spl4_52)),
  inference(backward_demodulation,[],[f668,f3672])).
fof(f3672,plain,(
  e2 = op(op(e4,e2),e1) | ~spl4_52),
  inference(backward_demodulation,[],[f828,f345])).
fof(f4293,plain,(
  ~spl4_9 | ~spl4_14),
  inference(avatar_split_clause,[],[f4289,f666,f645])).
fof(f4289,plain,(
  e3 != op(e4,e3) | ~spl4_14),
  inference(backward_demodulation,[],[f668,f453])).
fof(f4288,plain,(
  ~spl4_28 | ~spl4_40 | ~spl4_59),
  inference(avatar_contradiction_clause,[],[f4287])).
fof(f4287,plain,(
  $false | (~spl4_28 | ~spl4_40 | ~spl4_59)),
  inference(subsumption_resolution,[],[f4286,f463])).
fof(f4286,plain,(
  e2 = e3 | (~spl4_28 | ~spl4_40 | ~spl4_59)),
  inference(forward_demodulation,[],[f4284,f727])).
fof(f4284,plain,(
  e3 = op(e3,e4) | (~spl4_40 | ~spl4_59)),
  inference(backward_demodulation,[],[f777,f4266])).
fof(f4285,plain,(
  spl4_8 | ~spl4_40 | ~spl4_59),
  inference(avatar_split_clause,[],[f4283,f855,f775,f641])).
fof(f4283,plain,(
  e2 = op(e4,e3) | (~spl4_40 | ~spl4_59)),
  inference(backward_demodulation,[],[f777,f4265])).
fof(f4280,plain,(
  ~spl4_49 | ~spl4_106),
  inference(avatar_contradiction_clause,[],[f4279])).
fof(f4279,plain,(
  $false | (~spl4_49 | ~spl4_106)),
  inference(subsumption_resolution,[],[f4278,f458])).
fof(f4278,plain,(
  e0 = e3 | (~spl4_49 | ~spl4_106)),
  inference(forward_demodulation,[],[f4274,f1055])).
fof(f4274,plain,(
  e3 = op(e0,e3) | (~spl4_49 | ~spl4_106)),
  inference(backward_demodulation,[],[f815,f4014])).
fof(f4264,plain,(
  ~spl4_52 | ~spl4_70 | ~spl4_88),
  inference(avatar_contradiction_clause,[],[f4263])).
fof(f4263,plain,(
  $false | (~spl4_52 | ~spl4_70 | ~spl4_88)),
  inference(subsumption_resolution,[],[f4262,f460])).
fof(f4262,plain,(
  e1 = e2 | (~spl4_52 | ~spl4_70 | ~spl4_88)),
  inference(forward_demodulation,[],[f4258,f828])).
fof(f4258,plain,(
  e2 = op(e2,e4) | (~spl4_70 | ~spl4_88)),
  inference(backward_demodulation,[],[f903,f4217])).
fof(f4240,plain,(
  ~spl4_41 | ~spl4_85 | spl4_102),
  inference(avatar_contradiction_clause,[],[f4239])).
fof(f4239,plain,(
  $false | (~spl4_41 | ~spl4_85 | spl4_102)),
  inference(subsumption_resolution,[],[f4230,f1037])).
fof(f4230,plain,(
  e1 = op(e0,e4) | (~spl4_41 | ~spl4_85)),
  inference(backward_demodulation,[],[f966,f4169])).
fof(f4204,plain,(
  ~spl4_120 | ~spl4_125),
  inference(avatar_split_clause,[],[f4200,f1132,f1111])).
fof(f4200,plain,(
  e4 != op(e0,e1) | ~spl4_125),
  inference(backward_demodulation,[],[f1134,f406])).
fof(f4187,plain,(
  spl4_123 | ~spl4_61),
  inference(avatar_split_clause,[],[f3930,f864,f1124])).
fof(f3930,plain,(
  op(e0,e0) = e2 | ~spl4_61),
  inference(backward_demodulation,[],[f866,f343])).
fof(f4167,plain,(
  ~spl4_84 | ~spl4_94),
  inference(avatar_split_clause,[],[f4166,f1002,f960])).
fof(f4166,plain,(
  e3 != op(e1,e3) | ~spl4_94),
  inference(forward_demodulation,[],[f421,f1004])).
fof(f4160,plain,(
  ~spl4_75 | ~spl4_61 | spl4_187),
  inference(avatar_split_clause,[],[f4058,f1553,f864,f922])).
fof(f4058,plain,(
  e4 != op(e2,e0) | (~spl4_61 | spl4_187)),
  inference(forward_demodulation,[],[f1555,f866])).
fof(f4151,plain,(
  spl4_50 | ~spl4_21 | ~spl4_104),
  inference(avatar_split_clause,[],[f4122,f1044,f696,f817])).
fof(f4122,plain,(
  e4 = op(e3,e0) | (~spl4_21 | ~spl4_104)),
  inference(backward_demodulation,[],[f698,f3838])).
fof(f3838,plain,(
  e4 = op(e3,op(e4,e0)) | ~spl4_104),
  inference(forward_demodulation,[],[f351,f1046])).
fof(f4146,plain,(
  ~spl4_118 | ~spl4_32 | ~spl4_41 | spl4_134),
  inference(avatar_split_clause,[],[f4145,f1245,f780,f742,f1103])).
fof(f4145,plain,(
  e2 != op(e0,e1) | (~spl4_32 | ~spl4_41 | spl4_134)),
  inference(forward_demodulation,[],[f3875,f782])).
fof(f3875,plain,(
  e2 != op(op(e3,e1),e1) | (~spl4_32 | spl4_134)),
  inference(forward_demodulation,[],[f1247,f744])).
fof(f4113,plain,(
  ~spl4_41 | ~spl4_43),
  inference(avatar_contradiction_clause,[],[f4112])).
fof(f4112,plain,(
  $false | (~spl4_41 | ~spl4_43)),
  inference(subsumption_resolution,[],[f4111,f457])).
fof(f4111,plain,(
  e0 = e2 | (~spl4_41 | ~spl4_43)),
  inference(backward_demodulation,[],[f782,f790])).
fof(f4088,plain,(
  ~spl4_52 | ~spl4_53),
  inference(avatar_contradiction_clause,[],[f4087])).
fof(f4087,plain,(
  $false | (~spl4_52 | ~spl4_53)),
  inference(subsumption_resolution,[],[f4086,f460])).
fof(f4086,plain,(
  e1 = e2 | (~spl4_52 | ~spl4_53)),
  inference(forward_demodulation,[],[f832,f828])).
fof(f4072,plain,(
  ~spl4_60 | ~spl4_83 | ~spl4_94 | spl4_152),
  inference(avatar_contradiction_clause,[],[f4071])).
fof(f4071,plain,(
  $false | (~spl4_60 | ~spl4_83 | ~spl4_94 | spl4_152)),
  inference(subsumption_resolution,[],[f4070,f861])).
fof(f4070,plain,(
  e4 != op(e2,e3) | (~spl4_83 | ~spl4_94 | spl4_152)),
  inference(forward_demodulation,[],[f4069,f958])).
fof(f4069,plain,(
  e4 != op(op(e1,e3),e3) | (~spl4_94 | spl4_152)),
  inference(forward_demodulation,[],[f1329,f1004])).
fof(f4039,plain,(
  ~spl4_1 | ~spl4_76),
  inference(avatar_split_clause,[],[f4038,f927,f612])).
fof(f4038,plain,(
  e0 != op(e4,e4) | ~spl4_76),
  inference(forward_demodulation,[],[f402,f929])).
fof(f4026,plain,(
  ~spl4_70 | ~spl4_120),
  inference(avatar_split_clause,[],[f4025,f1111,f901])).
fof(f4025,plain,(
  e4 != op(e2,e1) | ~spl4_120),
  inference(forward_demodulation,[],[f367,f1113])).
fof(f4020,plain,(
  spl4_68 | ~spl4_13 | ~spl4_52),
  inference(avatar_split_clause,[],[f3734,f826,f662,f893])).
fof(f3734,plain,(
  e2 = op(e2,e1) | (~spl4_13 | ~spl4_52)),
  inference(backward_demodulation,[],[f664,f3672])).
fof(f4013,plain,(
  ~spl4_48 | ~spl4_123),
  inference(avatar_split_clause,[],[f4012,f1124,f809])).
fof(f4012,plain,(
  e2 != op(e3,e0) | ~spl4_123),
  inference(forward_demodulation,[],[f358,f1126])).
fof(f4011,plain,(
  ~spl4_43 | ~spl4_32 | spl4_175),
  inference(avatar_split_clause,[],[f3791,f1457,f742,f788])).
fof(f3791,plain,(
  e2 != op(e3,e1) | (~spl4_32 | spl4_175)),
  inference(forward_demodulation,[],[f1459,f744])).
fof(f4009,plain,(
  ~spl4_41 | ~spl4_32 | spl4_133),
  inference(avatar_split_clause,[],[f3794,f1241,f742,f780])).
fof(f3794,plain,(
  e0 != op(e3,e1) | (~spl4_32 | spl4_133)),
  inference(forward_demodulation,[],[f1243,f744])).
fof(f4008,plain,(
  spl4_41 | ~spl4_74 | ~spl4_112),
  inference(avatar_split_clause,[],[f3925,f1078,f918,f780])).
fof(f3925,plain,(
  e0 = op(e3,e1) | (~spl4_74 | ~spl4_112)),
  inference(backward_demodulation,[],[f920,f3862])).
fof(f3862,plain,(
  e0 = op(op(e2,e0),e1) | ~spl4_112),
  inference(forward_demodulation,[],[f333,f1080])).
fof(f4007,plain,(
  spl4_39 | ~spl4_9 | ~spl4_28),
  inference(avatar_split_clause,[],[f3984,f725,f645,f771])).
fof(f3984,plain,(
  e3 = op(e3,e2) | (~spl4_9 | ~spl4_28)),
  inference(backward_demodulation,[],[f647,f3474])).
fof(f3474,plain,(
  e3 = op(op(e4,e3),e2) | ~spl4_28),
  inference(forward_demodulation,[],[f350,f727])).
fof(f3999,plain,(
  ~spl4_4 | ~spl4_32),
  inference(avatar_contradiction_clause,[],[f3998])).
fof(f3998,plain,(
  $false | (~spl4_4 | ~spl4_32)),
  inference(subsumption_resolution,[],[f3997,f462])).
fof(f3997,plain,(
  e1 = e4 | (~spl4_4 | ~spl4_32)),
  inference(forward_demodulation,[],[f3989,f744])).
fof(f3989,plain,(
  e4 = op(e3,e3) | ~spl4_4),
  inference(backward_demodulation,[],[f626,f355])).
fof(f3981,plain,(
  ~spl4_16 | ~spl4_17),
  inference(avatar_contradiction_clause,[],[f3980])).
fof(f3980,plain,(
  $false | (~spl4_16 | ~spl4_17)),
  inference(subsumption_resolution,[],[f3979,f456])).
fof(f3979,plain,(
  e0 = e1 | (~spl4_16 | ~spl4_17)),
  inference(backward_demodulation,[],[f677,f681])).
fof(f3966,plain,(
  spl4_3 | ~spl4_40 | ~spl4_60),
  inference(avatar_contradiction_clause,[],[f3965])).
fof(f3965,plain,(
  $false | (spl4_3 | ~spl4_40 | ~spl4_60)),
  inference(subsumption_resolution,[],[f3963,f621])).
fof(f3963,plain,(
  e2 = op(e4,e4) | (~spl4_40 | ~spl4_60)),
  inference(backward_demodulation,[],[f777,f3506])).
fof(f3506,plain,(
  e2 = op(op(e3,e2),e4) | ~spl4_60),
  inference(forward_demodulation,[],[f344,f861])).
fof(f3964,plain,(
  spl4_4 | ~spl4_40 | ~spl4_60),
  inference(avatar_split_clause,[],[f3962,f859,f775,f624])).
fof(f3962,plain,(
  e3 = op(e4,e4) | (~spl4_40 | ~spl4_60)),
  inference(backward_demodulation,[],[f777,f3505])).
fof(f3505,plain,(
  e3 = op(e4,op(e3,e2)) | ~spl4_60),
  inference(forward_demodulation,[],[f348,f861])).
fof(f3958,plain,(
  ~spl4_42 | ~spl4_83 | spl4_87),
  inference(avatar_contradiction_clause,[],[f3957])).
fof(f3957,plain,(
  $false | (~spl4_42 | ~spl4_83 | spl4_87)),
  inference(subsumption_resolution,[],[f3953,f974])).
fof(f974,plain,(
  e1 != op(e1,e2) | spl4_87),
  inference(avatar_component_clause,[],[f973])).
fof(f3953,plain,(
  e1 = op(e1,e2) | (~spl4_42 | ~spl4_83)),
  inference(backward_demodulation,[],[f786,f3802])).
fof(f3802,plain,(
  e1 = op(op(e3,e1),e2) | ~spl4_83),
  inference(forward_demodulation,[],[f339,f958])).
fof(f3927,plain,(
  ~spl4_64 | ~spl4_74),
  inference(avatar_split_clause,[],[f3922,f918,f876])).
fof(f3922,plain,(
  e3 != op(e2,e2) | ~spl4_74),
  inference(backward_demodulation,[],[f920,f427])).
fof(f3926,plain,(
  ~spl4_24 | ~spl4_74),
  inference(avatar_split_clause,[],[f3921,f918,f708])).
fof(f3921,plain,(
  e3 != op(e4,e0) | ~spl4_74),
  inference(backward_demodulation,[],[f920,f364])).
fof(f3920,plain,(
  ~spl4_76 | ~spl4_80),
  inference(avatar_contradiction_clause,[],[f3919])).
fof(f3919,plain,(
  $false | (~spl4_76 | ~spl4_80)),
  inference(subsumption_resolution,[],[f3918,f459])).
fof(f3918,plain,(
  e0 = e4 | (~spl4_76 | ~spl4_80)),
  inference(backward_demodulation,[],[f929,f945])).
fof(f3914,plain,(
  ~spl4_40 | ~spl4_90),
  inference(avatar_split_clause,[],[f3911,f985,f775])).
fof(f3911,plain,(
  e4 != op(e3,e2) | ~spl4_90),
  inference(backward_demodulation,[],[f987,f381])).
fof(f3907,plain,(
  ~spl4_83 | ~spl4_94 | spl4_176),
  inference(avatar_contradiction_clause,[],[f3906])).
fof(f3906,plain,(
  $false | (~spl4_83 | ~spl4_94 | spl4_176)),
  inference(subsumption_resolution,[],[f3902,f958])).
fof(f3902,plain,(
  e2 != op(e1,e3) | (~spl4_94 | spl4_176)),
  inference(backward_demodulation,[],[f1004,f1465])).
fof(f3895,plain,(
  ~spl4_112 | ~spl4_114),
  inference(avatar_contradiction_clause,[],[f3894])).
fof(f3894,plain,(
  $false | (~spl4_112 | ~spl4_114)),
  inference(subsumption_resolution,[],[f3893,f461])).
fof(f3893,plain,(
  e1 = e3 | (~spl4_112 | ~spl4_114)),
  inference(forward_demodulation,[],[f1088,f1080])).
fof(f3887,plain,(
  ~spl4_112 | ~spl4_123 | spl4_164),
  inference(avatar_contradiction_clause,[],[f3886])).
fof(f3886,plain,(
  $false | (~spl4_112 | ~spl4_123 | spl4_164)),
  inference(subsumption_resolution,[],[f3881,f1080])).
fof(f3881,plain,(
  e1 != op(e0,e2) | (~spl4_123 | spl4_164)),
  inference(backward_demodulation,[],[f1126,f1399])).
fof(f3837,plain,(
  ~spl4_103 | ~spl4_28),
  inference(avatar_split_clause,[],[f3836,f725,f1040])).
fof(f3836,plain,(
  e2 != op(e0,e4) | ~spl4_28),
  inference(forward_demodulation,[],[f398,f727])).
fof(f3820,plain,(
  ~spl4_90 | ~spl4_80),
  inference(avatar_split_clause,[],[f3819,f943,f985])).
fof(f3819,plain,(
  e4 != op(e1,e2) | ~spl4_80),
  inference(forward_demodulation,[],[f424,f945])).
fof(f3814,plain,(
  spl4_90 | ~spl4_13 | ~spl4_52),
  inference(avatar_split_clause,[],[f3735,f826,f662,f985])).
fof(f3735,plain,(
  e4 = op(e1,e2) | (~spl4_13 | ~spl4_52)),
  inference(backward_demodulation,[],[f664,f3673])).
fof(f3673,plain,(
  e4 = op(e1,op(e4,e2)) | ~spl4_52),
  inference(backward_demodulation,[],[f828,f353])).
fof(f3799,plain,(
  spl4_94 | ~spl4_32),
  inference(avatar_split_clause,[],[f3798,f742,f1002])).
fof(f3798,plain,(
  e3 = op(e1,e1) | ~spl4_32),
  inference(forward_demodulation,[],[f349,f744])).
fof(f3783,plain,(
  ~spl4_17 | ~spl4_80),
  inference(avatar_contradiction_clause,[],[f3782])).
fof(f3782,plain,(
  $false | (~spl4_17 | ~spl4_80)),
  inference(subsumption_resolution,[],[f3781,f462])).
fof(f3781,plain,(
  e1 = e4 | (~spl4_17 | ~spl4_80)),
  inference(forward_demodulation,[],[f3780,f681])).
fof(f3780,plain,(
  e4 = op(e4,e1) | (~spl4_17 | ~spl4_80)),
  inference(forward_demodulation,[],[f3779,f945])).
fof(f3779,plain,(
  e4 = op(op(e1,e4),e1) | ~spl4_17),
  inference(forward_demodulation,[],[f352,f681])).
fof(f3771,plain,(
  ~spl4_4 | spl4_144 | ~spl4_187),
  inference(avatar_split_clause,[],[f3508,f1553,f1289,f624])).
fof(f1289,plain,(
  spl4_144 <=> e3 = op(op(e2,op(e2,e2)),op(e2,op(e2,e2)))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_144])])).
fof(f3508,plain,(
  e3 != op(e4,e4) | (spl4_144 | ~spl4_187)),
  inference(backward_demodulation,[],[f1554,f1291])).
fof(f1291,plain,(
  e3 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | spl4_144),
  inference(avatar_component_clause,[],[f1289])).
fof(f3764,plain,(
  ~spl4_1 | ~spl4_124),
  inference(avatar_contradiction_clause,[],[f3763])).
fof(f3763,plain,(
  $false | (~spl4_1 | ~spl4_124)),
  inference(subsumption_resolution,[],[f3762,f465])).
fof(f3762,plain,(
  e3 = e4 | (~spl4_1 | ~spl4_124)),
  inference(forward_demodulation,[],[f3754,f1130])).
fof(f3754,plain,(
  op(e0,e0) = e4 | ~spl4_1),
  inference(backward_demodulation,[],[f614,f355])).
fof(f3750,plain,(
  ~spl4_10 | ~spl4_11 | ~spl4_28),
  inference(avatar_contradiction_clause,[],[f3749])).
fof(f3749,plain,(
  $false | (~spl4_10 | ~spl4_11 | ~spl4_28)),
  inference(subsumption_resolution,[],[f3748,f458])).
fof(f3748,plain,(
  e0 = e3 | (~spl4_10 | ~spl4_11 | ~spl4_28)),
  inference(forward_demodulation,[],[f3745,f656])).
fof(f3745,plain,(
  e3 = op(e4,e2) | (~spl4_10 | ~spl4_28)),
  inference(backward_demodulation,[],[f651,f3474])).
fof(f3747,plain,(
  ~spl4_5 | ~spl4_10),
  inference(avatar_split_clause,[],[f3744,f649,f628])).
fof(f3744,plain,(
  e4 != op(e4,e4) | ~spl4_10),
  inference(backward_demodulation,[],[f651,f455])).
fof(f3729,plain,(
  ~spl4_17 | ~spl4_19),
  inference(avatar_contradiction_clause,[],[f3728])).
fof(f3728,plain,(
  $false | (~spl4_17 | ~spl4_19)),
  inference(subsumption_resolution,[],[f3727,f461])).
fof(f3727,plain,(
  e1 = e3 | (~spl4_17 | ~spl4_19)),
  inference(backward_demodulation,[],[f681,f689])).
fof(f3725,plain,(
  ~spl4_24 | ~spl4_83 | ~spl4_102),
  inference(avatar_contradiction_clause,[],[f3724])).
fof(f3724,plain,(
  $false | (~spl4_24 | ~spl4_83 | ~spl4_102)),
  inference(subsumption_resolution,[],[f3723,f464])).
fof(f3723,plain,(
  e2 = e4 | (~spl4_24 | ~spl4_83 | ~spl4_102)),
  inference(forward_demodulation,[],[f3719,f958])).
fof(f3719,plain,(
  e4 = op(e1,e3) | (~spl4_24 | ~spl4_102)),
  inference(backward_demodulation,[],[f710,f3258])).
fof(f3258,plain,(
  e4 = op(e1,op(e4,e0)) | ~spl4_102),
  inference(forward_demodulation,[],[f351,f1038])).
fof(f3722,plain,(
  ~spl4_24 | spl4_41 | ~spl4_102),
  inference(avatar_contradiction_clause,[],[f3721])).
fof(f3721,plain,(
  $false | (~spl4_24 | spl4_41 | ~spl4_102)),
  inference(subsumption_resolution,[],[f3718,f781])).
fof(f781,plain,(
  e0 != op(e3,e1) | spl4_41),
  inference(avatar_component_clause,[],[f780])).
fof(f3718,plain,(
  e0 = op(e3,e1) | (~spl4_24 | ~spl4_102)),
  inference(backward_demodulation,[],[f710,f3257])).
fof(f3257,plain,(
  e0 = op(op(e4,e0),e1) | ~spl4_102),
  inference(forward_demodulation,[],[f335,f1038])).
fof(f3713,plain,(
  ~spl4_32 | ~spl4_33),
  inference(avatar_contradiction_clause,[],[f3712])).
fof(f3712,plain,(
  $false | (~spl4_32 | ~spl4_33)),
  inference(subsumption_resolution,[],[f3711,f460])).
fof(f3711,plain,(
  e1 = e2 | (~spl4_32 | ~spl4_33)),
  inference(backward_demodulation,[],[f744,f748])).
fof(f3694,plain,(
  ~spl4_46 | ~spl4_109),
  inference(avatar_contradiction_clause,[],[f3693])).
fof(f3693,plain,(
  $false | (~spl4_46 | ~spl4_109)),
  inference(subsumption_resolution,[],[f3692,f458])).
fof(f3692,plain,(
  e0 = e3 | (~spl4_46 | ~spl4_109)),
  inference(forward_demodulation,[],[f3689,f803])).
fof(f3689,plain,(
  e3 = op(e3,e0) | (~spl4_46 | ~spl4_109)),
  inference(backward_demodulation,[],[f803,f3596])).
fof(f3596,plain,(
  e3 = op(e3,op(e3,e0)) | ~spl4_109),
  inference(backward_demodulation,[],[f1067,f346])).
fof(f3690,plain,(
  ~spl4_21 | ~spl4_46),
  inference(avatar_split_clause,[],[f3681,f801,f696])).
fof(f3681,plain,(
  e0 != op(e4,e0) | ~spl4_46),
  inference(backward_demodulation,[],[f803,f365])).
fof(f3680,plain,(
  ~spl4_51 | ~spl4_52),
  inference(avatar_contradiction_clause,[],[f3679])).
fof(f3679,plain,(
  $false | (~spl4_51 | ~spl4_52)),
  inference(subsumption_resolution,[],[f3678,f456])).
fof(f3678,plain,(
  e0 = e1 | (~spl4_51 | ~spl4_52)),
  inference(backward_demodulation,[],[f824,f828])).
fof(f3647,plain,(
  ~spl4_19 | spl4_27 | ~spl4_80),
  inference(avatar_contradiction_clause,[],[f3646])).
fof(f3646,plain,(
  $false | (~spl4_19 | spl4_27 | ~spl4_80)),
  inference(subsumption_resolution,[],[f3643,f722])).
fof(f722,plain,(
  e1 != op(e3,e4) | spl4_27),
  inference(avatar_component_clause,[],[f721])).
fof(f3643,plain,(
  e1 = op(e3,e4) | (~spl4_19 | ~spl4_80)),
  inference(backward_demodulation,[],[f945,f3524])).
fof(f3524,plain,(
  e1 = op(e3,op(e1,e4)) | ~spl4_19),
  inference(forward_demodulation,[],[f340,f689])).
fof(f3639,plain,(
  spl4_12 | ~spl4_45 | ~spl4_83),
  inference(avatar_split_clause,[],[f3635,f956,f796,f658])).
fof(f3635,plain,(
  e1 = op(e4,e2) | (~spl4_45 | ~spl4_83)),
  inference(backward_demodulation,[],[f958,f3273])).
fof(f3273,plain,(
  e1 = op(e4,op(e1,e3)) | ~spl4_45),
  inference(forward_demodulation,[],[f339,f798])).
fof(f3638,plain,(
  spl4_54 | ~spl4_45 | ~spl4_83),
  inference(avatar_split_clause,[],[f3634,f956,f796,f834])).
fof(f3634,plain,(
  e3 = op(e2,e4) | (~spl4_45 | ~spl4_83)),
  inference(backward_demodulation,[],[f958,f3272])).
fof(f3272,plain,(
  e3 = op(op(e1,e3),e4) | ~spl4_45),
  inference(forward_demodulation,[],[f347,f798])).
fof(f3636,plain,(
  ~spl4_33 | ~spl4_83),
  inference(avatar_split_clause,[],[f3629,f956,f746])).
fof(f3629,plain,(
  e2 != op(e3,e3) | ~spl4_83),
  inference(backward_demodulation,[],[f958,f391])).
fof(f3623,plain,(
  ~spl4_91 | spl4_122),
  inference(avatar_contradiction_clause,[],[f3622])).
fof(f3622,plain,(
  $false | (~spl4_91 | spl4_122)),
  inference(subsumption_resolution,[],[f3611,f1121])).
fof(f1121,plain,(
  op(e0,e0) != e1 | spl4_122),
  inference(avatar_component_clause,[],[f1120])).
fof(f3611,plain,(
  op(e0,e0) = e1 | ~spl4_91),
  inference(backward_demodulation,[],[f992,f337])).
fof(f3610,plain,(
  ~spl4_102 | ~spl4_104),
  inference(avatar_contradiction_clause,[],[f3609])).
fof(f3609,plain,(
  $false | (~spl4_102 | ~spl4_104)),
  inference(subsumption_resolution,[],[f3608,f461])).
fof(f3608,plain,(
  e1 = e3 | (~spl4_102 | ~spl4_104)),
  inference(forward_demodulation,[],[f1046,f1038])).
fof(f3601,plain,(
  ~spl4_84 | ~spl4_109),
  inference(avatar_split_clause,[],[f3597,f1065,f960])).
fof(f3597,plain,(
  e3 != op(e1,e3) | ~spl4_109),
  inference(backward_demodulation,[],[f1067,f386])).
fof(f3594,plain,(
  ~spl4_112 | ~spl4_115),
  inference(avatar_contradiction_clause,[],[f3593])).
fof(f3593,plain,(
  $false | (~spl4_112 | ~spl4_115)),
  inference(subsumption_resolution,[],[f3592,f462])).
fof(f3592,plain,(
  e1 = e4 | (~spl4_112 | ~spl4_115)),
  inference(backward_demodulation,[],[f1080,f1092])).
fof(f3583,plain,(
  ~spl4_118 | ~spl4_120),
  inference(avatar_contradiction_clause,[],[f3582])).
fof(f3582,plain,(
  $false | (~spl4_118 | ~spl4_120)),
  inference(subsumption_resolution,[],[f3580,f464])).
fof(f3580,plain,(
  e2 = e4 | (~spl4_118 | ~spl4_120)),
  inference(backward_demodulation,[],[f1105,f1113])).
fof(f3579,plain,(
  ~spl4_124 | ~spl4_125),
  inference(avatar_contradiction_clause,[],[f3578])).
fof(f3578,plain,(
  $false | (~spl4_124 | ~spl4_125)),
  inference(subsumption_resolution,[],[f3577,f465])).
fof(f3577,plain,(
  e3 = e4 | (~spl4_124 | ~spl4_125)),
  inference(forward_demodulation,[],[f1134,f1130])).
fof(f3560,plain,(
  ~spl4_109 | ~spl4_124),
  inference(avatar_split_clause,[],[f3559,f1128,f1065])).
fof(f3559,plain,(
  e3 != op(e0,e3) | ~spl4_124),
  inference(forward_demodulation,[],[f408,f1130])).
fof(f3548,plain,(
  ~spl4_118 | ~spl4_68),
  inference(avatar_split_clause,[],[f3325,f893,f1103])).
fof(f3325,plain,(
  e2 != op(e0,e1) | ~spl4_68),
  inference(forward_demodulation,[],[f367,f895])).
fof(f3541,plain,(
  ~spl4_112 | ~spl4_102),
  inference(avatar_split_clause,[],[f3316,f1036,f1078])).
fof(f3316,plain,(
  e1 != op(e0,e2) | ~spl4_102),
  inference(forward_demodulation,[],[f414,f1038])).
fof(f3515,plain,(
  ~spl4_75 | ~spl4_60),
  inference(avatar_split_clause,[],[f3514,f859,f922])).
fof(f3514,plain,(
  e4 != op(e2,e0) | ~spl4_60),
  inference(forward_demodulation,[],[f428,f861])).
fof(f3497,plain,(
  ~spl4_110 | ~spl4_60),
  inference(avatar_split_clause,[],[f3496,f859,f1069])).
fof(f3496,plain,(
  e4 != op(e0,e3) | ~spl4_60),
  inference(forward_demodulation,[],[f387,f861])).
fof(f3467,plain,(
  ~spl4_8 | ~spl4_9),
  inference(avatar_contradiction_clause,[],[f3466])).
fof(f3466,plain,(
  $false | (~spl4_8 | ~spl4_9)),
  inference(subsumption_resolution,[],[f3465,f463])).
fof(f3465,plain,(
  e2 = e3 | (~spl4_8 | ~spl4_9)),
  inference(backward_demodulation,[],[f643,f647])).
fof(f3460,plain,(
  ~spl4_19 | spl4_47 | ~spl4_76),
  inference(avatar_contradiction_clause,[],[f3459])).
fof(f3459,plain,(
  $false | (~spl4_19 | spl4_47 | ~spl4_76)),
  inference(subsumption_resolution,[],[f3457,f806])).
fof(f806,plain,(
  e1 != op(e3,e0) | spl4_47),
  inference(avatar_component_clause,[],[f805])).
fof(f3457,plain,(
  e1 = op(e3,e0) | (~spl4_19 | ~spl4_76)),
  inference(backward_demodulation,[],[f689,f2854])).
fof(f2854,plain,(
  e1 = op(op(e4,e1),e0) | ~spl4_76),
  inference(backward_demodulation,[],[f929,f340])).
fof(f3454,plain,(
  ~spl4_21 | ~spl4_97 | ~spl4_102),
  inference(avatar_contradiction_clause,[],[f3453])).
fof(f3453,plain,(
  $false | (~spl4_21 | ~spl4_97 | ~spl4_102)),
  inference(subsumption_resolution,[],[f3452,f462])).
fof(f3452,plain,(
  e1 = e4 | (~spl4_21 | ~spl4_97 | ~spl4_102)),
  inference(forward_demodulation,[],[f3451,f1017])).
fof(f3451,plain,(
  e4 = op(e1,e0) | (~spl4_21 | ~spl4_102)),
  inference(backward_demodulation,[],[f698,f3258])).
fof(f3440,plain,(
  ~spl4_9 | ~spl4_28 | ~spl4_57),
  inference(avatar_contradiction_clause,[],[f3439])).
fof(f3439,plain,(
  $false | (~spl4_9 | ~spl4_28 | ~spl4_57)),
  inference(subsumption_resolution,[],[f3438,f462])).
fof(f3438,plain,(
  e1 = e4 | (~spl4_9 | ~spl4_28 | ~spl4_57)),
  inference(forward_demodulation,[],[f3435,f849])).
fof(f3435,plain,(
  e4 = op(e2,e3) | (~spl4_9 | ~spl4_28)),
  inference(backward_demodulation,[],[f727,f3237])).
fof(f3237,plain,(
  e4 = op(op(e3,e4),e3) | ~spl4_9),
  inference(forward_demodulation,[],[f354,f647])).
fof(f3430,plain,(
  ~spl4_31 | ~spl4_123),
  inference(avatar_contradiction_clause,[],[f3429])).
fof(f3429,plain,(
  $false | (~spl4_31 | ~spl4_123)),
  inference(subsumption_resolution,[],[f3428,f463])).
fof(f3428,plain,(
  e2 = e3 | (~spl4_31 | ~spl4_123)),
  inference(forward_demodulation,[],[f3418,f1126])).
fof(f3418,plain,(
  op(e0,e0) = e3 | ~spl4_31),
  inference(backward_demodulation,[],[f740,f349])).
fof(f3409,plain,(
  ~spl4_42 | ~spl4_45),
  inference(avatar_contradiction_clause,[],[f3408])).
fof(f3408,plain,(
  $false | (~spl4_42 | ~spl4_45)),
  inference(subsumption_resolution,[],[f3407,f462])).
fof(f3407,plain,(
  e1 = e4 | (~spl4_42 | ~spl4_45)),
  inference(backward_demodulation,[],[f786,f798])).
fof(f3394,plain,(
  ~spl4_57 | ~spl4_60),
  inference(avatar_contradiction_clause,[],[f3393])).
fof(f3393,plain,(
  $false | (~spl4_57 | ~spl4_60)),
  inference(subsumption_resolution,[],[f3392,f462])).
fof(f3392,plain,(
  e1 = e4 | (~spl4_57 | ~spl4_60)),
  inference(forward_demodulation,[],[f861,f849])).
fof(f3383,plain,(
  ~spl4_89 | ~spl4_90),
  inference(avatar_contradiction_clause,[],[f3382])).
fof(f3382,plain,(
  $false | (~spl4_89 | ~spl4_90)),
  inference(subsumption_resolution,[],[f3381,f465])).
fof(f3381,plain,(
  e3 = e4 | (~spl4_89 | ~spl4_90)),
  inference(backward_demodulation,[],[f983,f987])).
fof(f3380,plain,(
  ~spl4_90 | ~spl4_93 | spl4_188),
  inference(avatar_contradiction_clause,[],[f3379])).
fof(f3379,plain,(
  $false | (~spl4_90 | ~spl4_93 | spl4_188)),
  inference(subsumption_resolution,[],[f3373,f987])).
fof(f3373,plain,(
  e4 != op(e1,e2) | (~spl4_93 | spl4_188)),
  inference(backward_demodulation,[],[f1000,f1561])).
fof(f3378,plain,(
  ~spl4_89 | ~spl4_93 | spl4_184),
  inference(avatar_split_clause,[],[f3372,f1519,f998,f981])).
fof(f3372,plain,(
  e3 != op(e1,e2) | (~spl4_93 | spl4_184)),
  inference(backward_demodulation,[],[f1000,f1521])).
fof(f3376,plain,(
  ~spl4_61 | ~spl4_93),
  inference(avatar_contradiction_clause,[],[f3375])).
fof(f3375,plain,(
  $false | (~spl4_61 | ~spl4_93)),
  inference(subsumption_resolution,[],[f3374,f456])).
fof(f3374,plain,(
  e0 = e1 | (~spl4_61 | ~spl4_93)),
  inference(forward_demodulation,[],[f3365,f866])).
fof(f3365,plain,(
  e1 = op(e2,e2) | ~spl4_93),
  inference(backward_demodulation,[],[f1000,f337])).
fof(f3341,plain,(
  ~spl4_123 | ~spl4_124),
  inference(avatar_contradiction_clause,[],[f3340])).
fof(f3340,plain,(
  $false | (~spl4_123 | ~spl4_124)),
  inference(subsumption_resolution,[],[f3339,f463])).
fof(f3339,plain,(
  e2 = e3 | (~spl4_123 | ~spl4_124)),
  inference(forward_demodulation,[],[f1130,f1126])).
fof(f3311,plain,(
  ~spl4_107 | ~spl4_102),
  inference(avatar_split_clause,[],[f3310,f1036,f1057])).
fof(f3310,plain,(
  e1 != op(e0,e3) | ~spl4_102),
  inference(forward_demodulation,[],[f415,f1038])).
fof(f3303,plain,(
  ~spl4_97 | ~spl4_116),
  inference(avatar_contradiction_clause,[],[f3302])).
fof(f3302,plain,(
  $false | (~spl4_97 | ~spl4_116)),
  inference(subsumption_resolution,[],[f3301,f456])).
fof(f3301,plain,(
  e0 = e1 | (~spl4_97 | ~spl4_116)),
  inference(forward_demodulation,[],[f3300,f1097])).
fof(f3300,plain,(
  e1 = op(e0,e1) | (~spl4_97 | ~spl4_116)),
  inference(forward_demodulation,[],[f3299,f1097])).
fof(f3299,plain,(
  e1 = op(op(e0,e1),e1) | ~spl4_97),
  inference(forward_demodulation,[],[f336,f1017])).
fof(f3251,plain,(
  spl4_83 | ~spl4_12 | ~spl4_54),
  inference(avatar_split_clause,[],[f3250,f834,f658,f956])).
fof(f3250,plain,(
  e2 = op(e1,e3) | (~spl4_12 | ~spl4_54)),
  inference(forward_demodulation,[],[f2632,f660])).
fof(f2632,plain,(
  e2 = op(op(e4,e2),e3) | ~spl4_54),
  inference(backward_demodulation,[],[f836,f345])).
fof(f3246,plain,(
  ~spl4_84 | ~spl4_9),
  inference(avatar_split_clause,[],[f3245,f645,f960])).
fof(f3245,plain,(
  e3 != op(e1,e3) | ~spl4_9),
  inference(forward_demodulation,[],[f392,f647])).
fof(f3229,plain,(
  ~spl4_3 | ~spl4_12 | ~spl4_94 | spl4_132),
  inference(avatar_contradiction_clause,[],[f3228])).
fof(f3228,plain,(
  $false | (~spl4_3 | ~spl4_12 | ~spl4_94 | spl4_132)),
  inference(subsumption_resolution,[],[f3227,f1004])).
fof(f3227,plain,(
  e3 != op(e1,e1) | (~spl4_3 | ~spl4_12 | spl4_132)),
  inference(forward_demodulation,[],[f3221,f660])).
fof(f3221,plain,(
  e3 != op(op(e4,e2),op(e4,e2)) | (~spl4_3 | spl4_132)),
  inference(backward_demodulation,[],[f622,f1238])).
fof(f3226,plain,(
  ~spl4_3 | ~spl4_61),
  inference(avatar_contradiction_clause,[],[f3225])).
fof(f3225,plain,(
  $false | (~spl4_3 | ~spl4_61)),
  inference(subsumption_resolution,[],[f3224,f459])).
fof(f3224,plain,(
  e0 = e4 | (~spl4_3 | ~spl4_61)),
  inference(forward_demodulation,[],[f3218,f866])).
fof(f3218,plain,(
  e4 = op(e2,e2) | ~spl4_3),
  inference(backward_demodulation,[],[f622,f355])).
fof(f3216,plain,(
  ~spl4_6 | ~spl4_9),
  inference(avatar_contradiction_clause,[],[f3215])).
fof(f3215,plain,(
  $false | (~spl4_6 | ~spl4_9)),
  inference(subsumption_resolution,[],[f3214,f458])).
fof(f3214,plain,(
  e0 = e3 | (~spl4_6 | ~spl4_9)),
  inference(forward_demodulation,[],[f647,f635])).
fof(f635,plain,(
  e0 = op(e4,e3) | ~spl4_6),
  inference(avatar_component_clause,[],[f633])).
fof(f3213,plain,(
  ~spl4_12 | ~spl4_15),
  inference(avatar_contradiction_clause,[],[f3212])).
fof(f3212,plain,(
  $false | (~spl4_12 | ~spl4_15)),
  inference(subsumption_resolution,[],[f3211,f462])).
fof(f3211,plain,(
  e1 = e4 | (~spl4_12 | ~spl4_15)),
  inference(backward_demodulation,[],[f660,f672])).
fof(f3187,plain,(
  ~spl4_27 | ~spl4_30),
  inference(avatar_contradiction_clause,[],[f3186])).
fof(f3186,plain,(
  $false | (~spl4_27 | ~spl4_30)),
  inference(subsumption_resolution,[],[f3185,f462])).
fof(f3185,plain,(
  e1 = e4 | (~spl4_27 | ~spl4_30)),
  inference(backward_demodulation,[],[f723,f735])).
fof(f3182,plain,(
  ~spl4_39 | ~spl4_57 | ~spl4_81),
  inference(avatar_contradiction_clause,[],[f3181])).
fof(f3181,plain,(
  $false | (~spl4_39 | ~spl4_57 | ~spl4_81)),
  inference(subsumption_resolution,[],[f3180,f458])).
fof(f3180,plain,(
  e0 = e3 | (~spl4_39 | ~spl4_57 | ~spl4_81)),
  inference(forward_demodulation,[],[f3176,f950])).
fof(f3176,plain,(
  e3 = op(e1,e3) | (~spl4_39 | ~spl4_57)),
  inference(backward_demodulation,[],[f773,f2887])).
fof(f2887,plain,(
  e3 = op(e1,op(e3,e2)) | ~spl4_57),
  inference(backward_demodulation,[],[f849,f348])).
fof(f3179,plain,(
  ~spl4_39 | spl4_43 | ~spl4_57),
  inference(avatar_contradiction_clause,[],[f3178])).
fof(f3178,plain,(
  $false | (~spl4_39 | spl4_43 | ~spl4_57)),
  inference(subsumption_resolution,[],[f3175,f789])).
fof(f789,plain,(
  e2 != op(e3,e1) | spl4_43),
  inference(avatar_component_clause,[],[f788])).
fof(f3175,plain,(
  e2 = op(e3,e1) | (~spl4_39 | ~spl4_57)),
  inference(backward_demodulation,[],[f773,f2886])).
fof(f2886,plain,(
  e2 = op(op(e3,e2),e1) | ~spl4_57),
  inference(backward_demodulation,[],[f849,f344])).
fof(f3133,plain,(
  spl4_62 | ~spl4_68 | ~spl4_88),
  inference(avatar_split_clause,[],[f3129,f977,f893,f868])).
fof(f3129,plain,(
  e1 = op(e2,e2) | (~spl4_68 | ~spl4_88)),
  inference(backward_demodulation,[],[f895,f3013])).
fof(f3013,plain,(
  e1 = op(op(e2,e1),e2) | ~spl4_88),
  inference(forward_demodulation,[],[f338,f979])).
fof(f3132,plain,(
  spl4_63 | ~spl4_68 | ~spl4_88),
  inference(avatar_contradiction_clause,[],[f3131])).
fof(f3131,plain,(
  $false | (spl4_63 | ~spl4_68 | ~spl4_88)),
  inference(subsumption_resolution,[],[f3128,f873])).
fof(f873,plain,(
  e2 != op(e2,e2) | spl4_63),
  inference(avatar_component_clause,[],[f872])).
fof(f3128,plain,(
  e2 = op(e2,e2) | (~spl4_68 | ~spl4_88)),
  inference(backward_demodulation,[],[f895,f3011])).
fof(f3011,plain,(
  e2 = op(e2,op(e2,e1)) | ~spl4_88),
  inference(backward_demodulation,[],[f979,f342])).
fof(f3116,plain,(
  ~spl4_88 | ~spl4_90),
  inference(avatar_contradiction_clause,[],[f3115])).
fof(f3115,plain,(
  $false | (~spl4_88 | ~spl4_90)),
  inference(subsumption_resolution,[],[f3114,f464])).
fof(f3114,plain,(
  e2 = e4 | (~spl4_88 | ~spl4_90)),
  inference(forward_demodulation,[],[f987,f979])).
fof(f3098,plain,(
  ~spl4_97 | ~spl4_99),
  inference(avatar_contradiction_clause,[],[f3097])).
fof(f3097,plain,(
  $false | (~spl4_97 | ~spl4_99)),
  inference(subsumption_resolution,[],[f3096,f461])).
fof(f3096,plain,(
  e1 = e3 | (~spl4_97 | ~spl4_99)),
  inference(backward_demodulation,[],[f1017,f1025])).
fof(f3093,plain,(
  spl4_66 | ~spl4_23 | ~spl4_102),
  inference(avatar_split_clause,[],[f3088,f1036,f704,f885])).
fof(f3088,plain,(
  e0 = op(e2,e1) | (~spl4_23 | ~spl4_102)),
  inference(backward_demodulation,[],[f1038,f2800])).
fof(f2800,plain,(
  e0 = op(e2,op(e0,e4)) | ~spl4_23),
  inference(forward_demodulation,[],[f335,f706])).
fof(f3092,plain,(
  ~spl4_23 | ~spl4_88 | ~spl4_102),
  inference(avatar_contradiction_clause,[],[f3091])).
fof(f3091,plain,(
  $false | (~spl4_23 | ~spl4_88 | ~spl4_102)),
  inference(subsumption_resolution,[],[f3090,f464])).
fof(f3090,plain,(
  e2 = e4 | (~spl4_23 | ~spl4_88 | ~spl4_102)),
  inference(forward_demodulation,[],[f3087,f979])).
fof(f3087,plain,(
  e4 = op(e1,e2) | (~spl4_23 | ~spl4_102)),
  inference(backward_demodulation,[],[f1038,f2799])).
fof(f2799,plain,(
  e4 = op(op(e0,e4),e2) | ~spl4_23),
  inference(forward_demodulation,[],[f351,f706])).
fof(f3089,plain,(
  ~spl4_2 | ~spl4_102),
  inference(avatar_split_clause,[],[f3083,f1036,f616])).
fof(f3083,plain,(
  e1 != op(e4,e4) | ~spl4_102),
  inference(backward_demodulation,[],[f1038,f399])).
fof(f3082,plain,(
  spl4_3 | ~spl4_75 | ~spl4_115),
  inference(avatar_split_clause,[],[f3077,f1090,f922,f620])).
fof(f3077,plain,(
  e2 = op(e4,e4) | (~spl4_75 | ~spl4_115)),
  inference(backward_demodulation,[],[f1092,f2767])).
fof(f2767,plain,(
  e2 = op(op(e0,e2),e4) | ~spl4_75),
  inference(forward_demodulation,[],[f341,f924])).
fof(f3065,plain,(
  ~spl4_114 | ~spl4_123 | spl4_185),
  inference(avatar_split_clause,[],[f3061,f1529,f1124,f1086])).
fof(f3061,plain,(
  e3 != op(e0,e2) | (~spl4_123 | spl4_185)),
  inference(backward_demodulation,[],[f1126,f1531])).
fof(f3044,plain,(
  spl4_31 | ~spl4_99 | ~spl4_119),
  inference(avatar_contradiction_clause,[],[f3043])).
fof(f3043,plain,(
  $false | (spl4_31 | ~spl4_99 | ~spl4_119)),
  inference(subsumption_resolution,[],[f3042,f739])).
fof(f3042,plain,(
  e0 = op(e3,e3) | (~spl4_99 | ~spl4_119)),
  inference(forward_demodulation,[],[f3041,f1025])).
fof(f3041,plain,(
  e0 = op(op(e1,e0),e3) | ~spl4_119),
  inference(forward_demodulation,[],[f332,f1109])).
fof(f3040,plain,(
  spl4_32 | ~spl4_99 | ~spl4_119),
  inference(avatar_contradiction_clause,[],[f3039])).
fof(f3039,plain,(
  $false | (spl4_32 | ~spl4_99 | ~spl4_119)),
  inference(subsumption_resolution,[],[f3038,f743])).
fof(f3038,plain,(
  e1 = op(e3,e3) | (~spl4_99 | ~spl4_119)),
  inference(forward_demodulation,[],[f3037,f1109])).
fof(f3037,plain,(
  e1 = op(op(e0,e1),e3) | ~spl4_99),
  inference(forward_demodulation,[],[f336,f1025])).
fof(f3036,plain,(
  ~spl4_94 | ~spl4_119),
  inference(avatar_split_clause,[],[f3035,f1107,f1002])).
fof(f3035,plain,(
  e3 != op(e1,e1) | ~spl4_119),
  inference(forward_demodulation,[],[f366,f1109])).
fof(f2995,plain,(
  spl4_64 | ~spl4_33),
  inference(avatar_split_clause,[],[f2914,f746,f876])).
fof(f2914,plain,(
  e3 = op(e2,e2) | ~spl4_33),
  inference(backward_demodulation,[],[f748,f349])).
fof(f2977,plain,(
  ~spl4_37 | ~spl4_33 | spl4_159),
  inference(avatar_split_clause,[],[f2919,f1369,f746,f763])).
fof(f2919,plain,(
  e1 != op(e3,e2) | (~spl4_33 | spl4_159)),
  inference(backward_demodulation,[],[f748,f1371])).
fof(f2975,plain,(
  ~spl4_40 | ~spl4_33 | spl4_186),
  inference(avatar_split_clause,[],[f2921,f1548,f746,f775])).
fof(f2921,plain,(
  e4 != op(e3,e2) | (~spl4_33 | spl4_186)),
  inference(backward_demodulation,[],[f748,f1550])).
fof(f2960,plain,(
  ~spl4_81 | ~spl4_6),
  inference(avatar_split_clause,[],[f2959,f633,f948])).
fof(f2959,plain,(
  e0 != op(e1,e3) | ~spl4_6),
  inference(forward_demodulation,[],[f392,f635])).
fof(f2952,plain,(
  ~spl4_4 | ~spl4_54),
  inference(avatar_split_clause,[],[f2951,f834,f624])).
fof(f2951,plain,(
  e3 != op(e4,e4) | ~spl4_54),
  inference(forward_demodulation,[],[f404,f836])).
fof(f2944,plain,(
  ~spl4_20 | ~spl4_76 | ~spl4_103),
  inference(avatar_contradiction_clause,[],[f2943])).
fof(f2943,plain,(
  $false | (~spl4_20 | ~spl4_76 | ~spl4_103)),
  inference(subsumption_resolution,[],[f2942,f464])).
fof(f2942,plain,(
  e2 = e4 | (~spl4_20 | ~spl4_76 | ~spl4_103)),
  inference(forward_demodulation,[],[f2938,f1042])).
fof(f2938,plain,(
  e4 = op(e0,e4) | (~spl4_20 | ~spl4_76)),
  inference(backward_demodulation,[],[f693,f2855])).
fof(f2855,plain,(
  e4 = op(e0,op(e4,e1)) | ~spl4_76),
  inference(backward_demodulation,[],[f929,f352])).
fof(f2941,plain,(
  ~spl4_20 | spl4_22 | ~spl4_76),
  inference(avatar_contradiction_clause,[],[f2940])).
fof(f2940,plain,(
  $false | (~spl4_20 | spl4_22 | ~spl4_76)),
  inference(subsumption_resolution,[],[f2937,f701])).
fof(f2937,plain,(
  e1 = op(e4,e0) | (~spl4_20 | ~spl4_76)),
  inference(backward_demodulation,[],[f693,f2854])).
fof(f2924,plain,(
  ~spl4_33 | ~spl4_62),
  inference(avatar_contradiction_clause,[],[f2923])).
fof(f2923,plain,(
  $false | (~spl4_33 | ~spl4_62)),
  inference(subsumption_resolution,[],[f2922,f461])).
fof(f2922,plain,(
  e1 = e3 | (~spl4_33 | ~spl4_62)),
  inference(forward_demodulation,[],[f2914,f870])).
fof(f2913,plain,(
  ~spl4_36 | ~spl4_40),
  inference(avatar_contradiction_clause,[],[f2912])).
fof(f2912,plain,(
  $false | (~spl4_36 | ~spl4_40)),
  inference(subsumption_resolution,[],[f2911,f459])).
fof(f2911,plain,(
  e0 = e4 | (~spl4_36 | ~spl4_40)),
  inference(backward_demodulation,[],[f761,f777])).
fof(f2907,plain,(
  ~spl4_40 | ~spl4_57 | ~spl4_76),
  inference(avatar_contradiction_clause,[],[f2906])).
fof(f2906,plain,(
  $false | (~spl4_40 | ~spl4_57 | ~spl4_76)),
  inference(subsumption_resolution,[],[f2905,f458])).
fof(f2905,plain,(
  e0 = e3 | (~spl4_40 | ~spl4_57 | ~spl4_76)),
  inference(forward_demodulation,[],[f2900,f929])).
fof(f2900,plain,(
  e3 = op(e1,e4) | (~spl4_40 | ~spl4_57)),
  inference(backward_demodulation,[],[f777,f2887])).
fof(f2895,plain,(
  ~spl4_47 | ~spl4_49),
  inference(avatar_contradiction_clause,[],[f2894])).
fof(f2894,plain,(
  $false | (~spl4_47 | ~spl4_49)),
  inference(subsumption_resolution,[],[f2893,f461])).
fof(f2893,plain,(
  e1 = e3 | (~spl4_47 | ~spl4_49)),
  inference(backward_demodulation,[],[f807,f815])).
fof(f2882,plain,(
  ~spl4_62 | ~spl4_66 | spl4_138),
  inference(avatar_contradiction_clause,[],[f2881])).
fof(f2881,plain,(
  $false | (~spl4_62 | ~spl4_66 | spl4_138)),
  inference(subsumption_resolution,[],[f2873,f887])).
fof(f2873,plain,(
  e0 != op(e2,e1) | (~spl4_62 | spl4_138)),
  inference(backward_demodulation,[],[f870,f1265])).
fof(f2878,plain,(
  ~spl4_62 | spl4_93),
  inference(avatar_contradiction_clause,[],[f2877])).
fof(f2877,plain,(
  $false | (~spl4_62 | spl4_93)),
  inference(subsumption_resolution,[],[f2867,f999])).
fof(f2867,plain,(
  e2 = op(e1,e1) | ~spl4_62),
  inference(backward_demodulation,[],[f870,f343])).
fof(f2852,plain,(
  ~spl4_87 | ~spl4_88),
  inference(avatar_contradiction_clause,[],[f2851])).
fof(f2851,plain,(
  $false | (~spl4_87 | ~spl4_88)),
  inference(subsumption_resolution,[],[f2850,f460])).
fof(f2850,plain,(
  e1 = e2 | (~spl4_87 | ~spl4_88)),
  inference(forward_demodulation,[],[f979,f975])).
fof(f2839,plain,(
  ~spl4_78 | ~spl4_103),
  inference(avatar_split_clause,[],[f2833,f1040,f935])).
fof(f2833,plain,(
  e2 != op(e1,e4) | ~spl4_103),
  inference(backward_demodulation,[],[f1042,f396])).
fof(f2809,plain,(
  ~spl4_117 | ~spl4_119),
  inference(avatar_contradiction_clause,[],[f2808])).
fof(f2808,plain,(
  $false | (~spl4_117 | ~spl4_119)),
  inference(subsumption_resolution,[],[f2807,f461])).
fof(f2807,plain,(
  e1 = e3 | (~spl4_117 | ~spl4_119)),
  inference(forward_demodulation,[],[f1109,f1101])).
fof(f2782,plain,(
  spl4_57 | ~spl4_43 | ~spl4_84),
  inference(avatar_split_clause,[],[f2779,f960,f788,f847])).
fof(f2779,plain,(
  e1 = op(e2,e3) | (~spl4_43 | ~spl4_84)),
  inference(backward_demodulation,[],[f962,f2736])).
fof(f2736,plain,(
  e1 = op(e2,op(e1,e3)) | ~spl4_43),
  inference(forward_demodulation,[],[f339,f790])).
fof(f2758,plain,(
  spl4_98 | ~spl4_66 | ~spl4_87),
  inference(avatar_split_clause,[],[f2757,f973,f885,f1019])).
fof(f2757,plain,(
  e2 = op(e1,e0) | (~spl4_66 | ~spl4_87)),
  inference(forward_demodulation,[],[f2756,f975])).
fof(f2756,plain,(
  e2 = op(op(e1,e2),e0) | ~spl4_66),
  inference(forward_demodulation,[],[f342,f887])).
fof(f2728,plain,(
  ~spl4_20 | ~spl4_95),
  inference(avatar_split_clause,[],[f2516,f1006,f691])).
fof(f2516,plain,(
  e4 != op(e4,e1) | ~spl4_95),
  inference(forward_demodulation,[],[f372,f1008])).
fof(f2708,plain,(
  ~spl4_7 | ~spl4_8),
  inference(avatar_contradiction_clause,[],[f2707])).
fof(f2707,plain,(
  $false | (~spl4_7 | ~spl4_8)),
  inference(subsumption_resolution,[],[f2706,f460])).
fof(f2706,plain,(
  e1 = e2 | (~spl4_7 | ~spl4_8)),
  inference(forward_demodulation,[],[f643,f639])).
fof(f2705,plain,(
  ~spl4_7 | ~spl4_10),
  inference(avatar_contradiction_clause,[],[f2704])).
fof(f2704,plain,(
  $false | (~spl4_7 | ~spl4_10)),
  inference(subsumption_resolution,[],[f2703,f462])).
fof(f2703,plain,(
  e1 = e4 | (~spl4_7 | ~spl4_10)),
  inference(forward_demodulation,[],[f651,f639])).
fof(f2702,plain,(
  ~spl4_11 | ~spl4_14),
  inference(avatar_contradiction_clause,[],[f2701])).
fof(f2701,plain,(
  $false | (~spl4_11 | ~spl4_14)),
  inference(subsumption_resolution,[],[f2699,f458])).
fof(f2699,plain,(
  e0 = e3 | (~spl4_11 | ~spl4_14)),
  inference(backward_demodulation,[],[f656,f668])).
fof(f2680,plain,(
  ~spl4_23 | spl4_51 | ~spl4_105),
  inference(avatar_contradiction_clause,[],[f2679])).
fof(f2679,plain,(
  $false | (~spl4_23 | spl4_51 | ~spl4_105)),
  inference(subsumption_resolution,[],[f2676,f823])).
fof(f823,plain,(
  e0 != op(e2,e4) | spl4_51),
  inference(avatar_component_clause,[],[f822])).
fof(f2676,plain,(
  e0 = op(e2,e4) | (~spl4_23 | ~spl4_105)),
  inference(backward_demodulation,[],[f706,f1654])).
fof(f1654,plain,(
  e0 = op(op(e4,e0),e4) | ~spl4_105),
  inference(backward_demodulation,[],[f1050,f335])).
fof(f2671,plain,(
  ~spl4_7 | ~spl4_26 | ~spl4_117),
  inference(avatar_contradiction_clause,[],[f2670])).
fof(f2670,plain,(
  $false | (~spl4_7 | ~spl4_26 | ~spl4_117)),
  inference(subsumption_resolution,[],[f2669,f462])).
fof(f2669,plain,(
  e1 = e4 | (~spl4_7 | ~spl4_26 | ~spl4_117)),
  inference(forward_demodulation,[],[f2665,f1101])).
fof(f2665,plain,(
  e4 = op(e0,e1) | (~spl4_7 | ~spl4_26)),
  inference(backward_demodulation,[],[f719,f2507])).
fof(f2507,plain,(
  e4 = op(op(e3,e4),e1) | ~spl4_7),
  inference(forward_demodulation,[],[f354,f639])).
fof(f2657,plain,(
  ~spl4_2 | ~spl4_35),
  inference(avatar_contradiction_clause,[],[f2656])).
fof(f2656,plain,(
  $false | (~spl4_2 | ~spl4_35)),
  inference(subsumption_resolution,[],[f2655,f461])).
fof(f2655,plain,(
  e1 = e3 | (~spl4_2 | ~spl4_35)),
  inference(forward_demodulation,[],[f2649,f618])).
fof(f2649,plain,(
  e3 = op(e4,e4) | ~spl4_35),
  inference(backward_demodulation,[],[f756,f349])).
fof(f2648,plain,(
  ~spl4_43 | ~spl4_44),
  inference(avatar_contradiction_clause,[],[f2647])).
fof(f2647,plain,(
  $false | (~spl4_43 | ~spl4_44)),
  inference(subsumption_resolution,[],[f2646,f463])).
fof(f2646,plain,(
  e2 = e3 | (~spl4_43 | ~spl4_44)),
  inference(backward_demodulation,[],[f790,f794])).
fof(f2645,plain,(
  spl4_34 | ~spl4_49 | ~spl4_109),
  inference(avatar_contradiction_clause,[],[f2644])).
fof(f2644,plain,(
  $false | (spl4_34 | ~spl4_49 | ~spl4_109)),
  inference(subsumption_resolution,[],[f2641,f751])).
fof(f751,plain,(
  e3 != op(e3,e3) | spl4_34),
  inference(avatar_component_clause,[],[f750])).
fof(f2641,plain,(
  e3 = op(e3,e3) | (~spl4_49 | ~spl4_109)),
  inference(backward_demodulation,[],[f815,f2156])).
fof(f2156,plain,(
  e3 = op(e3,op(e3,e0)) | ~spl4_109),
  inference(backward_demodulation,[],[f1067,f346])).
fof(f2643,plain,(
  spl4_31 | ~spl4_49 | ~spl4_109),
  inference(avatar_split_clause,[],[f2640,f1065,f813,f738])).
fof(f2640,plain,(
  e0 = op(e3,e3) | (~spl4_49 | ~spl4_109)),
  inference(backward_demodulation,[],[f815,f2155])).
fof(f2155,plain,(
  e0 = op(op(e3,e0),e3) | ~spl4_109),
  inference(backward_demodulation,[],[f1067,f334])).
fof(f2631,plain,(
  ~spl4_37 | ~spl4_59 | ~spl4_81),
  inference(avatar_contradiction_clause,[],[f2630])).
fof(f2630,plain,(
  $false | (~spl4_37 | ~spl4_59 | ~spl4_81)),
  inference(subsumption_resolution,[],[f2629,f457])).
fof(f2629,plain,(
  e0 = e2 | (~spl4_37 | ~spl4_59 | ~spl4_81)),
  inference(forward_demodulation,[],[f2627,f950])).
fof(f2627,plain,(
  e2 = op(e1,e3) | (~spl4_37 | ~spl4_59)),
  inference(backward_demodulation,[],[f857,f2529])).
fof(f2529,plain,(
  e2 = op(e1,op(e2,e3)) | ~spl4_37),
  inference(forward_demodulation,[],[f344,f765])).
fof(f2618,plain,(
  ~spl4_2 | ~spl4_65),
  inference(avatar_contradiction_clause,[],[f2617])).
fof(f2617,plain,(
  $false | (~spl4_2 | ~spl4_65)),
  inference(subsumption_resolution,[],[f2616,f460])).
fof(f2616,plain,(
  e1 = e2 | (~spl4_2 | ~spl4_65)),
  inference(forward_demodulation,[],[f2608,f618])).
fof(f2608,plain,(
  e2 = op(e4,e4) | ~spl4_65),
  inference(backward_demodulation,[],[f882,f343])).
fof(f2607,plain,(
  ~spl4_66 | ~spl4_68),
  inference(avatar_contradiction_clause,[],[f2606])).
fof(f2606,plain,(
  $false | (~spl4_66 | ~spl4_68)),
  inference(subsumption_resolution,[],[f2605,f457])).
fof(f2605,plain,(
  e0 = e2 | (~spl4_66 | ~spl4_68)),
  inference(backward_demodulation,[],[f887,f895])).
fof(f2600,plain,(
  ~spl4_72 | ~spl4_75),
  inference(avatar_contradiction_clause,[],[f2599])).
fof(f2599,plain,(
  $false | (~spl4_72 | ~spl4_75)),
  inference(subsumption_resolution,[],[f2598,f462])).
fof(f2598,plain,(
  e1 = e4 | (~spl4_72 | ~spl4_75)),
  inference(forward_demodulation,[],[f924,f912])).
fof(f2590,plain,(
  ~spl4_81 | ~spl4_84),
  inference(avatar_contradiction_clause,[],[f2589])).
fof(f2589,plain,(
  $false | (~spl4_81 | ~spl4_84)),
  inference(subsumption_resolution,[],[f2587,f458])).
fof(f2587,plain,(
  e0 = e3 | (~spl4_81 | ~spl4_84)),
  inference(backward_demodulation,[],[f950,f962])).
fof(f2585,plain,(
  ~spl4_84 | ~spl4_85),
  inference(avatar_contradiction_clause,[],[f2584])).
fof(f2584,plain,(
  $false | (~spl4_84 | ~spl4_85)),
  inference(subsumption_resolution,[],[f2583,f465])).
fof(f2583,plain,(
  e3 = e4 | (~spl4_84 | ~spl4_85)),
  inference(backward_demodulation,[],[f962,f966])).
fof(f2576,plain,(
  ~spl4_68 | ~spl4_87),
  inference(avatar_contradiction_clause,[],[f2575])).
fof(f2575,plain,(
  $false | (~spl4_68 | ~spl4_87)),
  inference(subsumption_resolution,[],[f2574,f460])).
fof(f2574,plain,(
  e1 = e2 | (~spl4_68 | ~spl4_87)),
  inference(forward_demodulation,[],[f2573,f975])).
fof(f2573,plain,(
  e2 = op(e1,e2) | (~spl4_68 | ~spl4_87)),
  inference(forward_demodulation,[],[f2572,f975])).
fof(f2572,plain,(
  e2 = op(op(e1,e2),e2) | ~spl4_68),
  inference(forward_demodulation,[],[f342,f895])).
fof(f2560,plain,(
  spl4_86 | ~spl4_72 | ~spl4_113),
  inference(avatar_split_clause,[],[f2420,f1082,f910,f969])).
fof(f2420,plain,(
  e0 = op(e1,e2) | (~spl4_72 | ~spl4_113)),
  inference(backward_demodulation,[],[f1084,f2419])).
fof(f2419,plain,(
  e0 = op(e1,op(e0,e2)) | ~spl4_72),
  inference(forward_demodulation,[],[f333,f912])).
fof(f2539,plain,(
  ~spl4_78 | ~spl4_98),
  inference(avatar_split_clause,[],[f2538,f1019,f935])).
fof(f2538,plain,(
  e2 != op(e1,e4) | ~spl4_98),
  inference(forward_demodulation,[],[f419,f1021])).
fof(f2535,plain,(
  ~spl4_78 | ~spl4_95 | spl4_176),
  inference(avatar_split_clause,[],[f2440,f1463,f1006,f935])).
fof(f2440,plain,(
  e2 != op(e1,e4) | (~spl4_95 | spl4_176)),
  inference(backward_demodulation,[],[f1008,f1465])).
fof(f2534,plain,(
  spl4_66 | ~spl4_98 | ~spl4_117),
  inference(avatar_split_clause,[],[f2429,f1099,f1019,f885])).
fof(f2429,plain,(
  e0 = op(e2,e1) | (~spl4_98 | ~spl4_117)),
  inference(backward_demodulation,[],[f1021,f2135])).
fof(f2135,plain,(
  e0 = op(op(e1,e0),e1) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f332])).
fof(f2526,plain,(
  ~spl4_53 | ~spl4_68),
  inference(avatar_split_clause,[],[f2525,f893,f830])).
fof(f2525,plain,(
  e2 != op(e2,e4) | ~spl4_68),
  inference(forward_demodulation,[],[f432,f895])).
fof(f2501,plain,(
  ~spl4_32 | ~spl4_7),
  inference(avatar_split_clause,[],[f2500,f637,f742])).
fof(f2500,plain,(
  e1 != op(e3,e3) | ~spl4_7),
  inference(forward_demodulation,[],[f395,f639])).
fof(f2497,plain,(
  ~spl4_6 | ~spl4_7),
  inference(avatar_contradiction_clause,[],[f2496])).
fof(f2496,plain,(
  $false | (~spl4_6 | ~spl4_7)),
  inference(subsumption_resolution,[],[f2495,f456])).
fof(f2495,plain,(
  e0 = e1 | (~spl4_6 | ~spl4_7)),
  inference(forward_demodulation,[],[f639,f635])).
fof(f2490,plain,(
  ~spl4_15 | spl4_23 | ~spl4_51),
  inference(avatar_contradiction_clause,[],[f2489])).
fof(f2489,plain,(
  $false | (~spl4_15 | spl4_23 | ~spl4_51)),
  inference(subsumption_resolution,[],[f2487,f705])).
fof(f705,plain,(
  e2 != op(e4,e0) | spl4_23),
  inference(avatar_component_clause,[],[f704])).
fof(f2487,plain,(
  e2 = op(e4,e0) | (~spl4_15 | ~spl4_51)),
  inference(backward_demodulation,[],[f672,f2356])).
fof(f2356,plain,(
  e2 = op(op(e4,e2),e0) | ~spl4_51),
  inference(forward_demodulation,[],[f345,f824])).
fof(f2484,plain,(
  ~spl4_31 | ~spl4_50 | spl4_186),
  inference(avatar_contradiction_clause,[],[f2483])).
fof(f2483,plain,(
  $false | (~spl4_31 | ~spl4_50 | spl4_186)),
  inference(subsumption_resolution,[],[f2479,f819])).
fof(f2479,plain,(
  e4 != op(e3,e0) | (~spl4_31 | spl4_186)),
  inference(backward_demodulation,[],[f740,f1550])).
fof(f2482,plain,(
  ~spl4_31 | ~spl4_121),
  inference(avatar_contradiction_clause,[],[f2481])).
fof(f2481,plain,(
  $false | (~spl4_31 | ~spl4_121)),
  inference(subsumption_resolution,[],[f2480,f458])).
fof(f2480,plain,(
  e0 = e3 | (~spl4_31 | ~spl4_121)),
  inference(forward_demodulation,[],[f2472,f1118])).
fof(f2472,plain,(
  op(e0,e0) = e3 | ~spl4_31),
  inference(backward_demodulation,[],[f740,f349])).
fof(f2449,plain,(
  ~spl4_86 | ~spl4_87),
  inference(avatar_contradiction_clause,[],[f2448])).
fof(f2448,plain,(
  $false | (~spl4_86 | ~spl4_87)),
  inference(subsumption_resolution,[],[f2447,f456])).
fof(f2447,plain,(
  e0 = e1 | (~spl4_86 | ~spl4_87)),
  inference(forward_demodulation,[],[f975,f971])).
fof(f2433,plain,(
  ~spl4_86 | ~spl4_98 | ~spl4_117),
  inference(avatar_contradiction_clause,[],[f2432])).
fof(f2432,plain,(
  $false | (~spl4_86 | ~spl4_98 | ~spl4_117)),
  inference(subsumption_resolution,[],[f2431,f456])).
fof(f2431,plain,(
  e0 = e1 | (~spl4_86 | ~spl4_98 | ~spl4_117)),
  inference(forward_demodulation,[],[f2428,f971])).
fof(f2428,plain,(
  e1 = op(e1,e2) | (~spl4_98 | ~spl4_117)),
  inference(backward_demodulation,[],[f1021,f2136])).
fof(f2136,plain,(
  e1 = op(e1,op(e1,e0)) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f336])).
fof(f2425,plain,(
  ~spl4_63 | ~spl4_113),
  inference(avatar_split_clause,[],[f2422,f1082,f872])).
fof(f2422,plain,(
  e2 != op(e2,e2) | ~spl4_113),
  inference(backward_demodulation,[],[f1084,f377])).
fof(f2424,plain,(
  spl4_68 | ~spl4_72 | ~spl4_113),
  inference(avatar_split_clause,[],[f2421,f1082,f910,f893])).
fof(f2421,plain,(
  e2 = op(e2,e1) | (~spl4_72 | ~spl4_113)),
  inference(backward_demodulation,[],[f1084,f2418])).
fof(f2418,plain,(
  e2 = op(op(e0,e2),e1) | ~spl4_72),
  inference(forward_demodulation,[],[f341,f912])).
fof(f2417,plain,(
  ~spl4_112 | ~spl4_37),
  inference(avatar_split_clause,[],[f2416,f763,f1078])).
fof(f2416,plain,(
  e1 != op(e0,e2) | ~spl4_37),
  inference(forward_demodulation,[],[f378,f765])).
fof(f2405,plain,(
  ~spl4_94 | ~spl4_44),
  inference(avatar_split_clause,[],[f2404,f792,f1002])).
fof(f2404,plain,(
  e3 != op(e1,e1) | ~spl4_44),
  inference(forward_demodulation,[],[f371,f794])).
fof(f2401,plain,(
  spl4_59 | ~spl4_44 | ~spl4_83),
  inference(avatar_split_clause,[],[f2400,f956,f792,f855])).
fof(f2400,plain,(
  e3 = op(e2,e3) | (~spl4_44 | ~spl4_83)),
  inference(forward_demodulation,[],[f2399,f958])).
fof(f2399,plain,(
  e3 = op(op(e1,e3),e3) | ~spl4_44),
  inference(forward_demodulation,[],[f347,f794])).
fof(f2386,plain,(
  ~spl4_83 | ~spl4_85),
  inference(avatar_contradiction_clause,[],[f2385])).
fof(f2385,plain,(
  $false | (~spl4_83 | ~spl4_85)),
  inference(subsumption_resolution,[],[f2384,f464])).
fof(f2384,plain,(
  e2 = e4 | (~spl4_83 | ~spl4_85)),
  inference(forward_demodulation,[],[f966,f958])).
fof(f2383,plain,(
  ~spl4_69 | ~spl4_44),
  inference(avatar_split_clause,[],[f2382,f792,f897])).
fof(f2382,plain,(
  e3 != op(e2,e1) | ~spl4_44),
  inference(forward_demodulation,[],[f373,f794])).
fof(f2379,plain,(
  ~spl4_37 | ~spl4_60 | spl4_78),
  inference(avatar_contradiction_clause,[],[f2378])).
fof(f2378,plain,(
  $false | (~spl4_37 | ~spl4_60 | spl4_78)),
  inference(subsumption_resolution,[],[f2377,f936])).
fof(f936,plain,(
  e2 != op(e1,e4) | spl4_78),
  inference(avatar_component_clause,[],[f935])).
fof(f2377,plain,(
  e2 = op(e1,e4) | (~spl4_37 | ~spl4_60)),
  inference(forward_demodulation,[],[f2376,f765])).
fof(f2376,plain,(
  e2 = op(op(e3,e2),e4) | ~spl4_60),
  inference(forward_demodulation,[],[f344,f861])).
fof(f2375,plain,(
  ~spl4_16 | ~spl4_37 | ~spl4_60),
  inference(avatar_contradiction_clause,[],[f2374])).
fof(f2374,plain,(
  $false | (~spl4_16 | ~spl4_37 | ~spl4_60)),
  inference(subsumption_resolution,[],[f2373,f458])).
fof(f2373,plain,(
  e0 = e3 | (~spl4_16 | ~spl4_37 | ~spl4_60)),
  inference(forward_demodulation,[],[f2372,f677])).
fof(f2372,plain,(
  e3 = op(e4,e1) | (~spl4_37 | ~spl4_60)),
  inference(forward_demodulation,[],[f2371,f861])).
fof(f2371,plain,(
  e3 = op(op(e2,e3),e1) | ~spl4_37),
  inference(forward_demodulation,[],[f348,f765])).
fof(f2368,plain,(
  ~spl4_56 | ~spl4_6),
  inference(avatar_split_clause,[],[f2367,f633,f843])).
fof(f2367,plain,(
  e0 != op(e2,e3) | ~spl4_6),
  inference(forward_demodulation,[],[f394,f635])).
fof(f2340,plain,(
  ~spl4_48 | ~spl4_28),
  inference(avatar_split_clause,[],[f2333,f725,f809])).
fof(f2333,plain,(
  e2 != op(e3,e0) | ~spl4_28),
  inference(backward_demodulation,[],[f727,f439])).
fof(f2338,plain,(
  ~spl4_6 | ~spl4_28 | ~spl4_72),
  inference(avatar_contradiction_clause,[],[f2337])).
fof(f2337,plain,(
  $false | (~spl4_6 | ~spl4_28 | ~spl4_72)),
  inference(subsumption_resolution,[],[f2336,f462])).
fof(f2336,plain,(
  e1 = e4 | (~spl4_6 | ~spl4_28 | ~spl4_72)),
  inference(forward_demodulation,[],[f2331,f912])).
fof(f2331,plain,(
  e4 = op(e2,e0) | (~spl4_6 | ~spl4_28)),
  inference(backward_demodulation,[],[f727,f2025])).
fof(f2025,plain,(
  e4 = op(op(e3,e4),e0) | ~spl4_6),
  inference(forward_demodulation,[],[f354,f635])).
fof(f2319,plain,(
  spl4_26 | ~spl4_24 | ~spl4_105),
  inference(avatar_split_clause,[],[f2314,f1048,f708,f717])).
fof(f2314,plain,(
  e0 = op(e3,e4) | (~spl4_24 | ~spl4_105)),
  inference(backward_demodulation,[],[f710,f1654])).
fof(f2318,plain,(
  ~spl4_14 | ~spl4_24),
  inference(avatar_split_clause,[],[f2312,f708,f666])).
fof(f2312,plain,(
  e3 != op(e4,e2) | ~spl4_24),
  inference(backward_demodulation,[],[f710,f447])).
fof(f2293,plain,(
  spl4_50 | ~spl4_16 | ~spl4_79),
  inference(avatar_split_clause,[],[f2292,f939,f675,f817])).
fof(f2292,plain,(
  e4 = op(e3,e0) | (~spl4_16 | ~spl4_79)),
  inference(forward_demodulation,[],[f2291,f941])).
fof(f2291,plain,(
  e4 = op(op(e1,e4),e0) | ~spl4_16),
  inference(forward_demodulation,[],[f352,f677])).
fof(f2290,plain,(
  spl4_107 | ~spl4_16 | ~spl4_79),
  inference(avatar_split_clause,[],[f2289,f939,f675,f1057])).
fof(f2289,plain,(
  e1 = op(e0,e3) | (~spl4_16 | ~spl4_79)),
  inference(forward_demodulation,[],[f2092,f677])).
fof(f2092,plain,(
  e1 = op(op(e4,e1),e3) | ~spl4_79),
  inference(forward_demodulation,[],[f340,f941])).
fof(f2284,plain,(
  ~spl4_12 | ~spl4_37),
  inference(avatar_split_clause,[],[f2283,f763,f658])).
fof(f2283,plain,(
  e1 != op(e4,e2) | ~spl4_37),
  inference(forward_demodulation,[],[f385,f765])).
fof(f2275,plain,(
  ~spl4_16 | ~spl4_18),
  inference(avatar_contradiction_clause,[],[f2274])).
fof(f2274,plain,(
  $false | (~spl4_16 | ~spl4_18)),
  inference(subsumption_resolution,[],[f2273,f457])).
fof(f2273,plain,(
  e0 = e2 | (~spl4_16 | ~spl4_18)),
  inference(backward_demodulation,[],[f677,f685])).
fof(f2266,plain,(
  ~spl4_37 | ~spl4_38),
  inference(avatar_contradiction_clause,[],[f2265])).
fof(f2265,plain,(
  $false | (~spl4_37 | ~spl4_38)),
  inference(subsumption_resolution,[],[f2264,f460])).
fof(f2264,plain,(
  e1 = e2 | (~spl4_37 | ~spl4_38)),
  inference(backward_demodulation,[],[f765,f769])).
fof(f2263,plain,(
  ~spl4_38 | ~spl4_40),
  inference(avatar_contradiction_clause,[],[f2262])).
fof(f2262,plain,(
  $false | (~spl4_38 | ~spl4_40)),
  inference(subsumption_resolution,[],[f2261,f464])).
fof(f2261,plain,(
  e2 = e4 | (~spl4_38 | ~spl4_40)),
  inference(backward_demodulation,[],[f769,f777])).
fof(f2252,plain,(
  ~spl4_44 | ~spl4_82),
  inference(avatar_contradiction_clause,[],[f2251])).
fof(f2251,plain,(
  $false | (~spl4_44 | ~spl4_82)),
  inference(subsumption_resolution,[],[f2250,f461])).
fof(f2250,plain,(
  e1 = e3 | (~spl4_44 | ~spl4_82)),
  inference(forward_demodulation,[],[f2247,f954])).
fof(f2247,plain,(
  e3 = op(e1,e3) | (~spl4_44 | ~spl4_82)),
  inference(backward_demodulation,[],[f794,f1730])).
fof(f1730,plain,(
  e3 = op(e1,op(e3,e1)) | ~spl4_82),
  inference(backward_demodulation,[],[f954,f347])).
fof(f2244,plain,(
  ~spl4_48 | ~spl4_50),
  inference(avatar_contradiction_clause,[],[f2243])).
fof(f2243,plain,(
  $false | (~spl4_48 | ~spl4_50)),
  inference(subsumption_resolution,[],[f2242,f464])).
fof(f2242,plain,(
  e2 = e4 | (~spl4_48 | ~spl4_50)),
  inference(backward_demodulation,[],[f811,f819])).
fof(f2232,plain,(
  ~spl4_51 | ~spl4_54),
  inference(avatar_contradiction_clause,[],[f2231])).
fof(f2231,plain,(
  $false | (~spl4_51 | ~spl4_54)),
  inference(subsumption_resolution,[],[f2230,f458])).
fof(f2230,plain,(
  e0 = e3 | (~spl4_51 | ~spl4_54)),
  inference(backward_demodulation,[],[f824,f836])).
fof(f2218,plain,(
  ~spl4_62 | ~spl4_92),
  inference(avatar_contradiction_clause,[],[f2217])).
fof(f2217,plain,(
  $false | (~spl4_62 | ~spl4_92)),
  inference(subsumption_resolution,[],[f2216,f460])).
fof(f2216,plain,(
  e1 = e2 | (~spl4_62 | ~spl4_92)),
  inference(forward_demodulation,[],[f2210,f996])).
fof(f2210,plain,(
  e2 = op(e1,e1) | ~spl4_62),
  inference(backward_demodulation,[],[f870,f343])).
fof(f2206,plain,(
  ~spl4_62 | ~spl4_67),
  inference(avatar_split_clause,[],[f2202,f889,f868])).
fof(f2202,plain,(
  e1 != op(e2,e2) | ~spl4_67),
  inference(backward_demodulation,[],[f891,f430])).
fof(f2199,plain,(
  ~spl4_67 | ~spl4_72),
  inference(avatar_split_clause,[],[f2191,f910,f889])).
fof(f2191,plain,(
  e1 != op(e2,e1) | ~spl4_72),
  inference(backward_demodulation,[],[f912,f426])).
fof(f2197,plain,(
  ~spl4_72 | spl4_91 | ~spl4_112),
  inference(avatar_contradiction_clause,[],[f2196])).
fof(f2196,plain,(
  $false | (~spl4_72 | spl4_91 | ~spl4_112)),
  inference(subsumption_resolution,[],[f2189,f991])).
fof(f2189,plain,(
  e0 = op(e1,e1) | (~spl4_72 | ~spl4_112)),
  inference(backward_demodulation,[],[f912,f2147])).
fof(f2147,plain,(
  e0 = op(op(e2,e0),e1) | ~spl4_112),
  inference(backward_demodulation,[],[f1080,f333])).
fof(f2186,plain,(
  ~spl4_82 | ~spl4_83),
  inference(avatar_contradiction_clause,[],[f2185])).
fof(f2185,plain,(
  $false | (~spl4_82 | ~spl4_83)),
  inference(subsumption_resolution,[],[f2184,f460])).
fof(f2184,plain,(
  e1 = e2 | (~spl4_82 | ~spl4_83)),
  inference(forward_demodulation,[],[f958,f954])).
fof(f2183,plain,(
  ~spl4_82 | ~spl4_85),
  inference(avatar_contradiction_clause,[],[f2182])).
fof(f2182,plain,(
  $false | (~spl4_82 | ~spl4_85)),
  inference(subsumption_resolution,[],[f2181,f462])).
fof(f2181,plain,(
  e1 = e4 | (~spl4_82 | ~spl4_85)),
  inference(forward_demodulation,[],[f966,f954])).
fof(f2177,plain,(
  ~spl4_91 | ~spl4_92 | spl4_148),
  inference(avatar_split_clause,[],[f2171,f1309,f994,f990])).
fof(f2171,plain,(
  e0 != op(e1,e1) | (~spl4_92 | spl4_148)),
  inference(backward_demodulation,[],[f996,f1311])).
fof(f2164,plain,(
  spl4_91 | ~spl4_97 | ~spl4_117),
  inference(avatar_split_clause,[],[f2159,f1099,f1015,f990])).
fof(f2159,plain,(
  e0 = op(e1,e1) | (~spl4_97 | ~spl4_117)),
  inference(backward_demodulation,[],[f1017,f2135])).
fof(f2145,plain,(
  ~spl4_112 | ~spl4_117),
  inference(avatar_split_clause,[],[f2140,f1099,f1078])).
fof(f2140,plain,(
  e1 != op(e0,e2) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f410])).
fof(f2142,plain,(
  ~spl4_92 | ~spl4_117),
  inference(avatar_split_clause,[],[f2137,f1099,f994])).
fof(f2137,plain,(
  e1 != op(e1,e1) | ~spl4_117),
  inference(backward_demodulation,[],[f1101,f366])).
fof(f2114,plain,(
  ~spl4_103 | ~spl4_105),
  inference(avatar_contradiction_clause,[],[f2113])).
fof(f2113,plain,(
  $false | (~spl4_103 | ~spl4_105)),
  inference(subsumption_resolution,[],[f2112,f464])).
fof(f2112,plain,(
  e2 = e4 | (~spl4_103 | ~spl4_105)),
  inference(forward_demodulation,[],[f1050,f1042])).
fof(f2111,plain,(
  ~spl4_97 | ~spl4_22),
  inference(avatar_split_clause,[],[f2110,f700,f1015])).
fof(f2110,plain,(
  e1 != op(e1,e0) | ~spl4_22),
  inference(forward_demodulation,[],[f362,f702])).
fof(f2103,plain,(
  ~spl4_93 | ~spl4_18),
  inference(avatar_split_clause,[],[f2102,f683,f998])).
fof(f2102,plain,(
  e2 != op(e1,e1) | ~spl4_18),
  inference(forward_demodulation,[],[f372,f685])).
fof(f2095,plain,(
  spl4_95 | ~spl4_2),
  inference(avatar_split_clause,[],[f1968,f616,f1006])).
fof(f1968,plain,(
  e4 = op(e1,e1) | ~spl4_2),
  inference(backward_demodulation,[],[f618,f355])).
fof(f2094,plain,(
  spl4_57 | ~spl4_18 | ~spl4_79),
  inference(avatar_split_clause,[],[f2093,f939,f683,f847])).
fof(f2093,plain,(
  e1 = op(e2,e3) | (~spl4_18 | ~spl4_79)),
  inference(forward_demodulation,[],[f2092,f685])).
fof(f2091,plain,(
  spl4_40 | ~spl4_18 | ~spl4_79),
  inference(avatar_split_clause,[],[f2090,f939,f683,f775])).
fof(f2090,plain,(
  e4 = op(e3,e2) | (~spl4_18 | ~spl4_79)),
  inference(forward_demodulation,[],[f2089,f941])).
fof(f2089,plain,(
  e4 = op(op(e1,e4),e2) | ~spl4_18),
  inference(forward_demodulation,[],[f352,f685])).
fof(f2076,plain,(
  ~spl4_68 | ~spl4_18),
  inference(avatar_split_clause,[],[f2075,f683,f893])).
fof(f2075,plain,(
  e2 != op(e2,e1) | ~spl4_18),
  inference(forward_demodulation,[],[f374,f685])).
fof(f2064,plain,(
  ~spl4_12 | ~spl4_54 | ~spl4_82),
  inference(avatar_contradiction_clause,[],[f2063])).
fof(f2063,plain,(
  $false | (~spl4_12 | ~spl4_54 | ~spl4_82)),
  inference(subsumption_resolution,[],[f2062,f460])).
fof(f2062,plain,(
  e1 = e2 | (~spl4_12 | ~spl4_54 | ~spl4_82)),
  inference(forward_demodulation,[],[f2061,f954])).
fof(f2061,plain,(
  e2 = op(e1,e3) | (~spl4_12 | ~spl4_54)),
  inference(forward_demodulation,[],[f2060,f660])).
fof(f2060,plain,(
  e2 = op(op(e4,e2),e3) | ~spl4_54),
  inference(forward_demodulation,[],[f345,f836])).
fof(f2059,plain,(
  spl4_45 | ~spl4_12 | ~spl4_54),
  inference(avatar_split_clause,[],[f2058,f834,f658,f796])).
fof(f2058,plain,(
  e4 = op(e3,e1) | (~spl4_12 | ~spl4_54)),
  inference(forward_demodulation,[],[f2057,f836])).
fof(f2057,plain,(
  e4 = op(op(e2,e4),e1) | ~spl4_12),
  inference(forward_demodulation,[],[f353,f660])).
fof(f2048,plain,(
  ~spl4_46 | ~spl4_26),
  inference(avatar_split_clause,[],[f2047,f717,f801])).
fof(f2047,plain,(
  e0 != op(e3,e0) | ~spl4_26),
  inference(forward_demodulation,[],[f439,f719])).
fof(f2042,plain,(
  ~spl4_41 | ~spl4_26),
  inference(avatar_split_clause,[],[f2041,f717,f780])).
fof(f2041,plain,(
  e0 != op(e3,e1) | ~spl4_26),
  inference(forward_demodulation,[],[f442,f719])).
fof(f2022,plain,(
  ~spl4_26 | ~spl4_28),
  inference(avatar_contradiction_clause,[],[f2021])).
fof(f2021,plain,(
  $false | (~spl4_26 | ~spl4_28)),
  inference(subsumption_resolution,[],[f2020,f457])).
fof(f2020,plain,(
  e0 = e2 | (~spl4_26 | ~spl4_28)),
  inference(forward_demodulation,[],[f727,f719])).
fof(f2009,plain,(
  spl4_76 | ~spl4_22 | ~spl4_105),
  inference(avatar_split_clause,[],[f2008,f1048,f700,f927])).
fof(f2008,plain,(
  e0 = op(e1,e4) | (~spl4_22 | ~spl4_105)),
  inference(forward_demodulation,[],[f1654,f702])).
fof(f2003,plain,(
  ~spl4_22 | ~spl4_24),
  inference(avatar_contradiction_clause,[],[f2002])).
fof(f2002,plain,(
  $false | (~spl4_22 | ~spl4_24)),
  inference(subsumption_resolution,[],[f2001,f461])).
fof(f2001,plain,(
  e1 = e3 | (~spl4_22 | ~spl4_24)),
  inference(forward_demodulation,[],[f710,f702])).
fof(f2000,plain,(
  ~spl4_18 | ~spl4_20),
  inference(avatar_contradiction_clause,[],[f1999])).
fof(f1999,plain,(
  $false | (~spl4_18 | ~spl4_20)),
  inference(subsumption_resolution,[],[f1998,f464])).
fof(f1998,plain,(
  e2 = e4 | (~spl4_18 | ~spl4_20)),
  inference(forward_demodulation,[],[f693,f685])).
fof(f1997,plain,(
  ~spl4_69 | ~spl4_2 | ~spl4_18 | spl4_136),
  inference(avatar_split_clause,[],[f1995,f1254,f683,f616,f897])).
fof(f1995,plain,(
  e3 != op(e2,e1) | (~spl4_2 | ~spl4_18 | spl4_136)),
  inference(backward_demodulation,[],[f685,f1973])).
fof(f1973,plain,(
  e3 != op(op(e4,e1),e1) | (~spl4_2 | spl4_136)),
  inference(backward_demodulation,[],[f618,f1256])).
fof(f1996,plain,(
  ~spl4_64 | ~spl4_2 | ~spl4_18 | spl4_132),
  inference(avatar_split_clause,[],[f1994,f1236,f683,f616,f876])).
fof(f1994,plain,(
  e3 != op(e2,e2) | (~spl4_2 | ~spl4_18 | spl4_132)),
  inference(backward_demodulation,[],[f685,f1972])).
fof(f1972,plain,(
  e3 != op(op(e4,e1),op(e4,e1)) | (~spl4_2 | spl4_132)),
  inference(backward_demodulation,[],[f618,f1238])).
fof(f1993,plain,(
  ~spl4_19 | ~spl4_2 | spl4_182),
  inference(avatar_split_clause,[],[f1974,f1508,f616,f687])).
fof(f1974,plain,(
  e3 != op(e4,e1) | (~spl4_2 | spl4_182)),
  inference(backward_demodulation,[],[f618,f1510])).
fof(f1992,plain,(
  ~spl4_11 | ~spl4_6),
  inference(avatar_split_clause,[],[f1991,f633,f654])).
fof(f1991,plain,(
  e0 != op(e4,e2) | ~spl4_6),
  inference(forward_demodulation,[],[f453,f635])).
fof(f1987,plain,(
  ~spl4_12 | ~spl4_14),
  inference(avatar_contradiction_clause,[],[f1986])).
fof(f1986,plain,(
  $false | (~spl4_12 | ~spl4_14)),
  inference(subsumption_resolution,[],[f1985,f461])).
fof(f1985,plain,(
  e1 = e3 | (~spl4_12 | ~spl4_14)),
  inference(forward_demodulation,[],[f668,f660])).
fof(f1977,plain,(
  ~spl4_2 | ~spl4_91),
  inference(avatar_contradiction_clause,[],[f1976])).
fof(f1976,plain,(
  $false | (~spl4_2 | ~spl4_91)),
  inference(subsumption_resolution,[],[f1975,f459])).
fof(f1975,plain,(
  e0 = e4 | (~spl4_2 | ~spl4_91)),
  inference(forward_demodulation,[],[f1968,f992])).
fof(f1967,plain,(
  ~spl4_6 | ~spl4_10),
  inference(avatar_contradiction_clause,[],[f1966])).
fof(f1966,plain,(
  $false | (~spl4_6 | ~spl4_10)),
  inference(subsumption_resolution,[],[f1965,f459])).
fof(f1965,plain,(
  e0 = e4 | (~spl4_6 | ~spl4_10)),
  inference(backward_demodulation,[],[f635,f651])).
fof(f1964,plain,(
  spl4_5 | ~spl4_10 | ~spl4_30),
  inference(avatar_split_clause,[],[f1961,f733,f649,f628])).
fof(f1961,plain,(
  e4 = op(e4,e4) | (~spl4_10 | ~spl4_30)),
  inference(backward_demodulation,[],[f651,f1889])).
fof(f1889,plain,(
  e4 = op(e4,op(e4,e3)) | ~spl4_30),
  inference(backward_demodulation,[],[f735,f354])).
fof(f1958,plain,(
  ~spl4_12 | ~spl4_13),
  inference(avatar_contradiction_clause,[],[f1957])).
fof(f1957,plain,(
  $false | (~spl4_12 | ~spl4_13)),
  inference(subsumption_resolution,[],[f1956,f460])).
fof(f1956,plain,(
  e1 = e2 | (~spl4_12 | ~spl4_13)),
  inference(backward_demodulation,[],[f660,f664])).
fof(f1955,plain,(
  ~spl4_13 | ~spl4_14),
  inference(avatar_contradiction_clause,[],[f1954])).
fof(f1954,plain,(
  $false | (~spl4_13 | ~spl4_14)),
  inference(subsumption_resolution,[],[f1953,f463])).
fof(f1953,plain,(
  e2 = e3 | (~spl4_13 | ~spl4_14)),
  inference(backward_demodulation,[],[f664,f668])).
fof(f1952,plain,(
  ~spl4_14 | ~spl4_15),
  inference(avatar_contradiction_clause,[],[f1951])).
fof(f1951,plain,(
  $false | (~spl4_14 | ~spl4_15)),
  inference(subsumption_resolution,[],[f1950,f465])).
fof(f1950,plain,(
  e3 = e4 | (~spl4_14 | ~spl4_15)),
  inference(backward_demodulation,[],[f668,f672])).
fof(f1949,plain,(
  spl4_5 | ~spl4_15 | ~spl4_55),
  inference(avatar_split_clause,[],[f1945,f838,f670,f628])).
fof(f1945,plain,(
  e4 = op(e4,e4) | (~spl4_15 | ~spl4_55)),
  inference(backward_demodulation,[],[f672,f1818])).
fof(f1818,plain,(
  e4 = op(e4,op(e4,e2)) | ~spl4_55),
  inference(backward_demodulation,[],[f840,f353])).
fof(f1941,plain,(
  ~spl4_17 | ~spl4_18),
  inference(avatar_contradiction_clause,[],[f1940])).
fof(f1940,plain,(
  $false | (~spl4_17 | ~spl4_18)),
  inference(subsumption_resolution,[],[f1939,f460])).
fof(f1939,plain,(
  e1 = e2 | (~spl4_17 | ~spl4_18)),
  inference(backward_demodulation,[],[f681,f685])).
fof(f1938,plain,(
  ~spl4_18 | ~spl4_19),
  inference(avatar_contradiction_clause,[],[f1937])).
fof(f1937,plain,(
  $false | (~spl4_18 | ~spl4_19)),
  inference(subsumption_resolution,[],[f1936,f463])).
fof(f1936,plain,(
  e2 = e3 | (~spl4_18 | ~spl4_19)),
  inference(backward_demodulation,[],[f685,f689])).
fof(f1932,plain,(
  spl4_5 | ~spl4_20 | ~spl4_80),
  inference(avatar_split_clause,[],[f1927,f943,f691,f628])).
fof(f1927,plain,(
  e4 = op(e4,e4) | (~spl4_20 | ~spl4_80)),
  inference(backward_demodulation,[],[f693,f1740])).
fof(f1740,plain,(
  e4 = op(e4,op(e4,e1)) | ~spl4_80),
  inference(backward_demodulation,[],[f945,f352])).
fof(f1922,plain,(
  ~spl4_22 | ~spl4_23),
  inference(avatar_contradiction_clause,[],[f1921])).
fof(f1921,plain,(
  $false | (~spl4_22 | ~spl4_23)),
  inference(subsumption_resolution,[],[f1920,f460])).
fof(f1920,plain,(
  e1 = e2 | (~spl4_22 | ~spl4_23)),
  inference(backward_demodulation,[],[f702,f706])).
fof(f1919,plain,(
  ~spl4_23 | ~spl4_24),
  inference(avatar_contradiction_clause,[],[f1918])).
fof(f1918,plain,(
  $false | (~spl4_23 | ~spl4_24)),
  inference(subsumption_resolution,[],[f1917,f463])).
fof(f1917,plain,(
  e2 = e3 | (~spl4_23 | ~spl4_24)),
  inference(backward_demodulation,[],[f706,f710])).
fof(f1913,plain,(
  spl4_5 | ~spl4_25 | ~spl4_105),
  inference(avatar_split_clause,[],[f1907,f1048,f712,f628])).
fof(f1907,plain,(
  e4 = op(e4,e4) | (~spl4_25 | ~spl4_105)),
  inference(backward_demodulation,[],[f714,f1655])).
fof(f1655,plain,(
  e4 = op(e4,op(e4,e0)) | ~spl4_105),
  inference(backward_demodulation,[],[f1050,f351])).
fof(f1901,plain,(
  ~spl4_26 | ~spl4_27),
  inference(avatar_contradiction_clause,[],[f1900])).
fof(f1900,plain,(
  $false | (~spl4_26 | ~spl4_27)),
  inference(subsumption_resolution,[],[f1899,f456])).
fof(f1899,plain,(
  e0 = e1 | (~spl4_26 | ~spl4_27)),
  inference(backward_demodulation,[],[f719,f723])).
fof(f1898,plain,(
  ~spl4_27 | ~spl4_28),
  inference(avatar_contradiction_clause,[],[f1897])).
fof(f1897,plain,(
  $false | (~spl4_27 | ~spl4_28)),
  inference(subsumption_resolution,[],[f1896,f460])).
fof(f1896,plain,(
  e1 = e2 | (~spl4_27 | ~spl4_28)),
  inference(backward_demodulation,[],[f723,f727])).
fof(f1895,plain,(
  ~spl4_28 | ~spl4_30),
  inference(avatar_contradiction_clause,[],[f1894])).
fof(f1894,plain,(
  $false | (~spl4_28 | ~spl4_30)),
  inference(subsumption_resolution,[],[f1893,f464])).
fof(f1893,plain,(
  e2 = e4 | (~spl4_28 | ~spl4_30)),
  inference(backward_demodulation,[],[f727,f735])).
fof(f1860,plain,(
  ~spl4_41 | ~spl4_82 | ~spl4_96),
  inference(avatar_contradiction_clause,[],[f1859])).
fof(f1859,plain,(
  $false | (~spl4_41 | ~spl4_82 | ~spl4_96)),
  inference(subsumption_resolution,[],[f1858,f458])).
fof(f1858,plain,(
  e0 = e3 | (~spl4_41 | ~spl4_82 | ~spl4_96)),
  inference(forward_demodulation,[],[f1850,f1013])).
fof(f1850,plain,(
  e3 = op(e1,e0) | (~spl4_41 | ~spl4_82)),
  inference(backward_demodulation,[],[f782,f1730])).
fof(f1838,plain,(
  ~spl4_41 | ~spl4_46),
  inference(avatar_split_clause,[],[f1831,f801,f780])).
fof(f1831,plain,(
  e0 != op(e3,e1) | ~spl4_46),
  inference(backward_demodulation,[],[f803,f436])).
fof(f1829,plain,(
  ~spl4_52 | ~spl4_54),
  inference(avatar_contradiction_clause,[],[f1828])).
fof(f1828,plain,(
  $false | (~spl4_52 | ~spl4_54)),
  inference(subsumption_resolution,[],[f1827,f461])).
fof(f1827,plain,(
  e1 = e3 | (~spl4_52 | ~spl4_54)),
  inference(backward_demodulation,[],[f828,f836])).
fof(f1825,plain,(
  ~spl4_54 | ~spl4_55),
  inference(avatar_contradiction_clause,[],[f1824])).
fof(f1824,plain,(
  $false | (~spl4_54 | ~spl4_55)),
  inference(subsumption_resolution,[],[f1823,f465])).
fof(f1823,plain,(
  e3 = e4 | (~spl4_54 | ~spl4_55)),
  inference(backward_demodulation,[],[f836,f840])).
fof(f1815,plain,(
  ~spl4_8 | ~spl4_58),
  inference(avatar_split_clause,[],[f1812,f851,f641])).
fof(f1812,plain,(
  e2 != op(e4,e3) | ~spl4_58),
  inference(backward_demodulation,[],[f853,f394])).
fof(f1799,plain,(
  ~spl4_61 | ~spl4_121),
  inference(avatar_contradiction_clause,[],[f1798])).
fof(f1798,plain,(
  $false | (~spl4_61 | ~spl4_121)),
  inference(subsumption_resolution,[],[f1797,f457])).
fof(f1797,plain,(
  e0 = e2 | (~spl4_61 | ~spl4_121)),
  inference(forward_demodulation,[],[f1788,f1118])).
fof(f1788,plain,(
  op(e0,e0) = e2 | ~spl4_61),
  inference(backward_demodulation,[],[f866,f343])).
fof(f1787,plain,(
  ~spl4_66 | ~spl4_86 | ~spl4_121),
  inference(avatar_contradiction_clause,[],[f1786])).
fof(f1786,plain,(
  $false | (~spl4_66 | ~spl4_86 | ~spl4_121)),
  inference(subsumption_resolution,[],[f1785,f457])).
fof(f1785,plain,(
  e0 = e2 | (~spl4_66 | ~spl4_86 | ~spl4_121)),
  inference(forward_demodulation,[],[f1776,f1118])).
fof(f1776,plain,(
  op(e0,e0) = e2 | (~spl4_66 | ~spl4_86)),
  inference(backward_demodulation,[],[f887,f1718])).
fof(f1718,plain,(
  e2 = op(e0,op(e2,e1)) | ~spl4_86),
  inference(backward_demodulation,[],[f971,f342])).
fof(f1781,plain,(
  ~spl4_51 | ~spl4_66),
  inference(avatar_split_clause,[],[f1774,f885,f822])).
fof(f1774,plain,(
  e0 != op(e2,e4) | ~spl4_66),
  inference(backward_demodulation,[],[f887,f432])).
fof(f1752,plain,(
  ~spl4_78 | ~spl4_79),
  inference(avatar_contradiction_clause,[],[f1751])).
fof(f1751,plain,(
  $false | (~spl4_78 | ~spl4_79)),
  inference(subsumption_resolution,[],[f1750,f463])).
fof(f1750,plain,(
  e2 = e3 | (~spl4_78 | ~spl4_79)),
  inference(backward_demodulation,[],[f937,f941])).
fof(f1726,plain,(
  ~spl4_11 | ~spl4_86),
  inference(avatar_split_clause,[],[f1721,f969,f654])).
fof(f1721,plain,(
  e0 != op(e4,e2) | ~spl4_86),
  inference(backward_demodulation,[],[f971,f382])).
fof(f1684,plain,(
  ~spl4_21 | ~spl4_96),
  inference(avatar_split_clause,[],[f1675,f1011,f696])).
fof(f1675,plain,(
  e0 != op(e4,e0) | ~spl4_96),
  inference(backward_demodulation,[],[f1013,f362])).
fof(f1672,plain,(
  ~spl4_102 | ~spl4_103),
  inference(avatar_contradiction_clause,[],[f1671])).
fof(f1671,plain,(
  $false | (~spl4_102 | ~spl4_103)),
  inference(subsumption_resolution,[],[f1670,f460])).
fof(f1670,plain,(
  e1 = e2 | (~spl4_102 | ~spl4_103)),
  inference(backward_demodulation,[],[f1038,f1042])).
fof(f1669,plain,(
  ~spl4_103 | ~spl4_104),
  inference(avatar_contradiction_clause,[],[f1668])).
fof(f1668,plain,(
  $false | (~spl4_103 | ~spl4_104)),
  inference(subsumption_resolution,[],[f1667,f463])).
fof(f1667,plain,(
  e2 = e3 | (~spl4_103 | ~spl4_104)),
  inference(backward_demodulation,[],[f1042,f1046])).
fof(f1651,plain,(
  ~spl4_31 | ~spl4_106),
  inference(avatar_split_clause,[],[f1646,f1053,f738])).
fof(f1646,plain,(
  e0 != op(e3,e3) | ~spl4_106),
  inference(backward_demodulation,[],[f1055,f388])).
fof(f1637,plain,(
  ~spl4_61 | ~spl4_111),
  inference(avatar_split_clause,[],[f1631,f1074,f864])).
fof(f1631,plain,(
  e0 != op(e2,e2) | ~spl4_111),
  inference(backward_demodulation,[],[f1076,f377])).
fof(f1623,plain,(
  ~spl4_41 | ~spl4_116),
  inference(avatar_split_clause,[],[f1616,f1095,f780])).
fof(f1616,plain,(
  e0 != op(e3,e1) | ~spl4_116),
  inference(backward_demodulation,[],[f1097,f368])).
fof(f1622,plain,(
  ~spl4_66 | ~spl4_116),
  inference(avatar_split_clause,[],[f1615,f1095,f885])).
fof(f1615,plain,(
  e0 != op(e2,e1) | ~spl4_116),
  inference(backward_demodulation,[],[f1097,f367])).
fof(f1605,plain,(
  ~spl4_111 | ~spl4_121),
  inference(avatar_split_clause,[],[f1593,f1116,f1074])).
fof(f1593,plain,(
  e0 != op(e0,e2) | ~spl4_121),
  inference(backward_demodulation,[],[f1118,f407])).
fof(f1602,plain,(
  ~spl4_46 | ~spl4_121),
  inference(avatar_split_clause,[],[f1590,f1116,f801])).
fof(f1590,plain,(
  e0 != op(e3,e0) | ~spl4_121),
  inference(backward_demodulation,[],[f1118,f358])).
fof(f1584,plain,(
  ~spl4_189 | ~spl4_124 | ~spl4_179 | ~spl4_173),
  inference(avatar_split_clause,[],[f608,f1441,f1487,f1128,f1569])).
fof(f608,plain,(
  e2 != op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) | e1 != op(op(e0,op(e0,e0)),op(e0,e0)) | op(e0,e0) != e3 | e4 != op(e0,op(e0,e0))),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  e2 != op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) | e1 != op(op(e0,op(e0,e0)),op(e0,e0)) | op(e0,e0) != e3 | e4 != op(e0,op(e0,e0))),
  inference(ennf_transformation,[],[f124])).
fof(f124,axiom,(
  ~(e2 = op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) & e1 = op(op(e0,op(e0,e0)),op(e0,e0)) & op(e0,e0) = e3 & e4 = op(e0,op(e0,e0)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1582,plain,(
  ~spl4_188 | ~spl4_94 | ~spl4_178 | ~spl4_157),
  inference(avatar_split_clause,[],[f606,f1353,f1481,f1002,f1559])).
fof(f606,plain,(
  e2 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e0 != op(op(e1,op(e1,e1)),op(e1,e1)) | e3 != op(e1,e1) | e4 != op(e1,op(e1,e1))),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  e2 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e0 != op(op(e1,op(e1,e1)),op(e1,e1)) | e3 != op(e1,e1) | e4 != op(e1,op(e1,e1))),
  inference(ennf_transformation,[],[f122])).
fof(f122,axiom,(
  ~(e2 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) & e0 = op(op(e1,op(e1,e1)),op(e1,e1)) & e3 = op(e1,e1) & e4 = op(e1,op(e1,e1)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1579,plain,(
  ~spl4_188 | ~spl4_93 | ~spl4_149 | ~spl4_180),
  inference(avatar_split_clause,[],[f603,f1492,f1313,f998,f1559])).
fof(f603,plain,(
  e0 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e3 != op(op(e1,op(e1,e1)),op(e1,e1)) | e2 != op(e1,e1) | e4 != op(e1,op(e1,e1))),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  e0 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e3 != op(op(e1,op(e1,e1)),op(e1,e1)) | e2 != op(e1,e1) | e4 != op(e1,op(e1,e1))),
  inference(ennf_transformation,[],[f119])).
fof(f119,axiom,(
  ~(e0 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) & e3 = op(op(e1,op(e1,e1)),op(e1,e1)) & e2 = op(e1,e1) & e4 = op(e1,op(e1,e1)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1577,plain,(
  ~spl4_186 | ~spl4_33 | ~spl4_146 | ~spl4_167),
  inference(avatar_split_clause,[],[f601,f1410,f1299,f746,f1548])).
fof(f601,plain,(
  e0 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e1 != op(op(e3,op(e3,e3)),op(e3,e3)) | e2 != op(e3,e3) | e4 != op(e3,op(e3,e3))),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  e0 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e1 != op(op(e3,op(e3,e3)),op(e3,e3)) | e2 != op(e3,e3) | e4 != op(e3,op(e3,e3))),
  inference(ennf_transformation,[],[f117])).
fof(f117,axiom,(
  ~(e0 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) & e1 = op(op(e3,op(e3,e3)),op(e3,e3)) & e2 = op(e3,e3) & e4 = op(e3,op(e3,e3)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1576,plain,(
  ~spl4_188 | ~spl4_93 | ~spl4_178 | ~spl4_153),
  inference(avatar_split_clause,[],[f600,f1331,f1481,f998,f1559])).
fof(f600,plain,(
  e3 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e0 != op(op(e1,op(e1,e1)),op(e1,e1)) | e2 != op(e1,e1) | e4 != op(e1,op(e1,e1))),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  e3 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e0 != op(op(e1,op(e1,e1)),op(e1,e1)) | e2 != op(e1,e1) | e4 != op(e1,op(e1,e1))),
  inference(ennf_transformation,[],[f116])).
fof(f116,axiom,(
  ~(e3 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) & e0 = op(op(e1,op(e1,e1)),op(e1,e1)) & e2 = op(e1,e1) & e4 = op(e1,op(e1,e1)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1565,plain,(
  ~spl4_186 | ~spl4_32 | ~spl4_162 | ~spl4_142),
  inference(avatar_split_clause,[],[f593,f1280,f1387,f742,f1548])).
fof(f593,plain,(
  e2 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e0 != op(op(e3,op(e3,e3)),op(e3,e3)) | e1 != op(e3,e3) | e4 != op(e3,op(e3,e3))),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  e2 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e0 != op(op(e3,op(e3,e3)),op(e3,e3)) | e1 != op(e3,e3) | e4 != op(e3,op(e3,e3))),
  inference(ennf_transformation,[],[f109])).
fof(f109,axiom,(
  ~(e2 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) & e0 = op(op(e3,op(e3,e3)),op(e3,e3)) & e1 = op(e3,e3) & e4 = op(e3,op(e3,e3)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1557,plain,(
  ~spl4_186 | ~spl4_31 | ~spl4_134 | ~spl4_151),
  inference(avatar_split_clause,[],[f589,f1322,f1245,f738,f1548])).
fof(f589,plain,(
  e1 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e2 != op(op(e3,op(e3,e3)),op(e3,e3)) | e0 != op(e3,e3) | e4 != op(e3,op(e3,e3))),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  e1 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e2 != op(op(e3,op(e3,e3)),op(e3,e3)) | e0 != op(e3,e3) | e4 != op(e3,op(e3,e3))),
  inference(ennf_transformation,[],[f105])).
fof(f105,axiom,(
  ~(e1 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) & e2 = op(op(e3,op(e3,e3)),op(e3,e3)) & e0 = op(e3,e3) & e4 = op(e3,op(e3,e3)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1556,plain,(
  ~spl4_187 | ~spl4_61 | ~spl4_154 | ~spl4_144),
  inference(avatar_split_clause,[],[f588,f1289,f1337,f864,f1553])).
fof(f588,plain,(
  e3 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | e1 != op(op(e2,op(e2,e2)),op(e2,e2)) | e0 != op(e2,e2) | e4 != op(e2,op(e2,e2))),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  e3 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | e1 != op(op(e2,op(e2,e2)),op(e2,e2)) | e0 != op(e2,e2) | e4 != op(e2,op(e2,e2))),
  inference(ennf_transformation,[],[f104])).
fof(f104,axiom,(
  ~(e3 = op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) & e1 = op(op(e2,op(e2,e2)),op(e2,e2)) & e0 = op(e2,e2) & e4 = op(e2,op(e2,e2)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1551,plain,(
  ~spl4_186 | ~spl4_31 | ~spl4_146 | ~spl4_142),
  inference(avatar_split_clause,[],[f587,f1280,f1299,f738,f1548])).
fof(f587,plain,(
  e2 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e1 != op(op(e3,op(e3,e3)),op(e3,e3)) | e0 != op(e3,e3) | e4 != op(e3,op(e3,e3))),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  e2 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e1 != op(op(e3,op(e3,e3)),op(e3,e3)) | e0 != op(e3,e3) | e4 != op(e3,op(e3,e3))),
  inference(ennf_transformation,[],[f103])).
fof(f103,axiom,(
  ~(e2 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) & e1 = op(op(e3,op(e3,e3)),op(e3,e3)) & e0 = op(e3,e3) & e4 = op(e3,op(e3,e3)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1544,plain,(
  ~spl4_185 | ~spl4_125 | ~spl4_179 | ~spl4_173),
  inference(avatar_split_clause,[],[f584,f1441,f1487,f1132,f1529])).
fof(f584,plain,(
  e2 != op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) | e1 != op(op(e0,op(e0,e0)),op(e0,e0)) | op(e0,e0) != e4 | e3 != op(e0,op(e0,e0))),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  e2 != op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) | e1 != op(op(e0,op(e0,e0)),op(e0,e0)) | op(e0,e0) != e4 | e3 != op(e0,op(e0,e0))),
  inference(ennf_transformation,[],[f100])).
fof(f100,axiom,(
  ~(e2 = op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) & e1 = op(op(e0,op(e0,e0)),op(e0,e0)) & op(e0,e0) = e4 & e3 = op(e0,op(e0,e0)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1539,plain,(
  ~spl4_184 | ~spl4_93 | ~spl4_152 | ~spl4_180),
  inference(avatar_split_clause,[],[f579,f1492,f1327,f998,f1519])).
fof(f579,plain,(
  e0 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e4 != op(op(e1,op(e1,e1)),op(e1,e1)) | e2 != op(e1,e1) | e3 != op(e1,op(e1,e1))),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  e0 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e4 != op(op(e1,op(e1,e1)),op(e1,e1)) | e2 != op(e1,e1) | e3 != op(e1,op(e1,e1))),
  inference(ennf_transformation,[],[f95])).
fof(f95,axiom,(
  ~(e0 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) & e4 = op(op(e1,op(e1,e1)),op(e1,e1)) & e2 = op(e1,e1) & e3 = op(e1,op(e1,e1)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1536,plain,(
  ~spl4_184 | ~spl4_93 | ~spl4_178 | ~spl4_150),
  inference(avatar_split_clause,[],[f576,f1317,f1481,f998,f1519])).
fof(f576,plain,(
  e4 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e0 != op(op(e1,op(e1,e1)),op(e1,e1)) | e2 != op(e1,e1) | e3 != op(e1,op(e1,e1))),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  e4 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e0 != op(op(e1,op(e1,e1)),op(e1,e1)) | e2 != op(e1,e1) | e3 != op(e1,op(e1,e1))),
  inference(ennf_transformation,[],[f92])).
fof(f92,axiom,(
  ~(e4 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) & e0 = op(op(e1,op(e1,e1)),op(e1,e1)) & e2 = op(e1,e1) & e3 = op(e1,op(e1,e1)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1533,plain,(
  ~spl4_183 | ~spl4_62 | ~spl4_143 | ~spl4_172),
  inference(avatar_split_clause,[],[f573,f1436,f1285,f868,f1513])).
fof(f573,plain,(
  e0 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | e4 != op(op(e2,op(e2,e2)),op(e2,e2)) | e1 != op(e2,e2) | e3 != op(e2,op(e2,e2))),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  e0 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | e4 != op(op(e2,op(e2,e2)),op(e2,e2)) | e1 != op(e2,e2) | e3 != op(e2,op(e2,e2))),
  inference(ennf_transformation,[],[f89])).
fof(f89,axiom,(
  ~(e0 = op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) & e4 = op(op(e2,op(e2,e2)),op(e2,e2)) & e1 = op(e2,e2) & e3 = op(e2,op(e2,e2)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1532,plain,(
  ~spl4_185 | ~spl4_122 | ~spl4_171 | ~spl4_166),
  inference(avatar_split_clause,[],[f572,f1405,f1431,f1120,f1529])).
fof(f572,plain,(
  e4 != op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) | e2 != op(op(e0,op(e0,e0)),op(e0,e0)) | op(e0,e0) != e1 | e3 != op(e0,op(e0,e0))),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  e4 != op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) | e2 != op(op(e0,op(e0,e0)),op(e0,e0)) | op(e0,e0) != e1 | e3 != op(e0,op(e0,e0))),
  inference(ennf_transformation,[],[f88])).
fof(f88,axiom,(
  ~(e4 = op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) & e2 = op(op(e0,op(e0,e0)),op(e0,e0)) & op(e0,e0) = e1 & e3 = op(e0,op(e0,e0)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1526,plain,(
  ~spl4_183 | ~spl4_62 | ~spl4_170 | ~spl4_140),
  inference(avatar_split_clause,[],[f570,f1271,f1425,f868,f1513])).
fof(f570,plain,(
  e4 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | e0 != op(op(e2,op(e2,e2)),op(e2,e2)) | e1 != op(e2,e2) | e3 != op(e2,op(e2,e2))),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  e4 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | e0 != op(op(e2,op(e2,e2)),op(e2,e2)) | e1 != op(e2,e2) | e3 != op(e2,op(e2,e2))),
  inference(ennf_transformation,[],[f86])).
fof(f86,axiom,(
  ~(e4 = op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) & e0 = op(op(e2,op(e2,e2)),op(e2,e2)) & e1 = op(e2,e2) & e3 = op(e2,op(e2,e2)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1525,plain,(
  ~spl4_182 | ~spl4_2 | ~spl4_161 | ~spl4_137),
  inference(avatar_split_clause,[],[f569,f1258,f1382,f616,f1508])).
fof(f569,plain,(
  e2 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e0 != op(op(e4,op(e4,e4)),op(e4,e4)) | e1 != op(e4,e4) | e3 != op(e4,op(e4,e4))),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  e2 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e0 != op(op(e4,op(e4,e4)),op(e4,e4)) | e1 != op(e4,e4) | e3 != op(e4,op(e4,e4))),
  inference(ennf_transformation,[],[f85])).
fof(f85,axiom,(
  ~(e2 = op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) & e0 = op(op(e4,op(e4,e4)),op(e4,e4)) & e1 = op(e4,e4) & e3 = op(e4,op(e4,e4)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1522,plain,(
  ~spl4_184 | ~spl4_91 | ~spl4_155 | ~spl4_150),
  inference(avatar_split_clause,[],[f566,f1317,f1343,f990,f1519])).
fof(f566,plain,(
  e4 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e2 != op(op(e1,op(e1,e1)),op(e1,e1)) | e0 != op(e1,e1) | e3 != op(e1,op(e1,e1))),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  e4 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e2 != op(op(e1,op(e1,e1)),op(e1,e1)) | e0 != op(e1,e1) | e3 != op(e1,op(e1,e1))),
  inference(ennf_transformation,[],[f82])).
fof(f82,axiom,(
  ~(e4 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) & e2 = op(op(e1,op(e1,e1)),op(e1,e1)) & e0 = op(e1,e1) & e3 = op(e1,op(e1,e1)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1502,plain,(
  ~spl4_176 | ~spl4_95 | ~spl4_178 | ~spl4_153),
  inference(avatar_split_clause,[],[f558,f1331,f1481,f1006,f1463])).
fof(f558,plain,(
  e3 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e0 != op(op(e1,op(e1,e1)),op(e1,e1)) | e4 != op(e1,e1) | e2 != op(e1,op(e1,e1))),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  e3 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e0 != op(op(e1,op(e1,e1)),op(e1,e1)) | e4 != op(e1,e1) | e2 != op(e1,op(e1,e1))),
  inference(ennf_transformation,[],[f74])).
fof(f74,axiom,(
  ~(e3 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) & e0 = op(op(e1,op(e1,e1)),op(e1,e1)) & e4 = op(e1,e1) & e2 = op(e1,op(e1,e1)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1495,plain,(
  ~spl4_176 | ~spl4_94 | ~spl4_152 | ~spl4_180),
  inference(avatar_split_clause,[],[f555,f1492,f1327,f1002,f1463])).
fof(f555,plain,(
  e0 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e4 != op(op(e1,op(e1,e1)),op(e1,e1)) | e3 != op(e1,e1) | e2 != op(e1,op(e1,e1))),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  e0 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e4 != op(op(e1,op(e1,e1)),op(e1,e1)) | e3 != op(e1,e1) | e2 != op(e1,op(e1,e1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,axiom,(
  ~(e0 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) & e4 = op(op(e1,op(e1,e1)),op(e1,e1)) & e3 = op(e1,e1) & e2 = op(e1,op(e1,e1)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1477,plain,(
  ~spl4_175 | ~spl4_32 | ~spl4_141 | ~spl4_167),
  inference(avatar_split_clause,[],[f549,f1410,f1276,f742,f1457])).
fof(f549,plain,(
  e0 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e4 != op(op(e3,op(e3,e3)),op(e3,e3)) | e1 != op(e3,e3) | e2 != op(e3,op(e3,e3))),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  e0 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e4 != op(op(e3,op(e3,e3)),op(e3,e3)) | e1 != op(e3,e3) | e2 != op(e3,op(e3,e3))),
  inference(ennf_transformation,[],[f65])).
fof(f65,axiom,(
  ~(e0 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) & e4 = op(op(e3,op(e3,e3)),op(e3,e3)) & e1 = op(e3,e3) & e2 = op(e3,op(e3,e3)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1476,plain,(
  ~spl4_177 | ~spl4_122 | ~spl4_165 | ~spl4_166),
  inference(avatar_split_clause,[],[f548,f1405,f1401,f1120,f1473])).
fof(f548,plain,(
  e4 != op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) | e3 != op(op(e0,op(e0,e0)),op(e0,e0)) | op(e0,e0) != e1 | e2 != op(e0,op(e0,e0))),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  e4 != op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) | e3 != op(op(e0,op(e0,e0)),op(e0,e0)) | op(e0,e0) != e1 | e2 != op(e0,op(e0,e0))),
  inference(ennf_transformation,[],[f64])).
fof(f64,axiom,(
  ~(e4 = op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) & e3 = op(op(e0,op(e0,e0)),op(e0,e0)) & op(e0,e0) = e1 & e2 = op(e0,op(e0,e0)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1471,plain,(
  ~spl4_174 | ~spl4_2 | ~spl4_136 | ~spl4_163),
  inference(avatar_split_clause,[],[f547,f1392,f1254,f616,f1452])).
fof(f547,plain,(
  e0 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e3 != op(op(e4,op(e4,e4)),op(e4,e4)) | e1 != op(e4,e4) | e2 != op(e4,op(e4,e4))),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  e0 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e3 != op(op(e4,op(e4,e4)),op(e4,e4)) | e1 != op(e4,e4) | e2 != op(e4,op(e4,e4))),
  inference(ennf_transformation,[],[f63])).
fof(f63,axiom,(
  ~(e0 = op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) & e3 = op(op(e4,op(e4,e4)),op(e4,e4)) & e1 = op(e4,e4) & e2 = op(e4,op(e4,e4)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1470,plain,(
  ~spl4_175 | ~spl4_32 | ~spl4_162 | ~spl4_135),
  inference(avatar_split_clause,[],[f546,f1249,f1387,f742,f1457])).
fof(f546,plain,(
  e4 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e0 != op(op(e3,op(e3,e3)),op(e3,e3)) | e1 != op(e3,e3) | e2 != op(e3,op(e3,e3))),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  e4 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e0 != op(op(e3,op(e3,e3)),op(e3,e3)) | e1 != op(e3,e3) | e2 != op(e3,op(e3,e3))),
  inference(ennf_transformation,[],[f62])).
fof(f62,axiom,(
  ~(e4 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) & e0 = op(op(e3,op(e3,e3)),op(e3,e3)) & e1 = op(e3,e3) & e2 = op(e3,op(e3,e3)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1469,plain,(
  ~spl4_174 | ~spl4_2 | ~spl4_161 | ~spl4_132),
  inference(avatar_split_clause,[],[f545,f1236,f1382,f616,f1452])).
fof(f545,plain,(
  e3 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e0 != op(op(e4,op(e4,e4)),op(e4,e4)) | e1 != op(e4,e4) | e2 != op(e4,op(e4,e4))),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  e3 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e0 != op(op(e4,op(e4,e4)),op(e4,e4)) | e1 != op(e4,e4) | e2 != op(e4,op(e4,e4))),
  inference(ennf_transformation,[],[f61])).
fof(f61,axiom,(
  ~(e3 = op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) & e0 = op(op(e4,op(e4,e4)),op(e4,e4)) & e1 = op(e4,e4) & e2 = op(e4,op(e4,e4)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1468,plain,(
  ~spl4_176 | ~spl4_91 | ~spl4_152 | ~spl4_153),
  inference(avatar_split_clause,[],[f544,f1331,f1327,f990,f1463])).
fof(f544,plain,(
  e3 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e4 != op(op(e1,op(e1,e1)),op(e1,e1)) | e0 != op(e1,e1) | e2 != op(e1,op(e1,e1))),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  e3 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e4 != op(op(e1,op(e1,e1)),op(e1,e1)) | e0 != op(e1,e1) | e2 != op(e1,op(e1,e1))),
  inference(ennf_transformation,[],[f60])).
fof(f60,axiom,(
  ~(e3 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) & e4 = op(op(e1,op(e1,e1)),op(e1,e1)) & e0 = op(e1,e1) & e2 = op(e1,op(e1,e1)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1439,plain,(
  ~spl4_160 | ~spl4_64 | ~spl4_143 | ~spl4_172),
  inference(avatar_split_clause,[],[f531,f1436,f1285,f876,f1375])).
fof(f531,plain,(
  e0 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | e4 != op(op(e2,op(e2,e2)),op(e2,e2)) | e3 != op(e2,e2) | e1 != op(e2,op(e2,e2))),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  e0 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | e4 != op(op(e2,op(e2,e2)),op(e2,e2)) | e3 != op(e2,e2) | e1 != op(e2,op(e2,e2))),
  inference(ennf_transformation,[],[f47])).
fof(f47,axiom,(
  ~(e0 = op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) & e4 = op(op(e2,op(e2,e2)),op(e2,e2)) & e3 = op(e2,e2) & e1 = op(e2,op(e2,e2)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1423,plain,(
  ~spl4_158 | ~spl4_4 | ~spl4_161 | ~spl4_137),
  inference(avatar_split_clause,[],[f527,f1258,f1382,f624,f1364])).
fof(f527,plain,(
  e2 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e0 != op(op(e4,op(e4,e4)),op(e4,e4)) | e3 != op(e4,e4) | e1 != op(e4,op(e4,e4))),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  e2 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e0 != op(op(e4,op(e4,e4)),op(e4,e4)) | e3 != op(e4,e4) | e1 != op(e4,op(e4,e4))),
  inference(ennf_transformation,[],[f43])).
fof(f43,axiom,(
  ~(e2 = op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) & e0 = op(op(e4,op(e4,e4)),op(e4,e4)) & e3 = op(e4,e4) & e1 = op(e4,op(e4,e4)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1408,plain,(
  ~spl4_164 | ~spl4_123 | ~spl4_165 | ~spl4_166),
  inference(avatar_split_clause,[],[f524,f1405,f1401,f1124,f1397])).
fof(f524,plain,(
  e4 != op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) | e3 != op(op(e0,op(e0,e0)),op(e0,e0)) | op(e0,e0) != e2 | e1 != op(e0,op(e0,e0))),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  e4 != op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) | e3 != op(op(e0,op(e0,e0)),op(e0,e0)) | op(e0,e0) != e2 | e1 != op(e0,op(e0,e0))),
  inference(ennf_transformation,[],[f40])).
fof(f40,axiom,(
  ~(e4 = op(op(e0,op(e0,e0)),op(e0,op(e0,e0))) & e3 = op(op(e0,op(e0,e0)),op(e0,e0)) & op(e0,e0) = e2 & e1 = op(e0,op(e0,e0)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1395,plain,(
  ~spl4_158 | ~spl4_3 | ~spl4_136 | ~spl4_163),
  inference(avatar_split_clause,[],[f523,f1392,f1254,f620,f1364])).
fof(f523,plain,(
  e0 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e3 != op(op(e4,op(e4,e4)),op(e4,e4)) | e2 != op(e4,e4) | e1 != op(e4,op(e4,e4))),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  e0 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e3 != op(op(e4,op(e4,e4)),op(e4,e4)) | e2 != op(e4,e4) | e1 != op(e4,op(e4,e4))),
  inference(ennf_transformation,[],[f39])).
fof(f39,axiom,(
  ~(e0 = op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) & e3 = op(op(e4,op(e4,e4)),op(e4,e4)) & e2 = op(e4,e4) & e1 = op(e4,op(e4,e4)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1385,plain,(
  ~spl4_158 | ~spl4_3 | ~spl4_161 | ~spl4_132),
  inference(avatar_split_clause,[],[f521,f1236,f1382,f620,f1364])).
fof(f521,plain,(
  e3 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e0 != op(op(e4,op(e4,e4)),op(e4,e4)) | e2 != op(e4,e4) | e1 != op(e4,op(e4,e4))),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  e3 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e0 != op(op(e4,op(e4,e4)),op(e4,e4)) | e2 != op(e4,e4) | e1 != op(e4,op(e4,e4))),
  inference(ennf_transformation,[],[f37])).
fof(f37,axiom,(
  ~(e3 = op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) & e0 = op(op(e4,op(e4,e4)),op(e4,e4)) & e2 = op(e4,e4) & e1 = op(e4,op(e4,e4)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1379,plain,(
  ~spl4_159 | ~spl4_31 | ~spl4_141 | ~spl4_142),
  inference(avatar_split_clause,[],[f519,f1280,f1276,f738,f1369])).
fof(f519,plain,(
  e2 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e4 != op(op(e3,op(e3,e3)),op(e3,e3)) | e0 != op(e3,e3) | e1 != op(e3,op(e3,e3))),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  e2 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e4 != op(op(e3,op(e3,e3)),op(e3,e3)) | e0 != op(e3,e3) | e1 != op(e3,op(e3,e3))),
  inference(ennf_transformation,[],[f35])).
fof(f35,axiom,(
  ~(e2 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) & e4 = op(op(e3,op(e3,e3)),op(e3,e3)) & e0 = op(e3,e3) & e1 = op(e3,op(e3,e3)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1378,plain,(
  ~spl4_160 | ~spl4_61 | ~spl4_139 | ~spl4_140),
  inference(avatar_split_clause,[],[f518,f1271,f1267,f864,f1375])).
fof(f518,plain,(
  e4 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | e3 != op(op(e2,op(e2,e2)),op(e2,e2)) | e0 != op(e2,e2) | e1 != op(e2,op(e2,e2))),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  e4 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | e3 != op(op(e2,op(e2,e2)),op(e2,e2)) | e0 != op(e2,e2) | e1 != op(e2,op(e2,e2))),
  inference(ennf_transformation,[],[f34])).
fof(f34,axiom,(
  ~(e4 = op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) & e3 = op(op(e2,op(e2,e2)),op(e2,e2)) & e0 = op(e2,e2) & e1 = op(e2,op(e2,e2)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1372,plain,(
  ~spl4_159 | ~spl4_31 | ~spl4_134 | ~spl4_135),
  inference(avatar_split_clause,[],[f516,f1249,f1245,f738,f1369])).
fof(f516,plain,(
  e4 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e2 != op(op(e3,op(e3,e3)),op(e3,e3)) | e0 != op(e3,e3) | e1 != op(e3,op(e3,e3))),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  e4 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e2 != op(op(e3,op(e3,e3)),op(e3,e3)) | e0 != op(e3,e3) | e1 != op(e3,op(e3,e3))),
  inference(ennf_transformation,[],[f32])).
fof(f32,axiom,(
  ~(e4 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) & e2 = op(op(e3,op(e3,e3)),op(e3,e3)) & e0 = op(e3,e3) & e1 = op(e3,op(e3,e3)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1357,plain,(
  ~spl4_133 | ~spl4_35 | ~spl4_146 | ~spl4_142),
  inference(avatar_split_clause,[],[f509,f1280,f1299,f754,f1241])).
fof(f509,plain,(
  e2 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e1 != op(op(e3,op(e3,e3)),op(e3,e3)) | e4 != op(e3,e3) | e0 != op(e3,op(e3,e3))),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  e2 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e1 != op(op(e3,op(e3,e3)),op(e3,e3)) | e4 != op(e3,e3) | e0 != op(e3,op(e3,e3))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ~(e2 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) & e1 = op(op(e3,op(e3,e3)),op(e3,e3)) & e4 = op(e3,e3) & e0 = op(e3,op(e3,e3)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1346,plain,(
  ~spl4_148 | ~spl4_94 | ~spl4_155 | ~spl4_150),
  inference(avatar_split_clause,[],[f506,f1317,f1343,f1002,f1309])).
fof(f506,plain,(
  e4 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e2 != op(op(e1,op(e1,e1)),op(e1,e1)) | e3 != op(e1,e1) | e0 != op(e1,op(e1,e1))),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  e4 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e2 != op(op(e1,op(e1,e1)),op(e1,e1)) | e3 != op(e1,e1) | e0 != op(e1,op(e1,e1))),
  inference(ennf_transformation,[],[f22])).
fof(f22,axiom,(
  ~(e4 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) & e2 = op(op(e1,op(e1,e1)),op(e1,e1)) & e3 = op(e1,e1) & e0 = op(e1,op(e1,e1)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1320,plain,(
  ~spl4_148 | ~spl4_93 | ~spl4_149 | ~spl4_150),
  inference(avatar_split_clause,[],[f500,f1317,f1313,f998,f1309])).
fof(f500,plain,(
  e4 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e3 != op(op(e1,op(e1,e1)),op(e1,e1)) | e2 != op(e1,e1) | e0 != op(e1,op(e1,e1))),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  e4 != op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) | e3 != op(op(e1,op(e1,e1)),op(e1,e1)) | e2 != op(e1,e1) | e0 != op(e1,op(e1,e1))),
  inference(ennf_transformation,[],[f16])).
fof(f16,axiom,(
  ~(e4 = op(op(e1,op(e1,e1)),op(e1,op(e1,e1))) & e3 = op(op(e1,op(e1,e1)),op(e1,e1)) & e2 = op(e1,e1) & e0 = op(e1,op(e1,e1)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1302,plain,(
  ~spl4_133 | ~spl4_33 | ~spl4_146 | ~spl4_135),
  inference(avatar_split_clause,[],[f498,f1249,f1299,f746,f1241])).
fof(f498,plain,(
  e4 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e1 != op(op(e3,op(e3,e3)),op(e3,e3)) | e2 != op(e3,e3) | e0 != op(e3,op(e3,e3))),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  e4 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e1 != op(op(e3,op(e3,e3)),op(e3,e3)) | e2 != op(e3,e3) | e0 != op(e3,op(e3,e3))),
  inference(ennf_transformation,[],[f14])).
fof(f14,axiom,(
  ~(e4 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) & e1 = op(op(e3,op(e3,e3)),op(e3,e3)) & e2 = op(e3,e3) & e0 = op(e3,op(e3,e3)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1274,plain,(
  ~spl4_138 | ~spl4_62 | ~spl4_139 | ~spl4_140),
  inference(avatar_split_clause,[],[f494,f1271,f1267,f868,f1263])).
fof(f494,plain,(
  e4 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | e3 != op(op(e2,op(e2,e2)),op(e2,e2)) | e1 != op(e2,e2) | e0 != op(e2,op(e2,e2))),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  e4 != op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) | e3 != op(op(e2,op(e2,e2)),op(e2,e2)) | e1 != op(e2,e2) | e0 != op(e2,op(e2,e2))),
  inference(ennf_transformation,[],[f10])).
fof(f10,axiom,(
  ~(e4 = op(op(e2,op(e2,e2)),op(e2,op(e2,e2))) & e3 = op(op(e2,op(e2,e2)),op(e2,e2)) & e1 = op(e2,e2) & e0 = op(e2,op(e2,e2)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1261,plain,(
  ~spl4_130 | ~spl4_2 | ~spl4_136 | ~spl4_137),
  inference(avatar_split_clause,[],[f493,f1258,f1254,f616,f1228])).
fof(f493,plain,(
  e2 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e3 != op(op(e4,op(e4,e4)),op(e4,e4)) | e1 != op(e4,e4) | e0 != op(e4,op(e4,e4))),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  e2 != op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) | e3 != op(op(e4,op(e4,e4)),op(e4,e4)) | e1 != op(e4,e4) | e0 != op(e4,op(e4,e4))),
  inference(ennf_transformation,[],[f9])).
fof(f9,axiom,(
  ~(e2 = op(op(e4,op(e4,e4)),op(e4,op(e4,e4))) & e3 = op(op(e4,op(e4,e4)),op(e4,e4)) & e1 = op(e4,e4) & e0 = op(e4,op(e4,e4)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1252,plain,(
  ~spl4_133 | ~spl4_32 | ~spl4_134 | ~spl4_135),
  inference(avatar_split_clause,[],[f492,f1249,f1245,f742,f1241])).
fof(f492,plain,(
  e4 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e2 != op(op(e3,op(e3,e3)),op(e3,e3)) | e1 != op(e3,e3) | e0 != op(e3,op(e3,e3))),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  e4 != op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) | e2 != op(op(e3,op(e3,e3)),op(e3,e3)) | e1 != op(e3,e3) | e0 != op(e3,op(e3,e3))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ~(e4 = op(op(e3,op(e3,e3)),op(e3,op(e3,e3))) & e2 = op(op(e3,op(e3,e3)),op(e3,e3)) & e1 = op(e3,e3) & e0 = op(e3,op(e3,e3)))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1226,plain,(
  spl4_129 | spl4_128 | spl4_127 | spl4_126 | ~spl4_101),
  inference(avatar_split_clause,[],[f486,f1032,f1187,f1196,f1205,f1214])).
fof(f1214,plain,(
  spl4_129 <=> sP0),
  introduced(avatar_definition,[new_symbols(naming,[spl4_129])])).
fof(f1205,plain,(
  spl4_128 <=> sP1),
  introduced(avatar_definition,[new_symbols(naming,[spl4_128])])).
fof(f1196,plain,(
  spl4_127 <=> sP2),
  introduced(avatar_definition,[new_symbols(naming,[spl4_127])])).
fof(f1187,plain,(
  spl4_126 <=> sP3),
  introduced(avatar_definition,[new_symbols(naming,[spl4_126])])).
fof(f486,plain,(
  e0 != op(e0,e4) | sP3 | sP2 | sP1 | sP0),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  (e4 != op(e4,e4) & e3 != op(e3,e4) & e2 != op(e2,e4) & e1 != op(e1,e4) & e0 != op(e0,e4)) | sP3 | sP2 | sP1 | sP0),
  inference(definition_folding,[],[f6,f250,f249,f248,f247])).
fof(f247,plain,(
  (e4 != op(e4,e0) & e3 != op(e3,e0) & e2 != op(e2,e0) & e1 != op(e1,e0) & e0 != op(e0,e0)) | ~sP0),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f248,plain,(
  (e4 != op(e4,e1) & e3 != op(e3,e1) & e2 != op(e2,e1) & e1 != op(e1,e1) & e0 != op(e0,e1)) | ~sP1),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f249,plain,(
  (e4 != op(e4,e2) & e3 != op(e3,e2) & e2 != op(e2,e2) & e1 != op(e1,e2) & e0 != op(e0,e2)) | ~sP2),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f250,plain,(
  (e4 != op(e4,e3) & e3 != op(e3,e3) & e2 != op(e2,e3) & e1 != op(e1,e3) & e0 != op(e0,e3)) | ~sP3),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f6,axiom,(
  (e4 != op(e4,e4) & e3 != op(e3,e4) & e2 != op(e2,e4) & e1 != op(e1,e4) & e0 != op(e0,e4)) | (e4 != op(e4,e3) & e3 != op(e3,e3) & e2 != op(e2,e3) & e1 != op(e1,e3) & e0 != op(e0,e3)) | (e4 != op(e4,e2) & e3 != op(e3,e2) & e2 != op(e2,e2) & e1 != op(e1,e2) & e0 != op(e0,e2)) | (e4 != op(e4,e1) & e3 != op(e3,e1) & e2 != op(e2,e1) & e1 != op(e1,e1) & e0 != op(e0,e1)) | (e4 != op(e4,e0) & e3 != op(e3,e0) & e2 != op(e2,e0) & e1 != op(e1,e0) & e0 != op(e0,e0))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1225,plain,(
  spl4_129 | spl4_128 | spl4_127 | spl4_126 | ~spl4_77),
  inference(avatar_split_clause,[],[f487,f931,f1187,f1196,f1205,f1214])).
fof(f487,plain,(
  e1 != op(e1,e4) | sP3 | sP2 | sP1 | sP0),
  inference(cnf_transformation,[],[f251])).
fof(f1222,plain,(
  spl4_129 | spl4_128 | spl4_127 | spl4_126 | ~spl4_5),
  inference(avatar_split_clause,[],[f490,f628,f1187,f1196,f1205,f1214])).
fof(f490,plain,(
  e4 != op(e4,e4) | sP3 | sP2 | sP1 | sP0),
  inference(cnf_transformation,[],[f251])).
fof(f1221,plain,(
  ~spl4_129 | ~spl4_121),
  inference(avatar_split_clause,[],[f481,f1116,f1214])).
fof(f481,plain,(
  e0 != op(e0,e0) | ~sP0),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  (e4 != op(e4,e0) & e3 != op(e3,e0) & e2 != op(e2,e0) & e1 != op(e1,e0) & e0 != op(e0,e0)) | ~sP0),
  inference(nnf_transformation,[],[f247])).
fof(f1220,plain,(
  ~spl4_129 | ~spl4_97),
  inference(avatar_split_clause,[],[f482,f1015,f1214])).
fof(f482,plain,(
  e1 != op(e1,e0) | ~sP0),
  inference(cnf_transformation,[],[f255])).
fof(f1219,plain,(
  ~spl4_129 | ~spl4_73),
  inference(avatar_split_clause,[],[f483,f914,f1214])).
fof(f483,plain,(
  e2 != op(e2,e0) | ~sP0),
  inference(cnf_transformation,[],[f255])).
fof(f1218,plain,(
  ~spl4_129 | ~spl4_49),
  inference(avatar_split_clause,[],[f484,f813,f1214])).
fof(f484,plain,(
  e3 != op(e3,e0) | ~sP0),
  inference(cnf_transformation,[],[f255])).
fof(f1212,plain,(
  ~spl4_128 | ~spl4_116),
  inference(avatar_split_clause,[],[f476,f1095,f1205])).
fof(f476,plain,(
  e0 != op(e0,e1) | ~sP1),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  (e4 != op(e4,e1) & e3 != op(e3,e1) & e2 != op(e2,e1) & e1 != op(e1,e1) & e0 != op(e0,e1)) | ~sP1),
  inference(nnf_transformation,[],[f248])).
fof(f1211,plain,(
  ~spl4_128 | ~spl4_92),
  inference(avatar_split_clause,[],[f477,f994,f1205])).
fof(f477,plain,(
  e1 != op(e1,e1) | ~sP1),
  inference(cnf_transformation,[],[f254])).
fof(f1210,plain,(
  ~spl4_128 | ~spl4_68),
  inference(avatar_split_clause,[],[f478,f893,f1205])).
fof(f478,plain,(
  e2 != op(e2,e1) | ~sP1),
  inference(cnf_transformation,[],[f254])).
fof(f1208,plain,(
  ~spl4_128 | ~spl4_20),
  inference(avatar_split_clause,[],[f480,f691,f1205])).
fof(f480,plain,(
  e4 != op(e4,e1) | ~sP1),
  inference(cnf_transformation,[],[f254])).
fof(f1201,plain,(
  ~spl4_127 | ~spl4_63),
  inference(avatar_split_clause,[],[f473,f872,f1196])).
fof(f473,plain,(
  e2 != op(e2,e2) | ~sP2),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  (e4 != op(e4,e2) & e3 != op(e3,e2) & e2 != op(e2,e2) & e1 != op(e1,e2) & e0 != op(e0,e2)) | ~sP2),
  inference(nnf_transformation,[],[f249])).
fof(f1200,plain,(
  ~spl4_127 | ~spl4_39),
  inference(avatar_split_clause,[],[f474,f771,f1196])).
fof(f474,plain,(
  e3 != op(e3,e2) | ~sP2),
  inference(cnf_transformation,[],[f253])).
fof(f1199,plain,(
  ~spl4_127 | ~spl4_15),
  inference(avatar_split_clause,[],[f475,f670,f1196])).
fof(f475,plain,(
  e4 != op(e4,e2) | ~sP2),
  inference(cnf_transformation,[],[f253])).
fof(f1194,plain,(
  ~spl4_126 | ~spl4_106),
  inference(avatar_split_clause,[],[f466,f1053,f1187])).
fof(f466,plain,(
  e0 != op(e0,e3) | ~sP3),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  (e4 != op(e4,e3) & e3 != op(e3,e3) & e2 != op(e2,e3) & e1 != op(e1,e3) & e0 != op(e0,e3)) | ~sP3),
  inference(nnf_transformation,[],[f250])).
fof(f1193,plain,(
  ~spl4_126 | ~spl4_82),
  inference(avatar_split_clause,[],[f467,f952,f1187])).
fof(f467,plain,(
  e1 != op(e1,e3) | ~sP3),
  inference(cnf_transformation,[],[f252])).
fof(f1191,plain,(
  ~spl4_126 | ~spl4_34),
  inference(avatar_split_clause,[],[f469,f750,f1187])).
fof(f469,plain,(
  e3 != op(e3,e3) | ~sP3),
  inference(cnf_transformation,[],[f252])).
fof(f1185,plain,(
  spl4_121 | spl4_116 | spl4_111 | spl4_106 | spl4_101),
  inference(avatar_split_clause,[],[f281,f1032,f1053,f1074,f1095,f1116])).
fof(f281,plain,(
  e0 = op(e0,e4) | e0 = op(e0,e3) | e0 = op(e0,e2) | e0 = op(e0,e1) | e0 = op(e0,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f2,axiom,(
  (e4 = op(e4,e4) | e4 = op(e3,e4) | e4 = op(e2,e4) | e4 = op(e1,e4) | e4 = op(e0,e4)) & (e4 = op(e4,e4) | e4 = op(e4,e3) | e4 = op(e4,e2) | e4 = op(e4,e1) | e4 = op(e4,e0)) & (e3 = op(e4,e4) | e3 = op(e3,e4) | e3 = op(e2,e4) | e3 = op(e1,e4) | e3 = op(e0,e4)) & (e3 = op(e4,e4) | e3 = op(e4,e3) | e3 = op(e4,e2) | e3 = op(e4,e1) | e3 = op(e4,e0)) & (e2 = op(e4,e4) | e2 = op(e3,e4) | e2 = op(e2,e4) | e2 = op(e1,e4) | e2 = op(e0,e4)) & (e2 = op(e4,e4) | e2 = op(e4,e3) | e2 = op(e4,e2) | e2 = op(e4,e1) | e2 = op(e4,e0)) & (e1 = op(e4,e4) | e1 = op(e3,e4) | e1 = op(e2,e4) | e1 = op(e1,e4) | e1 = op(e0,e4)) & (e1 = op(e4,e4) | e1 = op(e4,e3) | e1 = op(e4,e2) | e1 = op(e4,e1) | e1 = op(e4,e0)) & (e0 = op(e4,e4) | e0 = op(e3,e4) | e0 = op(e2,e4) | e0 = op(e1,e4) | e0 = op(e0,e4)) & (e0 = op(e4,e4) | e0 = op(e4,e3) | e0 = op(e4,e2) | e0 = op(e4,e1) | e0 = op(e4,e0)) & (e4 = op(e4,e3) | e4 = op(e3,e3) | e4 = op(e2,e3) | e4 = op(e1,e3) | e4 = op(e0,e3)) & (e4 = op(e3,e4) | e4 = op(e3,e3) | e4 = op(e3,e2) | e4 = op(e3,e1) | e4 = op(e3,e0)) & (e3 = op(e4,e3) | e3 = op(e3,e3) | e3 = op(e2,e3) | e3 = op(e1,e3) | e3 = op(e0,e3)) & (e3 = op(e3,e4) | e3 = op(e3,e3) | e3 = op(e3,e2) | e3 = op(e3,e1) | e3 = op(e3,e0)) & (e2 = op(e4,e3) | e2 = op(e3,e3) | e2 = op(e2,e3) | e2 = op(e1,e3) | e2 = op(e0,e3)) & (e2 = op(e3,e4) | e2 = op(e3,e3) | e2 = op(e3,e2) | e2 = op(e3,e1) | e2 = op(e3,e0)) & (e1 = op(e4,e3) | e1 = op(e3,e3) | e1 = op(e2,e3) | e1 = op(e1,e3) | e1 = op(e0,e3)) & (e1 = op(e3,e4) | e1 = op(e3,e3) | e1 = op(e3,e2) | e1 = op(e3,e1) | e1 = op(e3,e0)) & (e0 = op(e4,e3) | e0 = op(e3,e3) | e0 = op(e2,e3) | e0 = op(e1,e3) | e0 = op(e0,e3)) & (e0 = op(e3,e4) | e0 = op(e3,e3) | e0 = op(e3,e2) | e0 = op(e3,e1) | e0 = op(e3,e0)) & (e4 = op(e4,e2) | e4 = op(e3,e2) | e4 = op(e2,e2) | e4 = op(e1,e2) | e4 = op(e0,e2)) & (e4 = op(e2,e4) | e4 = op(e2,e3) | e4 = op(e2,e2) | e4 = op(e2,e1) | e4 = op(e2,e0)) & (e3 = op(e4,e2) | e3 = op(e3,e2) | e3 = op(e2,e2) | e3 = op(e1,e2) | e3 = op(e0,e2)) & (e3 = op(e2,e4) | e3 = op(e2,e3) | e3 = op(e2,e2) | e3 = op(e2,e1) | e3 = op(e2,e0)) & (e2 = op(e4,e2) | e2 = op(e3,e2) | e2 = op(e2,e2) | e2 = op(e1,e2) | e2 = op(e0,e2)) & (e2 = op(e2,e4) | e2 = op(e2,e3) | e2 = op(e2,e2) | e2 = op(e2,e1) | e2 = op(e2,e0)) & (e1 = op(e4,e2) | e1 = op(e3,e2) | e1 = op(e2,e2) | e1 = op(e1,e2) | e1 = op(e0,e2)) & (e1 = op(e2,e4) | e1 = op(e2,e3) | e1 = op(e2,e2) | e1 = op(e2,e1) | e1 = op(e2,e0)) & (e0 = op(e4,e2) | e0 = op(e3,e2) | e0 = op(e2,e2) | e0 = op(e1,e2) | e0 = op(e0,e2)) & (e0 = op(e2,e4) | e0 = op(e2,e3) | e0 = op(e2,e2) | e0 = op(e2,e1) | e0 = op(e2,e0)) & (e4 = op(e4,e1) | e4 = op(e3,e1) | e4 = op(e2,e1) | e4 = op(e1,e1) | e4 = op(e0,e1)) & (e4 = op(e1,e4) | e4 = op(e1,e3) | e4 = op(e1,e2) | e4 = op(e1,e1) | e4 = op(e1,e0)) & (e3 = op(e4,e1) | e3 = op(e3,e1) | e3 = op(e2,e1) | e3 = op(e1,e1) | e3 = op(e0,e1)) & (e3 = op(e1,e4) | e3 = op(e1,e3) | e3 = op(e1,e2) | e3 = op(e1,e1) | e3 = op(e1,e0)) & (e2 = op(e4,e1) | e2 = op(e3,e1) | e2 = op(e2,e1) | e2 = op(e1,e1) | e2 = op(e0,e1)) & (e2 = op(e1,e4) | e2 = op(e1,e3) | e2 = op(e1,e2) | e2 = op(e1,e1) | e2 = op(e1,e0)) & (e1 = op(e4,e1) | e1 = op(e3,e1) | e1 = op(e2,e1) | e1 = op(e1,e1) | e1 = op(e0,e1)) & (e1 = op(e1,e4) | e1 = op(e1,e3) | e1 = op(e1,e2) | e1 = op(e1,e1) | e1 = op(e1,e0)) & (e0 = op(e4,e1) | e0 = op(e3,e1) | e0 = op(e2,e1) | e0 = op(e1,e1) | e0 = op(e0,e1)) & (e0 = op(e1,e4) | e0 = op(e1,e3) | e0 = op(e1,e2) | e0 = op(e1,e1) | e0 = op(e1,e0)) & (e4 = op(e4,e0) | e4 = op(e3,e0) | e4 = op(e2,e0) | e4 = op(e1,e0) | op(e0,e0) = e4) & (e4 = op(e0,e4) | e4 = op(e0,e3) | e4 = op(e0,e2) | e4 = op(e0,e1) | op(e0,e0) = e4) & (e3 = op(e4,e0) | e3 = op(e3,e0) | e3 = op(e2,e0) | e3 = op(e1,e0) | op(e0,e0) = e3) & (e3 = op(e0,e4) | e3 = op(e0,e3) | e3 = op(e0,e2) | e3 = op(e0,e1) | op(e0,e0) = e3) & (e2 = op(e4,e0) | e2 = op(e3,e0) | e2 = op(e2,e0) | e2 = op(e1,e0) | op(e0,e0) = e2) & (e2 = op(e0,e4) | e2 = op(e0,e3) | e2 = op(e0,e2) | e2 = op(e0,e1) | op(e0,e0) = e2) & (e1 = op(e4,e0) | e1 = op(e3,e0) | e1 = op(e2,e0) | e1 = op(e1,e0) | op(e0,e0) = e1) & (e1 = op(e0,e4) | e1 = op(e0,e3) | e1 = op(e0,e2) | e1 = op(e0,e1) | op(e0,e0) = e1) & (e0 = op(e4,e0) | e0 = op(e3,e0) | e0 = op(e2,e0) | e0 = op(e1,e0) | e0 = op(e0,e0)) & (e0 = op(e0,e4) | e0 = op(e0,e3) | e0 = op(e0,e2) | e0 = op(e0,e1) | e0 = op(e0,e0))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1184,plain,(
  spl4_121 | spl4_96 | spl4_71 | spl4_46 | spl4_21),
  inference(avatar_split_clause,[],[f282,f696,f801,f906,f1011,f1116])).
fof(f282,plain,(
  e0 = op(e4,e0) | e0 = op(e3,e0) | e0 = op(e2,e0) | e0 = op(e1,e0) | e0 = op(e0,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1183,plain,(
  spl4_122 | spl4_117 | spl4_112 | spl4_107 | spl4_102),
  inference(avatar_split_clause,[],[f283,f1036,f1057,f1078,f1099,f1120])).
fof(f283,plain,(
  e1 = op(e0,e4) | e1 = op(e0,e3) | e1 = op(e0,e2) | e1 = op(e0,e1) | op(e0,e0) = e1),
  inference(cnf_transformation,[],[f2])).
fof(f1182,plain,(
  spl4_122 | spl4_97 | spl4_72 | spl4_47 | spl4_22),
  inference(avatar_split_clause,[],[f284,f700,f805,f910,f1015,f1120])).
fof(f284,plain,(
  e1 = op(e4,e0) | e1 = op(e3,e0) | e1 = op(e2,e0) | e1 = op(e1,e0) | op(e0,e0) = e1),
  inference(cnf_transformation,[],[f2])).
fof(f1181,plain,(
  spl4_123 | spl4_118 | spl4_113 | spl4_108 | spl4_103),
  inference(avatar_split_clause,[],[f285,f1040,f1061,f1082,f1103,f1124])).
fof(f285,plain,(
  e2 = op(e0,e4) | e2 = op(e0,e3) | e2 = op(e0,e2) | e2 = op(e0,e1) | op(e0,e0) = e2),
  inference(cnf_transformation,[],[f2])).
fof(f1180,plain,(
  spl4_123 | spl4_98 | spl4_73 | spl4_48 | spl4_23),
  inference(avatar_split_clause,[],[f286,f704,f809,f914,f1019,f1124])).
fof(f286,plain,(
  e2 = op(e4,e0) | e2 = op(e3,e0) | e2 = op(e2,e0) | e2 = op(e1,e0) | op(e0,e0) = e2),
  inference(cnf_transformation,[],[f2])).
fof(f1179,plain,(
  spl4_124 | spl4_119 | spl4_114 | spl4_109 | spl4_104),
  inference(avatar_split_clause,[],[f287,f1044,f1065,f1086,f1107,f1128])).
fof(f287,plain,(
  e3 = op(e0,e4) | e3 = op(e0,e3) | e3 = op(e0,e2) | e3 = op(e0,e1) | op(e0,e0) = e3),
  inference(cnf_transformation,[],[f2])).
fof(f1178,plain,(
  spl4_124 | spl4_99 | spl4_74 | spl4_49 | spl4_24),
  inference(avatar_split_clause,[],[f288,f708,f813,f918,f1023,f1128])).
fof(f288,plain,(
  e3 = op(e4,e0) | e3 = op(e3,e0) | e3 = op(e2,e0) | e3 = op(e1,e0) | op(e0,e0) = e3),
  inference(cnf_transformation,[],[f2])).
fof(f1177,plain,(
  spl4_125 | spl4_120 | spl4_115 | spl4_110 | spl4_105),
  inference(avatar_split_clause,[],[f289,f1048,f1069,f1090,f1111,f1132])).
fof(f289,plain,(
  e4 = op(e0,e4) | e4 = op(e0,e3) | e4 = op(e0,e2) | e4 = op(e0,e1) | op(e0,e0) = e4),
  inference(cnf_transformation,[],[f2])).
fof(f1176,plain,(
  spl4_125 | spl4_100 | spl4_75 | spl4_50 | spl4_25),
  inference(avatar_split_clause,[],[f290,f712,f817,f922,f1027,f1132])).
fof(f290,plain,(
  e4 = op(e4,e0) | e4 = op(e3,e0) | e4 = op(e2,e0) | e4 = op(e1,e0) | op(e0,e0) = e4),
  inference(cnf_transformation,[],[f2])).
fof(f1175,plain,(
  spl4_96 | spl4_91 | spl4_86 | spl4_81 | spl4_76),
  inference(avatar_split_clause,[],[f291,f927,f948,f969,f990,f1011])).
fof(f291,plain,(
  e0 = op(e1,e4) | e0 = op(e1,e3) | e0 = op(e1,e2) | e0 = op(e1,e1) | e0 = op(e1,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1174,plain,(
  spl4_116 | spl4_91 | spl4_66 | spl4_41 | spl4_16),
  inference(avatar_split_clause,[],[f292,f675,f780,f885,f990,f1095])).
fof(f292,plain,(
  e0 = op(e4,e1) | e0 = op(e3,e1) | e0 = op(e2,e1) | e0 = op(e1,e1) | e0 = op(e0,e1)),
  inference(cnf_transformation,[],[f2])).
fof(f1173,plain,(
  spl4_97 | spl4_92 | spl4_87 | spl4_82 | spl4_77),
  inference(avatar_split_clause,[],[f293,f931,f952,f973,f994,f1015])).
fof(f293,plain,(
  e1 = op(e1,e4) | e1 = op(e1,e3) | e1 = op(e1,e2) | e1 = op(e1,e1) | e1 = op(e1,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1172,plain,(
  spl4_117 | spl4_92 | spl4_67 | spl4_42 | spl4_17),
  inference(avatar_split_clause,[],[f294,f679,f784,f889,f994,f1099])).
fof(f294,plain,(
  e1 = op(e4,e1) | e1 = op(e3,e1) | e1 = op(e2,e1) | e1 = op(e1,e1) | e1 = op(e0,e1)),
  inference(cnf_transformation,[],[f2])).
fof(f1171,plain,(
  spl4_98 | spl4_93 | spl4_88 | spl4_83 | spl4_78),
  inference(avatar_split_clause,[],[f295,f935,f956,f977,f998,f1019])).
fof(f295,plain,(
  e2 = op(e1,e4) | e2 = op(e1,e3) | e2 = op(e1,e2) | e2 = op(e1,e1) | e2 = op(e1,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1170,plain,(
  spl4_118 | spl4_93 | spl4_68 | spl4_43 | spl4_18),
  inference(avatar_split_clause,[],[f296,f683,f788,f893,f998,f1103])).
fof(f296,plain,(
  e2 = op(e4,e1) | e2 = op(e3,e1) | e2 = op(e2,e1) | e2 = op(e1,e1) | e2 = op(e0,e1)),
  inference(cnf_transformation,[],[f2])).
fof(f1169,plain,(
  spl4_99 | spl4_94 | spl4_89 | spl4_84 | spl4_79),
  inference(avatar_split_clause,[],[f297,f939,f960,f981,f1002,f1023])).
fof(f297,plain,(
  e3 = op(e1,e4) | e3 = op(e1,e3) | e3 = op(e1,e2) | e3 = op(e1,e1) | e3 = op(e1,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1168,plain,(
  spl4_119 | spl4_94 | spl4_69 | spl4_44 | spl4_19),
  inference(avatar_split_clause,[],[f298,f687,f792,f897,f1002,f1107])).
fof(f298,plain,(
  e3 = op(e4,e1) | e3 = op(e3,e1) | e3 = op(e2,e1) | e3 = op(e1,e1) | e3 = op(e0,e1)),
  inference(cnf_transformation,[],[f2])).
fof(f1167,plain,(
  spl4_100 | spl4_95 | spl4_90 | spl4_85 | spl4_80),
  inference(avatar_split_clause,[],[f299,f943,f964,f985,f1006,f1027])).
fof(f299,plain,(
  e4 = op(e1,e4) | e4 = op(e1,e3) | e4 = op(e1,e2) | e4 = op(e1,e1) | e4 = op(e1,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1166,plain,(
  spl4_120 | spl4_95 | spl4_70 | spl4_45 | spl4_20),
  inference(avatar_split_clause,[],[f300,f691,f796,f901,f1006,f1111])).
fof(f300,plain,(
  e4 = op(e4,e1) | e4 = op(e3,e1) | e4 = op(e2,e1) | e4 = op(e1,e1) | e4 = op(e0,e1)),
  inference(cnf_transformation,[],[f2])).
fof(f1165,plain,(
  spl4_71 | spl4_66 | spl4_61 | spl4_56 | spl4_51),
  inference(avatar_split_clause,[],[f301,f822,f843,f864,f885,f906])).
fof(f301,plain,(
  e0 = op(e2,e4) | e0 = op(e2,e3) | e0 = op(e2,e2) | e0 = op(e2,e1) | e0 = op(e2,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1164,plain,(
  spl4_111 | spl4_86 | spl4_61 | spl4_36 | spl4_11),
  inference(avatar_split_clause,[],[f302,f654,f759,f864,f969,f1074])).
fof(f302,plain,(
  e0 = op(e4,e2) | e0 = op(e3,e2) | e0 = op(e2,e2) | e0 = op(e1,e2) | e0 = op(e0,e2)),
  inference(cnf_transformation,[],[f2])).
fof(f1163,plain,(
  spl4_72 | spl4_67 | spl4_62 | spl4_57 | spl4_52),
  inference(avatar_split_clause,[],[f303,f826,f847,f868,f889,f910])).
fof(f303,plain,(
  e1 = op(e2,e4) | e1 = op(e2,e3) | e1 = op(e2,e2) | e1 = op(e2,e1) | e1 = op(e2,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1162,plain,(
  spl4_112 | spl4_87 | spl4_62 | spl4_37 | spl4_12),
  inference(avatar_split_clause,[],[f304,f658,f763,f868,f973,f1078])).
fof(f304,plain,(
  e1 = op(e4,e2) | e1 = op(e3,e2) | e1 = op(e2,e2) | e1 = op(e1,e2) | e1 = op(e0,e2)),
  inference(cnf_transformation,[],[f2])).
fof(f1161,plain,(
  spl4_73 | spl4_68 | spl4_63 | spl4_58 | spl4_53),
  inference(avatar_split_clause,[],[f305,f830,f851,f872,f893,f914])).
fof(f305,plain,(
  e2 = op(e2,e4) | e2 = op(e2,e3) | e2 = op(e2,e2) | e2 = op(e2,e1) | e2 = op(e2,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1160,plain,(
  spl4_113 | spl4_88 | spl4_63 | spl4_38 | spl4_13),
  inference(avatar_split_clause,[],[f306,f662,f767,f872,f977,f1082])).
fof(f306,plain,(
  e2 = op(e4,e2) | e2 = op(e3,e2) | e2 = op(e2,e2) | e2 = op(e1,e2) | e2 = op(e0,e2)),
  inference(cnf_transformation,[],[f2])).
fof(f1159,plain,(
  spl4_74 | spl4_69 | spl4_64 | spl4_59 | spl4_54),
  inference(avatar_split_clause,[],[f307,f834,f855,f876,f897,f918])).
fof(f307,plain,(
  e3 = op(e2,e4) | e3 = op(e2,e3) | e3 = op(e2,e2) | e3 = op(e2,e1) | e3 = op(e2,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1158,plain,(
  spl4_114 | spl4_89 | spl4_64 | spl4_39 | spl4_14),
  inference(avatar_split_clause,[],[f308,f666,f771,f876,f981,f1086])).
fof(f308,plain,(
  e3 = op(e4,e2) | e3 = op(e3,e2) | e3 = op(e2,e2) | e3 = op(e1,e2) | e3 = op(e0,e2)),
  inference(cnf_transformation,[],[f2])).
fof(f1157,plain,(
  spl4_75 | spl4_70 | spl4_65 | spl4_60 | spl4_55),
  inference(avatar_split_clause,[],[f309,f838,f859,f880,f901,f922])).
fof(f309,plain,(
  e4 = op(e2,e4) | e4 = op(e2,e3) | e4 = op(e2,e2) | e4 = op(e2,e1) | e4 = op(e2,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1156,plain,(
  spl4_115 | spl4_90 | spl4_65 | spl4_40 | spl4_15),
  inference(avatar_split_clause,[],[f310,f670,f775,f880,f985,f1090])).
fof(f310,plain,(
  e4 = op(e4,e2) | e4 = op(e3,e2) | e4 = op(e2,e2) | e4 = op(e1,e2) | e4 = op(e0,e2)),
  inference(cnf_transformation,[],[f2])).
fof(f1155,plain,(
  spl4_46 | spl4_41 | spl4_36 | spl4_31 | spl4_26),
  inference(avatar_split_clause,[],[f311,f717,f738,f759,f780,f801])).
fof(f311,plain,(
  e0 = op(e3,e4) | e0 = op(e3,e3) | e0 = op(e3,e2) | e0 = op(e3,e1) | e0 = op(e3,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1154,plain,(
  spl4_106 | spl4_81 | spl4_56 | spl4_31 | spl4_6),
  inference(avatar_split_clause,[],[f312,f633,f738,f843,f948,f1053])).
fof(f312,plain,(
  e0 = op(e4,e3) | e0 = op(e3,e3) | e0 = op(e2,e3) | e0 = op(e1,e3) | e0 = op(e0,e3)),
  inference(cnf_transformation,[],[f2])).
fof(f1153,plain,(
  spl4_47 | spl4_42 | spl4_37 | spl4_32 | spl4_27),
  inference(avatar_split_clause,[],[f313,f721,f742,f763,f784,f805])).
fof(f313,plain,(
  e1 = op(e3,e4) | e1 = op(e3,e3) | e1 = op(e3,e2) | e1 = op(e3,e1) | e1 = op(e3,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1152,plain,(
  spl4_107 | spl4_82 | spl4_57 | spl4_32 | spl4_7),
  inference(avatar_split_clause,[],[f314,f637,f742,f847,f952,f1057])).
fof(f314,plain,(
  e1 = op(e4,e3) | e1 = op(e3,e3) | e1 = op(e2,e3) | e1 = op(e1,e3) | e1 = op(e0,e3)),
  inference(cnf_transformation,[],[f2])).
fof(f1151,plain,(
  spl4_48 | spl4_43 | spl4_38 | spl4_33 | spl4_28),
  inference(avatar_split_clause,[],[f315,f725,f746,f767,f788,f809])).
fof(f315,plain,(
  e2 = op(e3,e4) | e2 = op(e3,e3) | e2 = op(e3,e2) | e2 = op(e3,e1) | e2 = op(e3,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1150,plain,(
  spl4_108 | spl4_83 | spl4_58 | spl4_33 | spl4_8),
  inference(avatar_split_clause,[],[f316,f641,f746,f851,f956,f1061])).
fof(f316,plain,(
  e2 = op(e4,e3) | e2 = op(e3,e3) | e2 = op(e2,e3) | e2 = op(e1,e3) | e2 = op(e0,e3)),
  inference(cnf_transformation,[],[f2])).
fof(f1149,plain,(
  spl4_49 | spl4_44 | spl4_39 | spl4_34 | spl4_29),
  inference(avatar_split_clause,[],[f317,f729,f750,f771,f792,f813])).
fof(f317,plain,(
  e3 = op(e3,e4) | e3 = op(e3,e3) | e3 = op(e3,e2) | e3 = op(e3,e1) | e3 = op(e3,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1148,plain,(
  spl4_109 | spl4_84 | spl4_59 | spl4_34 | spl4_9),
  inference(avatar_split_clause,[],[f318,f645,f750,f855,f960,f1065])).
fof(f318,plain,(
  e3 = op(e4,e3) | e3 = op(e3,e3) | e3 = op(e2,e3) | e3 = op(e1,e3) | e3 = op(e0,e3)),
  inference(cnf_transformation,[],[f2])).
fof(f1147,plain,(
  spl4_50 | spl4_45 | spl4_40 | spl4_35 | spl4_30),
  inference(avatar_split_clause,[],[f319,f733,f754,f775,f796,f817])).
fof(f319,plain,(
  e4 = op(e3,e4) | e4 = op(e3,e3) | e4 = op(e3,e2) | e4 = op(e3,e1) | e4 = op(e3,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1146,plain,(
  spl4_110 | spl4_85 | spl4_60 | spl4_35 | spl4_10),
  inference(avatar_split_clause,[],[f320,f649,f754,f859,f964,f1069])).
fof(f320,plain,(
  e4 = op(e4,e3) | e4 = op(e3,e3) | e4 = op(e2,e3) | e4 = op(e1,e3) | e4 = op(e0,e3)),
  inference(cnf_transformation,[],[f2])).
fof(f1145,plain,(
  spl4_21 | spl4_16 | spl4_11 | spl4_6 | spl4_1),
  inference(avatar_split_clause,[],[f321,f612,f633,f654,f675,f696])).
fof(f321,plain,(
  e0 = op(e4,e4) | e0 = op(e4,e3) | e0 = op(e4,e2) | e0 = op(e4,e1) | e0 = op(e4,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1144,plain,(
  spl4_101 | spl4_76 | spl4_51 | spl4_26 | spl4_1),
  inference(avatar_split_clause,[],[f322,f612,f717,f822,f927,f1032])).
fof(f322,plain,(
  e0 = op(e4,e4) | e0 = op(e3,e4) | e0 = op(e2,e4) | e0 = op(e1,e4) | e0 = op(e0,e4)),
  inference(cnf_transformation,[],[f2])).
fof(f1143,plain,(
  spl4_22 | spl4_17 | spl4_12 | spl4_7 | spl4_2),
  inference(avatar_split_clause,[],[f323,f616,f637,f658,f679,f700])).
fof(f323,plain,(
  e1 = op(e4,e4) | e1 = op(e4,e3) | e1 = op(e4,e2) | e1 = op(e4,e1) | e1 = op(e4,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1142,plain,(
  spl4_102 | spl4_77 | spl4_52 | spl4_27 | spl4_2),
  inference(avatar_split_clause,[],[f324,f616,f721,f826,f931,f1036])).
fof(f324,plain,(
  e1 = op(e4,e4) | e1 = op(e3,e4) | e1 = op(e2,e4) | e1 = op(e1,e4) | e1 = op(e0,e4)),
  inference(cnf_transformation,[],[f2])).
fof(f1141,plain,(
  spl4_23 | spl4_18 | spl4_13 | spl4_8 | spl4_3),
  inference(avatar_split_clause,[],[f325,f620,f641,f662,f683,f704])).
fof(f325,plain,(
  e2 = op(e4,e4) | e2 = op(e4,e3) | e2 = op(e4,e2) | e2 = op(e4,e1) | e2 = op(e4,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1140,plain,(
  spl4_103 | spl4_78 | spl4_53 | spl4_28 | spl4_3),
  inference(avatar_split_clause,[],[f326,f620,f725,f830,f935,f1040])).
fof(f326,plain,(
  e2 = op(e4,e4) | e2 = op(e3,e4) | e2 = op(e2,e4) | e2 = op(e1,e4) | e2 = op(e0,e4)),
  inference(cnf_transformation,[],[f2])).
fof(f1139,plain,(
  spl4_24 | spl4_19 | spl4_14 | spl4_9 | spl4_4),
  inference(avatar_split_clause,[],[f327,f624,f645,f666,f687,f708])).
fof(f327,plain,(
  e3 = op(e4,e4) | e3 = op(e4,e3) | e3 = op(e4,e2) | e3 = op(e4,e1) | e3 = op(e4,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1138,plain,(
  spl4_104 | spl4_79 | spl4_54 | spl4_29 | spl4_4),
  inference(avatar_split_clause,[],[f328,f624,f729,f834,f939,f1044])).
fof(f328,plain,(
  e3 = op(e4,e4) | e3 = op(e3,e4) | e3 = op(e2,e4) | e3 = op(e1,e4) | e3 = op(e0,e4)),
  inference(cnf_transformation,[],[f2])).
fof(f1137,plain,(
  spl4_25 | spl4_20 | spl4_15 | spl4_10 | spl4_5),
  inference(avatar_split_clause,[],[f329,f628,f649,f670,f691,f712])).
fof(f329,plain,(
  e4 = op(e4,e4) | e4 = op(e4,e3) | e4 = op(e4,e2) | e4 = op(e4,e1) | e4 = op(e4,e0)),
  inference(cnf_transformation,[],[f2])).
fof(f1136,plain,(
  spl4_105 | spl4_80 | spl4_55 | spl4_30 | spl4_5),
  inference(avatar_split_clause,[],[f330,f628,f733,f838,f943,f1048])).
fof(f330,plain,(
  e4 = op(e4,e4) | e4 = op(e3,e4) | e4 = op(e2,e4) | e4 = op(e1,e4) | e4 = op(e0,e4)),
  inference(cnf_transformation,[],[f2])).
fof(f1135,plain,(
  spl4_121 | spl4_122 | spl4_123 | spl4_124 | spl4_125),
  inference(avatar_split_clause,[],[f256,f1132,f1128,f1124,f1120,f1116])).
fof(f256,plain,(
  op(e0,e0) = e4 | op(e0,e0) = e3 | op(e0,e0) = e2 | op(e0,e0) = e1 | e0 = op(e0,e0)),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  (e4 = op(e4,e4) | e3 = op(e4,e4) | e2 = op(e4,e4) | e1 = op(e4,e4) | e0 = op(e4,e4)) & (e4 = op(e4,e3) | e3 = op(e4,e3) | e2 = op(e4,e3) | e1 = op(e4,e3) | e0 = op(e4,e3)) & (e4 = op(e4,e2) | e3 = op(e4,e2) | e2 = op(e4,e2) | e1 = op(e4,e2) | e0 = op(e4,e2)) & (e4 = op(e4,e1) | e3 = op(e4,e1) | e2 = op(e4,e1) | e1 = op(e4,e1) | e0 = op(e4,e1)) & (e4 = op(e4,e0) | e3 = op(e4,e0) | e2 = op(e4,e0) | e1 = op(e4,e0) | e0 = op(e4,e0)) & (e4 = op(e3,e4) | e3 = op(e3,e4) | e2 = op(e3,e4) | e1 = op(e3,e4) | e0 = op(e3,e4)) & (e4 = op(e3,e3) | e3 = op(e3,e3) | e2 = op(e3,e3) | e1 = op(e3,e3) | e0 = op(e3,e3)) & (e4 = op(e3,e2) | e3 = op(e3,e2) | e2 = op(e3,e2) | e1 = op(e3,e2) | e0 = op(e3,e2)) & (e4 = op(e3,e1) | e3 = op(e3,e1) | e2 = op(e3,e1) | e1 = op(e3,e1) | e0 = op(e3,e1)) & (e4 = op(e3,e0) | e3 = op(e3,e0) | e2 = op(e3,e0) | e1 = op(e3,e0) | e0 = op(e3,e0)) & (e4 = op(e2,e4) | e3 = op(e2,e4) | e2 = op(e2,e4) | e1 = op(e2,e4) | e0 = op(e2,e4)) & (e4 = op(e2,e3) | e3 = op(e2,e3) | e2 = op(e2,e3) | e1 = op(e2,e3) | e0 = op(e2,e3)) & (e4 = op(e2,e2) | e3 = op(e2,e2) | e2 = op(e2,e2) | e1 = op(e2,e2) | e0 = op(e2,e2)) & (e4 = op(e2,e1) | e3 = op(e2,e1) | e2 = op(e2,e1) | e1 = op(e2,e1) | e0 = op(e2,e1)) & (e4 = op(e2,e0) | e3 = op(e2,e0) | e2 = op(e2,e0) | e1 = op(e2,e0) | e0 = op(e2,e0)) & (e4 = op(e1,e4) | e3 = op(e1,e4) | e2 = op(e1,e4) | e1 = op(e1,e4) | e0 = op(e1,e4)) & (e4 = op(e1,e3) | e3 = op(e1,e3) | e2 = op(e1,e3) | e1 = op(e1,e3) | e0 = op(e1,e3)) & (e4 = op(e1,e2) | e3 = op(e1,e2) | e2 = op(e1,e2) | e1 = op(e1,e2) | e0 = op(e1,e2)) & (e4 = op(e1,e1) | e3 = op(e1,e1) | e2 = op(e1,e1) | e1 = op(e1,e1) | e0 = op(e1,e1)) & (e4 = op(e1,e0) | e3 = op(e1,e0) | e2 = op(e1,e0) | e1 = op(e1,e0) | e0 = op(e1,e0)) & (e4 = op(e0,e4) | e3 = op(e0,e4) | e2 = op(e0,e4) | e1 = op(e0,e4) | e0 = op(e0,e4)) & (e4 = op(e0,e3) | e3 = op(e0,e3) | e2 = op(e0,e3) | e1 = op(e0,e3) | e0 = op(e0,e3)) & (e4 = op(e0,e2) | e3 = op(e0,e2) | e2 = op(e0,e2) | e1 = op(e0,e2) | e0 = op(e0,e2)) & (e4 = op(e0,e1) | e3 = op(e0,e1) | e2 = op(e0,e1) | e1 = op(e0,e1) | e0 = op(e0,e1)) & (op(e0,e0) = e4 | op(e0,e0) = e3 | op(e0,e0) = e2 | op(e0,e0) = e1 | e0 = op(e0,e0))),
  file('./p/ALG190+1.tptp',unknown)).
fof(f1114,plain,(
  spl4_116 | spl4_117 | spl4_118 | spl4_119 | spl4_120),
  inference(avatar_split_clause,[],[f257,f1111,f1107,f1103,f1099,f1095])).
fof(f257,plain,(
  e4 = op(e0,e1) | e3 = op(e0,e1) | e2 = op(e0,e1) | e1 = op(e0,e1) | e0 = op(e0,e1)),
  inference(cnf_transformation,[],[f1])).
fof(f1093,plain,(
  spl4_111 | spl4_112 | spl4_113 | spl4_114 | spl4_115),
  inference(avatar_split_clause,[],[f258,f1090,f1086,f1082,f1078,f1074])).
fof(f258,plain,(
  e4 = op(e0,e2) | e3 = op(e0,e2) | e2 = op(e0,e2) | e1 = op(e0,e2) | e0 = op(e0,e2)),
  inference(cnf_transformation,[],[f1])).
fof(f1072,plain,(
  spl4_106 | spl4_107 | spl4_108 | spl4_109 | spl4_110),
  inference(avatar_split_clause,[],[f259,f1069,f1065,f1061,f1057,f1053])).
fof(f259,plain,(
  e4 = op(e0,e3) | e3 = op(e0,e3) | e2 = op(e0,e3) | e1 = op(e0,e3) | e0 = op(e0,e3)),
  inference(cnf_transformation,[],[f1])).
fof(f1051,plain,(
  spl4_101 | spl4_102 | spl4_103 | spl4_104 | spl4_105),
  inference(avatar_split_clause,[],[f260,f1048,f1044,f1040,f1036,f1032])).
fof(f260,plain,(
  e4 = op(e0,e4) | e3 = op(e0,e4) | e2 = op(e0,e4) | e1 = op(e0,e4) | e0 = op(e0,e4)),
  inference(cnf_transformation,[],[f1])).
fof(f1030,plain,(
  spl4_96 | spl4_97 | spl4_98 | spl4_99 | spl4_100),
  inference(avatar_split_clause,[],[f261,f1027,f1023,f1019,f1015,f1011])).
fof(f261,plain,(
  e4 = op(e1,e0) | e3 = op(e1,e0) | e2 = op(e1,e0) | e1 = op(e1,e0) | e0 = op(e1,e0)),
  inference(cnf_transformation,[],[f1])).
fof(f1009,plain,(
  spl4_91 | spl4_92 | spl4_93 | spl4_94 | spl4_95),
  inference(avatar_split_clause,[],[f262,f1006,f1002,f998,f994,f990])).
fof(f262,plain,(
  e4 = op(e1,e1) | e3 = op(e1,e1) | e2 = op(e1,e1) | e1 = op(e1,e1) | e0 = op(e1,e1)),
  inference(cnf_transformation,[],[f1])).
fof(f988,plain,(
  spl4_86 | spl4_87 | spl4_88 | spl4_89 | spl4_90),
  inference(avatar_split_clause,[],[f263,f985,f981,f977,f973,f969])).
fof(f263,plain,(
  e4 = op(e1,e2) | e3 = op(e1,e2) | e2 = op(e1,e2) | e1 = op(e1,e2) | e0 = op(e1,e2)),
  inference(cnf_transformation,[],[f1])).
fof(f967,plain,(
  spl4_81 | spl4_82 | spl4_83 | spl4_84 | spl4_85),
  inference(avatar_split_clause,[],[f264,f964,f960,f956,f952,f948])).
fof(f264,plain,(
  e4 = op(e1,e3) | e3 = op(e1,e3) | e2 = op(e1,e3) | e1 = op(e1,e3) | e0 = op(e1,e3)),
  inference(cnf_transformation,[],[f1])).
fof(f946,plain,(
  spl4_76 | spl4_77 | spl4_78 | spl4_79 | spl4_80),
  inference(avatar_split_clause,[],[f265,f943,f939,f935,f931,f927])).
fof(f265,plain,(
  e4 = op(e1,e4) | e3 = op(e1,e4) | e2 = op(e1,e4) | e1 = op(e1,e4) | e0 = op(e1,e4)),
  inference(cnf_transformation,[],[f1])).
fof(f925,plain,(
  spl4_71 | spl4_72 | spl4_73 | spl4_74 | spl4_75),
  inference(avatar_split_clause,[],[f266,f922,f918,f914,f910,f906])).
fof(f266,plain,(
  e4 = op(e2,e0) | e3 = op(e2,e0) | e2 = op(e2,e0) | e1 = op(e2,e0) | e0 = op(e2,e0)),
  inference(cnf_transformation,[],[f1])).
fof(f904,plain,(
  spl4_66 | spl4_67 | spl4_68 | spl4_69 | spl4_70),
  inference(avatar_split_clause,[],[f267,f901,f897,f893,f889,f885])).
fof(f267,plain,(
  e4 = op(e2,e1) | e3 = op(e2,e1) | e2 = op(e2,e1) | e1 = op(e2,e1) | e0 = op(e2,e1)),
  inference(cnf_transformation,[],[f1])).
fof(f883,plain,(
  spl4_61 | spl4_62 | spl4_63 | spl4_64 | spl4_65),
  inference(avatar_split_clause,[],[f268,f880,f876,f872,f868,f864])).
fof(f268,plain,(
  e4 = op(e2,e2) | e3 = op(e2,e2) | e2 = op(e2,e2) | e1 = op(e2,e2) | e0 = op(e2,e2)),
  inference(cnf_transformation,[],[f1])).
fof(f862,plain,(
  spl4_56 | spl4_57 | spl4_58 | spl4_59 | spl4_60),
  inference(avatar_split_clause,[],[f269,f859,f855,f851,f847,f843])).
fof(f269,plain,(
  e4 = op(e2,e3) | e3 = op(e2,e3) | e2 = op(e2,e3) | e1 = op(e2,e3) | e0 = op(e2,e3)),
  inference(cnf_transformation,[],[f1])).
fof(f841,plain,(
  spl4_51 | spl4_52 | spl4_53 | spl4_54 | spl4_55),
  inference(avatar_split_clause,[],[f270,f838,f834,f830,f826,f822])).
fof(f270,plain,(
  e4 = op(e2,e4) | e3 = op(e2,e4) | e2 = op(e2,e4) | e1 = op(e2,e4) | e0 = op(e2,e4)),
  inference(cnf_transformation,[],[f1])).
fof(f820,plain,(
  spl4_46 | spl4_47 | spl4_48 | spl4_49 | spl4_50),
  inference(avatar_split_clause,[],[f271,f817,f813,f809,f805,f801])).
fof(f271,plain,(
  e4 = op(e3,e0) | e3 = op(e3,e0) | e2 = op(e3,e0) | e1 = op(e3,e0) | e0 = op(e3,e0)),
  inference(cnf_transformation,[],[f1])).
fof(f799,plain,(
  spl4_41 | spl4_42 | spl4_43 | spl4_44 | spl4_45),
  inference(avatar_split_clause,[],[f272,f796,f792,f788,f784,f780])).
fof(f272,plain,(
  e4 = op(e3,e1) | e3 = op(e3,e1) | e2 = op(e3,e1) | e1 = op(e3,e1) | e0 = op(e3,e1)),
  inference(cnf_transformation,[],[f1])).
fof(f778,plain,(
  spl4_36 | spl4_37 | spl4_38 | spl4_39 | spl4_40),
  inference(avatar_split_clause,[],[f273,f775,f771,f767,f763,f759])).
fof(f273,plain,(
  e4 = op(e3,e2) | e3 = op(e3,e2) | e2 = op(e3,e2) | e1 = op(e3,e2) | e0 = op(e3,e2)),
  inference(cnf_transformation,[],[f1])).
fof(f757,plain,(
  spl4_31 | spl4_32 | spl4_33 | spl4_34 | spl4_35),
  inference(avatar_split_clause,[],[f274,f754,f750,f746,f742,f738])).
fof(f274,plain,(
  e4 = op(e3,e3) | e3 = op(e3,e3) | e2 = op(e3,e3) | e1 = op(e3,e3) | e0 = op(e3,e3)),
  inference(cnf_transformation,[],[f1])).
fof(f736,plain,(
  spl4_26 | spl4_27 | spl4_28 | spl4_29 | spl4_30),
  inference(avatar_split_clause,[],[f275,f733,f729,f725,f721,f717])).
fof(f275,plain,(
  e4 = op(e3,e4) | e3 = op(e3,e4) | e2 = op(e3,e4) | e1 = op(e3,e4) | e0 = op(e3,e4)),
  inference(cnf_transformation,[],[f1])).
fof(f715,plain,(
  spl4_21 | spl4_22 | spl4_23 | spl4_24 | spl4_25),
  inference(avatar_split_clause,[],[f276,f712,f708,f704,f700,f696])).
fof(f276,plain,(
  e4 = op(e4,e0) | e3 = op(e4,e0) | e2 = op(e4,e0) | e1 = op(e4,e0) | e0 = op(e4,e0)),
  inference(cnf_transformation,[],[f1])).
fof(f694,plain,(
  spl4_16 | spl4_17 | spl4_18 | spl4_19 | spl4_20),
  inference(avatar_split_clause,[],[f277,f691,f687,f683,f679,f675])).
fof(f277,plain,(
  e4 = op(e4,e1) | e3 = op(e4,e1) | e2 = op(e4,e1) | e1 = op(e4,e1) | e0 = op(e4,e1)),
  inference(cnf_transformation,[],[f1])).
fof(f673,plain,(
  spl4_11 | spl4_12 | spl4_13 | spl4_14 | spl4_15),
  inference(avatar_split_clause,[],[f278,f670,f666,f662,f658,f654])).
fof(f278,plain,(
  e4 = op(e4,e2) | e3 = op(e4,e2) | e2 = op(e4,e2) | e1 = op(e4,e2) | e0 = op(e4,e2)),
  inference(cnf_transformation,[],[f1])).
fof(f652,plain,(
  spl4_6 | spl4_7 | spl4_8 | spl4_9 | spl4_10),
  inference(avatar_split_clause,[],[f279,f649,f645,f641,f637,f633])).
fof(f279,plain,(
  e4 = op(e4,e3) | e3 = op(e4,e3) | e2 = op(e4,e3) | e1 = op(e4,e3) | e0 = op(e4,e3)),
  inference(cnf_transformation,[],[f1])).
fof(f631,plain,(
  spl4_1 | spl4_2 | spl4_3 | spl4_4 | spl4_5),
  inference(avatar_split_clause,[],[f280,f628,f624,f620,f616,f612])).
fof(f280,plain,(
  e4 = op(e4,e4) | e3 = op(e4,e4) | e2 = op(e4,e4) | e1 = op(e4,e4) | e0 = op(e4,e4)),
  inference(cnf_transformation,[],[f1])).
% SZS output end Proof for ALG190+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 9466
% Time elapsed: 0.100 s
% ------------------------------
% ------------------------------
