% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for CAT002-3
% SZS output start Proof for CAT002-3
fof(f10313,plain,(
  $false),
  inference(subsumption_resolution,[],[f10305,f23])).
fof(f23,axiom,(
  h != g),
  file('./p/CAT002-3.tptp',unknown)).
fof(f10305,plain,(
  h = g),
  inference(equality_resolution,[],[f396])).
fof(f396,plain,(
  ( ! [X0] : (compose(b,X0) != compose(b,h) | g = X0) )),
  inference(superposition,[],[f27,f390])).
fof(f390,plain,(
  compose(b,g) = compose(b,h)),
  inference(equality_resolution,[],[f55])).
fof(f55,plain,(
  ( ! [X0] : (compose(a,X0) != compose(a,compose(b,h)) | compose(b,g) = X0) )),
  inference(superposition,[],[f26,f30])).
fof(f30,plain,(
  compose(a,compose(b,g)) = compose(a,compose(b,h))),
  inference(forward_demodulation,[],[f29,f9])).
fof(f9,axiom,(
  ( ! [X2,X0,X1] : (compose(X0,compose(X1,X2)) = compose(compose(X0,X1),X2)) )),
  file('./p/CAT002-3.tptp',unknown)).
fof(f29,plain,(
  compose(compose(a,b),h) = compose(a,compose(b,g))),
  inference(forward_demodulation,[],[f22,f9])).
fof(f22,axiom,(
  compose(compose(a,b),h) = compose(compose(a,b),g)),
  file('./p/CAT002-3.tptp',unknown)).
fof(f26,plain,(
  ( ! [X2,X0] : (compose(a,X0) != compose(a,X2) | X0 = X2) )),
  inference(equality_resolution,[],[f19])).
fof(f19,axiom,(
  ( ! [X2,X0,X1] : (X0 = X2 | compose(a,X2) != X1 | compose(a,X0) != X1) )),
  file('./p/CAT002-3.tptp',unknown)).
fof(f27,plain,(
  ( ! [X2,X0] : (compose(b,X0) != compose(b,X2) | X0 = X2) )),
  inference(equality_resolution,[],[f20])).
fof(f20,axiom,(
  ( ! [X2,X0,X1] : (X0 = X2 | compose(b,X2) != X1 | compose(b,X0) != X1) )),
  file('./p/CAT002-3.tptp',unknown)).
% SZS output end Proof for CAT002-3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 10362
% Time elapsed: 0.135 s
% ------------------------------
% ------------------------------
