% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for COL062-2
% SZS output start Proof for COL062-2
fof(f11,plain,(
  $false),
  inference(subsumption_resolution,[],[f10,f2])).
fof(f2,axiom,(
  ( ! [X0,X1] : (apply(apply(t,X0),X1) = apply(X1,X0)) )),
  file('./p/COL062-2.tptp',unknown)).
fof(f10,plain,(
  apply(apply(x,z),y) != apply(apply(t,y),apply(x,z))),
  inference(superposition,[],[f9,f1])).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (apply(apply(apply(b,X0),X1),X2) = apply(X0,apply(X1,X2))) )),
  file('./p/COL062-2.tptp',unknown)).
fof(f9,plain,(
  apply(apply(x,z),y) != apply(apply(apply(b,apply(t,y)),x),z)),
  inference(forward_demodulation,[],[f8,f1])).
fof(f8,plain,(
  apply(apply(x,z),y) != apply(apply(apply(apply(apply(b,b),t),y),x),z)),
  inference(forward_demodulation,[],[f7,f2])).
fof(f7,plain,(
  apply(apply(x,z),y) != apply(apply(apply(t,x),apply(apply(apply(b,b),t),y)),z)),
  inference(forward_demodulation,[],[f6,f1])).
fof(f6,plain,(
  apply(apply(x,z),y) != apply(apply(apply(apply(b,apply(t,x)),apply(apply(b,b),t)),y),z)),
  inference(forward_demodulation,[],[f5,f1])).
fof(f5,plain,(
  apply(apply(x,z),y) != apply(apply(apply(apply(apply(apply(b,b),t),x),apply(apply(b,b),t)),y),z)),
  inference(forward_demodulation,[],[f4,f2])).
fof(f4,plain,(
  apply(apply(x,z),y) != apply(apply(apply(apply(t,apply(apply(b,b),t)),apply(apply(apply(b,b),t),x)),y),z)),
  inference(forward_demodulation,[],[f3,f1])).
fof(f3,axiom,(
  apply(apply(apply(apply(apply(b,apply(t,apply(apply(b,b),t))),apply(apply(b,b),t)),x),y),z) != apply(apply(x,z),y)),
  file('./p/COL062-2.tptp',unknown)).
% SZS output end Proof for COL062-2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.0000 s
% ------------------------------
% ------------------------------
