% Refutation found. Thanks to Tanya!
% SZS status Theorem for CSR030+3
% SZS output start Proof for CSR030+3
fof(f63174,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f25592,f47468,f63173])).
fof(f63173,plain,(
  ~spl0_200),
  inference(avatar_contradiction_clause,[],[f63172])).
fof(f63172,plain,(
  $false | ~spl0_200),
  inference(subsumption_resolution,[],[f63169,f24407])).
fof(f24407,plain,(
  ( ! [X0] : (~tptptypes_5_387(X0,c_pushingwithopenhand)) )),
  inference(cnf_transformation,[],[f17119])).
fof(f17119,plain,(
  ! [X0] : (~tptptypes_5_387(X0,c_pushingwithopenhand) & mtvisible(c_tptp_member3393_mt))),
  inference(ennf_transformation,[],[f11726])).
fof(f11726,plain,(
  ~? [X0] : (mtvisible(c_tptp_member3393_mt) => tptptypes_5_387(X0,c_pushingwithopenhand))),
  inference(rectify,[],[f8007])).
fof(f8007,negated_conjecture,(
  ~? [X3] : (mtvisible(c_tptp_member3393_mt) => tptptypes_5_387(X3,c_pushingwithopenhand))),
  inference(negated_conjecture,[],[f8006])).
fof(f8006,conjecture,(
  ? [X3] : (mtvisible(c_tptp_member3393_mt) => tptptypes_5_387(X3,c_pushingwithopenhand))),
  file('./p/CSR030+3.tptp',unknown)).
fof(f63169,plain,(
  tptptypes_5_387(c_tptpcol_16_4451,c_pushingwithopenhand) | ~spl0_200),
  inference(resolution,[],[f48361,f18240])).
fof(f18240,plain,(
  ( ! [X0,X1] : (~tptptypes_6_388(X0,X1) | tptptypes_5_387(X0,X1)) )),
  inference(cnf_transformation,[],[f12908])).
fof(f12908,plain,(
  ! [X0,X1] : (tptptypes_5_387(X0,X1) | ~tptptypes_6_388(X0,X1))),
  inference(ennf_transformation,[],[f8067])).
fof(f8067,plain,(
  ! [X0,X1] : (tptptypes_6_388(X0,X1) => tptptypes_5_387(X0,X1))),
  inference(rectify,[],[f1143])).
fof(f1143,axiom,(
  ! [X2,X3] : (tptptypes_6_388(X2,X3) => tptptypes_5_387(X2,X3))),
  file('./p/CSR030+3.tptp',unknown)).
fof(f48361,plain,(
  tptptypes_6_388(c_tptpcol_16_4451,c_pushingwithopenhand) | ~spl0_200),
  inference(resolution,[],[f25316,f18551])).
fof(f18551,plain,(
  ( ! [X0,X1] : (~tptptypes_7_389(X0,X1) | tptptypes_6_388(X1,X0)) )),
  inference(cnf_transformation,[],[f13033])).
fof(f13033,plain,(
  ! [X0,X1] : (tptptypes_6_388(X1,X0) | ~tptptypes_7_389(X0,X1))),
  inference(ennf_transformation,[],[f8079])).
fof(f8079,plain,(
  ! [X0,X1] : (tptptypes_7_389(X0,X1) => tptptypes_6_388(X1,X0))),
  inference(rectify,[],[f1462])).
fof(f1462,axiom,(
  ! [X2,X3] : (tptptypes_7_389(X2,X3) => tptptypes_6_388(X3,X2))),
  file('./p/CSR030+3.tptp',unknown)).
fof(f25316,plain,(
  tptptypes_7_389(c_pushingwithopenhand,c_tptpcol_16_4451) | ~spl0_200),
  inference(avatar_component_clause,[],[f25314])).
fof(f25314,plain,(
  spl0_200 <=> tptptypes_7_389(c_pushingwithopenhand,c_tptpcol_16_4451)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_200])])).
fof(f47468,plain,(
  spl0_25),
  inference(avatar_split_clause,[],[f47467,f24515])).
fof(f24515,plain,(
  spl0_25 <=> mtvisible(c_tptp_spindleheadmt)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_25])])).
fof(f47467,plain,(
  mtvisible(c_tptp_spindleheadmt)),
  inference(subsumption_resolution,[],[f43985,f24406])).
fof(f24406,plain,(
  mtvisible(c_tptp_member3393_mt)),
  inference(cnf_transformation,[],[f17119])).
fof(f43985,plain,(
  mtvisible(c_tptp_spindleheadmt) | ~mtvisible(c_tptp_member3393_mt)),
  inference(resolution,[],[f24397,f19510])).
fof(f19510,plain,(
  genlmt(c_tptp_member3393_mt,c_tptp_spindleheadmt)),
  inference(cnf_transformation,[],[f2434])).
fof(f2434,axiom,(
  genlmt(c_tptp_member3393_mt,c_tptp_spindleheadmt)),
  file('./p/CSR030+3.tptp',unknown)).
fof(f24397,plain,(
  ( ! [X0,X1] : (~genlmt(X0,X1) | mtvisible(X1) | ~mtvisible(X0)) )),
  inference(cnf_transformation,[],[f17110])).
fof(f17110,plain,(
  ! [X0,X1] : (mtvisible(X1) | ~genlmt(X0,X1) | ~mtvisible(X0))),
  inference(flattening,[],[f17109])).
fof(f17109,plain,(
  ! [X0,X1] : (mtvisible(X1) | (~genlmt(X0,X1) | ~mtvisible(X0)))),
  inference(ennf_transformation,[],[f11718])).
fof(f11718,plain,(
  ! [X0,X1] : ((genlmt(X0,X1) & mtvisible(X0)) => mtvisible(X1))),
  inference(rectify,[],[f7997])).
fof(f7997,axiom,(
  ! [X20,X21] : ((genlmt(X20,X21) & mtvisible(X20)) => mtvisible(X21))),
  file('./p/CSR030+3.tptp',unknown)).
fof(f25592,plain,(
  ~spl0_25 | spl0_200),
  inference(avatar_split_clause,[],[f18303,f25314,f24515])).
fof(f18303,plain,(
  tptptypes_7_389(c_pushingwithopenhand,c_tptpcol_16_4451) | ~mtvisible(c_tptp_spindleheadmt)),
  inference(cnf_transformation,[],[f12932])).
fof(f12932,plain,(
  tptptypes_7_389(c_pushingwithopenhand,c_tptpcol_16_4451) | ~mtvisible(c_tptp_spindleheadmt)),
  inference(ennf_transformation,[],[f1207])).
fof(f1207,axiom,(
  mtvisible(c_tptp_spindleheadmt) => tptptypes_7_389(c_pushingwithopenhand,c_tptpcol_16_4451)),
  file('./p/CSR030+3.tptp',unknown)).
% SZS output end Proof for CSR030+3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 34413
% Time elapsed: 0.271 s
% ------------------------------
% ------------------------------
