% Refutation found. Thanks to Tanya!
% SZS status Theorem for CSR113+27
% SZS output start Proof for CSR113+27
fof(f500293,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f20918,f20978,f500291])).
fof(f500291,plain,(
  ~spl64_1),
  inference(avatar_contradiction_clause,[],[f500290])).
fof(f500290,plain,(
  $false | ~spl64_1),
  inference(subsumption_resolution,[],[f500288,f20872])).
fof(f20872,plain,(
  sub(c78538,name_1_1)),
  inference(cnf_transformation,[],[f10346])).
fof(f10346,plain,(
  fact(statue_1_1,real) & etype(statue_1_1,int0) & fact(c78466,real) & etype(c78466,int0) & fact(c79239,real) & etype(c79239,etype_c) & fact(name_1_1,real) & etype(name_1_1,int0) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & fact(c78538,real) & etype(c78538,int0) & fact(c78537,real) & etype(c78537,int0) & fact(freiheitsstatue_1_1,real) & etype(freiheitsstatue_1_1,int0) & fact(c78531,real) & etype(c78531,int0) & fact(wort_2_1,real) & etype(wort_2_1,int0) & fact(c78522,real) & etype(c78522,int1) & fact(c78520,real) & etype(c78520,int1) & fact(freiheit_1_1,real) & etype(freiheit_1_1,int0) & fact(c78495,real) & etype(c78495,int0) & fact(sich_1_1,real) & etype(sich_1_1,int0) & fact(c78489,real) & etype(c78489,int0) & fact(masse_1_1,real) & etype(masse_1_1,int1) & fact(c78483,real) & etype(c78483,int2) & sub(freiheitsstatue_1_1,statue_1_1) & assoc(freiheitsstatue_1_1,freiheit_1_1) & val(c78538,new_york_0) & sub(c78538,name_1_1) & sub(c78537,stadt__1_1) & attr(c78537,c78538) & sub(c78531,freiheitsstatue_1_1) & pred(c78522,wort_2_1) & prop(c78520,ber__374hmt_1_1) & sub(c78495,freiheit_1_1) & sub(c78489,sich_1_1) & pred(c78483,masse_1_1)),
  inference(pure_predicate_removal,[],[f10341])).
fof(f10341,plain,(
  gener(statue_1_1,ge) & fact(statue_1_1,real) & etype(statue_1_1,int0) & gener(c78466,sp) & fact(c78466,real) & etype(c78466,int0) & gener(c79239,gener_c) & fact(c79239,real) & etype(c79239,etype_c) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & gener(c78538,sp) & fact(c78538,real) & etype(c78538,int0) & gener(c78537,sp) & fact(c78537,real) & etype(c78537,int0) & gener(freiheitsstatue_1_1,ge) & fact(freiheitsstatue_1_1,real) & etype(freiheitsstatue_1_1,int0) & gener(c78531,sp) & fact(c78531,real) & etype(c78531,int0) & gener(wort_2_1,ge) & fact(wort_2_1,real) & etype(wort_2_1,int0) & gener(c78522,gener_c) & fact(c78522,real) & etype(c78522,int1) & gener(c78520,sp) & fact(c78520,real) & etype(c78520,int1) & gener(freiheit_1_1,ge) & fact(freiheit_1_1,real) & etype(freiheit_1_1,int0) & gener(c78495,sp) & fact(c78495,real) & etype(c78495,int0) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & gener(c78489,gener_c) & fact(c78489,real) & etype(c78489,int0) & gener(masse_1_1,ge) & fact(masse_1_1,real) & etype(masse_1_1,int1) & gener(c78483,gener_c) & fact(c78483,real) & etype(c78483,int2) & sub(freiheitsstatue_1_1,statue_1_1) & assoc(freiheitsstatue_1_1,freiheit_1_1) & val(c78538,new_york_0) & sub(c78538,name_1_1) & sub(c78537,stadt__1_1) & attr(c78537,c78538) & sub(c78531,freiheitsstatue_1_1) & pred(c78522,wort_2_1) & prop(c78520,ber__374hmt_1_1) & sub(c78495,freiheit_1_1) & sub(c78489,sich_1_1) & pred(c78483,masse_1_1)),
  inference(pure_predicate_removal,[],[f10336])).
fof(f10336,plain,(
  varia(statue_1_1,varia_c) & gener(statue_1_1,ge) & fact(statue_1_1,real) & etype(statue_1_1,int0) & varia(c78466,varia_c) & gener(c78466,sp) & fact(c78466,real) & etype(c78466,int0) & varia(c79239,varia_c) & gener(c79239,gener_c) & fact(c79239,real) & etype(c79239,etype_c) & varia(name_1_1,varia_c) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & varia(stadt__1_1,varia_c) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & varia(c78538,varia_c) & gener(c78538,sp) & fact(c78538,real) & etype(c78538,int0) & varia(c78537,con) & gener(c78537,sp) & fact(c78537,real) & etype(c78537,int0) & varia(freiheitsstatue_1_1,varia_c) & gener(freiheitsstatue_1_1,ge) & fact(freiheitsstatue_1_1,real) & etype(freiheitsstatue_1_1,int0) & varia(c78531,con) & gener(c78531,sp) & fact(c78531,real) & etype(c78531,int0) & varia(wort_2_1,varia_c) & gener(wort_2_1,ge) & fact(wort_2_1,real) & etype(wort_2_1,int0) & varia(c78522,varia_c) & gener(c78522,gener_c) & fact(c78522,real) & etype(c78522,int1) & varia(c78520,con) & gener(c78520,sp) & fact(c78520,real) & etype(c78520,int1) & varia(freiheit_1_1,varia_c) & gener(freiheit_1_1,ge) & fact(freiheit_1_1,real) & etype(freiheit_1_1,int0) & varia(c78495,con) & gener(c78495,sp) & fact(c78495,real) & etype(c78495,int0) & varia(sich_1_1,varia_c) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & varia(c78489,varia_c) & gener(c78489,gener_c) & fact(c78489,real) & etype(c78489,int0) & varia(masse_1_1,varia_c) & gener(masse_1_1,ge) & fact(masse_1_1,real) & etype(masse_1_1,int1) & varia(c78483,varia_c) & gener(c78483,gener_c) & fact(c78483,real) & etype(c78483,int2) & sub(freiheitsstatue_1_1,statue_1_1) & assoc(freiheitsstatue_1_1,freiheit_1_1) & val(c78538,new_york_0) & sub(c78538,name_1_1) & sub(c78537,stadt__1_1) & attr(c78537,c78538) & sub(c78531,freiheitsstatue_1_1) & pred(c78522,wort_2_1) & prop(c78520,ber__374hmt_1_1) & sub(c78495,freiheit_1_1) & sub(c78489,sich_1_1) & pred(c78483,masse_1_1)),
  inference(pure_predicate_removal,[],[f10333])).
fof(f10333,plain,(
  varia(statue_1_1,varia_c) & refer(statue_1_1,refer_c) & gener(statue_1_1,ge) & fact(statue_1_1,real) & etype(statue_1_1,int0) & varia(c78466,varia_c) & refer(c78466,det) & gener(c78466,sp) & fact(c78466,real) & etype(c78466,int0) & varia(c79239,varia_c) & refer(c79239,refer_c) & gener(c79239,gener_c) & fact(c79239,real) & etype(c79239,etype_c) & varia(name_1_1,varia_c) & refer(name_1_1,refer_c) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & varia(stadt__1_1,varia_c) & refer(stadt__1_1,refer_c) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & varia(c78538,varia_c) & refer(c78538,indet) & gener(c78538,sp) & fact(c78538,real) & etype(c78538,int0) & varia(c78537,con) & refer(c78537,det) & gener(c78537,sp) & fact(c78537,real) & etype(c78537,int0) & varia(freiheitsstatue_1_1,varia_c) & refer(freiheitsstatue_1_1,refer_c) & gener(freiheitsstatue_1_1,ge) & fact(freiheitsstatue_1_1,real) & etype(freiheitsstatue_1_1,int0) & varia(c78531,con) & refer(c78531,det) & gener(c78531,sp) & fact(c78531,real) & etype(c78531,int0) & varia(wort_2_1,varia_c) & refer(wort_2_1,refer_c) & gener(wort_2_1,ge) & fact(wort_2_1,real) & etype(wort_2_1,int0) & varia(c78522,varia_c) & refer(c78522,indet) & gener(c78522,gener_c) & fact(c78522,real) & etype(c78522,int1) & varia(c78520,con) & refer(c78520,det) & gener(c78520,sp) & fact(c78520,real) & etype(c78520,int1) & varia(freiheit_1_1,varia_c) & refer(freiheit_1_1,refer_c) & gener(freiheit_1_1,ge) & fact(freiheit_1_1,real) & etype(freiheit_1_1,int0) & varia(c78495,con) & refer(c78495,det) & gener(c78495,sp) & fact(c78495,real) & etype(c78495,int0) & varia(sich_1_1,varia_c) & refer(sich_1_1,refer_c) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & varia(c78489,varia_c) & refer(c78489,refer_c) & gener(c78489,gener_c) & fact(c78489,real) & etype(c78489,int0) & varia(masse_1_1,varia_c) & refer(masse_1_1,refer_c) & gener(masse_1_1,ge) & fact(masse_1_1,real) & etype(masse_1_1,int1) & varia(c78483,varia_c) & refer(c78483,indet) & gener(c78483,gener_c) & fact(c78483,real) & etype(c78483,int2) & sub(freiheitsstatue_1_1,statue_1_1) & assoc(freiheitsstatue_1_1,freiheit_1_1) & val(c78538,new_york_0) & sub(c78538,name_1_1) & sub(c78537,stadt__1_1) & attr(c78537,c78538) & sub(c78531,freiheitsstatue_1_1) & pred(c78522,wort_2_1) & prop(c78520,ber__374hmt_1_1) & sub(c78495,freiheit_1_1) & sub(c78489,sich_1_1) & pred(c78483,masse_1_1)),
  inference(pure_predicate_removal,[],[f10330])).
fof(f10330,plain,(
  varia(statue_1_1,varia_c) & refer(statue_1_1,refer_c) & gener(statue_1_1,ge) & fact(statue_1_1,real) & etype(statue_1_1,int0) & card(statue_1_1,int1) & varia(c78466,varia_c) & refer(c78466,det) & gener(c78466,sp) & fact(c78466,real) & etype(c78466,int0) & card(c78466,int1) & varia(c79239,varia_c) & refer(c79239,refer_c) & gener(c79239,gener_c) & fact(c79239,real) & etype(c79239,etype_c) & card(c79239,card_c) & varia(name_1_1,varia_c) & refer(name_1_1,refer_c) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & card(name_1_1,int1) & varia(stadt__1_1,varia_c) & refer(stadt__1_1,refer_c) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & card(stadt__1_1,int1) & varia(c78538,varia_c) & refer(c78538,indet) & gener(c78538,sp) & fact(c78538,real) & etype(c78538,int0) & card(c78538,int1) & varia(c78537,con) & refer(c78537,det) & gener(c78537,sp) & fact(c78537,real) & etype(c78537,int0) & card(c78537,int1) & varia(freiheitsstatue_1_1,varia_c) & refer(freiheitsstatue_1_1,refer_c) & gener(freiheitsstatue_1_1,ge) & fact(freiheitsstatue_1_1,real) & etype(freiheitsstatue_1_1,int0) & card(freiheitsstatue_1_1,int1) & varia(c78531,con) & refer(c78531,det) & gener(c78531,sp) & fact(c78531,real) & etype(c78531,int0) & card(c78531,int1) & varia(wort_2_1,varia_c) & refer(wort_2_1,refer_c) & gener(wort_2_1,ge) & fact(wort_2_1,real) & etype(wort_2_1,int0) & card(wort_2_1,int1) & varia(c78522,varia_c) & refer(c78522,indet) & gener(c78522,gener_c) & fact(c78522,real) & etype(c78522,int1) & card(c78522,cons(x_constant,cons(int1,nil))) & varia(c78520,con) & refer(c78520,det) & gener(c78520,sp) & fact(c78520,real) & etype(c78520,int1) & card(c78520,cons(x_constant,cons(int1,nil))) & varia(freiheit_1_1,varia_c) & refer(freiheit_1_1,refer_c) & gener(freiheit_1_1,ge) & fact(freiheit_1_1,real) & etype(freiheit_1_1,int0) & card(freiheit_1_1,int1) & varia(c78495,con) & refer(c78495,det) & gener(c78495,sp) & fact(c78495,real) & etype(c78495,int0) & card(c78495,int1) & varia(sich_1_1,varia_c) & refer(sich_1_1,refer_c) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & card(sich_1_1,int1) & varia(c78489,varia_c) & refer(c78489,refer_c) & gener(c78489,gener_c) & fact(c78489,real) & etype(c78489,int0) & card(c78489,int1) & varia(masse_1_1,varia_c) & refer(masse_1_1,refer_c) & gener(masse_1_1,ge) & fact(masse_1_1,real) & etype(masse_1_1,int1) & card(masse_1_1,card_c) & varia(c78483,varia_c) & refer(c78483,indet) & gener(c78483,gener_c) & fact(c78483,real) & etype(c78483,int2) & card(c78483,cons(x_constant,cons(int1,nil))) & sub(freiheitsstatue_1_1,statue_1_1) & assoc(freiheitsstatue_1_1,freiheit_1_1) & val(c78538,new_york_0) & sub(c78538,name_1_1) & sub(c78537,stadt__1_1) & attr(c78537,c78538) & sub(c78531,freiheitsstatue_1_1) & pred(c78522,wort_2_1) & prop(c78520,ber__374hmt_1_1) & sub(c78495,freiheit_1_1) & sub(c78489,sich_1_1) & pred(c78483,masse_1_1)),
  inference(pure_predicate_removal,[],[f10327])).
fof(f10327,plain,(
  varia(statue_1_1,varia_c) & refer(statue_1_1,refer_c) & quant(statue_1_1,one) & gener(statue_1_1,ge) & fact(statue_1_1,real) & etype(statue_1_1,int0) & card(statue_1_1,int1) & varia(c78466,varia_c) & refer(c78466,det) & quant(c78466,one) & gener(c78466,sp) & fact(c78466,real) & etype(c78466,int0) & card(c78466,int1) & varia(c79239,varia_c) & refer(c79239,refer_c) & quant(c79239,quant_c) & gener(c79239,gener_c) & fact(c79239,real) & etype(c79239,etype_c) & card(c79239,card_c) & varia(name_1_1,varia_c) & refer(name_1_1,refer_c) & quant(name_1_1,one) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & card(name_1_1,int1) & varia(stadt__1_1,varia_c) & refer(stadt__1_1,refer_c) & quant(stadt__1_1,one) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & card(stadt__1_1,int1) & varia(c78538,varia_c) & refer(c78538,indet) & quant(c78538,one) & gener(c78538,sp) & fact(c78538,real) & etype(c78538,int0) & card(c78538,int1) & varia(c78537,con) & refer(c78537,det) & quant(c78537,one) & gener(c78537,sp) & fact(c78537,real) & etype(c78537,int0) & card(c78537,int1) & varia(freiheitsstatue_1_1,varia_c) & refer(freiheitsstatue_1_1,refer_c) & quant(freiheitsstatue_1_1,one) & gener(freiheitsstatue_1_1,ge) & fact(freiheitsstatue_1_1,real) & etype(freiheitsstatue_1_1,int0) & card(freiheitsstatue_1_1,int1) & varia(c78531,con) & refer(c78531,det) & quant(c78531,one) & gener(c78531,sp) & fact(c78531,real) & etype(c78531,int0) & card(c78531,int1) & varia(wort_2_1,varia_c) & refer(wort_2_1,refer_c) & quant(wort_2_1,one) & gener(wort_2_1,ge) & fact(wort_2_1,real) & etype(wort_2_1,int0) & card(wort_2_1,int1) & varia(c78522,varia_c) & refer(c78522,indet) & quant(c78522,mult) & gener(c78522,gener_c) & fact(c78522,real) & etype(c78522,int1) & card(c78522,cons(x_constant,cons(int1,nil))) & varia(c78520,con) & refer(c78520,det) & quant(c78520,mult) & gener(c78520,sp) & fact(c78520,real) & etype(c78520,int1) & card(c78520,cons(x_constant,cons(int1,nil))) & varia(freiheit_1_1,varia_c) & refer(freiheit_1_1,refer_c) & quant(freiheit_1_1,one) & gener(freiheit_1_1,ge) & fact(freiheit_1_1,real) & etype(freiheit_1_1,int0) & card(freiheit_1_1,int1) & varia(c78495,con) & refer(c78495,det) & quant(c78495,one) & gener(c78495,sp) & fact(c78495,real) & etype(c78495,int0) & card(c78495,int1) & varia(sich_1_1,varia_c) & refer(sich_1_1,refer_c) & quant(sich_1_1,one) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & card(sich_1_1,int1) & varia(c78489,varia_c) & refer(c78489,refer_c) & quant(c78489,one) & gener(c78489,gener_c) & fact(c78489,real) & etype(c78489,int0) & card(c78489,int1) & varia(masse_1_1,varia_c) & refer(masse_1_1,refer_c) & quant(masse_1_1,quant_c) & gener(masse_1_1,ge) & fact(masse_1_1,real) & etype(masse_1_1,int1) & card(masse_1_1,card_c) & varia(c78483,varia_c) & refer(c78483,indet) & quant(c78483,mult) & gener(c78483,gener_c) & fact(c78483,real) & etype(c78483,int2) & card(c78483,cons(x_constant,cons(int1,nil))) & sub(freiheitsstatue_1_1,statue_1_1) & assoc(freiheitsstatue_1_1,freiheit_1_1) & val(c78538,new_york_0) & sub(c78538,name_1_1) & sub(c78537,stadt__1_1) & attr(c78537,c78538) & sub(c78531,freiheitsstatue_1_1) & pred(c78522,wort_2_1) & prop(c78520,ber__374hmt_1_1) & sub(c78495,freiheit_1_1) & sub(c78489,sich_1_1) & pred(c78483,masse_1_1)),
  inference(pure_predicate_removal,[],[f10191])).
fof(f10191,plain,(
  varia(statue_1_1,varia_c) & refer(statue_1_1,refer_c) & quant(statue_1_1,one) & gener(statue_1_1,ge) & fact(statue_1_1,real) & etype(statue_1_1,int0) & card(statue_1_1,int1) & sort(statue_1_1,d) & varia(c78466,varia_c) & refer(c78466,det) & quant(c78466,one) & gener(c78466,sp) & fact(c78466,real) & etype(c78466,int0) & card(c78466,int1) & sort(c78466,d) & varia(c79239,varia_c) & refer(c79239,refer_c) & quant(c79239,quant_c) & gener(c79239,gener_c) & fact(c79239,real) & etype(c79239,etype_c) & card(c79239,card_c) & sort(c79239,ent) & sort(new_york_0,fe) & varia(name_1_1,varia_c) & refer(name_1_1,refer_c) & quant(name_1_1,one) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & card(name_1_1,int1) & sort(name_1_1,na) & varia(stadt__1_1,varia_c) & refer(stadt__1_1,refer_c) & quant(stadt__1_1,one) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & card(stadt__1_1,int1) & sort(stadt__1_1,io) & sort(stadt__1_1,d) & varia(c78538,varia_c) & refer(c78538,indet) & quant(c78538,one) & gener(c78538,sp) & fact(c78538,real) & etype(c78538,int0) & card(c78538,int1) & sort(c78538,na) & varia(c78537,con) & refer(c78537,det) & quant(c78537,one) & gener(c78537,sp) & fact(c78537,real) & etype(c78537,int0) & card(c78537,int1) & sort(c78537,io) & sort(c78537,d) & varia(freiheitsstatue_1_1,varia_c) & refer(freiheitsstatue_1_1,refer_c) & quant(freiheitsstatue_1_1,one) & gener(freiheitsstatue_1_1,ge) & fact(freiheitsstatue_1_1,real) & etype(freiheitsstatue_1_1,int0) & card(freiheitsstatue_1_1,int1) & sort(freiheitsstatue_1_1,d) & varia(c78531,con) & refer(c78531,det) & quant(c78531,one) & gener(c78531,sp) & fact(c78531,real) & etype(c78531,int0) & card(c78531,int1) & sort(c78531,d) & varia(wort_2_1,varia_c) & refer(wort_2_1,refer_c) & quant(wort_2_1,one) & gener(wort_2_1,ge) & fact(wort_2_1,real) & etype(wort_2_1,int0) & card(wort_2_1,int1) & sort(wort_2_1,io) & sort(wort_2_1,d) & varia(c78522,varia_c) & refer(c78522,indet) & quant(c78522,mult) & gener(c78522,gener_c) & fact(c78522,real) & etype(c78522,int1) & card(c78522,cons(x_constant,cons(int1,nil))) & sort(c78522,io) & sort(c78522,d) & sort(ber__374hmt_1_1,nq) & varia(c78520,con) & refer(c78520,det) & quant(c78520,mult) & gener(c78520,sp) & fact(c78520,real) & etype(c78520,int1) & card(c78520,cons(x_constant,cons(int1,nil))) & sort(c78520,o) & varia(freiheit_1_1,varia_c) & refer(freiheit_1_1,refer_c) & quant(freiheit_1_1,one) & gener(freiheit_1_1,ge) & fact(freiheit_1_1,real) & etype(freiheit_1_1,int0) & card(freiheit_1_1,int1) & sort(freiheit_1_1,io) & sort(freiheit_1_1,as) & varia(c78495,con) & refer(c78495,det) & quant(c78495,one) & gener(c78495,sp) & fact(c78495,real) & etype(c78495,int0) & card(c78495,int1) & sort(c78495,io) & sort(c78495,as) & varia(sich_1_1,varia_c) & refer(sich_1_1,refer_c) & quant(sich_1_1,one) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & card(sich_1_1,int1) & sort(sich_1_1,o) & varia(c78489,varia_c) & refer(c78489,refer_c) & quant(c78489,one) & gener(c78489,gener_c) & fact(c78489,real) & etype(c78489,int0) & card(c78489,int1) & sort(c78489,o) & varia(masse_1_1,varia_c) & refer(masse_1_1,refer_c) & quant(masse_1_1,quant_c) & gener(masse_1_1,ge) & fact(masse_1_1,real) & etype(masse_1_1,int1) & card(masse_1_1,card_c) & sort(masse_1_1,io) & varia(c78483,varia_c) & refer(c78483,indet) & quant(c78483,mult) & gener(c78483,gener_c) & fact(c78483,real) & etype(c78483,int2) & card(c78483,cons(x_constant,cons(int1,nil))) & sort(c78483,io) & sub(freiheitsstatue_1_1,statue_1_1) & assoc(freiheitsstatue_1_1,freiheit_1_1) & val(c78538,new_york_0) & sub(c78538,name_1_1) & sub(c78537,stadt__1_1) & attr(c78537,c78538) & sub(c78531,freiheitsstatue_1_1) & pred(c78522,wort_2_1) & prop(c78520,ber__374hmt_1_1) & sub(c78495,freiheit_1_1) & sub(c78489,sich_1_1) & pred(c78483,masse_1_1)),
  inference(pure_predicate_removal,[],[f10190])).
fof(f10190,axiom,(
  varia(statue_1_1,varia_c) & refer(statue_1_1,refer_c) & quant(statue_1_1,one) & gener(statue_1_1,ge) & fact(statue_1_1,real) & etype(statue_1_1,int0) & card(statue_1_1,int1) & sort(statue_1_1,d) & varia(c78466,varia_c) & refer(c78466,det) & quant(c78466,one) & gener(c78466,sp) & fact(c78466,real) & etype(c78466,int0) & card(c78466,int1) & sort(c78466,d) & varia(c79239,varia_c) & refer(c79239,refer_c) & quant(c79239,quant_c) & gener(c79239,gener_c) & fact(c79239,real) & etype(c79239,etype_c) & card(c79239,card_c) & sort(c79239,ent) & sort(new_york_0,fe) & varia(name_1_1,varia_c) & refer(name_1_1,refer_c) & quant(name_1_1,one) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & card(name_1_1,int1) & sort(name_1_1,na) & varia(stadt__1_1,varia_c) & refer(stadt__1_1,refer_c) & quant(stadt__1_1,one) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & card(stadt__1_1,int1) & sort(stadt__1_1,io) & sort(stadt__1_1,d) & varia(c78538,varia_c) & refer(c78538,indet) & quant(c78538,one) & gener(c78538,sp) & fact(c78538,real) & etype(c78538,int0) & card(c78538,int1) & sort(c78538,na) & varia(c78537,con) & refer(c78537,det) & quant(c78537,one) & gener(c78537,sp) & fact(c78537,real) & etype(c78537,int0) & card(c78537,int1) & sort(c78537,io) & sort(c78537,d) & varia(freiheitsstatue_1_1,varia_c) & refer(freiheitsstatue_1_1,refer_c) & quant(freiheitsstatue_1_1,one) & gener(freiheitsstatue_1_1,ge) & fact(freiheitsstatue_1_1,real) & etype(freiheitsstatue_1_1,int0) & card(freiheitsstatue_1_1,int1) & sort(freiheitsstatue_1_1,d) & varia(c78531,con) & refer(c78531,det) & quant(c78531,one) & gener(c78531,sp) & fact(c78531,real) & etype(c78531,int0) & card(c78531,int1) & sort(c78531,d) & varia(wort_2_1,varia_c) & refer(wort_2_1,refer_c) & quant(wort_2_1,one) & gener(wort_2_1,ge) & fact(wort_2_1,real) & etype(wort_2_1,int0) & card(wort_2_1,int1) & sort(wort_2_1,io) & sort(wort_2_1,d) & varia(c78522,varia_c) & refer(c78522,indet) & quant(c78522,mult) & gener(c78522,gener_c) & fact(c78522,real) & etype(c78522,int1) & card(c78522,cons(x_constant,cons(int1,nil))) & sort(c78522,io) & sort(c78522,d) & sort(ber__374hmt_1_1,nq) & varia(c78520,con) & refer(c78520,det) & quant(c78520,mult) & gener(c78520,sp) & fact(c78520,real) & etype(c78520,int1) & card(c78520,cons(x_constant,cons(int1,nil))) & sort(c78520,o) & varia(freiheit_1_1,varia_c) & refer(freiheit_1_1,refer_c) & quant(freiheit_1_1,one) & gener(freiheit_1_1,ge) & fact(freiheit_1_1,real) & etype(freiheit_1_1,int0) & card(freiheit_1_1,int1) & sort(freiheit_1_1,io) & sort(freiheit_1_1,as) & varia(c78495,con) & refer(c78495,det) & quant(c78495,one) & gener(c78495,sp) & fact(c78495,real) & etype(c78495,int0) & card(c78495,int1) & sort(c78495,io) & sort(c78495,as) & varia(sich_1_1,varia_c) & refer(sich_1_1,refer_c) & quant(sich_1_1,one) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & card(sich_1_1,int1) & sort(sich_1_1,o) & varia(c78489,varia_c) & refer(c78489,refer_c) & quant(c78489,one) & gener(c78489,gener_c) & fact(c78489,real) & etype(c78489,int0) & card(c78489,int1) & sort(c78489,o) & varia(masse_1_1,varia_c) & refer(masse_1_1,refer_c) & quant(masse_1_1,quant_c) & gener(masse_1_1,ge) & fact(masse_1_1,real) & etype(masse_1_1,int1) & card(masse_1_1,card_c) & sort(masse_1_1,io) & varia(c78483,varia_c) & refer(c78483,indet) & quant(c78483,mult) & gener(c78483,gener_c) & fact(c78483,real) & etype(c78483,int2) & card(c78483,cons(x_constant,cons(int1,nil))) & sort(c78483,io) & sub(freiheitsstatue_1_1,statue_1_1) & assoc(freiheitsstatue_1_1,freiheit_1_1) & tupl_p9(c79239,c78466,c78483,c78489,c78495,c78520,c78522,c78531,c78537) & val(c78538,new_york_0) & sub(c78538,name_1_1) & sub(c78537,stadt__1_1) & attr(c78537,c78538) & sub(c78531,freiheitsstatue_1_1) & pred(c78522,wort_2_1) & prop(c78520,ber__374hmt_1_1) & sub(c78495,freiheit_1_1) & sub(c78489,sich_1_1) & pred(c78483,masse_1_1)),
  file('./p/CSR113+27.tptp',unknown)).
fof(f500288,plain,(
  ~sub(c78538,name_1_1) | ~spl64_1),
  inference(resolution,[],[f500286,f20870])).
fof(f20870,plain,(
  attr(c78537,c78538)),
  inference(cnf_transformation,[],[f10346])).
fof(f500286,plain,(
  ( ! [X0,X1] : (~attr(X1,X0) | ~sub(X0,name_1_1)) ) | ~spl64_1),
  inference(resolution,[],[f381231,f10652])).
fof(f10652,plain,(
  ( ! [X0,X1] : (member(X0,cons(X0,X1))) )),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1] : member(X0,cons(X0,X1))),
  file('./p/CSR113+27.tptp',unknown)).
fof(f381231,plain,(
  ( ! [X4,X2,X3] : (~member(X4,cons(name_1_1,nil)) | ~sub(X3,X4) | ~attr(X2,X3)) ) | ~spl64_1),
  inference(resolution,[],[f108834,f10653])).
fof(f10653,plain,(
  ( ! [X2,X0,X1] : (member(X0,cons(X1,X2)) | ~member(X0,X2)) )),
  inference(cnf_transformation,[],[f10349])).
fof(f10349,plain,(
  ! [X0,X1,X2] : (member(X0,cons(X1,X2)) | ~member(X0,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (member(X0,X2) => member(X0,cons(X1,X2)))),
  file('./p/CSR113+27.tptp',unknown)).
fof(f108834,plain,(
  ( ! [X4,X2,X3] : (~member(X3,cons(familiename_1_1,cons(name_1_1,nil))) | ~attr(X4,X2) | ~sub(X2,X3)) ) | ~spl64_1),
  inference(subsumption_resolution,[],[f108832,f20914])).
fof(f20914,plain,(
  ( ! [X2,X3] : (~scar(X2,X3)) ) | ~spl64_1),
  inference(avatar_component_clause,[],[f20913])).
fof(f20913,plain,(
  spl64_1 <=> ! [X3,X2] : ~scar(X2,X3)),
  introduced(avatar_definition,[new_symbols(naming,[spl64_1])])).
fof(f108832,plain,(
  ( ! [X4,X2,X3] : (~sub(X2,X3) | scar(sK52(X4),X4) | ~attr(X4,X2) | ~member(X3,cons(familiename_1_1,cons(name_1_1,nil)))) )),
  inference(resolution,[],[f10917,f10653])).
fof(f10917,plain,(
  ( ! [X2,X0,X1] : (~member(X1,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))) | ~sub(X0,X1) | scar(sK52(X2),X2) | ~attr(X2,X0)) )),
  inference(cnf_transformation,[],[f10633])).
fof(f10633,plain,(
  ! [X0,X1,X2] : ((subs(sK52(X2),stehen_1_b) & scar(sK52(X2),X2) & obj(sK52(X2),X2) & mcont(sK52(X2),X2)) | ~sub(X0,X1) | ~member(X1,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))) | ~attr(X2,X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK52])],[f10517,f10632])).
fof(f10632,plain,(
  ! [X2] : (? [X3] : (subs(X3,stehen_1_b) & scar(X3,X2) & obj(X3,X2) & mcont(X3,X2)) => (subs(sK52(X2),stehen_1_b) & scar(sK52(X2),X2) & obj(sK52(X2),X2) & mcont(sK52(X2),X2)))),
  introduced(choice_axiom,[])).
fof(f10517,plain,(
  ! [X0,X1,X2] : (? [X3] : (subs(X3,stehen_1_b) & scar(X3,X2) & obj(X3,X2) & mcont(X3,X2)) | ~sub(X0,X1) | ~member(X1,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))) | ~attr(X2,X0))),
  inference(flattening,[],[f10516])).
fof(f10516,plain,(
  ! [X0,X1,X2] : (? [X3] : (subs(X3,stehen_1_b) & scar(X3,X2) & obj(X3,X2) & mcont(X3,X2)) | (~sub(X0,X1) | ~member(X1,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))) | ~attr(X2,X0)))),
  inference(ennf_transformation,[],[f160])).
fof(f160,axiom,(
  ! [X0,X1,X2] : ((sub(X0,X1) & member(X1,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))) & attr(X2,X0)) => ? [X3] : (subs(X3,stehen_1_b) & scar(X3,X2) & obj(X3,X2) & mcont(X3,X2)))),
  file('./p/CSR113+27.tptp',unknown)).
fof(f20978,plain,(
  ~spl64_2),
  inference(avatar_contradiction_clause,[],[f20977])).
fof(f20977,plain,(
  $false | ~spl64_2),
  inference(resolution,[],[f20975,f20870])).
fof(f20975,plain,(
  ( ! [X0] : (~attr(X0,c78538)) ) | ~spl64_2),
  inference(subsumption_resolution,[],[f20973,f20872])).
fof(f20973,plain,(
  ( ! [X0] : (~attr(X0,c78538) | ~sub(c78538,name_1_1)) ) | ~spl64_2),
  inference(resolution,[],[f20917,f20873])).
fof(f20873,plain,(
  val(c78538,new_york_0)),
  inference(cnf_transformation,[],[f10346])).
fof(f20917,plain,(
  ( ! [X0,X1] : (~val(X0,new_york_0) | ~attr(X1,X0) | ~sub(X0,name_1_1)) ) | ~spl64_2),
  inference(avatar_component_clause,[],[f20916])).
fof(f20916,plain,(
  spl64_2 <=> ! [X1,X0] : (~val(X0,new_york_0) | ~attr(X1,X0) | ~sub(X0,name_1_1))),
  introduced(avatar_definition,[new_symbols(naming,[spl64_2])])).
fof(f20918,plain,(
  spl64_1 | spl64_2),
  inference(avatar_split_clause,[],[f20863,f20916,f20913])).
fof(f20863,plain,(
  ( ! [X2,X0,X3,X1] : (~val(X0,new_york_0) | ~sub(X0,name_1_1) | ~scar(X2,X3) | ~attr(X1,X0)) )),
  inference(cnf_transformation,[],[f10551])).
fof(f10551,plain,(
  ! [X0,X1,X2,X3] : (~val(X0,new_york_0) | ~sub(X0,name_1_1) | ~scar(X2,X3) | ~attr(X1,X0))),
  inference(ennf_transformation,[],[f10189])).
fof(f10189,negated_conjecture,(
  ~? [X0,X1,X2,X3] : (val(X0,new_york_0) & sub(X0,name_1_1) & scar(X2,X3) & attr(X1,X0))),
  inference(negated_conjecture,[],[f10188])).
fof(f10188,conjecture,(
  ? [X0,X1,X2,X3] : (val(X0,new_york_0) & sub(X0,name_1_1) & scar(X2,X3) & attr(X1,X0))),
  file('./p/CSR113+27.tptp',unknown)).
% SZS output end Proof for CSR113+27
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 317777
% Time elapsed: 4.649 s
% ------------------------------
% ------------------------------
