% Refutation found. Thanks to Tanya!
% SZS status Theorem for CSR114+28
% SZS output start Proof for CSR114+28
fof(f3550891,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f20910,f107383,f3550890])).
fof(f3550890,plain,(
  ~spl64_1),
  inference(avatar_contradiction_clause,[],[f3550889])).
fof(f3550889,plain,(
  $false | ~spl64_1),
  inference(subsumption_resolution,[],[f3550698,f20872])).
fof(f20872,plain,(
  sub(c621,name_1_1)),
  inference(cnf_transformation,[],[f10346])).
fof(f10346,plain,(
  fact(c895,real) & etype(c895,etype_c) & fact(name_1_1,real) & etype(name_1_1,int0) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & fact(c621,real) & etype(c621,int0) & fact(c620,real) & etype(c620,int0) & fact(kolosseum_1_1,real) & etype(kolosseum_1_1,int0) & fact(c614,real) & etype(c614,int0) & fact(sich_1_1,real) & etype(sich_1_1,int0) & fact(c609,real) & etype(c609,int0) & fact(c582,real) & etype(c582,int0) & fact(architektur_1_1,real) & etype(architektur_1_1,int0) & fact(c573,real) & etype(c573,int0) & fact(fassade_1_1,real) & etype(fassade_1_1,int0) & fact(au__337en_1_1,real) & etype(au__337en_1_1,int0) & fact(au__337enfassade_1_1,real) & etype(au__337enfassade_1_1,int0) & val(c621,rom_0) & sub(c621,name_1_1) & sub(c620,stadt__1_1) & attr(c620,c621) & sub(c614,kolosseum_1_1) & sub(c609,sich_1_1) & sub(c582,au__337enfassade_1_1) & sub(c573,architektur_1_1) & sub(au__337enfassade_1_1,fassade_1_1) & assoc(au__337enfassade_1_1,au__337en_1_1)),
  inference(pure_predicate_removal,[],[f10341])).
fof(f10341,plain,(
  gener(c895,gener_c) & fact(c895,real) & etype(c895,etype_c) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & gener(c621,sp) & fact(c621,real) & etype(c621,int0) & gener(c620,sp) & fact(c620,real) & etype(c620,int0) & gener(kolosseum_1_1,sp) & fact(kolosseum_1_1,real) & etype(kolosseum_1_1,int0) & gener(c614,sp) & fact(c614,real) & etype(c614,int0) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & gener(c609,gener_c) & fact(c609,real) & etype(c609,int0) & gener(c582,sp) & fact(c582,real) & etype(c582,int0) & gener(architektur_1_1,ge) & fact(architektur_1_1,real) & etype(architektur_1_1,int0) & gener(c573,sp) & fact(c573,real) & etype(c573,int0) & gener(fassade_1_1,ge) & fact(fassade_1_1,real) & etype(fassade_1_1,int0) & gener(au__337en_1_1,gener_c) & fact(au__337en_1_1,real) & etype(au__337en_1_1,int0) & gener(au__337enfassade_1_1,ge) & fact(au__337enfassade_1_1,real) & etype(au__337enfassade_1_1,int0) & val(c621,rom_0) & sub(c621,name_1_1) & sub(c620,stadt__1_1) & attr(c620,c621) & sub(c614,kolosseum_1_1) & sub(c609,sich_1_1) & sub(c582,au__337enfassade_1_1) & sub(c573,architektur_1_1) & sub(au__337enfassade_1_1,fassade_1_1) & assoc(au__337enfassade_1_1,au__337en_1_1)),
  inference(pure_predicate_removal,[],[f10336])).
fof(f10336,plain,(
  varia(c895,varia_c) & gener(c895,gener_c) & fact(c895,real) & etype(c895,etype_c) & varia(name_1_1,varia_c) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & varia(stadt__1_1,varia_c) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & varia(c621,varia_c) & gener(c621,sp) & fact(c621,real) & etype(c621,int0) & varia(c620,con) & gener(c620,sp) & fact(c620,real) & etype(c620,int0) & varia(kolosseum_1_1,con) & gener(kolosseum_1_1,sp) & fact(kolosseum_1_1,real) & etype(kolosseum_1_1,int0) & varia(c614,con) & gener(c614,sp) & fact(c614,real) & etype(c614,int0) & varia(sich_1_1,varia_c) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & varia(c609,varia_c) & gener(c609,gener_c) & fact(c609,real) & etype(c609,int0) & varia(c582,con) & gener(c582,sp) & fact(c582,real) & etype(c582,int0) & varia(architektur_1_1,varia_c) & gener(architektur_1_1,ge) & fact(architektur_1_1,real) & etype(architektur_1_1,int0) & varia(c573,con) & gener(c573,sp) & fact(c573,real) & etype(c573,int0) & varia(fassade_1_1,varia_c) & gener(fassade_1_1,ge) & fact(fassade_1_1,real) & etype(fassade_1_1,int0) & varia(au__337en_1_1,varia_c) & gener(au__337en_1_1,gener_c) & fact(au__337en_1_1,real) & etype(au__337en_1_1,int0) & varia(au__337enfassade_1_1,varia_c) & gener(au__337enfassade_1_1,ge) & fact(au__337enfassade_1_1,real) & etype(au__337enfassade_1_1,int0) & val(c621,rom_0) & sub(c621,name_1_1) & sub(c620,stadt__1_1) & attr(c620,c621) & sub(c614,kolosseum_1_1) & sub(c609,sich_1_1) & sub(c582,au__337enfassade_1_1) & sub(c573,architektur_1_1) & sub(au__337enfassade_1_1,fassade_1_1) & assoc(au__337enfassade_1_1,au__337en_1_1)),
  inference(pure_predicate_removal,[],[f10333])).
fof(f10333,plain,(
  varia(c895,varia_c) & refer(c895,refer_c) & gener(c895,gener_c) & fact(c895,real) & etype(c895,etype_c) & varia(name_1_1,varia_c) & refer(name_1_1,refer_c) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & varia(stadt__1_1,varia_c) & refer(stadt__1_1,refer_c) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & varia(c621,varia_c) & refer(c621,indet) & gener(c621,sp) & fact(c621,real) & etype(c621,int0) & varia(c620,con) & refer(c620,det) & gener(c620,sp) & fact(c620,real) & etype(c620,int0) & varia(kolosseum_1_1,con) & refer(kolosseum_1_1,det) & gener(kolosseum_1_1,sp) & fact(kolosseum_1_1,real) & etype(kolosseum_1_1,int0) & varia(c614,con) & refer(c614,det) & gener(c614,sp) & fact(c614,real) & etype(c614,int0) & varia(sich_1_1,varia_c) & refer(sich_1_1,refer_c) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & varia(c609,varia_c) & refer(c609,refer_c) & gener(c609,gener_c) & fact(c609,real) & etype(c609,int0) & varia(c582,con) & refer(c582,det) & gener(c582,sp) & fact(c582,real) & etype(c582,int0) & varia(architektur_1_1,varia_c) & refer(architektur_1_1,refer_c) & gener(architektur_1_1,ge) & fact(architektur_1_1,real) & etype(architektur_1_1,int0) & varia(c573,con) & refer(c573,det) & gener(c573,sp) & fact(c573,real) & etype(c573,int0) & varia(fassade_1_1,varia_c) & refer(fassade_1_1,refer_c) & gener(fassade_1_1,ge) & fact(fassade_1_1,real) & etype(fassade_1_1,int0) & varia(au__337en_1_1,varia_c) & refer(au__337en_1_1,refer_c) & gener(au__337en_1_1,gener_c) & fact(au__337en_1_1,real) & etype(au__337en_1_1,int0) & varia(au__337enfassade_1_1,varia_c) & refer(au__337enfassade_1_1,refer_c) & gener(au__337enfassade_1_1,ge) & fact(au__337enfassade_1_1,real) & etype(au__337enfassade_1_1,int0) & val(c621,rom_0) & sub(c621,name_1_1) & sub(c620,stadt__1_1) & attr(c620,c621) & sub(c614,kolosseum_1_1) & sub(c609,sich_1_1) & sub(c582,au__337enfassade_1_1) & sub(c573,architektur_1_1) & sub(au__337enfassade_1_1,fassade_1_1) & assoc(au__337enfassade_1_1,au__337en_1_1)),
  inference(pure_predicate_removal,[],[f10330])).
fof(f10330,plain,(
  varia(c895,varia_c) & refer(c895,refer_c) & gener(c895,gener_c) & fact(c895,real) & etype(c895,etype_c) & card(c895,card_c) & varia(name_1_1,varia_c) & refer(name_1_1,refer_c) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & card(name_1_1,int1) & varia(stadt__1_1,varia_c) & refer(stadt__1_1,refer_c) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & card(stadt__1_1,int1) & varia(c621,varia_c) & refer(c621,indet) & gener(c621,sp) & fact(c621,real) & etype(c621,int0) & card(c621,int1) & varia(c620,con) & refer(c620,det) & gener(c620,sp) & fact(c620,real) & etype(c620,int0) & card(c620,int1) & varia(kolosseum_1_1,con) & refer(kolosseum_1_1,det) & gener(kolosseum_1_1,sp) & fact(kolosseum_1_1,real) & etype(kolosseum_1_1,int0) & card(kolosseum_1_1,int1) & varia(c614,con) & refer(c614,det) & gener(c614,sp) & fact(c614,real) & etype(c614,int0) & card(c614,int1) & varia(sich_1_1,varia_c) & refer(sich_1_1,refer_c) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & card(sich_1_1,int1) & varia(c609,varia_c) & refer(c609,refer_c) & gener(c609,gener_c) & fact(c609,real) & etype(c609,int0) & card(c609,int1) & varia(c582,con) & refer(c582,det) & gener(c582,sp) & fact(c582,real) & etype(c582,int0) & card(c582,int1) & varia(architektur_1_1,varia_c) & refer(architektur_1_1,refer_c) & gener(architektur_1_1,ge) & fact(architektur_1_1,real) & etype(architektur_1_1,int0) & card(architektur_1_1,int1) & varia(c573,con) & refer(c573,det) & gener(c573,sp) & fact(c573,real) & etype(c573,int0) & card(c573,int1) & varia(fassade_1_1,varia_c) & refer(fassade_1_1,refer_c) & gener(fassade_1_1,ge) & fact(fassade_1_1,real) & etype(fassade_1_1,int0) & card(fassade_1_1,int1) & varia(au__337en_1_1,varia_c) & refer(au__337en_1_1,refer_c) & gener(au__337en_1_1,gener_c) & fact(au__337en_1_1,real) & etype(au__337en_1_1,int0) & card(au__337en_1_1,int1) & varia(au__337enfassade_1_1,varia_c) & refer(au__337enfassade_1_1,refer_c) & gener(au__337enfassade_1_1,ge) & fact(au__337enfassade_1_1,real) & etype(au__337enfassade_1_1,int0) & card(au__337enfassade_1_1,int1) & val(c621,rom_0) & sub(c621,name_1_1) & sub(c620,stadt__1_1) & attr(c620,c621) & sub(c614,kolosseum_1_1) & sub(c609,sich_1_1) & sub(c582,au__337enfassade_1_1) & sub(c573,architektur_1_1) & sub(au__337enfassade_1_1,fassade_1_1) & assoc(au__337enfassade_1_1,au__337en_1_1)),
  inference(pure_predicate_removal,[],[f10327])).
fof(f10327,plain,(
  varia(c895,varia_c) & refer(c895,refer_c) & quant(c895,quant_c) & gener(c895,gener_c) & fact(c895,real) & etype(c895,etype_c) & card(c895,card_c) & varia(name_1_1,varia_c) & refer(name_1_1,refer_c) & quant(name_1_1,one) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & card(name_1_1,int1) & varia(stadt__1_1,varia_c) & refer(stadt__1_1,refer_c) & quant(stadt__1_1,one) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & card(stadt__1_1,int1) & varia(c621,varia_c) & refer(c621,indet) & quant(c621,one) & gener(c621,sp) & fact(c621,real) & etype(c621,int0) & card(c621,int1) & varia(c620,con) & refer(c620,det) & quant(c620,one) & gener(c620,sp) & fact(c620,real) & etype(c620,int0) & card(c620,int1) & varia(kolosseum_1_1,con) & refer(kolosseum_1_1,det) & quant(kolosseum_1_1,one) & gener(kolosseum_1_1,sp) & fact(kolosseum_1_1,real) & etype(kolosseum_1_1,int0) & card(kolosseum_1_1,int1) & varia(c614,con) & refer(c614,det) & quant(c614,one) & gener(c614,sp) & fact(c614,real) & etype(c614,int0) & card(c614,int1) & varia(sich_1_1,varia_c) & refer(sich_1_1,refer_c) & quant(sich_1_1,one) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & card(sich_1_1,int1) & varia(c609,varia_c) & refer(c609,refer_c) & quant(c609,one) & gener(c609,gener_c) & fact(c609,real) & etype(c609,int0) & card(c609,int1) & varia(c582,con) & refer(c582,det) & quant(c582,one) & gener(c582,sp) & fact(c582,real) & etype(c582,int0) & card(c582,int1) & varia(architektur_1_1,varia_c) & refer(architektur_1_1,refer_c) & quant(architektur_1_1,one) & gener(architektur_1_1,ge) & fact(architektur_1_1,real) & etype(architektur_1_1,int0) & card(architektur_1_1,int1) & varia(c573,con) & refer(c573,det) & quant(c573,one) & gener(c573,sp) & fact(c573,real) & etype(c573,int0) & card(c573,int1) & varia(fassade_1_1,varia_c) & refer(fassade_1_1,refer_c) & quant(fassade_1_1,one) & gener(fassade_1_1,ge) & fact(fassade_1_1,real) & etype(fassade_1_1,int0) & card(fassade_1_1,int1) & varia(au__337en_1_1,varia_c) & refer(au__337en_1_1,refer_c) & quant(au__337en_1_1,one) & gener(au__337en_1_1,gener_c) & fact(au__337en_1_1,real) & etype(au__337en_1_1,int0) & card(au__337en_1_1,int1) & varia(au__337enfassade_1_1,varia_c) & refer(au__337enfassade_1_1,refer_c) & quant(au__337enfassade_1_1,one) & gener(au__337enfassade_1_1,ge) & fact(au__337enfassade_1_1,real) & etype(au__337enfassade_1_1,int0) & card(au__337enfassade_1_1,int1) & val(c621,rom_0) & sub(c621,name_1_1) & sub(c620,stadt__1_1) & attr(c620,c621) & sub(c614,kolosseum_1_1) & sub(c609,sich_1_1) & sub(c582,au__337enfassade_1_1) & sub(c573,architektur_1_1) & sub(au__337enfassade_1_1,fassade_1_1) & assoc(au__337enfassade_1_1,au__337en_1_1)),
  inference(pure_predicate_removal,[],[f10191])).
fof(f10191,plain,(
  varia(c895,varia_c) & refer(c895,refer_c) & quant(c895,quant_c) & gener(c895,gener_c) & fact(c895,real) & etype(c895,etype_c) & card(c895,card_c) & sort(c895,ent) & sort(rom_0,fe) & varia(name_1_1,varia_c) & refer(name_1_1,refer_c) & quant(name_1_1,one) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & card(name_1_1,int1) & sort(name_1_1,na) & varia(stadt__1_1,varia_c) & refer(stadt__1_1,refer_c) & quant(stadt__1_1,one) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & card(stadt__1_1,int1) & sort(stadt__1_1,io) & sort(stadt__1_1,d) & varia(c621,varia_c) & refer(c621,indet) & quant(c621,one) & gener(c621,sp) & fact(c621,real) & etype(c621,int0) & card(c621,int1) & sort(c621,na) & varia(c620,con) & refer(c620,det) & quant(c620,one) & gener(c620,sp) & fact(c620,real) & etype(c620,int0) & card(c620,int1) & sort(c620,io) & sort(c620,d) & varia(kolosseum_1_1,con) & refer(kolosseum_1_1,det) & quant(kolosseum_1_1,one) & gener(kolosseum_1_1,sp) & fact(kolosseum_1_1,real) & etype(kolosseum_1_1,int0) & card(kolosseum_1_1,int1) & sort(kolosseum_1_1,d) & varia(c614,con) & refer(c614,det) & quant(c614,one) & gener(c614,sp) & fact(c614,real) & etype(c614,int0) & card(c614,int1) & sort(c614,d) & varia(sich_1_1,varia_c) & refer(sich_1_1,refer_c) & quant(sich_1_1,one) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & card(sich_1_1,int1) & sort(sich_1_1,o) & varia(c609,varia_c) & refer(c609,refer_c) & quant(c609,one) & gener(c609,gener_c) & fact(c609,real) & etype(c609,int0) & card(c609,int1) & sort(c609,o) & varia(c582,con) & refer(c582,det) & quant(c582,one) & gener(c582,sp) & fact(c582,real) & etype(c582,int0) & card(c582,int1) & sort(c582,d) & varia(architektur_1_1,varia_c) & refer(architektur_1_1,refer_c) & quant(architektur_1_1,one) & gener(architektur_1_1,ge) & fact(architektur_1_1,real) & etype(architektur_1_1,int0) & card(architektur_1_1,int1) & sort(architektur_1_1,io) & varia(c573,con) & refer(c573,det) & quant(c573,one) & gener(c573,sp) & fact(c573,real) & etype(c573,int0) & card(c573,int1) & sort(c573,io) & varia(fassade_1_1,varia_c) & refer(fassade_1_1,refer_c) & quant(fassade_1_1,one) & gener(fassade_1_1,ge) & fact(fassade_1_1,real) & etype(fassade_1_1,int0) & card(fassade_1_1,int1) & sort(fassade_1_1,d) & varia(au__337en_1_1,varia_c) & refer(au__337en_1_1,refer_c) & quant(au__337en_1_1,one) & gener(au__337en_1_1,gener_c) & fact(au__337en_1_1,real) & etype(au__337en_1_1,int0) & card(au__337en_1_1,int1) & sort(au__337en_1_1,l) & varia(au__337enfassade_1_1,varia_c) & refer(au__337enfassade_1_1,refer_c) & quant(au__337enfassade_1_1,one) & gener(au__337enfassade_1_1,ge) & fact(au__337enfassade_1_1,real) & etype(au__337enfassade_1_1,int0) & card(au__337enfassade_1_1,int1) & sort(au__337enfassade_1_1,d) & val(c621,rom_0) & sub(c621,name_1_1) & sub(c620,stadt__1_1) & attr(c620,c621) & sub(c614,kolosseum_1_1) & sub(c609,sich_1_1) & sub(c582,au__337enfassade_1_1) & sub(c573,architektur_1_1) & sub(au__337enfassade_1_1,fassade_1_1) & assoc(au__337enfassade_1_1,au__337en_1_1)),
  inference(pure_predicate_removal,[],[f10190])).
fof(f10190,axiom,(
  varia(c895,varia_c) & refer(c895,refer_c) & quant(c895,quant_c) & gener(c895,gener_c) & fact(c895,real) & etype(c895,etype_c) & card(c895,card_c) & sort(c895,ent) & sort(rom_0,fe) & varia(name_1_1,varia_c) & refer(name_1_1,refer_c) & quant(name_1_1,one) & gener(name_1_1,ge) & fact(name_1_1,real) & etype(name_1_1,int0) & card(name_1_1,int1) & sort(name_1_1,na) & varia(stadt__1_1,varia_c) & refer(stadt__1_1,refer_c) & quant(stadt__1_1,one) & gener(stadt__1_1,ge) & fact(stadt__1_1,real) & etype(stadt__1_1,int0) & card(stadt__1_1,int1) & sort(stadt__1_1,io) & sort(stadt__1_1,d) & varia(c621,varia_c) & refer(c621,indet) & quant(c621,one) & gener(c621,sp) & fact(c621,real) & etype(c621,int0) & card(c621,int1) & sort(c621,na) & varia(c620,con) & refer(c620,det) & quant(c620,one) & gener(c620,sp) & fact(c620,real) & etype(c620,int0) & card(c620,int1) & sort(c620,io) & sort(c620,d) & varia(kolosseum_1_1,con) & refer(kolosseum_1_1,det) & quant(kolosseum_1_1,one) & gener(kolosseum_1_1,sp) & fact(kolosseum_1_1,real) & etype(kolosseum_1_1,int0) & card(kolosseum_1_1,int1) & sort(kolosseum_1_1,d) & varia(c614,con) & refer(c614,det) & quant(c614,one) & gener(c614,sp) & fact(c614,real) & etype(c614,int0) & card(c614,int1) & sort(c614,d) & varia(sich_1_1,varia_c) & refer(sich_1_1,refer_c) & quant(sich_1_1,one) & gener(sich_1_1,gener_c) & fact(sich_1_1,real) & etype(sich_1_1,int0) & card(sich_1_1,int1) & sort(sich_1_1,o) & varia(c609,varia_c) & refer(c609,refer_c) & quant(c609,one) & gener(c609,gener_c) & fact(c609,real) & etype(c609,int0) & card(c609,int1) & sort(c609,o) & varia(c582,con) & refer(c582,det) & quant(c582,one) & gener(c582,sp) & fact(c582,real) & etype(c582,int0) & card(c582,int1) & sort(c582,d) & varia(architektur_1_1,varia_c) & refer(architektur_1_1,refer_c) & quant(architektur_1_1,one) & gener(architektur_1_1,ge) & fact(architektur_1_1,real) & etype(architektur_1_1,int0) & card(architektur_1_1,int1) & sort(architektur_1_1,io) & varia(c573,con) & refer(c573,det) & quant(c573,one) & gener(c573,sp) & fact(c573,real) & etype(c573,int0) & card(c573,int1) & sort(c573,io) & varia(fassade_1_1,varia_c) & refer(fassade_1_1,refer_c) & quant(fassade_1_1,one) & gener(fassade_1_1,ge) & fact(fassade_1_1,real) & etype(fassade_1_1,int0) & card(fassade_1_1,int1) & sort(fassade_1_1,d) & varia(au__337en_1_1,varia_c) & refer(au__337en_1_1,refer_c) & quant(au__337en_1_1,one) & gener(au__337en_1_1,gener_c) & fact(au__337en_1_1,real) & etype(au__337en_1_1,int0) & card(au__337en_1_1,int1) & sort(au__337en_1_1,l) & varia(au__337enfassade_1_1,varia_c) & refer(au__337enfassade_1_1,refer_c) & quant(au__337enfassade_1_1,one) & gener(au__337enfassade_1_1,ge) & fact(au__337enfassade_1_1,real) & etype(au__337enfassade_1_1,int0) & card(au__337enfassade_1_1,int1) & sort(au__337enfassade_1_1,d) & tupl_p6(c895,c573,c582,c609,c614,c620) & val(c621,rom_0) & sub(c621,name_1_1) & sub(c620,stadt__1_1) & attr(c620,c621) & sub(c614,kolosseum_1_1) & sub(c609,sich_1_1) & sub(c582,au__337enfassade_1_1) & sub(c573,architektur_1_1) & sub(au__337enfassade_1_1,fassade_1_1) & assoc(au__337enfassade_1_1,au__337en_1_1)),
  file('./p/CSR114+28.tptp',unknown)).
fof(f3550698,plain,(
  ~sub(c621,name_1_1) | ~spl64_1),
  inference(resolution,[],[f3550691,f20870])).
fof(f20870,plain,(
  attr(c620,c621)),
  inference(cnf_transformation,[],[f10346])).
fof(f3550691,plain,(
  ( ! [X0,X1] : (~attr(X1,X0) | ~sub(X0,name_1_1)) ) | ~spl64_1),
  inference(resolution,[],[f3550491,f10652])).
fof(f10652,plain,(
  ( ! [X0,X1] : (member(X0,cons(X0,X1))) )),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1] : member(X0,cons(X0,X1))),
  file('./p/CSR114+28.tptp',unknown)).
fof(f3550491,plain,(
  ( ! [X4,X2,X3] : (~member(X4,cons(name_1_1,nil)) | ~sub(X3,X4) | ~attr(X2,X3)) ) | ~spl64_1),
  inference(resolution,[],[f108571,f10653])).
fof(f10653,plain,(
  ( ! [X2,X0,X1] : (member(X0,cons(X1,X2)) | ~member(X0,X2)) )),
  inference(cnf_transformation,[],[f10349])).
fof(f10349,plain,(
  ! [X0,X1,X2] : (member(X0,cons(X1,X2)) | ~member(X0,X2))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2] : (member(X0,X2) => member(X0,cons(X1,X2)))),
  file('./p/CSR114+28.tptp',unknown)).
fof(f108571,plain,(
  ( ! [X4,X2,X3] : (~member(X3,cons(familiename_1_1,cons(name_1_1,nil))) | ~attr(X4,X2) | ~sub(X2,X3)) ) | ~spl64_1),
  inference(subsumption_resolution,[],[f108569,f20906])).
fof(f20906,plain,(
  ( ! [X2,X3] : (~scar(X3,X2)) ) | ~spl64_1),
  inference(avatar_component_clause,[],[f20905])).
fof(f20905,plain,(
  spl64_1 <=> ! [X3,X2] : ~scar(X3,X2)),
  introduced(avatar_definition,[new_symbols(naming,[spl64_1])])).
fof(f108569,plain,(
  ( ! [X4,X2,X3] : (~sub(X2,X3) | scar(sK52(X4),X4) | ~attr(X4,X2) | ~member(X3,cons(familiename_1_1,cons(name_1_1,nil)))) )),
  inference(resolution,[],[f10917,f10653])).
fof(f10917,plain,(
  ( ! [X2,X0,X1] : (~member(X1,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))) | ~sub(X0,X1) | scar(sK52(X2),X2) | ~attr(X2,X0)) )),
  inference(cnf_transformation,[],[f10633])).
fof(f10633,plain,(
  ! [X0,X1,X2] : ((subs(sK52(X2),stehen_1_b) & scar(sK52(X2),X2) & obj(sK52(X2),X2) & mcont(sK52(X2),X2)) | ~sub(X0,X1) | ~member(X1,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))) | ~attr(X2,X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK52])],[f10517,f10632])).
fof(f10632,plain,(
  ! [X2] : (? [X3] : (subs(X3,stehen_1_b) & scar(X3,X2) & obj(X3,X2) & mcont(X3,X2)) => (subs(sK52(X2),stehen_1_b) & scar(sK52(X2),X2) & obj(sK52(X2),X2) & mcont(sK52(X2),X2)))),
  introduced(choice_axiom,[])).
fof(f10517,plain,(
  ! [X0,X1,X2] : (? [X3] : (subs(X3,stehen_1_b) & scar(X3,X2) & obj(X3,X2) & mcont(X3,X2)) | ~sub(X0,X1) | ~member(X1,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))) | ~attr(X2,X0))),
  inference(flattening,[],[f10516])).
fof(f10516,plain,(
  ! [X0,X1,X2] : (? [X3] : (subs(X3,stehen_1_b) & scar(X3,X2) & obj(X3,X2) & mcont(X3,X2)) | (~sub(X0,X1) | ~member(X1,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))) | ~attr(X2,X0)))),
  inference(ennf_transformation,[],[f160])).
fof(f160,axiom,(
  ! [X0,X1,X2] : ((sub(X0,X1) & member(X1,cons(eigenname_1_1,cons(familiename_1_1,cons(name_1_1,nil)))) & attr(X2,X0)) => ? [X3] : (subs(X3,stehen_1_b) & scar(X3,X2) & obj(X3,X2) & mcont(X3,X2)))),
  file('./p/CSR114+28.tptp',unknown)).
fof(f107383,plain,(
  ~spl64_2),
  inference(avatar_contradiction_clause,[],[f107382])).
fof(f107382,plain,(
  $false | ~spl64_2),
  inference(subsumption_resolution,[],[f107381,f20871])).
fof(f20871,plain,(
  sub(c620,stadt__1_1)),
  inference(cnf_transformation,[],[f10346])).
fof(f107381,plain,(
  ~sub(c620,stadt__1_1) | ~spl64_2),
  inference(resolution,[],[f107375,f20870])).
fof(f107375,plain,(
  ( ! [X0] : (~attr(X0,c621) | ~sub(X0,stadt__1_1)) ) | ~spl64_2),
  inference(subsumption_resolution,[],[f107373,f20872])).
fof(f107373,plain,(
  ( ! [X0] : (~sub(c621,name_1_1) | ~sub(X0,stadt__1_1) | ~attr(X0,c621)) ) | ~spl64_2),
  inference(resolution,[],[f20909,f20873])).
fof(f20873,plain,(
  val(c621,rom_0)),
  inference(cnf_transformation,[],[f10346])).
fof(f20909,plain,(
  ( ! [X0,X1] : (~val(X1,rom_0) | ~sub(X1,name_1_1) | ~sub(X0,stadt__1_1) | ~attr(X0,X1)) ) | ~spl64_2),
  inference(avatar_component_clause,[],[f20908])).
fof(f20908,plain,(
  spl64_2 <=> ! [X1,X0] : (~val(X1,rom_0) | ~sub(X1,name_1_1) | ~sub(X0,stadt__1_1) | ~attr(X0,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl64_2])])).
fof(f20910,plain,(
  spl64_1 | spl64_2),
  inference(avatar_split_clause,[],[f20863,f20908,f20905])).
fof(f20863,plain,(
  ( ! [X2,X0,X3,X1] : (~val(X1,rom_0) | ~sub(X0,stadt__1_1) | ~sub(X1,name_1_1) | ~scar(X3,X2) | ~attr(X0,X1)) )),
  inference(cnf_transformation,[],[f10551])).
fof(f10551,plain,(
  ! [X0,X1,X2,X3] : (~val(X1,rom_0) | ~sub(X0,stadt__1_1) | ~sub(X1,name_1_1) | ~scar(X3,X2) | ~attr(X0,X1))),
  inference(ennf_transformation,[],[f10189])).
fof(f10189,negated_conjecture,(
  ~? [X0,X1,X2,X3] : (val(X1,rom_0) & sub(X0,stadt__1_1) & sub(X1,name_1_1) & scar(X3,X2) & attr(X0,X1))),
  inference(negated_conjecture,[],[f10188])).
fof(f10188,conjecture,(
  ? [X0,X1,X2,X3] : (val(X1,rom_0) & sub(X0,stadt__1_1) & sub(X1,name_1_1) & scar(X3,X2) & attr(X0,X1))),
  file('./p/CSR114+28.tptp',unknown)).
% SZS output end Proof for CSR114+28
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 1765813
% Time elapsed: 11.569 s
% ------------------------------
% ------------------------------
