% Refutation found. Thanks to Tanya!
% SZS status Theorem for GEO187+1
% SZS output start Proof for GEO187+1
fof(f2064,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f68,f235,f258,f298,f1830,f1855,f1998,f2016,f2063])).
fof(f2063,plain,(
  ~spl4_26),
  inference(avatar_contradiction_clause,[],[f2062])).
fof(f2062,plain,(
  $false | ~spl4_26),
  inference(subsumption_resolution,[],[f2059,f56])).
fof(f56,plain,(
  distinct_points(sK2,sK3)),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  (apart_point_and_line(sK3,line_connecting(sK0,sK1)) | apart_point_and_line(sK2,line_connecting(sK0,sK1))) & ~apart_point_and_line(sK1,line_connecting(sK2,sK3)) & ~apart_point_and_line(sK0,line_connecting(sK2,sK3)) & distinct_points(sK2,sK3) & distinct_points(sK0,sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f38,f39])).
fof(f39,plain,(
  ? [X0,X1,X2,X3] : ((apart_point_and_line(X3,line_connecting(X0,X1)) | apart_point_and_line(X2,line_connecting(X0,X1))) & ~apart_point_and_line(X1,line_connecting(X2,X3)) & ~apart_point_and_line(X0,line_connecting(X2,X3)) & distinct_points(X2,X3) & distinct_points(X0,X1)) => ((apart_point_and_line(sK3,line_connecting(sK0,sK1)) | apart_point_and_line(sK2,line_connecting(sK0,sK1))) & ~apart_point_and_line(sK1,line_connecting(sK2,sK3)) & ~apart_point_and_line(sK0,line_connecting(sK2,sK3)) & distinct_points(sK2,sK3) & distinct_points(sK0,sK1))),
  introduced(choice_axiom,[])).
fof(f38,plain,(
  ? [X0,X1,X2,X3] : ((apart_point_and_line(X3,line_connecting(X0,X1)) | apart_point_and_line(X2,line_connecting(X0,X1))) & ~apart_point_and_line(X1,line_connecting(X2,X3)) & ~apart_point_and_line(X0,line_connecting(X2,X3)) & distinct_points(X2,X3) & distinct_points(X0,X1))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ? [X0,X1,X2,X3] : ((apart_point_and_line(X3,line_connecting(X0,X1)) | apart_point_and_line(X2,line_connecting(X0,X1))) & (~apart_point_and_line(X1,line_connecting(X2,X3)) & ~apart_point_and_line(X0,line_connecting(X2,X3)) & distinct_points(X2,X3) & distinct_points(X0,X1)))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1,X2,X3] : ((~apart_point_and_line(X1,line_connecting(X2,X3)) & ~apart_point_and_line(X0,line_connecting(X2,X3)) & distinct_points(X2,X3) & distinct_points(X0,X1)) => (~apart_point_and_line(X3,line_connecting(X0,X1)) & ~apart_point_and_line(X2,line_connecting(X0,X1))))),
  inference(rectify,[],[f16])).
fof(f16,negated_conjecture,(
  ~! [X0,X1,X3,X4] : ((~apart_point_and_line(X1,line_connecting(X3,X4)) & ~apart_point_and_line(X0,line_connecting(X3,X4)) & distinct_points(X3,X4) & distinct_points(X0,X1)) => (~apart_point_and_line(X4,line_connecting(X0,X1)) & ~apart_point_and_line(X3,line_connecting(X0,X1))))),
  inference(negated_conjecture,[],[f15])).
fof(f15,conjecture,(
  ! [X0,X1,X3,X4] : ((~apart_point_and_line(X1,line_connecting(X3,X4)) & ~apart_point_and_line(X0,line_connecting(X3,X4)) & distinct_points(X3,X4) & distinct_points(X0,X1)) => (~apart_point_and_line(X4,line_connecting(X0,X1)) & ~apart_point_and_line(X3,line_connecting(X0,X1))))),
  file('./p/GEO187+1.tptp',unknown)).
fof(f2059,plain,(
  ~distinct_points(sK2,sK3) | ~spl4_26),
  inference(resolution,[],[f1997,f47])).
fof(f47,plain,(
  ( ! [X0,X1] : (~apart_point_and_line(X0,line_connecting(X0,X1)) | ~distinct_points(X0,X1)) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : (~apart_point_and_line(X0,line_connecting(X0,X1)) | ~distinct_points(X0,X1))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1] : (distinct_points(X0,X1) => ~apart_point_and_line(X0,line_connecting(X0,X1)))),
  file('./p/GEO187+1.tptp',unknown)).
fof(f1997,plain,(
  apart_point_and_line(sK2,line_connecting(sK2,sK3)) | ~spl4_26),
  inference(avatar_component_clause,[],[f1995])).
fof(f1995,plain,(
  spl4_26 <=> apart_point_and_line(sK2,line_connecting(sK2,sK3))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_26])])).
fof(f2016,plain,(
  ~spl4_20),
  inference(avatar_contradiction_clause,[],[f2015])).
fof(f2015,plain,(
  $false | ~spl4_20),
  inference(subsumption_resolution,[],[f2012,f56])).
fof(f2012,plain,(
  ~distinct_points(sK2,sK3) | ~spl4_20),
  inference(resolution,[],[f1829,f48])).
fof(f48,plain,(
  ( ! [X0,X1] : (~apart_point_and_line(X1,line_connecting(X0,X1)) | ~distinct_points(X0,X1)) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : (~apart_point_and_line(X1,line_connecting(X0,X1)) | ~distinct_points(X0,X1))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1] : (distinct_points(X0,X1) => ~apart_point_and_line(X1,line_connecting(X0,X1)))),
  file('./p/GEO187+1.tptp',unknown)).
fof(f1829,plain,(
  apart_point_and_line(sK3,line_connecting(sK2,sK3)) | ~spl4_20),
  inference(avatar_component_clause,[],[f1827])).
fof(f1827,plain,(
  spl4_20 <=> apart_point_and_line(sK3,line_connecting(sK2,sK3))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_20])])).
fof(f1998,plain,(
  spl4_19 | spl4_26 | ~spl4_1 | spl4_4 | spl4_5),
  inference(avatar_split_clause,[],[f1993,f221,f216,f61,f1995,f1824])).
fof(f1824,plain,(
  spl4_19 <=> ! [X9] : distinct_points(sK1,X9)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_19])])).
fof(f61,plain,(
  spl4_1 <=> apart_point_and_line(sK2,line_connecting(sK0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_1])])).
fof(f216,plain,(
  spl4_4 <=> apart_point_and_line(sK1,line_connecting(sK0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_4])])).
fof(f221,plain,(
  spl4_5 <=> apart_point_and_line(sK0,line_connecting(sK0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_5])])).
fof(f1993,plain,(
  ( ! [X10] : (apart_point_and_line(sK2,line_connecting(sK2,sK3)) | distinct_points(sK1,X10)) ) | (~spl4_1 | spl4_4 | spl4_5)),
  inference(subsumption_resolution,[],[f1992,f57])).
fof(f57,plain,(
  ~apart_point_and_line(sK0,line_connecting(sK2,sK3))),
  inference(cnf_transformation,[],[f40])).
fof(f1992,plain,(
  ( ! [X10] : (apart_point_and_line(sK2,line_connecting(sK2,sK3)) | distinct_points(sK1,X10) | apart_point_and_line(sK0,line_connecting(sK2,sK3))) ) | (~spl4_1 | spl4_4 | spl4_5)),
  inference(subsumption_resolution,[],[f1978,f217])).
fof(f217,plain,(
  ~apart_point_and_line(sK1,line_connecting(sK0,sK1)) | spl4_4),
  inference(avatar_component_clause,[],[f216])).
fof(f1978,plain,(
  ( ! [X10] : (apart_point_and_line(sK1,line_connecting(sK0,sK1)) | apart_point_and_line(sK2,line_connecting(sK2,sK3)) | distinct_points(sK1,X10) | apart_point_and_line(sK0,line_connecting(sK2,sK3))) ) | (~spl4_1 | spl4_5)),
  inference(duplicate_literal_removal,[],[f1958])).
fof(f1958,plain,(
  ( ! [X10] : (apart_point_and_line(sK1,line_connecting(sK0,sK1)) | apart_point_and_line(sK2,line_connecting(sK2,sK3)) | distinct_points(sK1,X10) | distinct_points(sK1,X10) | apart_point_and_line(sK0,line_connecting(sK2,sK3))) ) | (~spl4_1 | spl4_5)),
  inference(resolution,[],[f358,f58])).
fof(f58,plain,(
  ~apart_point_and_line(sK1,line_connecting(sK2,sK3))),
  inference(cnf_transformation,[],[f40])).
fof(f358,plain,(
  ( ! [X17,X18,X16] : (apart_point_and_line(X16,line_connecting(sK0,sK1)) | apart_point_and_line(sK2,X17) | apart_point_and_line(X16,X17) | distinct_points(sK1,X18) | distinct_points(X16,X18) | apart_point_and_line(sK0,X17)) ) | (~spl4_1 | spl4_5)),
  inference(subsumption_resolution,[],[f328,f222])).
fof(f222,plain,(
  ~apart_point_and_line(sK0,line_connecting(sK0,sK1)) | spl4_5),
  inference(avatar_component_clause,[],[f221])).
fof(f328,plain,(
  ( ! [X17,X18,X16] : (apart_point_and_line(X16,line_connecting(sK0,sK1)) | apart_point_and_line(sK0,X17) | apart_point_and_line(sK0,line_connecting(sK0,sK1)) | apart_point_and_line(X16,X17) | apart_point_and_line(sK2,X17) | distinct_points(sK1,X18) | distinct_points(X16,X18)) ) | ~spl4_1),
  inference(resolution,[],[f252,f72])).
fof(f72,plain,(
  ( ! [X0,X1] : (distinct_points(sK0,X0) | distinct_points(sK1,X1) | distinct_points(X0,X1)) )),
  inference(resolution,[],[f69,f44])).
fof(f44,plain,(
  ( ! [X2,X0,X1] : (~distinct_points(X0,X1) | distinct_points(X0,X2) | distinct_points(X1,X2)) )),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1,X2] : (distinct_points(X1,X2) | distinct_points(X0,X2) | ~distinct_points(X0,X1))),
  inference(flattening,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2] : ((distinct_points(X1,X2) | distinct_points(X0,X2)) | ~distinct_points(X0,X1))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2] : (distinct_points(X0,X1) => (distinct_points(X1,X2) | distinct_points(X0,X2)))),
  file('./p/GEO187+1.tptp',unknown)).
fof(f69,plain,(
  ( ! [X0] : (distinct_points(sK1,X0) | distinct_points(sK0,X0)) )),
  inference(resolution,[],[f44,f55])).
fof(f55,plain,(
  distinct_points(sK0,sK1)),
  inference(cnf_transformation,[],[f40])).
fof(f252,plain,(
  ( ! [X2,X0,X1] : (~distinct_points(X2,X1) | apart_point_and_line(X1,line_connecting(sK0,sK1)) | apart_point_and_line(X2,X0) | apart_point_and_line(X2,line_connecting(sK0,sK1)) | apart_point_and_line(X1,X0) | apart_point_and_line(sK2,X0)) ) | ~spl4_1),
  inference(resolution,[],[f249,f51])).
fof(f51,plain,(
  ( ! [X2,X0,X3,X1] : (~distinct_lines(X2,X3) | apart_point_and_line(X1,X2) | apart_point_and_line(X0,X3) | apart_point_and_line(X0,X2) | apart_point_and_line(X1,X3) | ~distinct_points(X0,X1)) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2,X3] : (apart_point_and_line(X1,X3) | apart_point_and_line(X1,X2) | apart_point_and_line(X0,X3) | apart_point_and_line(X0,X2) | ~distinct_lines(X2,X3) | ~distinct_points(X0,X1))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2,X3] : ((apart_point_and_line(X1,X3) | apart_point_and_line(X1,X2) | apart_point_and_line(X0,X3) | apart_point_and_line(X0,X2)) | (~distinct_lines(X2,X3) | ~distinct_points(X0,X1)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1,X2,X3] : ((distinct_lines(X2,X3) & distinct_points(X0,X1)) => (apart_point_and_line(X1,X3) | apart_point_and_line(X1,X2) | apart_point_and_line(X0,X3) | apart_point_and_line(X0,X2)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X3,X4] : ((distinct_lines(X3,X4) & distinct_points(X0,X1)) => (apart_point_and_line(X1,X4) | apart_point_and_line(X1,X3) | apart_point_and_line(X0,X4) | apart_point_and_line(X0,X3)))),
  file('./p/GEO187+1.tptp',unknown)).
fof(f249,plain,(
  ( ! [X0] : (distinct_lines(line_connecting(sK0,sK1),X0) | apart_point_and_line(sK2,X0)) ) | ~spl4_1),
  inference(resolution,[],[f63,f53])).
fof(f53,plain,(
  ( ! [X2,X0,X1] : (~apart_point_and_line(X0,X1) | distinct_lines(X1,X2) | apart_point_and_line(X0,X2)) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : (apart_point_and_line(X0,X2) | distinct_lines(X1,X2) | ~apart_point_and_line(X0,X1))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2] : ((apart_point_and_line(X0,X2) | distinct_lines(X1,X2)) | ~apart_point_and_line(X0,X1))),
  inference(ennf_transformation,[],[f13])).
fof(f13,axiom,(
  ! [X0,X1,X2] : (apart_point_and_line(X0,X1) => (apart_point_and_line(X0,X2) | distinct_lines(X1,X2)))),
  file('./p/GEO187+1.tptp',unknown)).
fof(f63,plain,(
  apart_point_and_line(sK2,line_connecting(sK0,sK1)) | ~spl4_1),
  inference(avatar_component_clause,[],[f61])).
fof(f1855,plain,(
  ~spl4_19),
  inference(avatar_contradiction_clause,[],[f1852])).
fof(f1852,plain,(
  $false | ~spl4_19),
  inference(resolution,[],[f1825,f41])).
fof(f41,plain,(
  ( ! [X0] : (~distinct_points(X0,X0)) )),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0] : ~distinct_points(X0,X0)),
  file('./p/GEO187+1.tptp',unknown)).
fof(f1825,plain,(
  ( ! [X9] : (distinct_points(sK1,X9)) ) | ~spl4_19),
  inference(avatar_component_clause,[],[f1824])).
fof(f1830,plain,(
  spl4_19 | spl4_20 | spl4_4 | ~spl4_8),
  inference(avatar_split_clause,[],[f1822,f233,f216,f1827,f1824])).
fof(f233,plain,(
  spl4_8 <=> ! [X22,X21,X23] : (apart_point_and_line(X21,line_connecting(sK0,sK1)) | distinct_points(X21,X23) | distinct_points(sK1,X23) | apart_point_and_line(sK3,X22) | apart_point_and_line(X21,X22) | apart_point_and_line(sK0,X22))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_8])])).
fof(f1822,plain,(
  ( ! [X9] : (apart_point_and_line(sK3,line_connecting(sK2,sK3)) | distinct_points(sK1,X9)) ) | (spl4_4 | ~spl4_8)),
  inference(subsumption_resolution,[],[f1821,f57])).
fof(f1821,plain,(
  ( ! [X9] : (apart_point_and_line(sK3,line_connecting(sK2,sK3)) | distinct_points(sK1,X9) | apart_point_and_line(sK0,line_connecting(sK2,sK3))) ) | (spl4_4 | ~spl4_8)),
  inference(subsumption_resolution,[],[f1814,f217])).
fof(f1814,plain,(
  ( ! [X9] : (apart_point_and_line(sK1,line_connecting(sK0,sK1)) | apart_point_and_line(sK3,line_connecting(sK2,sK3)) | distinct_points(sK1,X9) | apart_point_and_line(sK0,line_connecting(sK2,sK3))) ) | ~spl4_8),
  inference(duplicate_literal_removal,[],[f1791])).
fof(f1791,plain,(
  ( ! [X9] : (apart_point_and_line(sK1,line_connecting(sK0,sK1)) | apart_point_and_line(sK3,line_connecting(sK2,sK3)) | distinct_points(sK1,X9) | distinct_points(sK1,X9) | apart_point_and_line(sK0,line_connecting(sK2,sK3))) ) | ~spl4_8),
  inference(resolution,[],[f234,f58])).
fof(f234,plain,(
  ( ! [X23,X21,X22] : (apart_point_and_line(X21,line_connecting(sK0,sK1)) | apart_point_and_line(sK3,X22) | apart_point_and_line(X21,X22) | distinct_points(sK1,X23) | distinct_points(X21,X23) | apart_point_and_line(sK0,X22)) ) | ~spl4_8),
  inference(avatar_component_clause,[],[f233])).
fof(f298,plain,(
  ~spl4_4),
  inference(avatar_contradiction_clause,[],[f297])).
fof(f297,plain,(
  $false | ~spl4_4),
  inference(subsumption_resolution,[],[f294,f55])).
fof(f294,plain,(
  ~distinct_points(sK0,sK1) | ~spl4_4),
  inference(resolution,[],[f218,f48])).
fof(f218,plain,(
  apart_point_and_line(sK1,line_connecting(sK0,sK1)) | ~spl4_4),
  inference(avatar_component_clause,[],[f216])).
fof(f258,plain,(
  ~spl4_5),
  inference(avatar_contradiction_clause,[],[f257])).
fof(f257,plain,(
  $false | ~spl4_5),
  inference(subsumption_resolution,[],[f254,f55])).
fof(f254,plain,(
  ~distinct_points(sK0,sK1) | ~spl4_5),
  inference(resolution,[],[f223,f47])).
fof(f223,plain,(
  apart_point_and_line(sK0,line_connecting(sK0,sK1)) | ~spl4_5),
  inference(avatar_component_clause,[],[f221])).
fof(f235,plain,(
  spl4_5 | spl4_8 | ~spl4_2),
  inference(avatar_split_clause,[],[f178,f65,f233,f221])).
fof(f65,plain,(
  spl4_2 <=> apart_point_and_line(sK3,line_connecting(sK0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_2])])).
fof(f178,plain,(
  ( ! [X23,X21,X22] : (apart_point_and_line(X21,line_connecting(sK0,sK1)) | apart_point_and_line(sK0,X22) | apart_point_and_line(sK0,line_connecting(sK0,sK1)) | apart_point_and_line(X21,X22) | apart_point_and_line(sK3,X22) | distinct_points(sK1,X23) | distinct_points(X21,X23)) ) | ~spl4_2),
  inference(resolution,[],[f90,f72])).
fof(f90,plain,(
  ( ! [X2,X0,X1] : (~distinct_points(X2,X1) | apart_point_and_line(X1,line_connecting(sK0,sK1)) | apart_point_and_line(X2,X0) | apart_point_and_line(X2,line_connecting(sK0,sK1)) | apart_point_and_line(X1,X0) | apart_point_and_line(sK3,X0)) ) | ~spl4_2),
  inference(resolution,[],[f83,f51])).
fof(f83,plain,(
  ( ! [X0] : (distinct_lines(line_connecting(sK0,sK1),X0) | apart_point_and_line(sK3,X0)) ) | ~spl4_2),
  inference(resolution,[],[f53,f67])).
fof(f67,plain,(
  apart_point_and_line(sK3,line_connecting(sK0,sK1)) | ~spl4_2),
  inference(avatar_component_clause,[],[f65])).
fof(f68,plain,(
  spl4_1 | spl4_2),
  inference(avatar_split_clause,[],[f59,f65,f61])).
fof(f59,plain,(
  apart_point_and_line(sK3,line_connecting(sK0,sK1)) | apart_point_and_line(sK2,line_connecting(sK0,sK1))),
  inference(cnf_transformation,[],[f40])).
% SZS output end Proof for GEO187+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5756
% Time elapsed: 0.042 s
% ------------------------------
% ------------------------------
