% Refutation found. Thanks to Tanya!
% SZS status Theorem for GEO171+3
% SZS output start Proof for GEO171+3
fof(f170,plain,(
  $false),
  inference(subsumption_resolution,[],[f167,f129])).
fof(f129,plain,(
  ~distinct_lines(sK0,sK1)),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ~distinct_lines(sK0,sK1) & convergent_lines(sK0,sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f89,f90])).
fof(f90,plain,(
  ? [X0,X1] : (~distinct_lines(X0,X1) & convergent_lines(X0,X1)) => (~distinct_lines(sK0,sK1) & convergent_lines(sK0,sK1))),
  introduced(choice_axiom,[])).
fof(f89,plain,(
  ? [X0,X1] : (~distinct_lines(X0,X1) & convergent_lines(X0,X1))),
  inference(ennf_transformation,[],[f37])).
fof(f37,negated_conjecture,(
  ~! [X0,X1] : (convergent_lines(X0,X1) => distinct_lines(X0,X1))),
  inference(negated_conjecture,[],[f36])).
fof(f36,conjecture,(
  ! [X0,X1] : (convergent_lines(X0,X1) => distinct_lines(X0,X1))),
  file('./p/GEO171+3.tptp',unknown)).
fof(f167,plain,(
  distinct_lines(sK0,sK1)),
  inference(resolution,[],[f152,f93])).
fof(f93,plain,(
  ( ! [X0] : (~distinct_lines(X0,X0)) )),
  inference(cnf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0] : ~distinct_lines(X0,X0)),
  file('./p/GEO171+3.tptp',unknown)).
fof(f152,plain,(
  ( ! [X0] : (distinct_lines(sK1,X0) | distinct_lines(sK0,X0)) )),
  inference(resolution,[],[f148,f96])).
fof(f96,plain,(
  ( ! [X2,X0,X1] : (~distinct_lines(X0,X1) | distinct_lines(X0,X2) | distinct_lines(X1,X2)) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : (distinct_lines(X1,X2) | distinct_lines(X0,X2) | ~distinct_lines(X0,X1))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : ((distinct_lines(X1,X2) | distinct_lines(X0,X2)) | ~distinct_lines(X0,X1))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2] : (distinct_lines(X0,X1) => (distinct_lines(X1,X2) | distinct_lines(X0,X2)))),
  file('./p/GEO171+3.tptp',unknown)).
fof(f148,plain,(
  distinct_lines(sK1,sK0)),
  inference(resolution,[],[f137,f94])).
fof(f94,plain,(
  ( ! [X0] : (~convergent_lines(X0,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : ~convergent_lines(X0,X0)),
  file('./p/GEO171+3.tptp',unknown)).
fof(f137,plain,(
  ( ! [X0] : (convergent_lines(sK0,X0) | distinct_lines(sK1,X0)) )),
  inference(resolution,[],[f105,f128])).
fof(f128,plain,(
  convergent_lines(sK0,sK1)),
  inference(cnf_transformation,[],[f91])).
fof(f105,plain,(
  ( ! [X2,X0,X1] : (~convergent_lines(X0,X1) | distinct_lines(X1,X2) | convergent_lines(X0,X2)) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : (convergent_lines(X0,X2) | distinct_lines(X1,X2) | ~convergent_lines(X0,X1))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : ((convergent_lines(X0,X2) | distinct_lines(X1,X2)) | ~convergent_lines(X0,X1))),
  inference(ennf_transformation,[],[f14])).
fof(f14,axiom,(
  ! [X0,X1,X2] : (convergent_lines(X0,X1) => (convergent_lines(X0,X2) | distinct_lines(X1,X2)))),
  file('./p/GEO171+3.tptp',unknown)).
% SZS output end Proof for GEO171+3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
