% Refutation found. Thanks to Tanya!
% SZS status Theorem for GEO181+3
% SZS output start Proof for GEO181+3
fof(f279,plain,(
  $false),
  inference(resolution,[],[f274,f109])).
fof(f109,plain,(
  ( ! [X0,X1] : (~apart_point_and_line(X0,parallel_through_point(X1,X0))) )),
  inference(cnf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0,X1] : ~apart_point_and_line(X0,parallel_through_point(X1,X0))),
  file('./p/GEO181+3.tptp',unknown)).
fof(f274,plain,(
  ( ! [X0] : (apart_point_and_line(sK2,parallel_through_point(line_connecting(sK0,sK1),X0))) )),
  inference(resolution,[],[f271,f108])).
fof(f108,plain,(
  ( ! [X0,X1] : (~convergent_lines(parallel_through_point(X1,X0),X1)) )),
  inference(cnf_transformation,[],[f16])).
fof(f16,axiom,(
  ! [X0,X1] : ~convergent_lines(parallel_through_point(X1,X0),X1)),
  file('./p/GEO181+3.tptp',unknown)).
fof(f271,plain,(
  ( ! [X8] : (convergent_lines(X8,line_connecting(sK0,sK1)) | apart_point_and_line(sK2,X8)) )),
  inference(resolution,[],[f256,f107])).
fof(f107,plain,(
  ( ! [X0,X1] : (~distinct_lines(X0,X1) | convergent_lines(X0,X1)) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (convergent_lines(X0,X1) | ~distinct_lines(X0,X1))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0,X1] : (distinct_lines(X0,X1) => convergent_lines(X0,X1))),
  file('./p/GEO181+3.tptp',unknown)).
fof(f256,plain,(
  ( ! [X0] : (distinct_lines(X0,line_connecting(sK0,sK1)) | apart_point_and_line(sK2,X0)) )),
  inference(resolution,[],[f160,f94])).
fof(f94,plain,(
  ( ! [X0] : (~distinct_lines(X0,X0)) )),
  inference(cnf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0] : ~distinct_lines(X0,X0)),
  file('./p/GEO181+3.tptp',unknown)).
fof(f160,plain,(
  ( ! [X0,X1] : (apart_point_and_line(sK2,X0) | distinct_lines(line_connecting(sK0,sK1),X1) | distinct_lines(X0,X1)) )),
  inference(resolution,[],[f145,f97])).
fof(f97,plain,(
  ( ! [X2,X0,X1] : (~distinct_lines(X0,X1) | distinct_lines(X0,X2) | distinct_lines(X1,X2)) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : (distinct_lines(X1,X2) | distinct_lines(X0,X2) | ~distinct_lines(X0,X1))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : ((distinct_lines(X1,X2) | distinct_lines(X0,X2)) | ~distinct_lines(X0,X1))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2] : (distinct_lines(X0,X1) => (distinct_lines(X1,X2) | distinct_lines(X0,X2)))),
  file('./p/GEO181+3.tptp',unknown)).
fof(f145,plain,(
  ( ! [X0] : (distinct_lines(line_connecting(sK0,sK1),X0) | apart_point_and_line(sK2,X0)) )),
  inference(resolution,[],[f105,f130])).
fof(f130,plain,(
  apart_point_and_line(sK2,line_connecting(sK0,sK1))),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ~apart_point_and_line(sK1,line_connecting(sK0,sK2)) & apart_point_and_line(sK2,line_connecting(sK0,sK1)) & distinct_points(sK0,sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f90,f91])).
fof(f91,plain,(
  ? [X0,X1,X2] : (~apart_point_and_line(X1,line_connecting(X0,X2)) & apart_point_and_line(X2,line_connecting(X0,X1)) & distinct_points(X0,X1)) => (~apart_point_and_line(sK1,line_connecting(sK0,sK2)) & apart_point_and_line(sK2,line_connecting(sK0,sK1)) & distinct_points(sK0,sK1))),
  introduced(choice_axiom,[])).
fof(f90,plain,(
  ? [X0,X1,X2] : (~apart_point_and_line(X1,line_connecting(X0,X2)) & apart_point_and_line(X2,line_connecting(X0,X1)) & distinct_points(X0,X1))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ? [X0,X1,X2] : ((~apart_point_and_line(X1,line_connecting(X0,X2)) & apart_point_and_line(X2,line_connecting(X0,X1))) & distinct_points(X0,X1))),
  inference(ennf_transformation,[],[f37])).
fof(f37,negated_conjecture,(
  ~! [X0,X1,X2] : (distinct_points(X0,X1) => (apart_point_and_line(X2,line_connecting(X0,X1)) => apart_point_and_line(X1,line_connecting(X0,X2))))),
  inference(negated_conjecture,[],[f36])).
fof(f36,conjecture,(
  ! [X0,X1,X2] : (distinct_points(X0,X1) => (apart_point_and_line(X2,line_connecting(X0,X1)) => apart_point_and_line(X1,line_connecting(X0,X2))))),
  file('./p/GEO181+3.tptp',unknown)).
fof(f105,plain,(
  ( ! [X2,X0,X1] : (~apart_point_and_line(X0,X1) | distinct_lines(X1,X2) | apart_point_and_line(X0,X2)) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2] : (apart_point_and_line(X0,X2) | distinct_lines(X1,X2) | ~apart_point_and_line(X0,X1))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : ((apart_point_and_line(X0,X2) | distinct_lines(X1,X2)) | ~apart_point_and_line(X0,X1))),
  inference(ennf_transformation,[],[f13])).
fof(f13,axiom,(
  ! [X0,X1,X2] : (apart_point_and_line(X0,X1) => (apart_point_and_line(X0,X2) | distinct_lines(X1,X2)))),
  file('./p/GEO181+3.tptp',unknown)).
% SZS output end Proof for GEO181+3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
