% Refutation found. Thanks to Tanya!
% SZS status Theorem for GEO200+3
% SZS output start Proof for GEO200+3
fof(f371,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f338,f355,f360,f365,f370])).
fof(f370,plain,(
  ~spl2_6),
  inference(avatar_contradiction_clause,[],[f369])).
fof(f369,plain,(
  $false | ~spl2_6),
  inference(subsumption_resolution,[],[f366,f206])).
fof(f206,plain,(
  distinct_points(sK1,sK0)),
  inference(resolution,[],[f195,f94])).
fof(f94,plain,(
  ( ! [X0] : (~distinct_points(X0,X0)) )),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0] : ~distinct_points(X0,X0)),
  file('./p/GEO200+3.tptp',unknown)).
fof(f195,plain,(
  ( ! [X0] : (distinct_points(X0,sK1) | distinct_points(X0,sK0)) )),
  inference(resolution,[],[f187,f94])).
fof(f187,plain,(
  ( ! [X0,X1] : (distinct_points(X0,sK1) | distinct_points(sK0,X1) | distinct_points(X0,X1)) )),
  inference(resolution,[],[f177,f97])).
fof(f97,plain,(
  ( ! [X2,X0,X1] : (~distinct_points(X0,X1) | distinct_points(X0,X2) | distinct_points(X1,X2)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : (distinct_points(X1,X2) | distinct_points(X0,X2) | ~distinct_points(X0,X1))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2] : ((distinct_points(X1,X2) | distinct_points(X0,X2)) | ~distinct_points(X0,X1))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2] : (distinct_points(X0,X1) => (distinct_points(X1,X2) | distinct_points(X0,X2)))),
  file('./p/GEO200+3.tptp',unknown)).
fof(f177,plain,(
  ( ! [X0] : (distinct_points(sK0,X0) | distinct_points(X0,sK1)) )),
  inference(resolution,[],[f138,f94])).
fof(f138,plain,(
  ( ! [X0,X1] : (distinct_points(sK0,X0) | distinct_points(sK1,X1) | distinct_points(X0,X1)) )),
  inference(resolution,[],[f136,f97])).
fof(f136,plain,(
  ( ! [X0] : (distinct_points(sK1,X0) | distinct_points(sK0,X0)) )),
  inference(resolution,[],[f97,f131])).
fof(f131,plain,(
  distinct_points(sK0,sK1)),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ~equal_lines(line_connecting(sK0,sK1),line_connecting(sK1,sK0)) & distinct_points(sK0,sK1)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f91,f92])).
fof(f92,plain,(
  ? [X0,X1] : (~equal_lines(line_connecting(X0,X1),line_connecting(X1,X0)) & distinct_points(X0,X1)) => (~equal_lines(line_connecting(sK0,sK1),line_connecting(sK1,sK0)) & distinct_points(sK0,sK1))),
  introduced(choice_axiom,[])).
fof(f91,plain,(
  ? [X0,X1] : (~equal_lines(line_connecting(X0,X1),line_connecting(X1,X0)) & distinct_points(X0,X1))),
  inference(ennf_transformation,[],[f37])).
fof(f37,negated_conjecture,(
  ~! [X0,X1] : (distinct_points(X0,X1) => equal_lines(line_connecting(X0,X1),line_connecting(X1,X0)))),
  inference(negated_conjecture,[],[f36])).
fof(f36,conjecture,(
  ! [X0,X1] : (distinct_points(X0,X1) => equal_lines(line_connecting(X0,X1),line_connecting(X1,X0)))),
  file('./p/GEO200+3.tptp',unknown)).
fof(f366,plain,(
  ~distinct_points(sK1,sK0) | ~spl2_6),
  inference(resolution,[],[f300,f101])).
fof(f101,plain,(
  ( ! [X0,X1] : (~apart_point_and_line(X1,line_connecting(X0,X1)) | ~distinct_points(X0,X1)) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : (~apart_point_and_line(X1,line_connecting(X0,X1)) | ~distinct_points(X0,X1))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1] : (distinct_points(X0,X1) => ~apart_point_and_line(X1,line_connecting(X0,X1)))),
  file('./p/GEO200+3.tptp',unknown)).
fof(f300,plain,(
  apart_point_and_line(sK0,line_connecting(sK1,sK0)) | ~spl2_6),
  inference(avatar_component_clause,[],[f298])).
fof(f298,plain,(
  spl2_6 <=> apart_point_and_line(sK0,line_connecting(sK1,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl2_6])])).
fof(f365,plain,(
  ~spl2_5),
  inference(avatar_contradiction_clause,[],[f364])).
fof(f364,plain,(
  $false | ~spl2_5),
  inference(subsumption_resolution,[],[f361,f131])).
fof(f361,plain,(
  ~distinct_points(sK0,sK1) | ~spl2_5),
  inference(resolution,[],[f296,f100])).
fof(f100,plain,(
  ( ! [X0,X1] : (~apart_point_and_line(X0,line_connecting(X0,X1)) | ~distinct_points(X0,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : (~apart_point_and_line(X0,line_connecting(X0,X1)) | ~distinct_points(X0,X1))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1] : (distinct_points(X0,X1) => ~apart_point_and_line(X0,line_connecting(X0,X1)))),
  file('./p/GEO200+3.tptp',unknown)).
fof(f296,plain,(
  apart_point_and_line(sK0,line_connecting(sK0,sK1)) | ~spl2_5),
  inference(avatar_component_clause,[],[f294])).
fof(f294,plain,(
  spl2_5 <=> apart_point_and_line(sK0,line_connecting(sK0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl2_5])])).
fof(f360,plain,(
  ~spl2_4),
  inference(avatar_contradiction_clause,[],[f359])).
fof(f359,plain,(
  $false | ~spl2_4),
  inference(subsumption_resolution,[],[f356,f206])).
fof(f356,plain,(
  ~distinct_points(sK1,sK0) | ~spl2_4),
  inference(resolution,[],[f292,f100])).
fof(f292,plain,(
  apart_point_and_line(sK1,line_connecting(sK1,sK0)) | ~spl2_4),
  inference(avatar_component_clause,[],[f290])).
fof(f290,plain,(
  spl2_4 <=> apart_point_and_line(sK1,line_connecting(sK1,sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl2_4])])).
fof(f355,plain,(
  ~spl2_3),
  inference(avatar_contradiction_clause,[],[f354])).
fof(f354,plain,(
  $false | ~spl2_3),
  inference(subsumption_resolution,[],[f351,f131])).
fof(f351,plain,(
  ~distinct_points(sK0,sK1) | ~spl2_3),
  inference(resolution,[],[f288,f101])).
fof(f288,plain,(
  apart_point_and_line(sK1,line_connecting(sK0,sK1)) | ~spl2_3),
  inference(avatar_component_clause,[],[f286])).
fof(f286,plain,(
  spl2_3 <=> apart_point_and_line(sK1,line_connecting(sK0,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl2_3])])).
fof(f338,plain,(
  spl2_5 | spl2_6 | spl2_3 | spl2_4),
  inference(avatar_split_clause,[],[f275,f290,f286,f298,f294])).
fof(f275,plain,(
  apart_point_and_line(sK1,line_connecting(sK1,sK0)) | apart_point_and_line(sK1,line_connecting(sK0,sK1)) | apart_point_and_line(sK0,line_connecting(sK1,sK0)) | apart_point_and_line(sK0,line_connecting(sK0,sK1))),
  inference(resolution,[],[f233,f206])).
fof(f233,plain,(
  ( ! [X0,X1] : (~distinct_points(X1,X0) | apart_point_and_line(X1,line_connecting(sK1,sK0)) | apart_point_and_line(X1,line_connecting(sK0,sK1)) | apart_point_and_line(X0,line_connecting(sK1,sK0)) | apart_point_and_line(X0,line_connecting(sK0,sK1))) )),
  inference(resolution,[],[f104,f134])).
fof(f134,plain,(
  distinct_lines(line_connecting(sK0,sK1),line_connecting(sK1,sK0))),
  inference(resolution,[],[f132,f130])).
fof(f130,plain,(
  ( ! [X0,X1] : (equal_lines(X0,X1) | distinct_lines(X0,X1)) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : (equal_lines(X0,X1) | distinct_lines(X0,X1))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : (~distinct_lines(X0,X1) => equal_lines(X0,X1))),
  inference(unused_predicate_definition_removal,[],[f32])).
fof(f32,axiom,(
  ! [X0,X1] : (equal_lines(X0,X1) <=> ~distinct_lines(X0,X1))),
  file('./p/GEO200+3.tptp',unknown)).
fof(f132,plain,(
  ~equal_lines(line_connecting(sK0,sK1),line_connecting(sK1,sK0))),
  inference(cnf_transformation,[],[f93])).
fof(f104,plain,(
  ( ! [X2,X0,X3,X1] : (~distinct_lines(X2,X3) | apart_point_and_line(X1,X2) | apart_point_and_line(X0,X3) | apart_point_and_line(X0,X2) | apart_point_and_line(X1,X3) | ~distinct_points(X0,X1)) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3] : (apart_point_and_line(X1,X3) | apart_point_and_line(X1,X2) | apart_point_and_line(X0,X3) | apart_point_and_line(X0,X2) | ~distinct_lines(X2,X3) | ~distinct_points(X0,X1))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : ((apart_point_and_line(X1,X3) | apart_point_and_line(X1,X2) | apart_point_and_line(X0,X3) | apart_point_and_line(X0,X2)) | (~distinct_lines(X2,X3) | ~distinct_points(X0,X1)))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : ((distinct_lines(X2,X3) & distinct_points(X0,X1)) => (apart_point_and_line(X1,X3) | apart_point_and_line(X1,X2) | apart_point_and_line(X0,X3) | apart_point_and_line(X0,X2)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1,X3,X4] : ((distinct_lines(X3,X4) & distinct_points(X0,X1)) => (apart_point_and_line(X1,X4) | apart_point_and_line(X1,X3) | apart_point_and_line(X0,X4) | apart_point_and_line(X0,X3)))),
  file('./p/GEO200+3.tptp',unknown)).
% SZS output end Proof for GEO200+3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
