% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GRP029-1
% SZS output start Proof for GRP029-1
fof(f537,plain,(
  $false),
  inference(resolution,[],[f400,f7])).
fof(f7,axiom,(
  ( ! [X6] : (~product(not_right_identity(X6),X6,not_right_identity(X6))) )),
  file('./p/GRP029-1.tptp',unknown)).
fof(f400,plain,(
  ( ! [X0] : (product(X0,identity,X0)) )),
  inference(backward_demodulation,[],[f387,f63])).
fof(f63,plain,(
  ( ! [X0] : (product(inverse(inverse(X0)),identity,X0)) )),
  inference(resolution,[],[f27,f6])).
fof(f6,axiom,(
  ( ! [X6] : (product(inverse(X6),X6,identity)) )),
  file('./p/GRP029-1.tptp',unknown)).
fof(f27,plain,(
  ( ! [X2,X3] : (~product(X2,inverse(X3),identity) | product(X2,identity,X3)) )),
  inference(resolution,[],[f11,f6])).
fof(f11,plain,(
  ( ! [X2,X0,X3,X1] : (~product(X3,X2,X1) | product(X0,X1,X2) | ~product(X0,X3,identity)) )),
  inference(resolution,[],[f3,f5])).
fof(f5,axiom,(
  ( ! [X6] : (product(identity,X6,X6)) )),
  file('./p/GRP029-1.tptp',unknown)).
fof(f3,axiom,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~product(X4,X2,X3) | product(X0,X5,X3) | ~product(X1,X2,X5) | ~product(X0,X1,X4)) )),
  file('./p/GRP029-1.tptp',unknown)).
fof(f387,plain,(
  ( ! [X0] : (inverse(inverse(X0)) = X0) )),
  inference(resolution,[],[f347,f8])).
fof(f8,plain,(
  ( ! [X0,X1] : (~product(identity,X1,X0) | X0 = X1) )),
  inference(resolution,[],[f2,f5])).
fof(f2,axiom,(
  ( ! [X2,X0,X3,X1] : (~product(X0,X1,X3) | X2 = X3 | ~product(X0,X1,X2)) )),
  file('./p/GRP029-1.tptp',unknown)).
fof(f347,plain,(
  ( ! [X0] : (product(identity,inverse(inverse(X0)),X0)) )),
  inference(resolution,[],[f296,f45])).
fof(f45,plain,(
  ( ! [X2,X3] : (~product(inverse(inverse(X2)),identity,X3) | product(identity,X2,X3)) )),
  inference(resolution,[],[f17,f6])).
fof(f17,plain,(
  ( ! [X6,X4,X7,X5] : (~product(X4,inverse(X7),X6) | product(X6,X7,X5) | ~product(X4,identity,X5)) )),
  inference(resolution,[],[f4,f6])).
fof(f4,axiom,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~product(X1,X2,X5) | ~product(X0,X5,X3) | product(X4,X2,X3) | ~product(X0,X1,X4)) )),
  file('./p/GRP029-1.tptp',unknown)).
fof(f296,plain,(
  ( ! [X22] : (product(inverse(inverse(inverse(inverse(X22)))),identity,X22)) )),
  inference(resolution,[],[f259,f27])).
fof(f259,plain,(
  ( ! [X0] : (product(inverse(inverse(inverse(X0))),X0,identity)) )),
  inference(resolution,[],[f213,f63])).
fof(f213,plain,(
  ( ! [X2,X3] : (~product(X3,identity,inverse(X2)) | product(X3,X2,identity)) )),
  inference(superposition,[],[f49,f21])).
fof(f21,plain,(
  ( ! [X0] : (identity = multiply(inverse(X0),X0)) )),
  inference(resolution,[],[f9,f1])).
fof(f1,axiom,(
  ( ! [X0,X1] : (product(X0,X1,multiply(X0,X1))) )),
  file('./p/GRP029-1.tptp',unknown)).
fof(f9,plain,(
  ( ! [X2,X3] : (~product(inverse(X3),X3,X2) | identity = X2) )),
  inference(resolution,[],[f2,f6])).
fof(f49,plain,(
  ( ! [X2,X0,X1] : (product(X0,X1,multiply(X2,X1)) | ~product(X0,identity,X2)) )),
  inference(resolution,[],[f13,f5])).
fof(f13,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~product(X12,X11,X9) | product(X8,X9,multiply(X10,X11)) | ~product(X8,X12,X10)) )),
  inference(resolution,[],[f3,f1])).
% SZS output end Proof for GRP029-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
