% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GRP226-1
% SZS output start Proof for GRP226-1
fof(f696,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f34,f39,f44,f49,f54,f55,f56,f57,f62,f63,f64,f65,f70,f71,f72,f73,f78,f79,f80,f81,f91,f101,f126,f161,f200,f205,f263,f302,f410,f515,f520,f632,f651,f655,f693,f695])).
fof(f695,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_11 | ~spl0_14 | ~spl0_17),
  inference(avatar_split_clause,[],[f694,f123,f107,f86,f36,f31])).
fof(f31,plain,(
  spl0_2 <=> sk_c7 = inverse(sk_c4)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f36,plain,(
  spl0_3 <=> sk_c7 = multiply(sk_c4,sk_c6)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f86,plain,(
  spl0_11 <=> ! [X5] : (sk_c7 != inverse(X5) | sk_c6 != multiply(sk_c7,multiply(X5,sk_c7)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_11])])).
fof(f107,plain,(
  spl0_14 <=> sk_c7 = sk_c6),
  introduced(avatar_definition,[new_symbols(naming,[spl0_14])])).
fof(f123,plain,(
  spl0_17 <=> sk_c6 = multiply(sk_c7,sk_c7)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_17])])).
fof(f694,plain,(
  sk_c7 != inverse(sk_c4) | (~spl0_3 | ~spl0_11 | ~spl0_14 | ~spl0_17)),
  inference(subsumption_resolution,[],[f687,f660])).
fof(f660,plain,(
  sk_c7 = multiply(sk_c7,sk_c7) | (~spl0_14 | ~spl0_17)),
  inference(backward_demodulation,[],[f108,f124])).
fof(f124,plain,(
  sk_c6 = multiply(sk_c7,sk_c7) | ~spl0_17),
  inference(avatar_component_clause,[],[f123])).
fof(f108,plain,(
  sk_c7 = sk_c6 | ~spl0_14),
  inference(avatar_component_clause,[],[f107])).
fof(f687,plain,(
  sk_c7 != multiply(sk_c7,sk_c7) | sk_c7 != inverse(sk_c4) | (~spl0_3 | ~spl0_11 | ~spl0_14)),
  inference(superposition,[],[f664,f657])).
fof(f657,plain,(
  sk_c7 = multiply(sk_c4,sk_c7) | (~spl0_3 | ~spl0_14)),
  inference(backward_demodulation,[],[f108,f38])).
fof(f38,plain,(
  sk_c7 = multiply(sk_c4,sk_c6) | ~spl0_3),
  inference(avatar_component_clause,[],[f36])).
fof(f664,plain,(
  ( ! [X5] : (sk_c7 != multiply(sk_c7,multiply(X5,sk_c7)) | sk_c7 != inverse(X5)) ) | (~spl0_11 | ~spl0_14)),
  inference(forward_demodulation,[],[f87,f108])).
fof(f87,plain,(
  ( ! [X5] : (sk_c6 != multiply(sk_c7,multiply(X5,sk_c7)) | sk_c7 != inverse(X5)) ) | ~spl0_11),
  inference(avatar_component_clause,[],[f86])).
fof(f693,plain,(
  ~spl0_1 | ~spl0_6 | ~spl0_11 | ~spl0_14 | ~spl0_17),
  inference(avatar_contradiction_clause,[],[f692])).
fof(f692,plain,(
  $false | (~spl0_1 | ~spl0_6 | ~spl0_11 | ~spl0_14 | ~spl0_17)),
  inference(subsumption_resolution,[],[f691,f29])).
fof(f29,plain,(
  inverse(sk_c1) = sk_c7 | ~spl0_1),
  inference(avatar_component_clause,[],[f27])).
fof(f27,plain,(
  spl0_1 <=> inverse(sk_c1) = sk_c7),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f691,plain,(
  inverse(sk_c1) != sk_c7 | (~spl0_6 | ~spl0_11 | ~spl0_14 | ~spl0_17)),
  inference(subsumption_resolution,[],[f686,f660])).
fof(f686,plain,(
  sk_c7 != multiply(sk_c7,sk_c7) | inverse(sk_c1) != sk_c7 | (~spl0_6 | ~spl0_11 | ~spl0_14)),
  inference(superposition,[],[f664,f658])).
fof(f658,plain,(
  sk_c7 = multiply(sk_c1,sk_c7) | (~spl0_6 | ~spl0_14)),
  inference(backward_demodulation,[],[f108,f53])).
fof(f53,plain,(
  sk_c7 = multiply(sk_c1,sk_c6) | ~spl0_6),
  inference(avatar_component_clause,[],[f51])).
fof(f51,plain,(
  spl0_6 <=> sk_c7 = multiply(sk_c1,sk_c6)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).
fof(f655,plain,(
  spl0_14 | ~spl0_7 | ~spl0_8 | ~spl0_9),
  inference(avatar_split_clause,[],[f654,f75,f67,f59,f107])).
fof(f59,plain,(
  spl0_7 <=> sk_c6 = multiply(sk_c7,sk_c3)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).
fof(f67,plain,(
  spl0_8 <=> sk_c3 = multiply(sk_c2,sk_c7)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_8])])).
fof(f75,plain,(
  spl0_9 <=> sk_c7 = inverse(sk_c2)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_9])])).
fof(f654,plain,(
  sk_c7 = sk_c6 | (~spl0_7 | ~spl0_8 | ~spl0_9)),
  inference(forward_demodulation,[],[f61,f568])).
fof(f568,plain,(
  sk_c7 = multiply(sk_c7,sk_c3) | (~spl0_8 | ~spl0_9)),
  inference(forward_demodulation,[],[f566,f77])).
fof(f77,plain,(
  sk_c7 = inverse(sk_c2) | ~spl0_9),
  inference(avatar_component_clause,[],[f75])).
fof(f566,plain,(
  sk_c7 = multiply(inverse(sk_c2),sk_c3) | ~spl0_8),
  inference(superposition,[],[f144,f69])).
fof(f69,plain,(
  sk_c3 = multiply(sk_c2,sk_c7) | ~spl0_8),
  inference(avatar_component_clause,[],[f67])).
fof(f144,plain,(
  ( ! [X2,X3] : (multiply(inverse(X2),multiply(X2,X3)) = X3) )),
  inference(forward_demodulation,[],[f137,f1])).
fof(f1,axiom,(
  ( ! [X0] : (multiply(identity,X0) = X0) )),
  file('./p/GRP226-1.tptp',unknown)).
fof(f137,plain,(
  ( ! [X2,X3] : (multiply(inverse(X2),multiply(X2,X3)) = multiply(identity,X3)) )),
  inference(superposition,[],[f3,f2])).
fof(f2,axiom,(
  ( ! [X0] : (identity = multiply(inverse(X0),X0)) )),
  file('./p/GRP226-1.tptp',unknown)).
fof(f3,axiom,(
  ( ! [X2,X0,X1] : (multiply(multiply(X0,X1),X2) = multiply(X0,multiply(X1,X2))) )),
  file('./p/GRP226-1.tptp',unknown)).
fof(f61,plain,(
  sk_c6 = multiply(sk_c7,sk_c3) | ~spl0_7),
  inference(avatar_component_clause,[],[f59])).
fof(f651,plain,(
  ~spl0_1 | ~spl0_6 | ~spl0_10 | ~spl0_14),
  inference(avatar_contradiction_clause,[],[f650])).
fof(f650,plain,(
  $false | (~spl0_1 | ~spl0_6 | ~spl0_10 | ~spl0_14)),
  inference(subsumption_resolution,[],[f649,f29])).
fof(f649,plain,(
  inverse(sk_c1) != sk_c7 | (~spl0_6 | ~spl0_10 | ~spl0_14)),
  inference(trivial_inequality_removal,[],[f639])).
fof(f639,plain,(
  sk_c7 != sk_c7 | inverse(sk_c1) != sk_c7 | (~spl0_6 | ~spl0_10 | ~spl0_14)),
  inference(superposition,[],[f638,f554])).
fof(f554,plain,(
  sk_c7 = multiply(sk_c1,sk_c7) | (~spl0_6 | ~spl0_14)),
  inference(backward_demodulation,[],[f108,f53])).
fof(f638,plain,(
  ( ! [X3] : (sk_c7 != multiply(X3,sk_c7) | sk_c7 != inverse(X3)) ) | (~spl0_10 | ~spl0_14)),
  inference(forward_demodulation,[],[f84,f108])).
fof(f84,plain,(
  ( ! [X3] : (sk_c7 != multiply(X3,sk_c6) | sk_c7 != inverse(X3)) ) | ~spl0_10),
  inference(avatar_component_clause,[],[f83])).
fof(f83,plain,(
  spl0_10 <=> ! [X3] : (sk_c7 != multiply(X3,sk_c6) | sk_c7 != inverse(X3))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f632,plain,(
  ~spl0_1 | ~spl0_6 | ~spl0_12 | ~spl0_14),
  inference(avatar_contradiction_clause,[],[f631])).
fof(f631,plain,(
  $false | (~spl0_1 | ~spl0_6 | ~spl0_12 | ~spl0_14)),
  inference(subsumption_resolution,[],[f630,f29])).
fof(f630,plain,(
  inverse(sk_c1) != sk_c7 | (~spl0_6 | ~spl0_12 | ~spl0_14)),
  inference(trivial_inequality_removal,[],[f620])).
fof(f620,plain,(
  sk_c7 != sk_c7 | inverse(sk_c1) != sk_c7 | (~spl0_6 | ~spl0_12 | ~spl0_14)),
  inference(superposition,[],[f523,f554])).
fof(f523,plain,(
  ( ! [X7] : (sk_c7 != multiply(X7,sk_c7) | sk_c7 != inverse(X7)) ) | (~spl0_12 | ~spl0_14)),
  inference(forward_demodulation,[],[f521,f108])).
fof(f521,plain,(
  ( ! [X7] : (sk_c7 != multiply(X7,sk_c7) | sk_c6 != inverse(X7)) ) | (~spl0_12 | ~spl0_14)),
  inference(backward_demodulation,[],[f108,f90])).
fof(f90,plain,(
  ( ! [X7] : (sk_c6 != inverse(X7) | sk_c7 != multiply(X7,sk_c6)) ) | ~spl0_12),
  inference(avatar_component_clause,[],[f89])).
fof(f89,plain,(
  spl0_12 <=> ! [X7] : (sk_c6 != inverse(X7) | sk_c7 != multiply(X7,sk_c6))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_12])])).
fof(f520,plain,(
  ~spl0_12 | ~spl0_16 | ~spl0_17 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f519])).
fof(f519,plain,(
  $false | (~spl0_12 | ~spl0_16 | ~spl0_17 | ~spl0_20)),
  inference(subsumption_resolution,[],[f518,f420])).
fof(f420,plain,(
  identity = inverse(identity) | (~spl0_16 | ~spl0_20)),
  inference(backward_demodulation,[],[f117,f169])).
fof(f169,plain,(
  sk_c7 = inverse(sk_c7) | ~spl0_20),
  inference(avatar_component_clause,[],[f168])).
fof(f168,plain,(
  spl0_20 <=> sk_c7 = inverse(sk_c7)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_20])])).
fof(f117,plain,(
  identity = sk_c7 | ~spl0_16),
  inference(avatar_component_clause,[],[f116])).
fof(f116,plain,(
  spl0_16 <=> identity = sk_c7),
  introduced(avatar_definition,[new_symbols(naming,[spl0_16])])).
fof(f518,plain,(
  identity != inverse(identity) | (~spl0_12 | ~spl0_16 | ~spl0_17 | ~spl0_20)),
  inference(forward_demodulation,[],[f507,f420])).
fof(f507,plain,(
  identity != inverse(inverse(identity)) | (~spl0_12 | ~spl0_16 | ~spl0_17)),
  inference(trivial_inequality_removal,[],[f505])).
fof(f505,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_12 | ~spl0_16 | ~spl0_17)),
  inference(superposition,[],[f448,f2])).
fof(f448,plain,(
  ( ! [X7] : (identity != multiply(X7,identity) | identity != inverse(X7)) ) | (~spl0_12 | ~spl0_16 | ~spl0_17)),
  inference(forward_demodulation,[],[f447,f117])).
fof(f447,plain,(
  ( ! [X7] : (sk_c7 != multiply(X7,identity) | identity != inverse(X7)) ) | (~spl0_12 | ~spl0_16 | ~spl0_17)),
  inference(forward_demodulation,[],[f446,f427])).
fof(f427,plain,(
  identity = sk_c6 | (~spl0_16 | ~spl0_17)),
  inference(forward_demodulation,[],[f416,f1])).
fof(f416,plain,(
  sk_c6 = multiply(identity,identity) | (~spl0_16 | ~spl0_17)),
  inference(backward_demodulation,[],[f117,f124])).
fof(f446,plain,(
  ( ! [X7] : (identity != inverse(X7) | sk_c7 != multiply(X7,sk_c6)) ) | (~spl0_12 | ~spl0_16 | ~spl0_17)),
  inference(forward_demodulation,[],[f90,f427])).
fof(f515,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_12 | ~spl0_16 | ~spl0_17),
  inference(avatar_contradiction_clause,[],[f514])).
fof(f514,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_12 | ~spl0_16 | ~spl0_17)),
  inference(subsumption_resolution,[],[f509,f454])).
fof(f454,plain,(
  identity = inverse(sk_c4) | (~spl0_2 | ~spl0_16)),
  inference(forward_demodulation,[],[f33,f117])).
fof(f33,plain,(
  sk_c7 = inverse(sk_c4) | ~spl0_2),
  inference(avatar_component_clause,[],[f31])).
fof(f509,plain,(
  identity != inverse(sk_c4) | (~spl0_3 | ~spl0_12 | ~spl0_16 | ~spl0_17)),
  inference(trivial_inequality_removal,[],[f503])).
fof(f503,plain,(
  identity != identity | identity != inverse(sk_c4) | (~spl0_3 | ~spl0_12 | ~spl0_16 | ~spl0_17)),
  inference(superposition,[],[f448,f453])).
fof(f453,plain,(
  identity = multiply(sk_c4,identity) | (~spl0_3 | ~spl0_16 | ~spl0_17)),
  inference(forward_demodulation,[],[f452,f117])).
fof(f452,plain,(
  sk_c7 = multiply(sk_c4,identity) | (~spl0_3 | ~spl0_16 | ~spl0_17)),
  inference(forward_demodulation,[],[f38,f427])).
fof(f410,plain,(
  ~spl0_1 | ~spl0_6 | spl0_17),
  inference(avatar_contradiction_clause,[],[f409])).
fof(f409,plain,(
  $false | (~spl0_1 | ~spl0_6 | spl0_17)),
  inference(subsumption_resolution,[],[f371,f125])).
fof(f125,plain,(
  sk_c6 != multiply(sk_c7,sk_c7) | spl0_17),
  inference(avatar_component_clause,[],[f123])).
fof(f371,plain,(
  sk_c6 = multiply(sk_c7,sk_c7) | (~spl0_1 | ~spl0_6)),
  inference(forward_demodulation,[],[f369,f29])).
fof(f369,plain,(
  sk_c6 = multiply(inverse(sk_c1),sk_c7) | ~spl0_6),
  inference(superposition,[],[f144,f53])).
fof(f302,plain,(
  ~spl0_4 | ~spl0_5 | ~spl0_14 | ~spl0_16 | spl0_20),
  inference(avatar_contradiction_clause,[],[f301])).
fof(f301,plain,(
  $false | (~spl0_4 | ~spl0_5 | ~spl0_14 | ~spl0_16 | spl0_20)),
  inference(subsumption_resolution,[],[f300,f222])).
fof(f222,plain,(
  identity != inverse(identity) | (~spl0_16 | spl0_20)),
  inference(backward_demodulation,[],[f117,f170])).
fof(f170,plain,(
  sk_c7 != inverse(sk_c7) | spl0_20),
  inference(avatar_component_clause,[],[f168])).
fof(f300,plain,(
  identity = inverse(identity) | (~spl0_4 | ~spl0_5 | ~spl0_14 | ~spl0_16)),
  inference(forward_demodulation,[],[f299,f278])).
fof(f278,plain,(
  identity = sk_c6 | (~spl0_14 | ~spl0_16)),
  inference(forward_demodulation,[],[f108,f117])).
fof(f299,plain,(
  sk_c6 = inverse(identity) | (~spl0_4 | ~spl0_5 | ~spl0_14 | ~spl0_16)),
  inference(forward_demodulation,[],[f48,f285])).
fof(f285,plain,(
  identity = sk_c5 | (~spl0_4 | ~spl0_5 | ~spl0_14 | ~spl0_16)),
  inference(backward_demodulation,[],[f284,f270])).
fof(f270,plain,(
  sk_c5 = multiply(sk_c5,identity) | (~spl0_4 | ~spl0_5 | ~spl0_16)),
  inference(forward_demodulation,[],[f224,f1])).
fof(f224,plain,(
  multiply(sk_c5,identity) = multiply(identity,sk_c5) | (~spl0_4 | ~spl0_5 | ~spl0_16)),
  inference(backward_demodulation,[],[f117,f182])).
fof(f182,plain,(
  multiply(sk_c7,sk_c5) = multiply(sk_c5,identity) | (~spl0_4 | ~spl0_5)),
  inference(superposition,[],[f141,f93])).
fof(f93,plain,(
  identity = multiply(sk_c6,sk_c5) | ~spl0_5),
  inference(superposition,[],[f2,f48])).
fof(f141,plain,(
  ( ! [X10] : (multiply(sk_c5,multiply(sk_c6,X10)) = multiply(sk_c7,X10)) ) | ~spl0_4),
  inference(superposition,[],[f3,f43])).
fof(f43,plain,(
  sk_c7 = multiply(sk_c5,sk_c6) | ~spl0_4),
  inference(avatar_component_clause,[],[f41])).
fof(f41,plain,(
  spl0_4 <=> sk_c7 = multiply(sk_c5,sk_c6)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f284,plain,(
  ( ! [X10] : (multiply(sk_c5,X10) = X10) ) | (~spl0_4 | ~spl0_14 | ~spl0_16)),
  inference(forward_demodulation,[],[f283,f1])).
fof(f283,plain,(
  ( ! [X10] : (multiply(sk_c5,multiply(identity,X10)) = X10) ) | (~spl0_4 | ~spl0_14 | ~spl0_16)),
  inference(forward_demodulation,[],[f250,f278])).
fof(f250,plain,(
  ( ! [X10] : (multiply(sk_c5,multiply(sk_c6,X10)) = X10) ) | (~spl0_4 | ~spl0_16)),
  inference(forward_demodulation,[],[f215,f1])).
fof(f215,plain,(
  ( ! [X10] : (multiply(sk_c5,multiply(sk_c6,X10)) = multiply(identity,X10)) ) | (~spl0_4 | ~spl0_16)),
  inference(backward_demodulation,[],[f117,f141])).
fof(f48,plain,(
  sk_c6 = inverse(sk_c5) | ~spl0_5),
  inference(avatar_component_clause,[],[f46])).
fof(f46,plain,(
  spl0_5 <=> sk_c6 = inverse(sk_c5)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f263,plain,(
  ~spl0_2 | spl0_13 | ~spl0_16),
  inference(avatar_contradiction_clause,[],[f262])).
fof(f262,plain,(
  $false | (~spl0_2 | spl0_13 | ~spl0_16)),
  inference(subsumption_resolution,[],[f261,f211])).
fof(f211,plain,(
  identity != inverse(identity) | (spl0_13 | ~spl0_16)),
  inference(backward_demodulation,[],[f117,f105])).
fof(f105,plain,(
  sk_c7 != inverse(identity) | spl0_13),
  inference(avatar_component_clause,[],[f103])).
fof(f103,plain,(
  spl0_13 <=> sk_c7 = inverse(identity)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_13])])).
fof(f261,plain,(
  identity = inverse(identity) | (~spl0_2 | ~spl0_16)),
  inference(backward_demodulation,[],[f257,f206])).
fof(f206,plain,(
  identity = inverse(sk_c4) | (~spl0_2 | ~spl0_16)),
  inference(backward_demodulation,[],[f117,f33])).
fof(f257,plain,(
  identity = sk_c4 | (~spl0_2 | ~spl0_16)),
  inference(forward_demodulation,[],[f218,f2])).
fof(f218,plain,(
  sk_c4 = multiply(inverse(identity),identity) | (~spl0_2 | ~spl0_16)),
  inference(backward_demodulation,[],[f117,f155])).
fof(f155,plain,(
  sk_c4 = multiply(inverse(sk_c7),identity) | ~spl0_2),
  inference(superposition,[],[f144,f92])).
fof(f92,plain,(
  identity = multiply(sk_c7,sk_c4) | ~spl0_2),
  inference(superposition,[],[f2,f33])).
fof(f205,plain,(
  spl0_16 | ~spl0_4 | ~spl0_5),
  inference(avatar_split_clause,[],[f204,f46,f41,f116])).
fof(f204,plain,(
  identity = sk_c7 | (~spl0_4 | ~spl0_5)),
  inference(forward_demodulation,[],[f191,f2])).
fof(f191,plain,(
  sk_c7 = multiply(inverse(sk_c6),sk_c6) | (~spl0_4 | ~spl0_5)),
  inference(superposition,[],[f144,f162])).
fof(f162,plain,(
  sk_c6 = multiply(sk_c6,sk_c7) | (~spl0_4 | ~spl0_5)),
  inference(forward_demodulation,[],[f157,f48])).
fof(f157,plain,(
  sk_c6 = multiply(inverse(sk_c5),sk_c7) | ~spl0_4),
  inference(superposition,[],[f144,f43])).
fof(f200,plain,(
  spl0_14 | ~spl0_4 | ~spl0_5 | ~spl0_17),
  inference(avatar_split_clause,[],[f199,f123,f46,f41,f107])).
fof(f199,plain,(
  sk_c7 = sk_c6 | (~spl0_4 | ~spl0_5 | ~spl0_17)),
  inference(forward_demodulation,[],[f198,f43])).
fof(f198,plain,(
  sk_c6 = multiply(sk_c5,sk_c6) | (~spl0_4 | ~spl0_5 | ~spl0_17)),
  inference(forward_demodulation,[],[f189,f124])).
fof(f189,plain,(
  multiply(sk_c5,sk_c6) = multiply(sk_c7,sk_c7) | (~spl0_4 | ~spl0_5)),
  inference(superposition,[],[f141,f162])).
fof(f161,plain,(
  spl0_17 | ~spl0_2 | ~spl0_3),
  inference(avatar_split_clause,[],[f160,f36,f31,f123])).
fof(f160,plain,(
  sk_c6 = multiply(sk_c7,sk_c7) | (~spl0_2 | ~spl0_3)),
  inference(forward_demodulation,[],[f156,f33])).
fof(f156,plain,(
  sk_c6 = multiply(inverse(sk_c4),sk_c7) | ~spl0_3),
  inference(superposition,[],[f144,f38])).
fof(f126,plain,(
  ~spl0_13 | ~spl0_17 | ~spl0_11),
  inference(avatar_split_clause,[],[f120,f86,f123,f103])).
fof(f120,plain,(
  sk_c6 != multiply(sk_c7,sk_c7) | sk_c7 != inverse(identity) | ~spl0_11),
  inference(superposition,[],[f87,f1])).
fof(f101,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_10),
  inference(avatar_contradiction_clause,[],[f100])).
fof(f100,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_10)),
  inference(subsumption_resolution,[],[f99,f33])).
fof(f99,plain,(
  sk_c7 != inverse(sk_c4) | (~spl0_3 | ~spl0_10)),
  inference(trivial_inequality_removal,[],[f94])).
fof(f94,plain,(
  sk_c7 != sk_c7 | sk_c7 != inverse(sk_c4) | (~spl0_3 | ~spl0_10)),
  inference(superposition,[],[f84,f38])).
fof(f91,plain,(
  spl0_10 | spl0_11 | spl0_10 | spl0_12),
  inference(avatar_split_clause,[],[f25,f89,f83,f86,f83])).
fof(f25,plain,(
  ( ! [X6,X7,X5,X3] : (sk_c6 != inverse(X7) | sk_c7 != multiply(X7,sk_c6) | sk_c7 != multiply(X6,sk_c6) | sk_c7 != inverse(X6) | sk_c7 != inverse(X5) | sk_c6 != multiply(sk_c7,multiply(X5,sk_c7)) | sk_c7 != multiply(X3,sk_c6) | sk_c7 != inverse(X3)) )),
  inference(equality_resolution,[],[f24])).
fof(f24,axiom,(
  ( ! [X6,X4,X7,X5,X3] : (sk_c6 != inverse(X7) | sk_c7 != multiply(X7,sk_c6) | sk_c7 != multiply(X6,sk_c6) | sk_c7 != inverse(X6) | sk_c7 != inverse(X5) | multiply(X5,sk_c7) != X4 | sk_c6 != multiply(sk_c7,X4) | sk_c7 != multiply(X3,sk_c6) | sk_c7 != inverse(X3)) )),
  file('./p/GRP226-1.tptp',unknown)).
fof(f81,plain,(
  spl0_9 | spl0_5),
  inference(avatar_split_clause,[],[f23,f46,f75])).
fof(f23,axiom,(
  sk_c6 = inverse(sk_c5) | sk_c7 = inverse(sk_c2)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f80,plain,(
  spl0_9 | spl0_4),
  inference(avatar_split_clause,[],[f22,f41,f75])).
fof(f22,axiom,(
  sk_c7 = multiply(sk_c5,sk_c6) | sk_c7 = inverse(sk_c2)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f79,plain,(
  spl0_9 | spl0_3),
  inference(avatar_split_clause,[],[f21,f36,f75])).
fof(f21,axiom,(
  sk_c7 = multiply(sk_c4,sk_c6) | sk_c7 = inverse(sk_c2)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f78,plain,(
  spl0_9 | spl0_2),
  inference(avatar_split_clause,[],[f20,f31,f75])).
fof(f20,axiom,(
  sk_c7 = inverse(sk_c4) | sk_c7 = inverse(sk_c2)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f73,plain,(
  spl0_8 | spl0_5),
  inference(avatar_split_clause,[],[f19,f46,f67])).
fof(f19,axiom,(
  sk_c6 = inverse(sk_c5) | sk_c3 = multiply(sk_c2,sk_c7)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f72,plain,(
  spl0_8 | spl0_4),
  inference(avatar_split_clause,[],[f18,f41,f67])).
fof(f18,axiom,(
  sk_c7 = multiply(sk_c5,sk_c6) | sk_c3 = multiply(sk_c2,sk_c7)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f71,plain,(
  spl0_8 | spl0_3),
  inference(avatar_split_clause,[],[f17,f36,f67])).
fof(f17,axiom,(
  sk_c7 = multiply(sk_c4,sk_c6) | sk_c3 = multiply(sk_c2,sk_c7)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f70,plain,(
  spl0_8 | spl0_2),
  inference(avatar_split_clause,[],[f16,f31,f67])).
fof(f16,axiom,(
  sk_c7 = inverse(sk_c4) | sk_c3 = multiply(sk_c2,sk_c7)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f65,plain,(
  spl0_7 | spl0_5),
  inference(avatar_split_clause,[],[f15,f46,f59])).
fof(f15,axiom,(
  sk_c6 = inverse(sk_c5) | sk_c6 = multiply(sk_c7,sk_c3)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f64,plain,(
  spl0_7 | spl0_4),
  inference(avatar_split_clause,[],[f14,f41,f59])).
fof(f14,axiom,(
  sk_c7 = multiply(sk_c5,sk_c6) | sk_c6 = multiply(sk_c7,sk_c3)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f63,plain,(
  spl0_7 | spl0_3),
  inference(avatar_split_clause,[],[f13,f36,f59])).
fof(f13,axiom,(
  sk_c7 = multiply(sk_c4,sk_c6) | sk_c6 = multiply(sk_c7,sk_c3)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f62,plain,(
  spl0_7 | spl0_2),
  inference(avatar_split_clause,[],[f12,f31,f59])).
fof(f12,axiom,(
  sk_c7 = inverse(sk_c4) | sk_c6 = multiply(sk_c7,sk_c3)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f57,plain,(
  spl0_6 | spl0_5),
  inference(avatar_split_clause,[],[f11,f46,f51])).
fof(f11,axiom,(
  sk_c6 = inverse(sk_c5) | sk_c7 = multiply(sk_c1,sk_c6)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f56,plain,(
  spl0_6 | spl0_4),
  inference(avatar_split_clause,[],[f10,f41,f51])).
fof(f10,axiom,(
  sk_c7 = multiply(sk_c5,sk_c6) | sk_c7 = multiply(sk_c1,sk_c6)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f55,plain,(
  spl0_6 | spl0_3),
  inference(avatar_split_clause,[],[f9,f36,f51])).
fof(f9,axiom,(
  sk_c7 = multiply(sk_c4,sk_c6) | sk_c7 = multiply(sk_c1,sk_c6)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f54,plain,(
  spl0_6 | spl0_2),
  inference(avatar_split_clause,[],[f8,f31,f51])).
fof(f8,axiom,(
  sk_c7 = inverse(sk_c4) | sk_c7 = multiply(sk_c1,sk_c6)),
  file('./p/GRP226-1.tptp',unknown)).
fof(f49,plain,(
  spl0_1 | spl0_5),
  inference(avatar_split_clause,[],[f7,f46,f27])).
fof(f7,axiom,(
  sk_c6 = inverse(sk_c5) | inverse(sk_c1) = sk_c7),
  file('./p/GRP226-1.tptp',unknown)).
fof(f44,plain,(
  spl0_1 | spl0_4),
  inference(avatar_split_clause,[],[f6,f41,f27])).
fof(f6,axiom,(
  sk_c7 = multiply(sk_c5,sk_c6) | inverse(sk_c1) = sk_c7),
  file('./p/GRP226-1.tptp',unknown)).
fof(f39,plain,(
  spl0_1 | spl0_3),
  inference(avatar_split_clause,[],[f5,f36,f27])).
fof(f5,axiom,(
  sk_c7 = multiply(sk_c4,sk_c6) | inverse(sk_c1) = sk_c7),
  file('./p/GRP226-1.tptp',unknown)).
fof(f34,plain,(
  spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f4,f31,f27])).
fof(f4,axiom,(
  sk_c7 = inverse(sk_c4) | inverse(sk_c1) = sk_c7),
  file('./p/GRP226-1.tptp',unknown)).
% SZS output end Proof for GRP226-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
