% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GRP212-1
% SZS output start Proof for GRP212-1
fof(f1003,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f45,f50,f55,f60,f65,f70,f75,f76,f77,f78,f79,f80,f85,f86,f87,f88,f89,f90,f95,f96,f97,f98,f99,f100,f105,f106,f107,f108,f109,f110,f120,f150,f156,f171,f213,f225,f242,f267,f281,f398,f429,f476,f847,f917,f923,f942,f949,f955,f966,f1002])).
fof(f1002,plain,(
  ~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_13 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f1001])).
fof(f1001,plain,(
  $false | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_13 | ~spl0_20)),
  inference(subsumption_resolution,[],[f1000,f909])).
fof(f909,plain,(
  identity = inverse(identity) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f908,f895])).
fof(f895,plain,(
  sk_c2 = inverse(identity) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f894,f74])).
fof(f74,plain,(
  sk_c2 = inverse(sk_c1) | ~spl0_8),
  inference(avatar_component_clause,[],[f72])).
fof(f72,plain,(
  spl0_8 <=> sk_c2 = inverse(sk_c1)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_8])])).
fof(f894,plain,(
  identity = sk_c1 | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f893,f881])).
fof(f881,plain,(
  sk_c1 = sk_c7 | (~spl0_8 | ~spl0_9 | ~spl0_20)),
  inference(forward_demodulation,[],[f864,f815])).
fof(f815,plain,(
  sk_c1 = multiply(inverse(sk_c2),identity) | ~spl0_8),
  inference(superposition,[],[f133,f343])).
fof(f343,plain,(
  identity = multiply(sk_c2,sk_c1) | ~spl0_8),
  inference(superposition,[],[f2,f74])).
fof(f2,axiom,(
  ( ! [X0] : (identity = multiply(inverse(X0),X0)) )),
  file('./p/GRP212-1.tptp',unknown)).
fof(f133,plain,(
  ( ! [X6,X7] : (multiply(inverse(X6),multiply(X6,X7)) = X7) )),
  inference(forward_demodulation,[],[f126,f1])).
fof(f1,axiom,(
  ( ! [X0] : (multiply(identity,X0) = X0) )),
  file('./p/GRP212-1.tptp',unknown)).
fof(f126,plain,(
  ( ! [X6,X7] : (multiply(inverse(X6),multiply(X6,X7)) = multiply(identity,X7)) )),
  inference(superposition,[],[f3,f2])).
fof(f3,axiom,(
  ( ! [X2,X0,X1] : (multiply(multiply(X0,X1),X2) = multiply(X0,multiply(X1,X2))) )),
  file('./p/GRP212-1.tptp',unknown)).
fof(f864,plain,(
  sk_c7 = multiply(inverse(sk_c2),identity) | (~spl0_9 | ~spl0_20)),
  inference(backward_demodulation,[],[f169,f452])).
fof(f452,plain,(
  sk_c7 = multiply(inverse(sk_c2),sk_c8) | ~spl0_9),
  inference(superposition,[],[f133,f84])).
fof(f84,plain,(
  sk_c8 = multiply(sk_c2,sk_c7) | ~spl0_9),
  inference(avatar_component_clause,[],[f82])).
fof(f82,plain,(
  spl0_9 <=> sk_c8 = multiply(sk_c2,sk_c7)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_9])])).
fof(f169,plain,(
  identity = sk_c8 | ~spl0_20),
  inference(avatar_component_clause,[],[f168])).
fof(f168,plain,(
  spl0_20 <=> identity = sk_c8),
  introduced(avatar_definition,[new_symbols(naming,[spl0_20])])).
fof(f893,plain,(
  identity = sk_c7 | (~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(forward_demodulation,[],[f867,f1])).
fof(f867,plain,(
  sk_c7 = multiply(identity,identity) | (~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f169,f456])).
fof(f456,plain,(
  sk_c7 = multiply(sk_c8,sk_c8) | (~spl0_10 | ~spl0_11)),
  inference(forward_demodulation,[],[f454,f94])).
fof(f94,plain,(
  sk_c8 = inverse(sk_c3) | ~spl0_10),
  inference(avatar_component_clause,[],[f92])).
fof(f92,plain,(
  spl0_10 <=> sk_c8 = inverse(sk_c3)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f454,plain,(
  sk_c7 = multiply(inverse(sk_c3),sk_c8) | ~spl0_11),
  inference(superposition,[],[f133,f104])).
fof(f104,plain,(
  sk_c8 = multiply(sk_c3,sk_c7) | ~spl0_11),
  inference(avatar_component_clause,[],[f102])).
fof(f102,plain,(
  spl0_11 <=> sk_c8 = multiply(sk_c3,sk_c7)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_11])])).
fof(f908,plain,(
  identity = sk_c2 | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f907,f863])).
fof(f863,plain,(
  sk_c2 = multiply(sk_c2,identity) | (~spl0_1 | ~spl0_8 | ~spl0_20)),
  inference(backward_demodulation,[],[f169,f443])).
fof(f443,plain,(
  sk_c2 = multiply(sk_c2,sk_c8) | (~spl0_1 | ~spl0_8)),
  inference(forward_demodulation,[],[f441,f74])).
fof(f441,plain,(
  sk_c2 = multiply(inverse(sk_c1),sk_c8) | ~spl0_1),
  inference(superposition,[],[f133,f40])).
fof(f40,plain,(
  multiply(sk_c1,sk_c2) = sk_c8 | ~spl0_1),
  inference(avatar_component_clause,[],[f38])).
fof(f38,plain,(
  spl0_1 <=> multiply(sk_c1,sk_c2) = sk_c8),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f907,plain,(
  ( ! [X0] : (multiply(sk_c2,X0) = X0) ) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f895,f177])).
fof(f177,plain,(
  ( ! [X0] : (multiply(inverse(identity),X0) = X0) )),
  inference(superposition,[],[f133,f1])).
fof(f1000,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f996,f909])).
fof(f996,plain,(
  identity != inverse(inverse(identity)) | (~spl0_10 | ~spl0_11 | ~spl0_13 | ~spl0_20)),
  inference(trivial_inequality_removal,[],[f994])).
fof(f994,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_10 | ~spl0_11 | ~spl0_13 | ~spl0_20)),
  inference(superposition,[],[f980,f2])).
fof(f980,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | identity != inverse(X5)) ) | (~spl0_10 | ~spl0_11 | ~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f979,f169])).
fof(f979,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | sk_c8 != inverse(X5)) ) | (~spl0_10 | ~spl0_11 | ~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f978,f169])).
fof(f978,plain,(
  ( ! [X5] : (sk_c8 != multiply(X5,identity) | sk_c8 != inverse(X5)) ) | (~spl0_10 | ~spl0_11 | ~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f116,f893])).
fof(f116,plain,(
  ( ! [X5] : (sk_c8 != multiply(X5,sk_c7) | sk_c8 != inverse(X5)) ) | ~spl0_13),
  inference(avatar_component_clause,[],[f115])).
fof(f115,plain,(
  spl0_13 <=> ! [X5] : (sk_c8 != multiply(X5,sk_c7) | sk_c8 != inverse(X5))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_13])])).
fof(f966,plain,(
  ~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20 | spl0_22),
  inference(avatar_contradiction_clause,[],[f965])).
fof(f965,plain,(
  $false | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20 | spl0_22)),
  inference(subsumption_resolution,[],[f964,f169])).
fof(f964,plain,(
  identity != sk_c8 | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20 | spl0_22)),
  inference(forward_demodulation,[],[f963,f909])).
fof(f963,plain,(
  sk_c8 != inverse(identity) | (~spl0_10 | ~spl0_11 | ~spl0_20 | spl0_22)),
  inference(forward_demodulation,[],[f474,f893])).
fof(f474,plain,(
  sk_c8 != inverse(sk_c7) | spl0_22),
  inference(avatar_component_clause,[],[f472])).
fof(f472,plain,(
  spl0_22 <=> sk_c8 = inverse(sk_c7)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_22])])).
fof(f955,plain,(
  ~spl0_10 | ~spl0_11 | spl0_15 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f954])).
fof(f954,plain,(
  $false | (~spl0_10 | ~spl0_11 | spl0_15 | ~spl0_20)),
  inference(subsumption_resolution,[],[f953,f169])).
fof(f953,plain,(
  identity != sk_c8 | (~spl0_10 | ~spl0_11 | spl0_15 | ~spl0_20)),
  inference(forward_demodulation,[],[f952,f1])).
fof(f952,plain,(
  sk_c8 != multiply(identity,identity) | (~spl0_10 | ~spl0_11 | spl0_15 | ~spl0_20)),
  inference(forward_demodulation,[],[f145,f893])).
fof(f145,plain,(
  sk_c8 != multiply(sk_c7,sk_c7) | spl0_15),
  inference(avatar_component_clause,[],[f143])).
fof(f143,plain,(
  spl0_15 <=> sk_c8 = multiply(sk_c7,sk_c7)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_15])])).
fof(f949,plain,(
  ~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | spl0_19 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f948])).
fof(f948,plain,(
  $false | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | spl0_19 | ~spl0_20)),
  inference(subsumption_resolution,[],[f947,f169])).
fof(f947,plain,(
  identity != sk_c8 | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | spl0_19 | ~spl0_20)),
  inference(forward_demodulation,[],[f946,f1])).
fof(f946,plain,(
  sk_c8 != multiply(identity,identity) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | spl0_19 | ~spl0_20)),
  inference(forward_demodulation,[],[f945,f909])).
fof(f945,plain,(
  sk_c8 != multiply(identity,inverse(identity)) | (~spl0_10 | ~spl0_11 | spl0_19 | ~spl0_20)),
  inference(forward_demodulation,[],[f166,f893])).
fof(f166,plain,(
  sk_c8 != multiply(sk_c7,inverse(sk_c7)) | spl0_19),
  inference(avatar_component_clause,[],[f164])).
fof(f164,plain,(
  spl0_19 <=> sk_c8 = multiply(sk_c7,inverse(sk_c7))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_19])])).
fof(f942,plain,(
  ~spl0_10 | ~spl0_11 | spl0_16 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f941])).
fof(f941,plain,(
  $false | (~spl0_10 | ~spl0_11 | spl0_16 | ~spl0_20)),
  inference(subsumption_resolution,[],[f940,f1])).
fof(f940,plain,(
  identity != multiply(identity,identity) | (~spl0_10 | ~spl0_11 | spl0_16 | ~spl0_20)),
  inference(forward_demodulation,[],[f939,f169])).
fof(f939,plain,(
  sk_c8 != multiply(sk_c8,identity) | (~spl0_10 | ~spl0_11 | spl0_16 | ~spl0_20)),
  inference(forward_demodulation,[],[f149,f893])).
fof(f149,plain,(
  sk_c8 != multiply(sk_c8,sk_c7) | spl0_16),
  inference(avatar_component_clause,[],[f147])).
fof(f147,plain,(
  spl0_16 <=> sk_c8 = multiply(sk_c8,sk_c7)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_16])])).
fof(f923,plain,(
  ~spl0_10 | ~spl0_11 | ~spl0_20 | spl0_21),
  inference(avatar_contradiction_clause,[],[f922])).
fof(f922,plain,(
  $false | (~spl0_10 | ~spl0_11 | ~spl0_20 | spl0_21)),
  inference(subsumption_resolution,[],[f921,f893])).
fof(f921,plain,(
  identity != sk_c7 | (~spl0_20 | spl0_21)),
  inference(forward_demodulation,[],[f868,f1])).
fof(f868,plain,(
  sk_c7 != multiply(identity,identity) | (~spl0_20 | spl0_21)),
  inference(backward_demodulation,[],[f169,f470])).
fof(f470,plain,(
  sk_c7 != multiply(sk_c8,identity) | spl0_21),
  inference(avatar_component_clause,[],[f468])).
fof(f468,plain,(
  spl0_21 <=> sk_c7 = multiply(sk_c8,identity)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_21])])).
fof(f917,plain,(
  ~spl0_1 | spl0_2 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f916])).
fof(f916,plain,(
  $false | (~spl0_1 | spl0_2 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(subsumption_resolution,[],[f900,f909])).
fof(f900,plain,(
  identity != inverse(identity) | (spl0_2 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f894,f882])).
fof(f882,plain,(
  sk_c1 != inverse(identity) | (spl0_2 | ~spl0_8 | ~spl0_9 | ~spl0_20)),
  inference(backward_demodulation,[],[f881,f849])).
fof(f849,plain,(
  sk_c7 != inverse(identity) | (spl0_2 | ~spl0_20)),
  inference(backward_demodulation,[],[f169,f43])).
fof(f43,plain,(
  inverse(sk_c8) != sk_c7 | spl0_2),
  inference(avatar_component_clause,[],[f42])).
fof(f42,plain,(
  spl0_2 <=> inverse(sk_c8) = sk_c7),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f847,plain,(
  spl0_20 | ~spl0_1 | ~spl0_8),
  inference(avatar_split_clause,[],[f842,f72,f38,f168])).
fof(f842,plain,(
  identity = sk_c8 | (~spl0_1 | ~spl0_8)),
  inference(forward_demodulation,[],[f840,f2])).
fof(f840,plain,(
  sk_c8 = multiply(inverse(sk_c2),sk_c2) | (~spl0_1 | ~spl0_8)),
  inference(superposition,[],[f133,f443])).
fof(f476,plain,(
  ~spl0_22 | ~spl0_21 | ~spl0_2 | ~spl0_14),
  inference(avatar_split_clause,[],[f463,f118,f42,f468,f472])).
fof(f118,plain,(
  spl0_14 <=> ! [X8] : (sk_c8 != inverse(X8) | sk_c7 != multiply(sk_c8,multiply(X8,sk_c8)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_14])])).
fof(f463,plain,(
  sk_c7 != multiply(sk_c8,identity) | sk_c8 != inverse(sk_c7) | (~spl0_2 | ~spl0_14)),
  inference(superposition,[],[f119,f121])).
fof(f121,plain,(
  identity = multiply(sk_c7,sk_c8) | ~spl0_2),
  inference(superposition,[],[f2,f44])).
fof(f44,plain,(
  inverse(sk_c8) = sk_c7 | ~spl0_2),
  inference(avatar_component_clause,[],[f42])).
fof(f119,plain,(
  ( ! [X8] : (sk_c7 != multiply(sk_c8,multiply(X8,sk_c8)) | sk_c8 != inverse(X8)) ) | ~spl0_14),
  inference(avatar_component_clause,[],[f118])).
fof(f429,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_14 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f428])).
fof(f428,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_14 | ~spl0_20)),
  inference(subsumption_resolution,[],[f427,f292])).
fof(f292,plain,(
  identity = inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_20)),
  inference(backward_demodulation,[],[f290,f284])).
fof(f284,plain,(
  identity = inverse(sk_c4) | (~spl0_3 | ~spl0_20)),
  inference(forward_demodulation,[],[f49,f169])).
fof(f49,plain,(
  sk_c8 = inverse(sk_c4) | ~spl0_3),
  inference(avatar_component_clause,[],[f47])).
fof(f47,plain,(
  spl0_3 <=> sk_c8 = inverse(sk_c4)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f290,plain,(
  identity = sk_c4 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_20)),
  inference(forward_demodulation,[],[f289,f1])).
fof(f289,plain,(
  sk_c4 = multiply(identity,identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_20)),
  inference(forward_demodulation,[],[f189,f287])).
fof(f287,plain,(
  identity = sk_c7 | (~spl0_3 | ~spl0_4 | ~spl0_20)),
  inference(forward_demodulation,[],[f286,f1])).
fof(f286,plain,(
  sk_c7 = multiply(identity,identity) | (~spl0_3 | ~spl0_4 | ~spl0_20)),
  inference(forward_demodulation,[],[f197,f169])).
fof(f197,plain,(
  sk_c7 = multiply(sk_c8,sk_c8) | (~spl0_3 | ~spl0_4)),
  inference(forward_demodulation,[],[f185,f49])).
fof(f185,plain,(
  sk_c7 = multiply(inverse(sk_c4),sk_c8) | ~spl0_4),
  inference(superposition,[],[f133,f54])).
fof(f54,plain,(
  sk_c8 = multiply(sk_c4,sk_c7) | ~spl0_4),
  inference(avatar_component_clause,[],[f52])).
fof(f52,plain,(
  spl0_4 <=> sk_c8 = multiply(sk_c4,sk_c7)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f189,plain,(
  sk_c4 = multiply(sk_c7,identity) | (~spl0_2 | ~spl0_3)),
  inference(forward_demodulation,[],[f182,f44])).
fof(f182,plain,(
  sk_c4 = multiply(inverse(sk_c8),identity) | ~spl0_3),
  inference(superposition,[],[f133,f122])).
fof(f122,plain,(
  identity = multiply(sk_c8,sk_c4) | ~spl0_3),
  inference(superposition,[],[f2,f49])).
fof(f427,plain,(
  identity != inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_14 | ~spl0_20)),
  inference(forward_demodulation,[],[f426,f292])).
fof(f426,plain,(
  identity != inverse(inverse(identity)) | (~spl0_3 | ~spl0_4 | ~spl0_14 | ~spl0_20)),
  inference(subsumption_resolution,[],[f418,f1])).
fof(f418,plain,(
  identity != multiply(identity,identity) | identity != inverse(inverse(identity)) | (~spl0_3 | ~spl0_4 | ~spl0_14 | ~spl0_20)),
  inference(superposition,[],[f405,f2])).
fof(f405,plain,(
  ( ! [X8] : (identity != multiply(identity,multiply(X8,identity)) | identity != inverse(X8)) ) | (~spl0_3 | ~spl0_4 | ~spl0_14 | ~spl0_20)),
  inference(forward_demodulation,[],[f404,f287])).
fof(f404,plain,(
  ( ! [X8] : (sk_c7 != multiply(identity,multiply(X8,identity)) | identity != inverse(X8)) ) | (~spl0_14 | ~spl0_20)),
  inference(forward_demodulation,[],[f403,f169])).
fof(f403,plain,(
  ( ! [X8] : (identity != inverse(X8) | sk_c7 != multiply(sk_c8,multiply(X8,sk_c8))) ) | (~spl0_14 | ~spl0_20)),
  inference(forward_demodulation,[],[f119,f169])).
fof(f398,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_13 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f397])).
fof(f397,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_13 | ~spl0_20)),
  inference(subsumption_resolution,[],[f396,f292])).
fof(f396,plain,(
  identity != inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f392,f292])).
fof(f392,plain,(
  identity != inverse(inverse(identity)) | (~spl0_3 | ~spl0_4 | ~spl0_13 | ~spl0_20)),
  inference(trivial_inequality_removal,[],[f383])).
fof(f383,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_3 | ~spl0_4 | ~spl0_13 | ~spl0_20)),
  inference(superposition,[],[f334,f2])).
fof(f334,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | identity != inverse(X5)) ) | (~spl0_3 | ~spl0_4 | ~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f333,f169])).
fof(f333,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | sk_c8 != inverse(X5)) ) | (~spl0_3 | ~spl0_4 | ~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f332,f169])).
fof(f332,plain,(
  ( ! [X5] : (sk_c8 != multiply(X5,identity) | sk_c8 != inverse(X5)) ) | (~spl0_3 | ~spl0_4 | ~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f116,f287])).
fof(f281,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_14 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f280])).
fof(f280,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_14 | ~spl0_20)),
  inference(subsumption_resolution,[],[f279,f244])).
fof(f244,plain,(
  identity = inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_20)),
  inference(backward_demodulation,[],[f169,f201])).
fof(f201,plain,(
  sk_c8 = inverse(sk_c8) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7)),
  inference(backward_demodulation,[],[f200,f44])).
fof(f200,plain,(
  sk_c8 = sk_c7 | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7)),
  inference(backward_demodulation,[],[f199,f59])).
fof(f59,plain,(
  sk_c7 = multiply(sk_c8,sk_c6) | ~spl0_5),
  inference(avatar_component_clause,[],[f57])).
fof(f57,plain,(
  spl0_5 <=> sk_c7 = multiply(sk_c8,sk_c6)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f199,plain,(
  sk_c8 = multiply(sk_c8,sk_c6) | (~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7)),
  inference(forward_demodulation,[],[f198,f49])).
fof(f198,plain,(
  sk_c8 = multiply(inverse(sk_c4),sk_c6) | (~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7)),
  inference(forward_demodulation,[],[f186,f191])).
fof(f191,plain,(
  sk_c4 = sk_c5 | (~spl0_2 | ~spl0_3 | ~spl0_7)),
  inference(forward_demodulation,[],[f190,f189])).
fof(f190,plain,(
  sk_c5 = multiply(sk_c7,identity) | (~spl0_2 | ~spl0_7)),
  inference(forward_demodulation,[],[f183,f44])).
fof(f183,plain,(
  sk_c5 = multiply(inverse(sk_c8),identity) | ~spl0_7),
  inference(superposition,[],[f133,f123])).
fof(f123,plain,(
  identity = multiply(sk_c8,sk_c5) | ~spl0_7),
  inference(superposition,[],[f2,f69])).
fof(f69,plain,(
  sk_c8 = inverse(sk_c5) | ~spl0_7),
  inference(avatar_component_clause,[],[f67])).
fof(f67,plain,(
  spl0_7 <=> sk_c8 = inverse(sk_c5)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).
fof(f186,plain,(
  sk_c8 = multiply(inverse(sk_c5),sk_c6) | ~spl0_6),
  inference(superposition,[],[f133,f64])).
fof(f64,plain,(
  sk_c6 = multiply(sk_c5,sk_c8) | ~spl0_6),
  inference(avatar_component_clause,[],[f62])).
fof(f62,plain,(
  spl0_6 <=> sk_c6 = multiply(sk_c5,sk_c8)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).
fof(f279,plain,(
  identity != inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_14 | ~spl0_20)),
  inference(forward_demodulation,[],[f278,f244])).
fof(f278,plain,(
  identity != inverse(inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_14 | ~spl0_20)),
  inference(subsumption_resolution,[],[f273,f1])).
fof(f273,plain,(
  identity != multiply(identity,identity) | identity != inverse(inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_14 | ~spl0_20)),
  inference(superposition,[],[f270,f2])).
fof(f270,plain,(
  ( ! [X8] : (identity != multiply(identity,multiply(X8,identity)) | identity != inverse(X8)) ) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_14 | ~spl0_20)),
  inference(forward_demodulation,[],[f269,f243])).
fof(f243,plain,(
  identity = sk_c7 | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_20)),
  inference(backward_demodulation,[],[f169,f200])).
fof(f269,plain,(
  ( ! [X8] : (sk_c7 != multiply(identity,multiply(X8,identity)) | identity != inverse(X8)) ) | (~spl0_14 | ~spl0_20)),
  inference(forward_demodulation,[],[f268,f169])).
fof(f268,plain,(
  ( ! [X8] : (identity != inverse(X8) | sk_c7 != multiply(sk_c8,multiply(X8,sk_c8))) ) | (~spl0_14 | ~spl0_20)),
  inference(forward_demodulation,[],[f119,f169])).
fof(f267,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_13 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f266])).
fof(f266,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_13 | ~spl0_20)),
  inference(subsumption_resolution,[],[f265,f244])).
fof(f265,plain,(
  identity != inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f261,f244])).
fof(f261,plain,(
  identity != inverse(inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_13 | ~spl0_20)),
  inference(trivial_inequality_removal,[],[f260])).
fof(f260,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_13 | ~spl0_20)),
  inference(superposition,[],[f254,f2])).
fof(f254,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | identity != inverse(X5)) ) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f253,f169])).
fof(f253,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | sk_c8 != inverse(X5)) ) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f252,f169])).
fof(f252,plain,(
  ( ! [X5] : (sk_c8 != multiply(X5,identity) | sk_c8 != inverse(X5)) ) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f116,f243])).
fof(f242,plain,(
  spl0_20 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7),
  inference(avatar_split_clause,[],[f227,f67,f62,f57,f52,f47,f42,f168])).
fof(f227,plain,(
  identity = sk_c8 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7)),
  inference(forward_demodulation,[],[f226,f216])).
fof(f216,plain,(
  ( ! [X8] : (multiply(sk_c8,X8) = X8) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7)),
  inference(backward_demodulation,[],[f207,f215])).
fof(f215,plain,(
  ( ! [X8] : (multiply(sk_c8,X8) = multiply(sk_c8,multiply(sk_c8,X8))) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7)),
  inference(backward_demodulation,[],[f214,f205])).
fof(f205,plain,(
  ( ! [X8] : (multiply(sk_c8,multiply(sk_c6,X8)) = multiply(sk_c8,X8)) ) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7)),
  inference(backward_demodulation,[],[f200,f127])).
fof(f127,plain,(
  ( ! [X8] : (multiply(sk_c8,multiply(sk_c6,X8)) = multiply(sk_c7,X8)) ) | ~spl0_5),
  inference(superposition,[],[f3,f59])).
fof(f214,plain,(
  ( ! [X11] : (multiply(sk_c8,X11) = multiply(sk_c6,X11)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7)),
  inference(backward_demodulation,[],[f206,f193])).
fof(f193,plain,(
  ( ! [X11] : (multiply(sk_c6,X11) = multiply(sk_c4,multiply(sk_c8,X11))) ) | (~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7)),
  inference(backward_demodulation,[],[f191,f130])).
fof(f130,plain,(
  ( ! [X11] : (multiply(sk_c5,multiply(sk_c8,X11)) = multiply(sk_c6,X11)) ) | ~spl0_6),
  inference(superposition,[],[f3,f64])).
fof(f206,plain,(
  ( ! [X10] : (multiply(sk_c8,X10) = multiply(sk_c4,multiply(sk_c8,X10))) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7)),
  inference(backward_demodulation,[],[f200,f129])).
fof(f129,plain,(
  ( ! [X10] : (multiply(sk_c4,multiply(sk_c7,X10)) = multiply(sk_c8,X10)) ) | ~spl0_4),
  inference(superposition,[],[f3,f54])).
fof(f207,plain,(
  ( ! [X9] : (multiply(sk_c8,multiply(sk_c8,X9)) = X9) ) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7)),
  inference(backward_demodulation,[],[f200,f134])).
fof(f134,plain,(
  ( ! [X9] : (multiply(sk_c7,multiply(sk_c8,X9)) = X9) ) | ~spl0_2),
  inference(forward_demodulation,[],[f128,f1])).
fof(f128,plain,(
  ( ! [X9] : (multiply(sk_c7,multiply(sk_c8,X9)) = multiply(identity,X9)) ) | ~spl0_2),
  inference(superposition,[],[f3,f121])).
fof(f226,plain,(
  sk_c8 = multiply(sk_c8,identity) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7)),
  inference(forward_demodulation,[],[f209,f201])).
fof(f209,plain,(
  sk_c8 = multiply(inverse(sk_c8),identity) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7)),
  inference(backward_demodulation,[],[f200,f184])).
fof(f184,plain,(
  sk_c8 = multiply(inverse(sk_c7),identity) | ~spl0_2),
  inference(superposition,[],[f133,f121])).
fof(f225,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7 | spl0_16),
  inference(avatar_contradiction_clause,[],[f224])).
fof(f224,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7 | spl0_16)),
  inference(subsumption_resolution,[],[f208,f216])).
fof(f208,plain,(
  sk_c8 != multiply(sk_c8,sk_c8) | (~spl0_2 | ~spl0_3 | ~spl0_5 | ~spl0_6 | ~spl0_7 | spl0_16)),
  inference(backward_demodulation,[],[f200,f149])).
fof(f213,plain,(
  spl0_17 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7),
  inference(avatar_split_clause,[],[f202,f67,f62,f57,f52,f47,f42,f153])).
fof(f153,plain,(
  spl0_17 <=> sk_c8 = multiply(sk_c4,sk_c8)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_17])])).
fof(f202,plain,(
  sk_c8 = multiply(sk_c4,sk_c8) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7)),
  inference(backward_demodulation,[],[f200,f54])).
fof(f171,plain,(
  ~spl0_19 | ~spl0_20 | ~spl0_12),
  inference(avatar_split_clause,[],[f140,f112,f168,f164])).
fof(f112,plain,(
  spl0_12 <=> ! [X3] : (sk_c8 != multiply(inverse(X3),sk_c7) | sk_c8 != multiply(X3,inverse(X3)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_12])])).
fof(f140,plain,(
  identity != sk_c8 | sk_c8 != multiply(sk_c7,inverse(sk_c7)) | ~spl0_12),
  inference(superposition,[],[f113,f2])).
fof(f113,plain,(
  ( ! [X3] : (sk_c8 != multiply(inverse(X3),sk_c7) | sk_c8 != multiply(X3,inverse(X3))) ) | ~spl0_12),
  inference(avatar_component_clause,[],[f112])).
fof(f156,plain,(
  ~spl0_16 | ~spl0_17 | ~spl0_3 | ~spl0_12),
  inference(avatar_split_clause,[],[f151,f112,f47,f153,f147])).
fof(f151,plain,(
  sk_c8 != multiply(sk_c4,sk_c8) | sk_c8 != multiply(sk_c8,sk_c7) | (~spl0_3 | ~spl0_12)),
  inference(forward_demodulation,[],[f138,f49])).
fof(f138,plain,(
  sk_c8 != multiply(sk_c8,sk_c7) | sk_c8 != multiply(sk_c4,inverse(sk_c4)) | (~spl0_3 | ~spl0_12)),
  inference(superposition,[],[f113,f49])).
fof(f150,plain,(
  ~spl0_15 | ~spl0_16 | ~spl0_2 | ~spl0_12),
  inference(avatar_split_clause,[],[f141,f112,f42,f147,f143])).
fof(f141,plain,(
  sk_c8 != multiply(sk_c8,sk_c7) | sk_c8 != multiply(sk_c7,sk_c7) | (~spl0_2 | ~spl0_12)),
  inference(forward_demodulation,[],[f137,f44])).
fof(f137,plain,(
  sk_c8 != multiply(sk_c7,sk_c7) | sk_c8 != multiply(sk_c8,inverse(sk_c8)) | (~spl0_2 | ~spl0_12)),
  inference(superposition,[],[f113,f44])).
fof(f120,plain,(
  spl0_12 | spl0_13 | ~spl0_2 | spl0_13 | spl0_14),
  inference(avatar_split_clause,[],[f36,f118,f115,f42,f115,f112])).
fof(f36,plain,(
  ( ! [X6,X8,X5,X3] : (sk_c8 != inverse(X8) | sk_c7 != multiply(sk_c8,multiply(X8,sk_c8)) | sk_c8 != multiply(X6,sk_c7) | sk_c8 != inverse(X6) | inverse(sk_c8) != sk_c7 | sk_c8 != multiply(X5,sk_c7) | sk_c8 != inverse(X5) | sk_c8 != multiply(inverse(X3),sk_c7) | sk_c8 != multiply(X3,inverse(X3))) )),
  inference(equality_resolution,[],[f35])).
fof(f35,plain,(
  ( ! [X6,X4,X8,X5,X3] : (sk_c8 != inverse(X8) | sk_c7 != multiply(sk_c8,multiply(X8,sk_c8)) | sk_c8 != multiply(X6,sk_c7) | sk_c8 != inverse(X6) | inverse(sk_c8) != sk_c7 | sk_c8 != multiply(X5,sk_c7) | sk_c8 != inverse(X5) | sk_c8 != multiply(X4,sk_c7) | inverse(X3) != X4 | sk_c8 != multiply(X3,X4)) )),
  inference(equality_resolution,[],[f34])).
fof(f34,axiom,(
  ( ! [X6,X4,X8,X7,X5,X3] : (sk_c8 != inverse(X8) | multiply(X8,sk_c8) != X7 | sk_c7 != multiply(sk_c8,X7) | sk_c8 != multiply(X6,sk_c7) | sk_c8 != inverse(X6) | inverse(sk_c8) != sk_c7 | sk_c8 != multiply(X5,sk_c7) | sk_c8 != inverse(X5) | sk_c8 != multiply(X4,sk_c7) | inverse(X3) != X4 | sk_c8 != multiply(X3,X4)) )),
  file('./p/GRP212-1.tptp',unknown)).
fof(f110,plain,(
  spl0_11 | spl0_7),
  inference(avatar_split_clause,[],[f33,f67,f102])).
fof(f33,axiom,(
  sk_c8 = inverse(sk_c5) | sk_c8 = multiply(sk_c3,sk_c7)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f109,plain,(
  spl0_11 | spl0_6),
  inference(avatar_split_clause,[],[f32,f62,f102])).
fof(f32,axiom,(
  sk_c6 = multiply(sk_c5,sk_c8) | sk_c8 = multiply(sk_c3,sk_c7)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f108,plain,(
  spl0_11 | spl0_5),
  inference(avatar_split_clause,[],[f31,f57,f102])).
fof(f31,axiom,(
  sk_c7 = multiply(sk_c8,sk_c6) | sk_c8 = multiply(sk_c3,sk_c7)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f107,plain,(
  spl0_11 | spl0_4),
  inference(avatar_split_clause,[],[f30,f52,f102])).
fof(f30,axiom,(
  sk_c8 = multiply(sk_c4,sk_c7) | sk_c8 = multiply(sk_c3,sk_c7)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f106,plain,(
  spl0_11 | spl0_3),
  inference(avatar_split_clause,[],[f29,f47,f102])).
fof(f29,axiom,(
  sk_c8 = inverse(sk_c4) | sk_c8 = multiply(sk_c3,sk_c7)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f105,plain,(
  spl0_11 | spl0_2),
  inference(avatar_split_clause,[],[f28,f42,f102])).
fof(f28,axiom,(
  inverse(sk_c8) = sk_c7 | sk_c8 = multiply(sk_c3,sk_c7)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f100,plain,(
  spl0_10 | spl0_7),
  inference(avatar_split_clause,[],[f27,f67,f92])).
fof(f27,axiom,(
  sk_c8 = inverse(sk_c5) | sk_c8 = inverse(sk_c3)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f99,plain,(
  spl0_10 | spl0_6),
  inference(avatar_split_clause,[],[f26,f62,f92])).
fof(f26,axiom,(
  sk_c6 = multiply(sk_c5,sk_c8) | sk_c8 = inverse(sk_c3)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f98,plain,(
  spl0_10 | spl0_5),
  inference(avatar_split_clause,[],[f25,f57,f92])).
fof(f25,axiom,(
  sk_c7 = multiply(sk_c8,sk_c6) | sk_c8 = inverse(sk_c3)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f97,plain,(
  spl0_10 | spl0_4),
  inference(avatar_split_clause,[],[f24,f52,f92])).
fof(f24,axiom,(
  sk_c8 = multiply(sk_c4,sk_c7) | sk_c8 = inverse(sk_c3)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f96,plain,(
  spl0_10 | spl0_3),
  inference(avatar_split_clause,[],[f23,f47,f92])).
fof(f23,axiom,(
  sk_c8 = inverse(sk_c4) | sk_c8 = inverse(sk_c3)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f95,plain,(
  spl0_10 | spl0_2),
  inference(avatar_split_clause,[],[f22,f42,f92])).
fof(f22,axiom,(
  inverse(sk_c8) = sk_c7 | sk_c8 = inverse(sk_c3)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f90,plain,(
  spl0_9 | spl0_7),
  inference(avatar_split_clause,[],[f21,f67,f82])).
fof(f21,axiom,(
  sk_c8 = inverse(sk_c5) | sk_c8 = multiply(sk_c2,sk_c7)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f89,plain,(
  spl0_9 | spl0_6),
  inference(avatar_split_clause,[],[f20,f62,f82])).
fof(f20,axiom,(
  sk_c6 = multiply(sk_c5,sk_c8) | sk_c8 = multiply(sk_c2,sk_c7)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f88,plain,(
  spl0_9 | spl0_5),
  inference(avatar_split_clause,[],[f19,f57,f82])).
fof(f19,axiom,(
  sk_c7 = multiply(sk_c8,sk_c6) | sk_c8 = multiply(sk_c2,sk_c7)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f87,plain,(
  spl0_9 | spl0_4),
  inference(avatar_split_clause,[],[f18,f52,f82])).
fof(f18,axiom,(
  sk_c8 = multiply(sk_c4,sk_c7) | sk_c8 = multiply(sk_c2,sk_c7)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f86,plain,(
  spl0_9 | spl0_3),
  inference(avatar_split_clause,[],[f17,f47,f82])).
fof(f17,axiom,(
  sk_c8 = inverse(sk_c4) | sk_c8 = multiply(sk_c2,sk_c7)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f85,plain,(
  spl0_9 | spl0_2),
  inference(avatar_split_clause,[],[f16,f42,f82])).
fof(f16,axiom,(
  inverse(sk_c8) = sk_c7 | sk_c8 = multiply(sk_c2,sk_c7)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f80,plain,(
  spl0_8 | spl0_7),
  inference(avatar_split_clause,[],[f15,f67,f72])).
fof(f15,axiom,(
  sk_c8 = inverse(sk_c5) | sk_c2 = inverse(sk_c1)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f79,plain,(
  spl0_8 | spl0_6),
  inference(avatar_split_clause,[],[f14,f62,f72])).
fof(f14,axiom,(
  sk_c6 = multiply(sk_c5,sk_c8) | sk_c2 = inverse(sk_c1)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f78,plain,(
  spl0_8 | spl0_5),
  inference(avatar_split_clause,[],[f13,f57,f72])).
fof(f13,axiom,(
  sk_c7 = multiply(sk_c8,sk_c6) | sk_c2 = inverse(sk_c1)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f77,plain,(
  spl0_8 | spl0_4),
  inference(avatar_split_clause,[],[f12,f52,f72])).
fof(f12,axiom,(
  sk_c8 = multiply(sk_c4,sk_c7) | sk_c2 = inverse(sk_c1)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f76,plain,(
  spl0_8 | spl0_3),
  inference(avatar_split_clause,[],[f11,f47,f72])).
fof(f11,axiom,(
  sk_c8 = inverse(sk_c4) | sk_c2 = inverse(sk_c1)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f75,plain,(
  spl0_8 | spl0_2),
  inference(avatar_split_clause,[],[f10,f42,f72])).
fof(f10,axiom,(
  inverse(sk_c8) = sk_c7 | sk_c2 = inverse(sk_c1)),
  file('./p/GRP212-1.tptp',unknown)).
fof(f70,plain,(
  spl0_1 | spl0_7),
  inference(avatar_split_clause,[],[f9,f67,f38])).
fof(f9,axiom,(
  sk_c8 = inverse(sk_c5) | multiply(sk_c1,sk_c2) = sk_c8),
  file('./p/GRP212-1.tptp',unknown)).
fof(f65,plain,(
  spl0_1 | spl0_6),
  inference(avatar_split_clause,[],[f8,f62,f38])).
fof(f8,axiom,(
  sk_c6 = multiply(sk_c5,sk_c8) | multiply(sk_c1,sk_c2) = sk_c8),
  file('./p/GRP212-1.tptp',unknown)).
fof(f60,plain,(
  spl0_1 | spl0_5),
  inference(avatar_split_clause,[],[f7,f57,f38])).
fof(f7,axiom,(
  sk_c7 = multiply(sk_c8,sk_c6) | multiply(sk_c1,sk_c2) = sk_c8),
  file('./p/GRP212-1.tptp',unknown)).
fof(f55,plain,(
  spl0_1 | spl0_4),
  inference(avatar_split_clause,[],[f6,f52,f38])).
fof(f6,axiom,(
  sk_c8 = multiply(sk_c4,sk_c7) | multiply(sk_c1,sk_c2) = sk_c8),
  file('./p/GRP212-1.tptp',unknown)).
fof(f50,plain,(
  spl0_1 | spl0_3),
  inference(avatar_split_clause,[],[f5,f47,f38])).
fof(f5,axiom,(
  sk_c8 = inverse(sk_c4) | multiply(sk_c1,sk_c2) = sk_c8),
  file('./p/GRP212-1.tptp',unknown)).
fof(f45,plain,(
  spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f4,f42,f38])).
fof(f4,axiom,(
  inverse(sk_c8) = sk_c7 | multiply(sk_c1,sk_c2) = sk_c8),
  file('./p/GRP212-1.tptp',unknown)).
% SZS output end Proof for GRP212-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
