% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GRP255-1
% SZS output start Proof for GRP255-1
fof(f1043,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f43,f48,f53,f58,f63,f68,f69,f70,f71,f72,f77,f78,f79,f80,f81,f86,f87,f88,f89,f90,f95,f96,f97,f98,f99,f104,f105,f106,f107,f108,f124,f138,f156,f193,f222,f249,f266,f288,f343,f502,f578,f627,f647,f686,f700,f725,f746,f759,f772,f967,f969,f988,f990,f1006,f1040,f1042])).
fof(f1042,plain,(
  ~spl0_3 | ~spl0_4 | ~spl0_14 | ~spl0_18 | ~spl0_20 | ~spl0_23),
  inference(avatar_contradiction_clause,[],[f1041])).
fof(f1041,plain,(
  $false | (~spl0_3 | ~spl0_4 | ~spl0_14 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(subsumption_resolution,[],[f1033,f882])).
fof(f882,plain,(
  identity = inverse(sk_c4) | (~spl0_3 | ~spl0_18 | ~spl0_20)),
  inference(backward_demodulation,[],[f727,f47])).
fof(f47,plain,(
  sk_c8 = inverse(sk_c4) | ~spl0_3),
  inference(avatar_component_clause,[],[f45])).
fof(f45,plain,(
  spl0_3 <=> sk_c8 = inverse(sk_c4)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f727,plain,(
  identity = sk_c8 | (~spl0_18 | ~spl0_20)),
  inference(forward_demodulation,[],[f136,f145])).
fof(f145,plain,(
  identity = sk_c7 | ~spl0_20),
  inference(avatar_component_clause,[],[f144])).
fof(f144,plain,(
  spl0_20 <=> identity = sk_c7),
  introduced(avatar_definition,[new_symbols(naming,[spl0_20])])).
fof(f136,plain,(
  sk_c8 = sk_c7 | ~spl0_18),
  inference(avatar_component_clause,[],[f135])).
fof(f135,plain,(
  spl0_18 <=> sk_c8 = sk_c7),
  introduced(avatar_definition,[new_symbols(naming,[spl0_18])])).
fof(f1033,plain,(
  identity != inverse(sk_c4) | (~spl0_4 | ~spl0_14 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(trivial_inequality_removal,[],[f1030])).
fof(f1030,plain,(
  identity != identity | identity != inverse(sk_c4) | (~spl0_4 | ~spl0_14 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(superposition,[],[f1024,f888])).
fof(f888,plain,(
  identity = multiply(sk_c4,identity) | (~spl0_4 | ~spl0_18 | ~spl0_20)),
  inference(backward_demodulation,[],[f727,f629])).
fof(f629,plain,(
  sk_c8 = multiply(sk_c4,identity) | (~spl0_4 | ~spl0_20)),
  inference(backward_demodulation,[],[f145,f52])).
fof(f52,plain,(
  sk_c8 = multiply(sk_c4,sk_c7) | ~spl0_4),
  inference(avatar_component_clause,[],[f50])).
fof(f50,plain,(
  spl0_4 <=> sk_c8 = multiply(sk_c4,sk_c7)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f1024,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | identity != inverse(X5)) ) | (~spl0_14 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(forward_demodulation,[],[f1023,f498])).
fof(f498,plain,(
  identity = sk_c6 | ~spl0_23),
  inference(avatar_component_clause,[],[f497])).
fof(f497,plain,(
  spl0_23 <=> identity = sk_c6),
  introduced(avatar_definition,[new_symbols(naming,[spl0_23])])).
fof(f1023,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | sk_c6 != inverse(X5)) ) | (~spl0_14 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(forward_demodulation,[],[f1022,f498])).
fof(f1022,plain,(
  ( ! [X5] : (sk_c6 != multiply(X5,identity) | sk_c6 != inverse(X5)) ) | (~spl0_14 | ~spl0_18 | ~spl0_20)),
  inference(forward_demodulation,[],[f117,f727])).
fof(f117,plain,(
  ( ! [X5] : (sk_c6 != multiply(X5,sk_c8) | sk_c6 != inverse(X5)) ) | ~spl0_14),
  inference(avatar_component_clause,[],[f116])).
fof(f116,plain,(
  spl0_14 <=> ! [X5] : (sk_c6 != multiply(X5,sk_c8) | sk_c6 != inverse(X5))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_14])])).
fof(f1040,plain,(
  ~spl0_7 | ~spl0_14 | ~spl0_18 | ~spl0_20 | ~spl0_23),
  inference(avatar_contradiction_clause,[],[f1039])).
fof(f1039,plain,(
  $false | (~spl0_7 | ~spl0_14 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(subsumption_resolution,[],[f1038,f900])).
fof(f900,plain,(
  identity = inverse(identity) | (~spl0_7 | ~spl0_18 | ~spl0_20)),
  inference(backward_demodulation,[],[f897,f883])).
fof(f883,plain,(
  identity = inverse(sk_c1) | (~spl0_7 | ~spl0_18 | ~spl0_20)),
  inference(backward_demodulation,[],[f727,f67])).
fof(f67,plain,(
  sk_c8 = inverse(sk_c1) | ~spl0_7),
  inference(avatar_component_clause,[],[f65])).
fof(f65,plain,(
  spl0_7 <=> sk_c8 = inverse(sk_c1)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).
fof(f897,plain,(
  identity = sk_c1 | (~spl0_7 | ~spl0_18 | ~spl0_20)),
  inference(forward_demodulation,[],[f886,f2])).
fof(f2,axiom,(
  ( ! [X0] : (identity = multiply(inverse(X0),X0)) )),
  file('./p/GRP255-1.tptp',unknown)).
fof(f886,plain,(
  sk_c1 = multiply(inverse(identity),identity) | (~spl0_7 | ~spl0_18 | ~spl0_20)),
  inference(backward_demodulation,[],[f727,f503])).
fof(f503,plain,(
  sk_c1 = multiply(inverse(sk_c8),identity) | ~spl0_7),
  inference(superposition,[],[f166,f458])).
fof(f458,plain,(
  identity = multiply(sk_c8,sk_c1) | ~spl0_7),
  inference(superposition,[],[f2,f67])).
fof(f166,plain,(
  ( ! [X6,X7] : (multiply(inverse(X6),multiply(X6,X7)) = X7) )),
  inference(forward_demodulation,[],[f159,f1])).
fof(f1,axiom,(
  ( ! [X0] : (multiply(identity,X0) = X0) )),
  file('./p/GRP255-1.tptp',unknown)).
fof(f159,plain,(
  ( ! [X6,X7] : (multiply(inverse(X6),multiply(X6,X7)) = multiply(identity,X7)) )),
  inference(superposition,[],[f3,f2])).
fof(f3,axiom,(
  ( ! [X2,X0,X1] : (multiply(multiply(X0,X1),X2) = multiply(X0,multiply(X1,X2))) )),
  file('./p/GRP255-1.tptp',unknown)).
fof(f1038,plain,(
  identity != inverse(identity) | (~spl0_7 | ~spl0_14 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(forward_demodulation,[],[f1034,f900])).
fof(f1034,plain,(
  identity != inverse(inverse(identity)) | (~spl0_14 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(trivial_inequality_removal,[],[f1029])).
fof(f1029,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_14 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(superposition,[],[f1024,f2])).
fof(f1006,plain,(
  spl0_23 | ~spl0_2 | ~spl0_18 | ~spl0_20),
  inference(avatar_split_clause,[],[f1005,f144,f135,f40,f497])).
fof(f40,plain,(
  spl0_2 <=> multiply(sk_c7,sk_c8) = sk_c6),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f1005,plain,(
  identity = sk_c6 | (~spl0_2 | ~spl0_18 | ~spl0_20)),
  inference(forward_demodulation,[],[f1004,f1])).
fof(f1004,plain,(
  sk_c6 = multiply(identity,identity) | (~spl0_2 | ~spl0_18 | ~spl0_20)),
  inference(forward_demodulation,[],[f1003,f145])).
fof(f1003,plain,(
  sk_c6 = multiply(sk_c7,identity) | (~spl0_2 | ~spl0_18 | ~spl0_20)),
  inference(forward_demodulation,[],[f42,f727])).
fof(f42,plain,(
  multiply(sk_c7,sk_c8) = sk_c6 | ~spl0_2),
  inference(avatar_component_clause,[],[f40])).
fof(f990,plain,(
  ~spl0_3 | ~spl0_4 | ~spl0_16 | ~spl0_18 | ~spl0_20 | ~spl0_23),
  inference(avatar_contradiction_clause,[],[f989])).
fof(f989,plain,(
  $false | (~spl0_3 | ~spl0_4 | ~spl0_16 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(subsumption_resolution,[],[f981,f882])).
fof(f981,plain,(
  identity != inverse(sk_c4) | (~spl0_4 | ~spl0_16 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(trivial_inequality_removal,[],[f976])).
fof(f976,plain,(
  identity != identity | identity != inverse(sk_c4) | (~spl0_4 | ~spl0_16 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(superposition,[],[f972,f888])).
fof(f972,plain,(
  ( ! [X7] : (identity != multiply(X7,identity) | identity != inverse(X7)) ) | (~spl0_16 | ~spl0_20 | ~spl0_23)),
  inference(forward_demodulation,[],[f971,f145])).
fof(f971,plain,(
  ( ! [X7] : (identity != multiply(X7,identity) | sk_c7 != inverse(X7)) ) | (~spl0_16 | ~spl0_20 | ~spl0_23)),
  inference(forward_demodulation,[],[f970,f145])).
fof(f970,plain,(
  ( ! [X7] : (sk_c7 != multiply(X7,identity) | sk_c7 != inverse(X7)) ) | (~spl0_16 | ~spl0_23)),
  inference(forward_demodulation,[],[f123,f498])).
fof(f123,plain,(
  ( ! [X7] : (sk_c7 != multiply(X7,sk_c6) | sk_c7 != inverse(X7)) ) | ~spl0_16),
  inference(avatar_component_clause,[],[f122])).
fof(f122,plain,(
  spl0_16 <=> ! [X7] : (sk_c7 != multiply(X7,sk_c6) | sk_c7 != inverse(X7))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_16])])).
fof(f988,plain,(
  ~spl0_7 | ~spl0_16 | ~spl0_18 | ~spl0_20 | ~spl0_23),
  inference(avatar_contradiction_clause,[],[f987])).
fof(f987,plain,(
  $false | (~spl0_7 | ~spl0_16 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(subsumption_resolution,[],[f986,f900])).
fof(f986,plain,(
  identity != inverse(identity) | (~spl0_7 | ~spl0_16 | ~spl0_18 | ~spl0_20 | ~spl0_23)),
  inference(forward_demodulation,[],[f982,f900])).
fof(f982,plain,(
  identity != inverse(inverse(identity)) | (~spl0_16 | ~spl0_20 | ~spl0_23)),
  inference(trivial_inequality_removal,[],[f975])).
fof(f975,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_16 | ~spl0_20 | ~spl0_23)),
  inference(superposition,[],[f972,f2])).
fof(f969,plain,(
  ~spl0_3 | ~spl0_4 | ~spl0_15 | ~spl0_18 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f968])).
fof(f968,plain,(
  $false | (~spl0_3 | ~spl0_4 | ~spl0_15 | ~spl0_18 | ~spl0_20)),
  inference(subsumption_resolution,[],[f960,f882])).
fof(f960,plain,(
  identity != inverse(sk_c4) | (~spl0_4 | ~spl0_15 | ~spl0_18 | ~spl0_20)),
  inference(trivial_inequality_removal,[],[f955])).
fof(f955,plain,(
  identity != identity | identity != inverse(sk_c4) | (~spl0_4 | ~spl0_15 | ~spl0_18 | ~spl0_20)),
  inference(superposition,[],[f927,f888])).
fof(f927,plain,(
  ( ! [X6] : (identity != multiply(X6,identity) | identity != inverse(X6)) ) | (~spl0_15 | ~spl0_18 | ~spl0_20)),
  inference(forward_demodulation,[],[f926,f727])).
fof(f926,plain,(
  ( ! [X6] : (identity != multiply(X6,identity) | sk_c8 != inverse(X6)) ) | (~spl0_15 | ~spl0_18 | ~spl0_20)),
  inference(forward_demodulation,[],[f925,f727])).
fof(f925,plain,(
  ( ! [X6] : (sk_c8 != multiply(X6,identity) | sk_c8 != inverse(X6)) ) | (~spl0_15 | ~spl0_20)),
  inference(forward_demodulation,[],[f120,f145])).
fof(f120,plain,(
  ( ! [X6] : (sk_c8 != multiply(X6,sk_c7) | sk_c8 != inverse(X6)) ) | ~spl0_15),
  inference(avatar_component_clause,[],[f119])).
fof(f119,plain,(
  spl0_15 <=> ! [X6] : (sk_c8 != multiply(X6,sk_c7) | sk_c8 != inverse(X6))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_15])])).
fof(f967,plain,(
  ~spl0_7 | ~spl0_15 | ~spl0_18 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f966])).
fof(f966,plain,(
  $false | (~spl0_7 | ~spl0_15 | ~spl0_18 | ~spl0_20)),
  inference(subsumption_resolution,[],[f965,f900])).
fof(f965,plain,(
  identity != inverse(identity) | (~spl0_7 | ~spl0_15 | ~spl0_18 | ~spl0_20)),
  inference(forward_demodulation,[],[f961,f900])).
fof(f961,plain,(
  identity != inverse(inverse(identity)) | (~spl0_15 | ~spl0_18 | ~spl0_20)),
  inference(trivial_inequality_removal,[],[f954])).
fof(f954,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_15 | ~spl0_18 | ~spl0_20)),
  inference(superposition,[],[f927,f2])).
fof(f772,plain,(
  ~spl0_9 | ~spl0_16 | ~spl0_20 | ~spl0_23),
  inference(avatar_contradiction_clause,[],[f771])).
fof(f771,plain,(
  $false | (~spl0_9 | ~spl0_16 | ~spl0_20 | ~spl0_23)),
  inference(subsumption_resolution,[],[f770,f654])).
fof(f654,plain,(
  identity = inverse(identity) | (~spl0_9 | ~spl0_20)),
  inference(backward_demodulation,[],[f651,f631])).
fof(f631,plain,(
  identity = inverse(sk_c2) | (~spl0_9 | ~spl0_20)),
  inference(backward_demodulation,[],[f145,f85])).
fof(f85,plain,(
  sk_c7 = inverse(sk_c2) | ~spl0_9),
  inference(avatar_component_clause,[],[f83])).
fof(f83,plain,(
  spl0_9 <=> sk_c7 = inverse(sk_c2)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_9])])).
fof(f651,plain,(
  identity = sk_c2 | (~spl0_9 | ~spl0_20)),
  inference(forward_demodulation,[],[f642,f2])).
fof(f642,plain,(
  sk_c2 = multiply(inverse(identity),identity) | (~spl0_9 | ~spl0_20)),
  inference(backward_demodulation,[],[f145,f593])).
fof(f593,plain,(
  sk_c2 = multiply(inverse(sk_c7),identity) | ~spl0_9),
  inference(superposition,[],[f166,f567])).
fof(f567,plain,(
  identity = multiply(sk_c7,sk_c2) | ~spl0_9),
  inference(superposition,[],[f2,f85])).
fof(f770,plain,(
  identity != inverse(identity) | (~spl0_9 | ~spl0_16 | ~spl0_20 | ~spl0_23)),
  inference(forward_demodulation,[],[f766,f654])).
fof(f766,plain,(
  identity != inverse(inverse(identity)) | (~spl0_16 | ~spl0_20 | ~spl0_23)),
  inference(trivial_inequality_removal,[],[f765])).
fof(f765,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_16 | ~spl0_20 | ~spl0_23)),
  inference(superposition,[],[f762,f2])).
fof(f762,plain,(
  ( ! [X7] : (identity != multiply(X7,identity) | identity != inverse(X7)) ) | (~spl0_16 | ~spl0_20 | ~spl0_23)),
  inference(forward_demodulation,[],[f761,f145])).
fof(f761,plain,(
  ( ! [X7] : (identity != multiply(X7,identity) | sk_c7 != inverse(X7)) ) | (~spl0_16 | ~spl0_20 | ~spl0_23)),
  inference(forward_demodulation,[],[f760,f145])).
fof(f760,plain,(
  ( ! [X7] : (sk_c7 != multiply(X7,identity) | sk_c7 != inverse(X7)) ) | (~spl0_16 | ~spl0_23)),
  inference(forward_demodulation,[],[f123,f498])).
fof(f759,plain,(
  ~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_15 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f758])).
fof(f758,plain,(
  $false | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_15 | ~spl0_20)),
  inference(subsumption_resolution,[],[f757,f654])).
fof(f757,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_15 | ~spl0_20)),
  inference(forward_demodulation,[],[f753,f654])).
fof(f753,plain,(
  identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_15 | ~spl0_20)),
  inference(trivial_inequality_removal,[],[f752])).
fof(f752,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_15 | ~spl0_20)),
  inference(superposition,[],[f749,f2])).
fof(f749,plain,(
  ( ! [X6] : (identity != multiply(X6,identity) | identity != inverse(X6)) ) | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_15 | ~spl0_20)),
  inference(forward_demodulation,[],[f748,f714])).
fof(f714,plain,(
  identity = sk_c8 | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(forward_demodulation,[],[f689,f654])).
fof(f689,plain,(
  sk_c8 = inverse(identity) | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f688,f67])).
fof(f688,plain,(
  identity = sk_c1 | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(forward_demodulation,[],[f687,f661])).
fof(f661,plain,(
  ( ! [X0] : (multiply(sk_c1,X0) = X0) ) | (~spl0_1 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f660,f649])).
fof(f649,plain,(
  ( ! [X0] : (multiply(sk_c1,multiply(sk_c8,X0)) = X0) ) | (~spl0_1 | ~spl0_20)),
  inference(forward_demodulation,[],[f638,f1])).
fof(f638,plain,(
  ( ! [X0] : (multiply(identity,X0) = multiply(sk_c1,multiply(sk_c8,X0))) ) | (~spl0_1 | ~spl0_20)),
  inference(backward_demodulation,[],[f145,f569])).
fof(f569,plain,(
  ( ! [X0] : (multiply(sk_c7,X0) = multiply(sk_c1,multiply(sk_c8,X0))) ) | ~spl0_1),
  inference(superposition,[],[f3,f38])).
fof(f38,plain,(
  multiply(sk_c1,sk_c8) = sk_c7 | ~spl0_1),
  inference(avatar_component_clause,[],[f36])).
fof(f36,plain,(
  spl0_1 <=> multiply(sk_c1,sk_c8) = sk_c7),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f660,plain,(
  ( ! [X0] : (multiply(sk_c8,X0) = X0) ) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(forward_demodulation,[],[f657,f656])).
fof(f656,plain,(
  ( ! [X0] : (multiply(sk_c6,X0) = X0) ) | (~spl0_8 | ~spl0_9 | ~spl0_20)),
  inference(forward_demodulation,[],[f652,f1])).
fof(f652,plain,(
  ( ! [X0] : (multiply(identity,X0) = multiply(sk_c6,X0)) ) | (~spl0_8 | ~spl0_9 | ~spl0_20)),
  inference(backward_demodulation,[],[f651,f650])).
fof(f650,plain,(
  ( ! [X0] : (multiply(sk_c6,X0) = multiply(sk_c2,X0)) ) | (~spl0_8 | ~spl0_20)),
  inference(forward_demodulation,[],[f640,f1])).
fof(f640,plain,(
  ( ! [X0] : (multiply(sk_c6,X0) = multiply(sk_c2,multiply(identity,X0))) ) | (~spl0_8 | ~spl0_20)),
  inference(backward_demodulation,[],[f145,f582])).
fof(f582,plain,(
  ( ! [X0] : (multiply(sk_c6,X0) = multiply(sk_c2,multiply(sk_c7,X0))) ) | ~spl0_8),
  inference(superposition,[],[f3,f76])).
fof(f76,plain,(
  sk_c6 = multiply(sk_c2,sk_c7) | ~spl0_8),
  inference(avatar_component_clause,[],[f74])).
fof(f74,plain,(
  spl0_8 <=> sk_c6 = multiply(sk_c2,sk_c7)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_8])])).
fof(f657,plain,(
  ( ! [X0] : (multiply(sk_c6,X0) = multiply(sk_c8,X0)) ) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f656,f612])).
fof(f612,plain,(
  ( ! [X0] : (multiply(sk_c8,X0) = multiply(sk_c6,multiply(sk_c6,X0))) ) | (~spl0_10 | ~spl0_11)),
  inference(superposition,[],[f3,f467])).
fof(f467,plain,(
  sk_c8 = multiply(sk_c6,sk_c6) | (~spl0_10 | ~spl0_11)),
  inference(forward_demodulation,[],[f465,f94])).
fof(f94,plain,(
  sk_c6 = inverse(sk_c3) | ~spl0_10),
  inference(avatar_component_clause,[],[f92])).
fof(f92,plain,(
  spl0_10 <=> sk_c6 = inverse(sk_c3)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f465,plain,(
  sk_c8 = multiply(inverse(sk_c3),sk_c6) | ~spl0_11),
  inference(superposition,[],[f166,f103])).
fof(f103,plain,(
  sk_c6 = multiply(sk_c3,sk_c8) | ~spl0_11),
  inference(avatar_component_clause,[],[f101])).
fof(f101,plain,(
  spl0_11 <=> sk_c6 = multiply(sk_c3,sk_c8)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_11])])).
fof(f687,plain,(
  sk_c1 = multiply(sk_c1,identity) | (~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(forward_demodulation,[],[f659,f680])).
fof(f680,plain,(
  sk_c1 = sk_c3 | (~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(forward_demodulation,[],[f676,f503])).
fof(f676,plain,(
  sk_c3 = multiply(inverse(sk_c8),identity) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f672,f511])).
fof(f511,plain,(
  sk_c3 = multiply(inverse(sk_c6),identity) | ~spl0_10),
  inference(superposition,[],[f166,f460])).
fof(f460,plain,(
  identity = multiply(sk_c6,sk_c3) | ~spl0_10),
  inference(superposition,[],[f2,f94])).
fof(f672,plain,(
  sk_c8 = sk_c6 | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f670,f103])).
fof(f670,plain,(
  ( ! [X0] : (multiply(sk_c3,X0) = X0) ) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(forward_demodulation,[],[f658,f660])).
fof(f658,plain,(
  ( ! [X0] : (multiply(sk_c3,multiply(sk_c8,X0)) = X0) ) | (~spl0_8 | ~spl0_9 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f656,f466])).
fof(f466,plain,(
  ( ! [X0] : (multiply(sk_c6,X0) = multiply(sk_c3,multiply(sk_c8,X0))) ) | ~spl0_11),
  inference(superposition,[],[f3,f103])).
fof(f659,plain,(
  sk_c1 = multiply(sk_c3,identity) | (~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f656,f613])).
fof(f613,plain,(
  multiply(sk_c3,identity) = multiply(sk_c6,sk_c1) | (~spl0_7 | ~spl0_11)),
  inference(superposition,[],[f466,f458])).
fof(f748,plain,(
  ( ! [X6] : (identity != multiply(X6,identity) | sk_c8 != inverse(X6)) ) | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_15 | ~spl0_20)),
  inference(forward_demodulation,[],[f747,f714])).
fof(f747,plain,(
  ( ! [X6] : (sk_c8 != multiply(X6,identity) | sk_c8 != inverse(X6)) ) | (~spl0_15 | ~spl0_20)),
  inference(forward_demodulation,[],[f120,f145])).
fof(f746,plain,(
  ~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_14 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f745])).
fof(f745,plain,(
  $false | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_14 | ~spl0_20)),
  inference(subsumption_resolution,[],[f744,f654])).
fof(f744,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_14 | ~spl0_20)),
  inference(forward_demodulation,[],[f740,f654])).
fof(f740,plain,(
  identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_14 | ~spl0_20)),
  inference(trivial_inequality_removal,[],[f739])).
fof(f739,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_14 | ~spl0_20)),
  inference(superposition,[],[f722,f2])).
fof(f722,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | identity != inverse(X5)) ) | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_14 | ~spl0_20)),
  inference(forward_demodulation,[],[f720,f714])).
fof(f720,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | sk_c8 != inverse(X5)) ) | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_14 | ~spl0_20)),
  inference(backward_demodulation,[],[f714,f712])).
fof(f712,plain,(
  ( ! [X5] : (sk_c8 != inverse(X5) | sk_c8 != multiply(X5,sk_c8)) ) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_14 | ~spl0_20)),
  inference(forward_demodulation,[],[f711,f672])).
fof(f711,plain,(
  ( ! [X5] : (sk_c8 != multiply(X5,sk_c8) | sk_c6 != inverse(X5)) ) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_14 | ~spl0_20)),
  inference(forward_demodulation,[],[f117,f672])).
fof(f725,plain,(
  ~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20 | spl0_21),
  inference(avatar_contradiction_clause,[],[f724])).
fof(f724,plain,(
  $false | (~spl0_1 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20 | spl0_21)),
  inference(subsumption_resolution,[],[f723,f714])).
fof(f723,plain,(
  identity != sk_c8 | (~spl0_9 | ~spl0_20 | spl0_21)),
  inference(forward_demodulation,[],[f633,f654])).
fof(f633,plain,(
  sk_c8 != inverse(identity) | (~spl0_20 | spl0_21)),
  inference(backward_demodulation,[],[f145,f151])).
fof(f151,plain,(
  sk_c8 != inverse(sk_c7) | spl0_21),
  inference(avatar_component_clause,[],[f149])).
fof(f149,plain,(
  spl0_21 <=> sk_c8 = inverse(sk_c7)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_21])])).
fof(f700,plain,(
  spl0_23 | ~spl0_8 | ~spl0_9 | ~spl0_20),
  inference(avatar_split_clause,[],[f699,f144,f83,f74,f497])).
fof(f699,plain,(
  identity = sk_c6 | (~spl0_8 | ~spl0_9 | ~spl0_20)),
  inference(forward_demodulation,[],[f655,f1])).
fof(f655,plain,(
  sk_c6 = multiply(identity,identity) | (~spl0_8 | ~spl0_9 | ~spl0_20)),
  inference(backward_demodulation,[],[f651,f630])).
fof(f630,plain,(
  sk_c6 = multiply(sk_c2,identity) | (~spl0_8 | ~spl0_20)),
  inference(backward_demodulation,[],[f145,f76])).
fof(f686,plain,(
  ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | spl0_18 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f685])).
fof(f685,plain,(
  $false | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | spl0_18 | ~spl0_20)),
  inference(subsumption_resolution,[],[f684,f632])).
fof(f632,plain,(
  identity != sk_c8 | (spl0_18 | ~spl0_20)),
  inference(backward_demodulation,[],[f145,f137])).
fof(f137,plain,(
  sk_c8 != sk_c7 | spl0_18),
  inference(avatar_component_clause,[],[f135])).
fof(f684,plain,(
  identity = sk_c8 | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(forward_demodulation,[],[f677,f2])).
fof(f677,plain,(
  sk_c8 = multiply(inverse(sk_c8),sk_c8) | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_20)),
  inference(backward_demodulation,[],[f672,f611])).
fof(f611,plain,(
  sk_c6 = multiply(inverse(sk_c6),sk_c8) | (~spl0_10 | ~spl0_11)),
  inference(superposition,[],[f166,f467])).
fof(f647,plain,(
  ~spl0_23 | ~spl0_20 | spl0_22),
  inference(avatar_split_clause,[],[f634,f153,f144,f497])).
fof(f153,plain,(
  spl0_22 <=> sk_c7 = sk_c6),
  introduced(avatar_definition,[new_symbols(naming,[spl0_22])])).
fof(f634,plain,(
  identity != sk_c6 | (~spl0_20 | spl0_22)),
  inference(backward_demodulation,[],[f145,f155])).
fof(f155,plain,(
  sk_c7 != sk_c6 | spl0_22),
  inference(avatar_component_clause,[],[f153])).
fof(f627,plain,(
  spl0_20 | ~spl0_1 | ~spl0_7),
  inference(avatar_split_clause,[],[f623,f65,f36,f144])).
fof(f623,plain,(
  identity = sk_c7 | (~spl0_1 | ~spl0_7)),
  inference(forward_demodulation,[],[f620,f2])).
fof(f620,plain,(
  sk_c7 = multiply(inverse(sk_c8),sk_c8) | (~spl0_1 | ~spl0_7)),
  inference(superposition,[],[f166,f570])).
fof(f570,plain,(
  sk_c8 = multiply(sk_c8,sk_c7) | (~spl0_1 | ~spl0_7)),
  inference(forward_demodulation,[],[f568,f67])).
fof(f568,plain,(
  sk_c8 = multiply(inverse(sk_c1),sk_c7) | ~spl0_1),
  inference(superposition,[],[f166,f38])).
fof(f578,plain,(
  ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | spl0_18 | ~spl0_22),
  inference(avatar_contradiction_clause,[],[f577])).
fof(f577,plain,(
  $false | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | spl0_18 | ~spl0_22)),
  inference(subsumption_resolution,[],[f576,f137])).
fof(f576,plain,(
  sk_c8 = sk_c7 | (~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11 | ~spl0_22)),
  inference(forward_demodulation,[],[f575,f549])).
fof(f549,plain,(
  sk_c8 = multiply(sk_c7,sk_c7) | (~spl0_10 | ~spl0_11 | ~spl0_22)),
  inference(forward_demodulation,[],[f467,f154])).
fof(f154,plain,(
  sk_c7 = sk_c6 | ~spl0_22),
  inference(avatar_component_clause,[],[f153])).
fof(f575,plain,(
  sk_c7 = multiply(sk_c7,sk_c7) | (~spl0_8 | ~spl0_9 | ~spl0_22)),
  inference(forward_demodulation,[],[f573,f85])).
fof(f573,plain,(
  sk_c7 = multiply(inverse(sk_c2),sk_c7) | (~spl0_8 | ~spl0_22)),
  inference(superposition,[],[f166,f557])).
fof(f557,plain,(
  sk_c7 = multiply(sk_c2,sk_c7) | (~spl0_8 | ~spl0_22)),
  inference(forward_demodulation,[],[f76,f154])).
fof(f502,plain,(
  ~spl0_8 | ~spl0_9 | ~spl0_13 | ~spl0_18),
  inference(avatar_contradiction_clause,[],[f501])).
fof(f501,plain,(
  $false | (~spl0_8 | ~spl0_9 | ~spl0_13 | ~spl0_18)),
  inference(subsumption_resolution,[],[f495,f412])).
fof(f412,plain,(
  sk_c8 = inverse(sk_c2) | (~spl0_9 | ~spl0_18)),
  inference(backward_demodulation,[],[f136,f85])).
fof(f495,plain,(
  sk_c8 != inverse(sk_c2) | (~spl0_8 | ~spl0_13 | ~spl0_18)),
  inference(trivial_inequality_removal,[],[f492])).
fof(f492,plain,(
  sk_c6 != sk_c6 | sk_c8 != inverse(sk_c2) | (~spl0_8 | ~spl0_13 | ~spl0_18)),
  inference(superposition,[],[f447,f448])).
fof(f448,plain,(
  sk_c6 = multiply(sk_c2,sk_c8) | (~spl0_8 | ~spl0_18)),
  inference(forward_demodulation,[],[f76,f136])).
fof(f447,plain,(
  ( ! [X4] : (sk_c6 != multiply(X4,sk_c8) | sk_c8 != inverse(X4)) ) | (~spl0_13 | ~spl0_18)),
  inference(forward_demodulation,[],[f445,f136])).
fof(f445,plain,(
  ( ! [X4] : (sk_c6 != multiply(X4,sk_c8) | sk_c7 != inverse(X4)) ) | (~spl0_13 | ~spl0_18)),
  inference(backward_demodulation,[],[f136,f114])).
fof(f114,plain,(
  ( ! [X4] : (sk_c7 != inverse(X4) | sk_c6 != multiply(X4,sk_c7)) ) | ~spl0_13),
  inference(avatar_component_clause,[],[f113])).
fof(f113,plain,(
  spl0_13 <=> ! [X4] : (sk_c7 != inverse(X4) | sk_c6 != multiply(X4,sk_c7))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_13])])).
fof(f343,plain,(
  spl0_2 | ~spl0_18 | ~spl0_20 | ~spl0_22),
  inference(avatar_contradiction_clause,[],[f342])).
fof(f342,plain,(
  $false | (spl0_2 | ~spl0_18 | ~spl0_20 | ~spl0_22)),
  inference(subsumption_resolution,[],[f339,f1])).
fof(f339,plain,(
  identity != multiply(identity,identity) | (spl0_2 | ~spl0_18 | ~spl0_20 | ~spl0_22)),
  inference(backward_demodulation,[],[f309,f332])).
fof(f332,plain,(
  identity != multiply(identity,sk_c8) | (spl0_2 | ~spl0_20 | ~spl0_22)),
  inference(forward_demodulation,[],[f325,f145])).
fof(f325,plain,(
  identity != multiply(sk_c7,sk_c8) | (spl0_2 | ~spl0_20 | ~spl0_22)),
  inference(forward_demodulation,[],[f41,f291])).
fof(f291,plain,(
  identity = sk_c6 | (~spl0_20 | ~spl0_22)),
  inference(forward_demodulation,[],[f154,f145])).
fof(f41,plain,(
  multiply(sk_c7,sk_c8) != sk_c6 | spl0_2),
  inference(avatar_component_clause,[],[f40])).
fof(f309,plain,(
  identity = sk_c8 | (~spl0_18 | ~spl0_20)),
  inference(forward_demodulation,[],[f136,f145])).
fof(f288,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13 | ~spl0_18 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f287])).
fof(f287,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13 | ~spl0_18 | ~spl0_20)),
  inference(subsumption_resolution,[],[f286,f245])).
fof(f245,plain,(
  identity = inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_18)),
  inference(backward_demodulation,[],[f243,f228])).
fof(f228,plain,(
  identity = inverse(sk_c4) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_18)),
  inference(backward_demodulation,[],[f227,f47])).
fof(f227,plain,(
  identity = sk_c8 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_18)),
  inference(forward_demodulation,[],[f225,f2])).
fof(f225,plain,(
  sk_c8 = multiply(inverse(sk_c8),sk_c8) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_18)),
  inference(backward_demodulation,[],[f220,f207])).
fof(f207,plain,(
  sk_c8 = multiply(inverse(sk_c8),sk_c6) | (~spl0_2 | ~spl0_18)),
  inference(backward_demodulation,[],[f136,f182])).
fof(f182,plain,(
  sk_c8 = multiply(inverse(sk_c7),sk_c6) | ~spl0_2),
  inference(superposition,[],[f166,f42])).
fof(f220,plain,(
  sk_c8 = sk_c6 | (~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_18)),
  inference(forward_demodulation,[],[f209,f206])).
fof(f206,plain,(
  sk_c8 = multiply(sk_c8,sk_c8) | (~spl0_3 | ~spl0_4 | ~spl0_18)),
  inference(backward_demodulation,[],[f136,f171])).
fof(f171,plain,(
  sk_c7 = multiply(sk_c8,sk_c8) | (~spl0_3 | ~spl0_4)),
  inference(superposition,[],[f168,f52])).
fof(f168,plain,(
  ( ! [X0] : (multiply(sk_c8,multiply(sk_c4,X0)) = X0) ) | ~spl0_3),
  inference(forward_demodulation,[],[f167,f1])).
fof(f167,plain,(
  ( ! [X0] : (multiply(identity,X0) = multiply(sk_c8,multiply(sk_c4,X0))) ) | ~spl0_3),
  inference(superposition,[],[f3,f125])).
fof(f125,plain,(
  identity = multiply(sk_c8,sk_c4) | ~spl0_3),
  inference(superposition,[],[f2,f47])).
fof(f209,plain,(
  sk_c6 = multiply(sk_c8,sk_c8) | (~spl0_5 | ~spl0_6 | ~spl0_18)),
  inference(backward_demodulation,[],[f136,f188])).
fof(f188,plain,(
  sk_c6 = multiply(sk_c7,sk_c7) | (~spl0_5 | ~spl0_6)),
  inference(forward_demodulation,[],[f185,f57])).
fof(f57,plain,(
  sk_c7 = inverse(sk_c5) | ~spl0_5),
  inference(avatar_component_clause,[],[f55])).
fof(f55,plain,(
  spl0_5 <=> sk_c7 = inverse(sk_c5)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f185,plain,(
  sk_c6 = multiply(inverse(sk_c5),sk_c7) | ~spl0_6),
  inference(superposition,[],[f166,f62])).
fof(f62,plain,(
  sk_c7 = multiply(sk_c5,sk_c6) | ~spl0_6),
  inference(avatar_component_clause,[],[f60])).
fof(f60,plain,(
  spl0_6 <=> sk_c7 = multiply(sk_c5,sk_c6)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).
fof(f243,plain,(
  identity = sk_c4 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_18)),
  inference(forward_demodulation,[],[f233,f2])).
fof(f233,plain,(
  sk_c4 = multiply(inverse(identity),identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_18)),
  inference(backward_demodulation,[],[f227,f180])).
fof(f180,plain,(
  sk_c4 = multiply(inverse(sk_c8),identity) | ~spl0_3),
  inference(superposition,[],[f166,f125])).
fof(f286,plain,(
  identity != inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13 | ~spl0_18 | ~spl0_20)),
  inference(forward_demodulation,[],[f282,f245])).
fof(f282,plain,(
  identity != inverse(inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13 | ~spl0_18 | ~spl0_20)),
  inference(trivial_inequality_removal,[],[f281])).
fof(f281,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13 | ~spl0_18 | ~spl0_20)),
  inference(superposition,[],[f273,f2])).
fof(f273,plain,(
  ( ! [X4] : (identity != multiply(X4,identity) | identity != inverse(X4)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_13 | ~spl0_18 | ~spl0_20)),
  inference(forward_demodulation,[],[f272,f240])).
fof(f240,plain,(
  identity = sk_c6 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_18)),
  inference(backward_demodulation,[],[f227,f220])).
fof(f272,plain,(
  ( ! [X4] : (sk_c6 != multiply(X4,identity) | identity != inverse(X4)) ) | (~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f271,f145])).
fof(f271,plain,(
  ( ! [X4] : (identity != inverse(X4) | sk_c6 != multiply(X4,sk_c7)) ) | (~spl0_13 | ~spl0_20)),
  inference(forward_demodulation,[],[f114,f145])).
fof(f266,plain,(
  spl0_20 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_18),
  inference(avatar_split_clause,[],[f231,f135,f60,f55,f50,f45,f40,f144])).
fof(f231,plain,(
  identity = sk_c7 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_18)),
  inference(backward_demodulation,[],[f227,f136])).
fof(f249,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | spl0_17 | ~spl0_18),
  inference(avatar_contradiction_clause,[],[f248])).
fof(f248,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | spl0_17 | ~spl0_18)),
  inference(subsumption_resolution,[],[f245,f230])).
fof(f230,plain,(
  identity != inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | spl0_17 | ~spl0_18)),
  inference(backward_demodulation,[],[f227,f133])).
fof(f133,plain,(
  sk_c8 != inverse(identity) | spl0_17),
  inference(avatar_component_clause,[],[f131])).
fof(f131,plain,(
  spl0_17 <=> sk_c8 = inverse(identity)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_17])])).
fof(f222,plain,(
  ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_18 | spl0_22),
  inference(avatar_contradiction_clause,[],[f221])).
fof(f221,plain,(
  $false | (~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_18 | spl0_22)),
  inference(subsumption_resolution,[],[f220,f200])).
fof(f200,plain,(
  sk_c8 != sk_c6 | (~spl0_18 | spl0_22)),
  inference(backward_demodulation,[],[f136,f155])).
fof(f193,plain,(
  spl0_18 | ~spl0_2 | ~spl0_5 | ~spl0_6),
  inference(avatar_split_clause,[],[f192,f60,f55,f40,f135])).
fof(f192,plain,(
  sk_c8 = sk_c7 | (~spl0_2 | ~spl0_5 | ~spl0_6)),
  inference(forward_demodulation,[],[f190,f182])).
fof(f190,plain,(
  sk_c7 = multiply(inverse(sk_c7),sk_c6) | (~spl0_5 | ~spl0_6)),
  inference(superposition,[],[f166,f188])).
fof(f156,plain,(
  ~spl0_21 | ~spl0_22 | ~spl0_2 | ~spl0_12),
  inference(avatar_split_clause,[],[f129,f110,f40,f153,f149])).
fof(f110,plain,(
  spl0_12 <=> ! [X3] : (sk_c8 != inverse(X3) | sk_c7 != multiply(X3,sk_c8))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_12])])).
fof(f129,plain,(
  sk_c7 != sk_c6 | sk_c8 != inverse(sk_c7) | (~spl0_2 | ~spl0_12)),
  inference(superposition,[],[f111,f42])).
fof(f111,plain,(
  ( ! [X3] : (sk_c7 != multiply(X3,sk_c8) | sk_c8 != inverse(X3)) ) | ~spl0_12),
  inference(avatar_component_clause,[],[f110])).
fof(f138,plain,(
  ~spl0_17 | ~spl0_18 | ~spl0_12),
  inference(avatar_split_clause,[],[f127,f110,f135,f131])).
fof(f127,plain,(
  sk_c8 != sk_c7 | sk_c8 != inverse(identity) | ~spl0_12),
  inference(superposition,[],[f111,f1])).
fof(f124,plain,(
  spl0_12 | spl0_13 | spl0_14 | ~spl0_2 | spl0_15 | spl0_16),
  inference(avatar_split_clause,[],[f34,f122,f119,f40,f116,f113,f110])).
fof(f34,axiom,(
  ( ! [X6,X4,X7,X5,X3] : (sk_c7 != multiply(X7,sk_c6) | sk_c7 != inverse(X7) | sk_c8 != multiply(X6,sk_c7) | sk_c8 != inverse(X6) | multiply(sk_c7,sk_c8) != sk_c6 | sk_c6 != multiply(X5,sk_c8) | sk_c6 != inverse(X5) | sk_c7 != inverse(X4) | sk_c6 != multiply(X4,sk_c7) | sk_c8 != inverse(X3) | sk_c7 != multiply(X3,sk_c8)) )),
  file('./p/GRP255-1.tptp',unknown)).
fof(f108,plain,(
  spl0_11 | spl0_6),
  inference(avatar_split_clause,[],[f33,f60,f101])).
fof(f33,axiom,(
  sk_c7 = multiply(sk_c5,sk_c6) | sk_c6 = multiply(sk_c3,sk_c8)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f107,plain,(
  spl0_11 | spl0_5),
  inference(avatar_split_clause,[],[f32,f55,f101])).
fof(f32,axiom,(
  sk_c7 = inverse(sk_c5) | sk_c6 = multiply(sk_c3,sk_c8)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f106,plain,(
  spl0_11 | spl0_4),
  inference(avatar_split_clause,[],[f31,f50,f101])).
fof(f31,axiom,(
  sk_c8 = multiply(sk_c4,sk_c7) | sk_c6 = multiply(sk_c3,sk_c8)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f105,plain,(
  spl0_11 | spl0_3),
  inference(avatar_split_clause,[],[f30,f45,f101])).
fof(f30,axiom,(
  sk_c8 = inverse(sk_c4) | sk_c6 = multiply(sk_c3,sk_c8)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f104,plain,(
  spl0_11 | spl0_2),
  inference(avatar_split_clause,[],[f29,f40,f101])).
fof(f29,axiom,(
  multiply(sk_c7,sk_c8) = sk_c6 | sk_c6 = multiply(sk_c3,sk_c8)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f99,plain,(
  spl0_10 | spl0_6),
  inference(avatar_split_clause,[],[f28,f60,f92])).
fof(f28,axiom,(
  sk_c7 = multiply(sk_c5,sk_c6) | sk_c6 = inverse(sk_c3)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f98,plain,(
  spl0_10 | spl0_5),
  inference(avatar_split_clause,[],[f27,f55,f92])).
fof(f27,axiom,(
  sk_c7 = inverse(sk_c5) | sk_c6 = inverse(sk_c3)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f97,plain,(
  spl0_10 | spl0_4),
  inference(avatar_split_clause,[],[f26,f50,f92])).
fof(f26,axiom,(
  sk_c8 = multiply(sk_c4,sk_c7) | sk_c6 = inverse(sk_c3)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f96,plain,(
  spl0_10 | spl0_3),
  inference(avatar_split_clause,[],[f25,f45,f92])).
fof(f25,axiom,(
  sk_c8 = inverse(sk_c4) | sk_c6 = inverse(sk_c3)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f95,plain,(
  spl0_10 | spl0_2),
  inference(avatar_split_clause,[],[f24,f40,f92])).
fof(f24,axiom,(
  multiply(sk_c7,sk_c8) = sk_c6 | sk_c6 = inverse(sk_c3)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f90,plain,(
  spl0_9 | spl0_6),
  inference(avatar_split_clause,[],[f23,f60,f83])).
fof(f23,axiom,(
  sk_c7 = multiply(sk_c5,sk_c6) | sk_c7 = inverse(sk_c2)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f89,plain,(
  spl0_9 | spl0_5),
  inference(avatar_split_clause,[],[f22,f55,f83])).
fof(f22,axiom,(
  sk_c7 = inverse(sk_c5) | sk_c7 = inverse(sk_c2)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f88,plain,(
  spl0_9 | spl0_4),
  inference(avatar_split_clause,[],[f21,f50,f83])).
fof(f21,axiom,(
  sk_c8 = multiply(sk_c4,sk_c7) | sk_c7 = inverse(sk_c2)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f87,plain,(
  spl0_9 | spl0_3),
  inference(avatar_split_clause,[],[f20,f45,f83])).
fof(f20,axiom,(
  sk_c8 = inverse(sk_c4) | sk_c7 = inverse(sk_c2)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f86,plain,(
  spl0_9 | spl0_2),
  inference(avatar_split_clause,[],[f19,f40,f83])).
fof(f19,axiom,(
  multiply(sk_c7,sk_c8) = sk_c6 | sk_c7 = inverse(sk_c2)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f81,plain,(
  spl0_8 | spl0_6),
  inference(avatar_split_clause,[],[f18,f60,f74])).
fof(f18,axiom,(
  sk_c7 = multiply(sk_c5,sk_c6) | sk_c6 = multiply(sk_c2,sk_c7)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f80,plain,(
  spl0_8 | spl0_5),
  inference(avatar_split_clause,[],[f17,f55,f74])).
fof(f17,axiom,(
  sk_c7 = inverse(sk_c5) | sk_c6 = multiply(sk_c2,sk_c7)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f79,plain,(
  spl0_8 | spl0_4),
  inference(avatar_split_clause,[],[f16,f50,f74])).
fof(f16,axiom,(
  sk_c8 = multiply(sk_c4,sk_c7) | sk_c6 = multiply(sk_c2,sk_c7)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f78,plain,(
  spl0_8 | spl0_3),
  inference(avatar_split_clause,[],[f15,f45,f74])).
fof(f15,axiom,(
  sk_c8 = inverse(sk_c4) | sk_c6 = multiply(sk_c2,sk_c7)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f77,plain,(
  spl0_8 | spl0_2),
  inference(avatar_split_clause,[],[f14,f40,f74])).
fof(f14,axiom,(
  multiply(sk_c7,sk_c8) = sk_c6 | sk_c6 = multiply(sk_c2,sk_c7)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f72,plain,(
  spl0_7 | spl0_6),
  inference(avatar_split_clause,[],[f13,f60,f65])).
fof(f13,axiom,(
  sk_c7 = multiply(sk_c5,sk_c6) | sk_c8 = inverse(sk_c1)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f71,plain,(
  spl0_7 | spl0_5),
  inference(avatar_split_clause,[],[f12,f55,f65])).
fof(f12,axiom,(
  sk_c7 = inverse(sk_c5) | sk_c8 = inverse(sk_c1)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f70,plain,(
  spl0_7 | spl0_4),
  inference(avatar_split_clause,[],[f11,f50,f65])).
fof(f11,axiom,(
  sk_c8 = multiply(sk_c4,sk_c7) | sk_c8 = inverse(sk_c1)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f69,plain,(
  spl0_7 | spl0_3),
  inference(avatar_split_clause,[],[f10,f45,f65])).
fof(f10,axiom,(
  sk_c8 = inverse(sk_c4) | sk_c8 = inverse(sk_c1)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f68,plain,(
  spl0_7 | spl0_2),
  inference(avatar_split_clause,[],[f9,f40,f65])).
fof(f9,axiom,(
  multiply(sk_c7,sk_c8) = sk_c6 | sk_c8 = inverse(sk_c1)),
  file('./p/GRP255-1.tptp',unknown)).
fof(f63,plain,(
  spl0_1 | spl0_6),
  inference(avatar_split_clause,[],[f8,f60,f36])).
fof(f8,axiom,(
  sk_c7 = multiply(sk_c5,sk_c6) | multiply(sk_c1,sk_c8) = sk_c7),
  file('./p/GRP255-1.tptp',unknown)).
fof(f58,plain,(
  spl0_1 | spl0_5),
  inference(avatar_split_clause,[],[f7,f55,f36])).
fof(f7,axiom,(
  sk_c7 = inverse(sk_c5) | multiply(sk_c1,sk_c8) = sk_c7),
  file('./p/GRP255-1.tptp',unknown)).
fof(f53,plain,(
  spl0_1 | spl0_4),
  inference(avatar_split_clause,[],[f6,f50,f36])).
fof(f6,axiom,(
  sk_c8 = multiply(sk_c4,sk_c7) | multiply(sk_c1,sk_c8) = sk_c7),
  file('./p/GRP255-1.tptp',unknown)).
fof(f48,plain,(
  spl0_1 | spl0_3),
  inference(avatar_split_clause,[],[f5,f45,f36])).
fof(f5,axiom,(
  sk_c8 = inverse(sk_c4) | multiply(sk_c1,sk_c8) = sk_c7),
  file('./p/GRP255-1.tptp',unknown)).
fof(f43,plain,(
  spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f4,f40,f36])).
fof(f4,axiom,(
  multiply(sk_c7,sk_c8) = sk_c6 | multiply(sk_c1,sk_c8) = sk_c7),
  file('./p/GRP255-1.tptp',unknown)).
% SZS output end Proof for GRP255-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.015 s
% ------------------------------
% ------------------------------
