% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GRP338-1
% SZS output start Proof for GRP338-1
fof(f1151,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f65,f70,f75,f80,f85,f90,f95,f100,f105,f110,f115,f116,f117,f118,f129,f130,f131,f132,f143,f144,f145,f146,f157,f158,f159,f160,f182,f342,f409,f552,f571,f621,f660,f746,f1013,f1047,f1068,f1091,f1150])).
fof(f1150,plain,(
  ~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_19),
  inference(avatar_contradiction_clause,[],[f1149])).
fof(f1149,plain,(
  $false | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_19)),
  inference(subsumption_resolution,[],[f1148,f843])).
fof(f843,plain,(
  identity = inverse(identity) | (~spl0_1 | ~spl0_12 | ~spl0_14 | ~spl0_15)),
  inference(backward_demodulation,[],[f837,f701])).
fof(f701,plain,(
  identity = inverse(sk_c1) | (~spl0_1 | ~spl0_12 | ~spl0_14 | ~spl0_15)),
  inference(backward_demodulation,[],[f696,f114])).
fof(f114,plain,(
  sk_c11 = inverse(sk_c1) | ~spl0_12),
  inference(avatar_component_clause,[],[f112])).
fof(f112,plain,(
  spl0_12 <=> sk_c11 = inverse(sk_c1)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_12])])).
fof(f696,plain,(
  identity = sk_c11 | (~spl0_1 | ~spl0_14 | ~spl0_15)),
  inference(forward_demodulation,[],[f692,f2])).
fof(f2,axiom,(
  ( ! [X0] : (identity = multiply(inverse(X0),X0)) )),
  file('./p/GRP338-1.tptp',unknown)).
fof(f692,plain,(
  sk_c11 = multiply(inverse(sk_c10),sk_c10) | (~spl0_1 | ~spl0_14 | ~spl0_15)),
  inference(backward_demodulation,[],[f690,f674])).
fof(f674,plain,(
  sk_c11 = multiply(inverse(sk_c10),sk_c9) | ~spl0_1),
  inference(superposition,[],[f198,f60])).
fof(f60,plain,(
  multiply(sk_c10,sk_c11) = sk_c9 | ~spl0_1),
  inference(avatar_component_clause,[],[f58])).
fof(f58,plain,(
  spl0_1 <=> multiply(sk_c10,sk_c11) = sk_c9),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f198,plain,(
  ( ! [X6,X7] : (multiply(inverse(X6),multiply(X6,X7)) = X7) )),
  inference(forward_demodulation,[],[f190,f1])).
fof(f1,axiom,(
  ( ! [X0] : (multiply(identity,X0) = X0) )),
  file('./p/GRP338-1.tptp',unknown)).
fof(f190,plain,(
  ( ! [X6,X7] : (multiply(inverse(X6),multiply(X6,X7)) = multiply(identity,X7)) )),
  inference(superposition,[],[f3,f2])).
fof(f3,axiom,(
  ( ! [X2,X0,X1] : (multiply(multiply(X0,X1),X2) = multiply(X0,multiply(X1,X2))) )),
  file('./p/GRP338-1.tptp',unknown)).
fof(f690,plain,(
  sk_c10 = sk_c9 | (~spl0_1 | ~spl0_14 | ~spl0_15)),
  inference(backward_demodulation,[],[f689,f60])).
fof(f689,plain,(
  sk_c10 = multiply(sk_c10,sk_c11) | (~spl0_14 | ~spl0_15)),
  inference(forward_demodulation,[],[f687,f156])).
fof(f156,plain,(
  sk_c10 = inverse(sk_c2) | ~spl0_15),
  inference(avatar_component_clause,[],[f154])).
fof(f154,plain,(
  spl0_15 <=> sk_c10 = inverse(sk_c2)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_15])])).
fof(f687,plain,(
  sk_c10 = multiply(inverse(sk_c2),sk_c11) | ~spl0_14),
  inference(superposition,[],[f198,f142])).
fof(f142,plain,(
  sk_c11 = multiply(sk_c2,sk_c10) | ~spl0_14),
  inference(avatar_component_clause,[],[f140])).
fof(f140,plain,(
  spl0_14 <=> sk_c11 = multiply(sk_c2,sk_c10)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_14])])).
fof(f837,plain,(
  identity = sk_c1 | (~spl0_1 | ~spl0_12 | ~spl0_14 | ~spl0_15)),
  inference(superposition,[],[f711,f1])).
fof(f711,plain,(
  identity = multiply(identity,sk_c1) | (~spl0_1 | ~spl0_12 | ~spl0_14 | ~spl0_15)),
  inference(backward_demodulation,[],[f696,f432])).
fof(f432,plain,(
  identity = multiply(sk_c11,sk_c1) | ~spl0_12),
  inference(superposition,[],[f2,f114])).
fof(f1148,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_19)),
  inference(forward_demodulation,[],[f1144,f843])).
fof(f1144,plain,(
  identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_19)),
  inference(trivial_inequality_removal,[],[f1143])).
fof(f1143,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_19)),
  inference(superposition,[],[f1100,f2])).
fof(f1100,plain,(
  ( ! [X6] : (identity != multiply(X6,identity) | identity != inverse(X6)) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_19)),
  inference(forward_demodulation,[],[f1099,f882])).
fof(f882,plain,(
  identity = sk_c9 | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15)),
  inference(forward_demodulation,[],[f690,f744])).
fof(f744,plain,(
  identity = sk_c10 | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15)),
  inference(forward_demodulation,[],[f713,f1])).
fof(f713,plain,(
  sk_c10 = multiply(identity,identity) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15)),
  inference(backward_demodulation,[],[f696,f686])).
fof(f686,plain,(
  sk_c10 = multiply(sk_c11,sk_c11) | (~spl0_12 | ~spl0_13)),
  inference(forward_demodulation,[],[f684,f114])).
fof(f684,plain,(
  sk_c10 = multiply(inverse(sk_c1),sk_c11) | ~spl0_13),
  inference(superposition,[],[f198,f128])).
fof(f128,plain,(
  sk_c11 = multiply(sk_c1,sk_c10) | ~spl0_13),
  inference(avatar_component_clause,[],[f126])).
fof(f126,plain,(
  spl0_13 <=> sk_c11 = multiply(sk_c1,sk_c10)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_13])])).
fof(f1099,plain,(
  ( ! [X6] : (sk_c9 != multiply(X6,identity) | identity != inverse(X6)) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_19)),
  inference(forward_demodulation,[],[f1098,f744])).
fof(f1098,plain,(
  ( ! [X6] : (identity != inverse(X6) | sk_c9 != multiply(X6,sk_c10)) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_19)),
  inference(forward_demodulation,[],[f178,f744])).
fof(f178,plain,(
  ( ! [X6] : (sk_c10 != inverse(X6) | sk_c9 != multiply(X6,sk_c10)) ) | ~spl0_19),
  inference(avatar_component_clause,[],[f177])).
fof(f177,plain,(
  spl0_19 <=> ! [X6] : (sk_c10 != inverse(X6) | sk_c9 != multiply(X6,sk_c10))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_19])])).
fof(f1091,plain,(
  ~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_16),
  inference(avatar_contradiction_clause,[],[f1090])).
fof(f1090,plain,(
  $false | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_16)),
  inference(subsumption_resolution,[],[f1089,f843])).
fof(f1089,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_16)),
  inference(forward_demodulation,[],[f1085,f843])).
fof(f1085,plain,(
  identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_16)),
  inference(trivial_inequality_removal,[],[f1080])).
fof(f1080,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_16)),
  inference(superposition,[],[f1072,f2])).
fof(f1072,plain,(
  ( ! [X3] : (identity != multiply(X3,identity) | identity != inverse(X3)) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_16)),
  inference(forward_demodulation,[],[f1071,f696])).
fof(f1071,plain,(
  ( ! [X3] : (identity != multiply(X3,identity) | sk_c11 != inverse(X3)) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_16)),
  inference(forward_demodulation,[],[f1070,f696])).
fof(f1070,plain,(
  ( ! [X3] : (sk_c11 != multiply(X3,identity) | sk_c11 != inverse(X3)) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_16)),
  inference(forward_demodulation,[],[f169,f744])).
fof(f169,plain,(
  ( ! [X3] : (sk_c11 != multiply(X3,sk_c10) | sk_c11 != inverse(X3)) ) | ~spl0_16),
  inference(avatar_component_clause,[],[f168])).
fof(f168,plain,(
  spl0_16 <=> ! [X3] : (sk_c11 != multiply(X3,sk_c10) | sk_c11 != inverse(X3))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_16])])).
fof(f1068,plain,(
  ~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_17),
  inference(avatar_contradiction_clause,[],[f1067])).
fof(f1067,plain,(
  $false | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_17)),
  inference(subsumption_resolution,[],[f1066,f843])).
fof(f1066,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_17)),
  inference(forward_demodulation,[],[f1062,f843])).
fof(f1062,plain,(
  identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_17)),
  inference(trivial_inequality_removal,[],[f1058])).
fof(f1058,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_17)),
  inference(superposition,[],[f1055,f2])).
fof(f1055,plain,(
  ( ! [X4] : (identity != multiply(X4,identity) | identity != inverse(X4)) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_17)),
  inference(forward_demodulation,[],[f1054,f696])).
fof(f1054,plain,(
  ( ! [X4] : (sk_c11 != multiply(X4,identity) | identity != inverse(X4)) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_17)),
  inference(forward_demodulation,[],[f1053,f744])).
fof(f1053,plain,(
  ( ! [X4] : (identity != inverse(X4) | sk_c11 != multiply(X4,sk_c10)) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_17)),
  inference(forward_demodulation,[],[f172,f744])).
fof(f172,plain,(
  ( ! [X4] : (sk_c10 != inverse(X4) | sk_c11 != multiply(X4,sk_c10)) ) | ~spl0_17),
  inference(avatar_component_clause,[],[f171])).
fof(f171,plain,(
  spl0_17 <=> ! [X4] : (sk_c10 != inverse(X4) | sk_c11 != multiply(X4,sk_c10))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_17])])).
fof(f1047,plain,(
  ~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_18),
  inference(avatar_contradiction_clause,[],[f1046])).
fof(f1046,plain,(
  $false | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_18)),
  inference(subsumption_resolution,[],[f1045,f843])).
fof(f1045,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_18)),
  inference(forward_demodulation,[],[f1041,f843])).
fof(f1041,plain,(
  identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_18)),
  inference(trivial_inequality_removal,[],[f1037])).
fof(f1037,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_18)),
  inference(superposition,[],[f1028,f2])).
fof(f1028,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | identity != inverse(X5)) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_18)),
  inference(forward_demodulation,[],[f1027,f744])).
fof(f1027,plain,(
  ( ! [X5] : (sk_c10 != multiply(X5,identity) | identity != inverse(X5)) ) | (~spl0_1 | ~spl0_14 | ~spl0_15 | ~spl0_18)),
  inference(forward_demodulation,[],[f1026,f696])).
fof(f1026,plain,(
  ( ! [X5] : (identity != inverse(X5) | sk_c10 != multiply(X5,sk_c11)) ) | (~spl0_1 | ~spl0_14 | ~spl0_15 | ~spl0_18)),
  inference(forward_demodulation,[],[f175,f696])).
fof(f175,plain,(
  ( ! [X5] : (sk_c11 != inverse(X5) | sk_c10 != multiply(X5,sk_c11)) ) | ~spl0_18),
  inference(avatar_component_clause,[],[f174])).
fof(f174,plain,(
  spl0_18 <=> ! [X5] : (sk_c11 != inverse(X5) | sk_c10 != multiply(X5,sk_c11))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_18])])).
fof(f1013,plain,(
  ~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f1012])).
fof(f1012,plain,(
  $false | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_20)),
  inference(subsumption_resolution,[],[f1011,f843])).
fof(f1011,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_20)),
  inference(forward_demodulation,[],[f1010,f2])).
fof(f1010,plain,(
  identity != inverse(multiply(inverse(identity),identity)) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_20)),
  inference(subsumption_resolution,[],[f1009,f843])).
fof(f1009,plain,(
  identity != inverse(identity) | identity != inverse(multiply(inverse(identity),identity)) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_20)),
  inference(forward_demodulation,[],[f1001,f843])).
fof(f1001,plain,(
  identity != inverse(inverse(identity)) | identity != inverse(multiply(inverse(identity),identity)) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_20)),
  inference(superposition,[],[f957,f2])).
fof(f957,plain,(
  ( ! [X1] : (multiply(X1,identity) != inverse(X1) | identity != inverse(multiply(X1,identity))) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_20)),
  inference(forward_demodulation,[],[f956,f843])).
fof(f956,plain,(
  ( ! [X1] : (identity != inverse(multiply(X1,identity)) | inverse(X1) != multiply(X1,inverse(identity))) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_20)),
  inference(subsumption_resolution,[],[f955,f1])).
fof(f955,plain,(
  ( ! [X1] : (identity != multiply(identity,identity) | identity != inverse(multiply(X1,identity)) | inverse(X1) != multiply(X1,inverse(identity))) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_20)),
  inference(forward_demodulation,[],[f954,f843])).
fof(f954,plain,(
  ( ! [X1] : (identity != inverse(multiply(X1,identity)) | identity != multiply(identity,inverse(identity)) | inverse(X1) != multiply(X1,inverse(identity))) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_20)),
  inference(subsumption_resolution,[],[f933,f2])).
fof(f933,plain,(
  ( ! [X1] : (identity != inverse(multiply(X1,identity)) | identity != multiply(identity,inverse(identity)) | identity != multiply(inverse(identity),identity) | inverse(X1) != multiply(X1,inverse(identity))) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_20)),
  inference(superposition,[],[f878,f843])).
fof(f878,plain,(
  ( ! [X7,X9] : (inverse(X7) != inverse(multiply(X9,inverse(X7))) | identity != multiply(X7,inverse(X7)) | identity != multiply(inverse(X7),identity) | inverse(X9) != multiply(X9,inverse(X7))) ) | (~spl0_1 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_20)),
  inference(forward_demodulation,[],[f716,f744])).
fof(f716,plain,(
  ( ! [X7,X9] : (identity != multiply(X7,inverse(X7)) | identity != multiply(inverse(X7),sk_c10) | inverse(X7) != inverse(multiply(X9,inverse(X7))) | inverse(X9) != multiply(X9,inverse(X7))) ) | (~spl0_1 | ~spl0_14 | ~spl0_15 | ~spl0_20)),
  inference(forward_demodulation,[],[f704,f696])).
fof(f704,plain,(
  ( ! [X7,X9] : (identity != multiply(inverse(X7),sk_c10) | inverse(X7) != inverse(multiply(X9,inverse(X7))) | sk_c11 != multiply(X7,inverse(X7)) | inverse(X9) != multiply(X9,inverse(X7))) ) | (~spl0_1 | ~spl0_14 | ~spl0_15 | ~spl0_20)),
  inference(backward_demodulation,[],[f696,f181])).
fof(f181,plain,(
  ( ! [X7,X9] : (inverse(X7) != inverse(multiply(X9,inverse(X7))) | sk_c11 != multiply(X7,inverse(X7)) | sk_c11 != multiply(inverse(X7),sk_c10) | inverse(X9) != multiply(X9,inverse(X7))) ) | ~spl0_20),
  inference(avatar_component_clause,[],[f180])).
fof(f180,plain,(
  spl0_20 <=> ! [X9,X7] : (inverse(X7) != inverse(multiply(X9,inverse(X7))) | sk_c11 != multiply(X7,inverse(X7)) | sk_c11 != multiply(inverse(X7),sk_c10) | inverse(X9) != multiply(X9,inverse(X7)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_20])])).
fof(f746,plain,(
  ~spl0_1 | spl0_2 | ~spl0_3 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15),
  inference(avatar_contradiction_clause,[],[f745])).
fof(f745,plain,(
  $false | (~spl0_1 | spl0_2 | ~spl0_3 | ~spl0_12 | ~spl0_13 | ~spl0_14 | ~spl0_15)),
  inference(subsumption_resolution,[],[f744,f725])).
fof(f725,plain,(
  identity != sk_c10 | (~spl0_1 | spl0_2 | ~spl0_3 | ~spl0_14 | ~spl0_15)),
  inference(forward_demodulation,[],[f722,f1])).
fof(f722,plain,(
  sk_c10 != multiply(identity,identity) | (~spl0_1 | spl0_2 | ~spl0_3 | ~spl0_14 | ~spl0_15)),
  inference(backward_demodulation,[],[f718,f697])).
fof(f697,plain,(
  sk_c10 != multiply(sk_c3,identity) | (~spl0_1 | spl0_2 | ~spl0_14 | ~spl0_15)),
  inference(backward_demodulation,[],[f696,f63])).
fof(f63,plain,(
  sk_c10 != multiply(sk_c3,sk_c11) | spl0_2),
  inference(avatar_component_clause,[],[f62])).
fof(f62,plain,(
  spl0_2 <=> sk_c10 = multiply(sk_c3,sk_c11)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f718,plain,(
  identity = sk_c3 | (~spl0_1 | ~spl0_3 | ~spl0_14 | ~spl0_15)),
  inference(forward_demodulation,[],[f708,f2])).
fof(f708,plain,(
  sk_c3 = multiply(inverse(identity),identity) | (~spl0_1 | ~spl0_3 | ~spl0_14 | ~spl0_15)),
  inference(backward_demodulation,[],[f696,f219])).
fof(f219,plain,(
  sk_c3 = multiply(inverse(sk_c11),identity) | ~spl0_3),
  inference(superposition,[],[f198,f183])).
fof(f183,plain,(
  identity = multiply(sk_c11,sk_c3) | ~spl0_3),
  inference(superposition,[],[f2,f69])).
fof(f69,plain,(
  sk_c11 = inverse(sk_c3) | ~spl0_3),
  inference(avatar_component_clause,[],[f67])).
fof(f67,plain,(
  spl0_3 <=> sk_c11 = inverse(sk_c3)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f660,plain,(
  ~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_18),
  inference(avatar_contradiction_clause,[],[f659])).
fof(f659,plain,(
  $false | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_18)),
  inference(subsumption_resolution,[],[f658,f336])).
fof(f336,plain,(
  identity = inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_5)),
  inference(backward_demodulation,[],[f333,f261])).
fof(f261,plain,(
  identity = inverse(sk_c4) | (~spl0_2 | ~spl0_3 | ~spl0_5)),
  inference(backward_demodulation,[],[f257,f79])).
fof(f79,plain,(
  sk_c10 = inverse(sk_c4) | ~spl0_5),
  inference(avatar_component_clause,[],[f77])).
fof(f77,plain,(
  spl0_5 <=> sk_c10 = inverse(sk_c4)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f257,plain,(
  identity = sk_c10 | (~spl0_2 | ~spl0_3)),
  inference(forward_demodulation,[],[f255,f2])).
fof(f255,plain,(
  sk_c10 = multiply(inverse(sk_c11),sk_c11) | (~spl0_2 | ~spl0_3)),
  inference(superposition,[],[f198,f229])).
fof(f229,plain,(
  sk_c11 = multiply(sk_c11,sk_c10) | (~spl0_2 | ~spl0_3)),
  inference(forward_demodulation,[],[f220,f69])).
fof(f220,plain,(
  sk_c11 = multiply(inverse(sk_c3),sk_c10) | ~spl0_2),
  inference(superposition,[],[f198,f64])).
fof(f64,plain,(
  sk_c10 = multiply(sk_c3,sk_c11) | ~spl0_2),
  inference(avatar_component_clause,[],[f62])).
fof(f333,plain,(
  identity = sk_c4 | (~spl0_2 | ~spl0_3 | ~spl0_5)),
  inference(forward_demodulation,[],[f268,f2])).
fof(f268,plain,(
  sk_c4 = multiply(inverse(identity),identity) | (~spl0_2 | ~spl0_3 | ~spl0_5)),
  inference(backward_demodulation,[],[f257,f218])).
fof(f218,plain,(
  sk_c4 = multiply(inverse(sk_c10),identity) | ~spl0_5),
  inference(superposition,[],[f198,f184])).
fof(f184,plain,(
  identity = multiply(sk_c10,sk_c4) | ~spl0_5),
  inference(superposition,[],[f2,f79])).
fof(f658,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_18)),
  inference(forward_demodulation,[],[f654,f336])).
fof(f654,plain,(
  identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_18)),
  inference(trivial_inequality_removal,[],[f653])).
fof(f653,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_18)),
  inference(superposition,[],[f626,f2])).
fof(f626,plain,(
  ( ! [X5] : (identity != multiply(X5,identity) | identity != inverse(X5)) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_18)),
  inference(forward_demodulation,[],[f625,f257])).
fof(f625,plain,(
  ( ! [X5] : (sk_c10 != multiply(X5,identity) | identity != inverse(X5)) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_18)),
  inference(forward_demodulation,[],[f624,f450])).
fof(f450,plain,(
  identity = sk_c11 | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5)),
  inference(superposition,[],[f407,f1])).
fof(f407,plain,(
  identity = multiply(identity,sk_c11) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5)),
  inference(forward_demodulation,[],[f352,f257])).
fof(f352,plain,(
  identity = multiply(sk_c10,sk_c11) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5)),
  inference(forward_demodulation,[],[f60,f340])).
fof(f340,plain,(
  identity = sk_c9 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5)),
  inference(forward_demodulation,[],[f337,f1])).
fof(f337,plain,(
  sk_c9 = multiply(identity,identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5)),
  inference(backward_demodulation,[],[f333,f260])).
fof(f260,plain,(
  sk_c9 = multiply(sk_c4,identity) | (~spl0_2 | ~spl0_3 | ~spl0_4)),
  inference(backward_demodulation,[],[f257,f74])).
fof(f74,plain,(
  sk_c9 = multiply(sk_c4,sk_c10) | ~spl0_4),
  inference(avatar_component_clause,[],[f72])).
fof(f72,plain,(
  spl0_4 <=> sk_c9 = multiply(sk_c4,sk_c10)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f624,plain,(
  ( ! [X5] : (identity != inverse(X5) | sk_c10 != multiply(X5,sk_c11)) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_18)),
  inference(forward_demodulation,[],[f175,f450])).
fof(f621,plain,(
  ~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20),
  inference(avatar_contradiction_clause,[],[f620])).
fof(f620,plain,(
  $false | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20)),
  inference(subsumption_resolution,[],[f619,f336])).
fof(f619,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20)),
  inference(forward_demodulation,[],[f618,f2])).
fof(f618,plain,(
  identity != inverse(multiply(inverse(identity),identity)) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20)),
  inference(subsumption_resolution,[],[f617,f336])).
fof(f617,plain,(
  identity != inverse(identity) | identity != inverse(multiply(inverse(identity),identity)) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20)),
  inference(forward_demodulation,[],[f608,f336])).
fof(f608,plain,(
  identity != inverse(inverse(identity)) | identity != inverse(multiply(inverse(identity),identity)) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20)),
  inference(superposition,[],[f594,f2])).
fof(f594,plain,(
  ( ! [X0] : (inverse(X0) != multiply(X0,identity) | identity != inverse(multiply(X0,identity))) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20)),
  inference(forward_demodulation,[],[f593,f336])).
fof(f593,plain,(
  ( ! [X0] : (identity != inverse(multiply(X0,identity)) | inverse(X0) != multiply(X0,inverse(identity))) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20)),
  inference(subsumption_resolution,[],[f592,f1])).
fof(f592,plain,(
  ( ! [X0] : (identity != multiply(identity,identity) | identity != inverse(multiply(X0,identity)) | inverse(X0) != multiply(X0,inverse(identity))) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20)),
  inference(forward_demodulation,[],[f591,f336])).
fof(f591,plain,(
  ( ! [X0] : (identity != inverse(multiply(X0,identity)) | identity != multiply(identity,inverse(identity)) | inverse(X0) != multiply(X0,inverse(identity))) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20)),
  inference(subsumption_resolution,[],[f587,f2])).
fof(f587,plain,(
  ( ! [X0] : (identity != inverse(multiply(X0,identity)) | identity != multiply(identity,inverse(identity)) | identity != multiply(inverse(identity),identity) | inverse(X0) != multiply(X0,inverse(identity))) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20)),
  inference(superposition,[],[f576,f336])).
fof(f576,plain,(
  ( ! [X7,X9] : (inverse(X7) != inverse(multiply(X9,inverse(X7))) | identity != multiply(X7,inverse(X7)) | identity != multiply(inverse(X7),identity) | inverse(X9) != multiply(X9,inverse(X7))) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20)),
  inference(forward_demodulation,[],[f575,f450])).
fof(f575,plain,(
  ( ! [X7,X9] : (sk_c11 != multiply(inverse(X7),identity) | identity != multiply(X7,inverse(X7)) | inverse(X7) != inverse(multiply(X9,inverse(X7))) | inverse(X9) != multiply(X9,inverse(X7))) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20)),
  inference(forward_demodulation,[],[f574,f257])).
fof(f574,plain,(
  ( ! [X7,X9] : (identity != multiply(X7,inverse(X7)) | inverse(X7) != inverse(multiply(X9,inverse(X7))) | sk_c11 != multiply(inverse(X7),sk_c10) | inverse(X9) != multiply(X9,inverse(X7))) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_20)),
  inference(forward_demodulation,[],[f181,f450])).
fof(f571,plain,(
  ~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_16),
  inference(avatar_contradiction_clause,[],[f570])).
fof(f570,plain,(
  $false | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_16)),
  inference(subsumption_resolution,[],[f569,f336])).
fof(f569,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_16)),
  inference(forward_demodulation,[],[f565,f336])).
fof(f565,plain,(
  identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_16)),
  inference(trivial_inequality_removal,[],[f560])).
fof(f560,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_16)),
  inference(superposition,[],[f557,f2])).
fof(f557,plain,(
  ( ! [X3] : (identity != multiply(X3,identity) | identity != inverse(X3)) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_16)),
  inference(forward_demodulation,[],[f556,f450])).
fof(f556,plain,(
  ( ! [X3] : (identity != multiply(X3,identity) | sk_c11 != inverse(X3)) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_16)),
  inference(forward_demodulation,[],[f555,f450])).
fof(f555,plain,(
  ( ! [X3] : (sk_c11 != multiply(X3,identity) | sk_c11 != inverse(X3)) ) | (~spl0_2 | ~spl0_3 | ~spl0_16)),
  inference(forward_demodulation,[],[f169,f257])).
fof(f552,plain,(
  ~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_17),
  inference(avatar_contradiction_clause,[],[f551])).
fof(f551,plain,(
  $false | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_17)),
  inference(subsumption_resolution,[],[f550,f336])).
fof(f550,plain,(
  identity != inverse(identity) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_17)),
  inference(forward_demodulation,[],[f546,f336])).
fof(f546,plain,(
  identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_17)),
  inference(trivial_inequality_removal,[],[f541])).
fof(f541,plain,(
  identity != identity | identity != inverse(inverse(identity)) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_17)),
  inference(superposition,[],[f476,f2])).
fof(f476,plain,(
  ( ! [X4] : (identity != multiply(X4,identity) | identity != inverse(X4)) ) | (~spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_17)),
  inference(backward_demodulation,[],[f450,f426])).
fof(f426,plain,(
  ( ! [X4] : (sk_c11 != multiply(X4,identity) | identity != inverse(X4)) ) | (~spl0_2 | ~spl0_3 | ~spl0_17)),
  inference(forward_demodulation,[],[f425,f257])).
fof(f425,plain,(
  ( ! [X4] : (identity != inverse(X4) | sk_c11 != multiply(X4,sk_c10)) ) | (~spl0_2 | ~spl0_3 | ~spl0_17)),
  inference(forward_demodulation,[],[f172,f257])).
fof(f409,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_19),
  inference(avatar_contradiction_clause,[],[f408])).
fof(f408,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_19)),
  inference(subsumption_resolution,[],[f336,f401])).
fof(f401,plain,(
  identity != inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_19)),
  inference(trivial_inequality_removal,[],[f397])).
fof(f397,plain,(
  identity != identity | identity != inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_19)),
  inference(superposition,[],[f396,f1])).
fof(f396,plain,(
  ( ! [X6] : (identity != multiply(X6,identity) | identity != inverse(X6)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_19)),
  inference(forward_demodulation,[],[f395,f340])).
fof(f395,plain,(
  ( ! [X6] : (sk_c9 != multiply(X6,identity) | identity != inverse(X6)) ) | (~spl0_2 | ~spl0_3 | ~spl0_19)),
  inference(forward_demodulation,[],[f394,f257])).
fof(f394,plain,(
  ( ! [X6] : (identity != inverse(X6) | sk_c9 != multiply(X6,sk_c10)) ) | (~spl0_2 | ~spl0_3 | ~spl0_19)),
  inference(forward_demodulation,[],[f178,f257])).
fof(f342,plain,(
  spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11),
  inference(avatar_contradiction_clause,[],[f341])).
fof(f341,plain,(
  $false | (spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(subsumption_resolution,[],[f340,f301])).
fof(f301,plain,(
  identity != sk_c9 | (spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(backward_demodulation,[],[f294,f272])).
fof(f272,plain,(
  sk_c11 != sk_c9 | (spl0_1 | ~spl0_2 | ~spl0_3)),
  inference(forward_demodulation,[],[f258,f1])).
fof(f258,plain,(
  sk_c9 != multiply(identity,sk_c11) | (spl0_1 | ~spl0_2 | ~spl0_3)),
  inference(backward_demodulation,[],[f257,f59])).
fof(f59,plain,(
  multiply(sk_c10,sk_c11) != sk_c9 | spl0_1),
  inference(avatar_component_clause,[],[f58])).
fof(f294,plain,(
  identity = sk_c11 | (~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(backward_demodulation,[],[f291,f262])).
fof(f262,plain,(
  sk_c11 = multiply(sk_c8,identity) | (~spl0_2 | ~spl0_3 | ~spl0_8)),
  inference(backward_demodulation,[],[f257,f94])).
fof(f94,plain,(
  sk_c11 = multiply(sk_c8,sk_c10) | ~spl0_8),
  inference(avatar_component_clause,[],[f92])).
fof(f92,plain,(
  spl0_8 <=> sk_c11 = multiply(sk_c8,sk_c10)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_8])])).
fof(f291,plain,(
  ( ! [X0] : (multiply(sk_c8,X0) = X0) ) | (~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(backward_demodulation,[],[f290,f204])).
fof(f204,plain,(
  ( ! [X0] : (multiply(sk_c8,multiply(sk_c5,X0)) = X0) ) | ~spl0_7),
  inference(forward_demodulation,[],[f203,f1])).
fof(f203,plain,(
  ( ! [X0] : (multiply(identity,X0) = multiply(sk_c8,multiply(sk_c5,X0))) ) | ~spl0_7),
  inference(superposition,[],[f3,f185])).
fof(f185,plain,(
  identity = multiply(sk_c8,sk_c5) | ~spl0_7),
  inference(superposition,[],[f2,f89])).
fof(f89,plain,(
  sk_c8 = inverse(sk_c5) | ~spl0_7),
  inference(avatar_component_clause,[],[f87])).
fof(f87,plain,(
  spl0_7 <=> sk_c8 = inverse(sk_c5)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).
fof(f290,plain,(
  ( ! [X12] : (multiply(sk_c5,X12) = X12) ) | (~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(forward_demodulation,[],[f286,f285])).
fof(f285,plain,(
  ( ! [X10] : (multiply(sk_c11,X10) = X10) ) | (~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f284,f274])).
fof(f274,plain,(
  ( ! [X0] : (multiply(sk_c5,multiply(sk_c11,X0)) = X0) ) | (~spl0_2 | ~spl0_3 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f273,f239])).
fof(f239,plain,(
  ( ! [X0] : (multiply(sk_c5,multiply(sk_c7,X0)) = X0) ) | (~spl0_7 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f232,f206])).
fof(f206,plain,(
  ( ! [X0] : (multiply(sk_c6,multiply(sk_c7,X0)) = X0) ) | ~spl0_9),
  inference(forward_demodulation,[],[f205,f1])).
fof(f205,plain,(
  ( ! [X0] : (multiply(identity,X0) = multiply(sk_c6,multiply(sk_c7,X0))) ) | ~spl0_9),
  inference(superposition,[],[f3,f186])).
fof(f186,plain,(
  identity = multiply(sk_c6,sk_c7) | ~spl0_9),
  inference(superposition,[],[f2,f99])).
fof(f99,plain,(
  inverse(sk_c7) = sk_c6 | ~spl0_9),
  inference(avatar_component_clause,[],[f97])).
fof(f97,plain,(
  spl0_9 <=> inverse(sk_c7) = sk_c6),
  introduced(avatar_definition,[new_symbols(naming,[spl0_9])])).
fof(f232,plain,(
  sk_c5 = sk_c6 | (~spl0_7 | ~spl0_10)),
  inference(forward_demodulation,[],[f225,f224])).
fof(f224,plain,(
  sk_c5 = multiply(inverse(sk_c8),identity) | ~spl0_7),
  inference(superposition,[],[f198,f185])).
fof(f225,plain,(
  sk_c6 = multiply(inverse(sk_c8),identity) | ~spl0_10),
  inference(superposition,[],[f198,f187])).
fof(f187,plain,(
  identity = multiply(sk_c8,sk_c6) | ~spl0_10),
  inference(superposition,[],[f2,f104])).
fof(f104,plain,(
  sk_c8 = inverse(sk_c6) | ~spl0_10),
  inference(avatar_component_clause,[],[f102])).
fof(f102,plain,(
  spl0_10 <=> sk_c8 = inverse(sk_c6)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f273,plain,(
  sk_c11 = sk_c7 | (~spl0_2 | ~spl0_3 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f262,f244])).
fof(f244,plain,(
  sk_c7 = multiply(sk_c8,identity) | (~spl0_7 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f243,f89])).
fof(f243,plain,(
  sk_c7 = multiply(inverse(sk_c5),identity) | (~spl0_7 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f227,f232])).
fof(f227,plain,(
  sk_c7 = multiply(inverse(sk_c6),identity) | ~spl0_9),
  inference(superposition,[],[f198,f186])).
fof(f284,plain,(
  ( ! [X10] : (multiply(sk_c11,X10) = multiply(sk_c5,multiply(sk_c11,X10))) ) | (~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_8)),
  inference(backward_demodulation,[],[f282,f193])).
fof(f193,plain,(
  ( ! [X10] : (multiply(sk_c5,multiply(sk_c8,X10)) = multiply(sk_c11,X10)) ) | ~spl0_6),
  inference(superposition,[],[f3,f84])).
fof(f84,plain,(
  sk_c11 = multiply(sk_c5,sk_c8) | ~spl0_6),
  inference(avatar_component_clause,[],[f82])).
fof(f82,plain,(
  spl0_6 <=> sk_c11 = multiply(sk_c5,sk_c8)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).
fof(f282,plain,(
  ( ! [X11] : (multiply(sk_c11,X11) = multiply(sk_c8,X11)) ) | (~spl0_2 | ~spl0_3 | ~spl0_8)),
  inference(forward_demodulation,[],[f266,f1])).
fof(f266,plain,(
  ( ! [X11] : (multiply(sk_c11,X11) = multiply(sk_c8,multiply(identity,X11))) ) | (~spl0_2 | ~spl0_3 | ~spl0_8)),
  inference(backward_demodulation,[],[f257,f194])).
fof(f194,plain,(
  ( ! [X11] : (multiply(sk_c8,multiply(sk_c10,X11)) = multiply(sk_c11,X11)) ) | ~spl0_8),
  inference(superposition,[],[f3,f94])).
fof(f286,plain,(
  ( ! [X12] : (multiply(sk_c5,X12) = multiply(sk_c11,X12)) ) | (~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(backward_demodulation,[],[f285,f283])).
fof(f283,plain,(
  ( ! [X12] : (multiply(sk_c5,X12) = multiply(sk_c11,multiply(sk_c11,X12))) ) | (~spl0_2 | ~spl0_3 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(backward_demodulation,[],[f282,f275])).
fof(f275,plain,(
  ( ! [X12] : (multiply(sk_c5,X12) = multiply(sk_c11,multiply(sk_c8,X12))) ) | (~spl0_2 | ~spl0_3 | ~spl0_7 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_11)),
  inference(backward_demodulation,[],[f273,f238])).
fof(f238,plain,(
  ( ! [X12] : (multiply(sk_c7,multiply(sk_c8,X12)) = multiply(sk_c5,X12)) ) | (~spl0_7 | ~spl0_10 | ~spl0_11)),
  inference(backward_demodulation,[],[f232,f195])).
fof(f195,plain,(
  ( ! [X12] : (multiply(sk_c7,multiply(sk_c8,X12)) = multiply(sk_c6,X12)) ) | ~spl0_11),
  inference(superposition,[],[f3,f109])).
fof(f109,plain,(
  sk_c6 = multiply(sk_c7,sk_c8) | ~spl0_11),
  inference(avatar_component_clause,[],[f107])).
fof(f107,plain,(
  spl0_11 <=> sk_c6 = multiply(sk_c7,sk_c8)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_11])])).
fof(f182,plain,(
  ~spl0_1 | spl0_16 | spl0_17 | spl0_18 | spl0_19 | spl0_20),
  inference(avatar_split_clause,[],[f56,f180,f177,f174,f171,f168,f58])).
fof(f56,plain,(
  ( ! [X6,X4,X7,X5,X3,X9] : (inverse(X7) != inverse(multiply(X9,inverse(X7))) | inverse(X9) != multiply(X9,inverse(X7)) | sk_c11 != multiply(inverse(X7),sk_c10) | sk_c11 != multiply(X7,inverse(X7)) | sk_c10 != inverse(X6) | sk_c9 != multiply(X6,sk_c10) | sk_c11 != inverse(X5) | sk_c10 != multiply(X5,sk_c11) | sk_c10 != inverse(X4) | sk_c11 != multiply(X4,sk_c10) | sk_c11 != multiply(X3,sk_c10) | sk_c11 != inverse(X3) | multiply(sk_c10,sk_c11) != sk_c9) )),
  inference(equality_resolution,[],[f55])).
fof(f55,plain,(
  ( ! [X6,X4,X8,X7,X5,X3,X9] : (inverse(multiply(X9,X8)) != X8 | inverse(X9) != multiply(X9,X8) | sk_c11 != multiply(X8,sk_c10) | inverse(X7) != X8 | sk_c11 != multiply(X7,X8) | sk_c10 != inverse(X6) | sk_c9 != multiply(X6,sk_c10) | sk_c11 != inverse(X5) | sk_c10 != multiply(X5,sk_c11) | sk_c10 != inverse(X4) | sk_c11 != multiply(X4,sk_c10) | sk_c11 != multiply(X3,sk_c10) | sk_c11 != inverse(X3) | multiply(sk_c10,sk_c11) != sk_c9) )),
  inference(equality_resolution,[],[f54])).
fof(f54,axiom,(
  ( ! [X6,X4,X10,X8,X7,X5,X3,X9] : (multiply(X9,X8) != X10 | inverse(X10) != X8 | inverse(X9) != X10 | sk_c11 != multiply(X8,sk_c10) | inverse(X7) != X8 | sk_c11 != multiply(X7,X8) | sk_c10 != inverse(X6) | sk_c9 != multiply(X6,sk_c10) | sk_c11 != inverse(X5) | sk_c10 != multiply(X5,sk_c11) | sk_c10 != inverse(X4) | sk_c11 != multiply(X4,sk_c10) | sk_c11 != multiply(X3,sk_c10) | sk_c11 != inverse(X3) | multiply(sk_c10,sk_c11) != sk_c9) )),
  file('./p/GRP338-1.tptp',unknown)).
fof(f160,plain,(
  spl0_15 | spl0_5),
  inference(avatar_split_clause,[],[f47,f77,f154])).
fof(f47,axiom,(
  sk_c10 = inverse(sk_c4) | sk_c10 = inverse(sk_c2)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f159,plain,(
  spl0_15 | spl0_4),
  inference(avatar_split_clause,[],[f46,f72,f154])).
fof(f46,axiom,(
  sk_c9 = multiply(sk_c4,sk_c10) | sk_c10 = inverse(sk_c2)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f158,plain,(
  spl0_15 | spl0_3),
  inference(avatar_split_clause,[],[f45,f67,f154])).
fof(f45,axiom,(
  sk_c11 = inverse(sk_c3) | sk_c10 = inverse(sk_c2)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f157,plain,(
  spl0_15 | spl0_2),
  inference(avatar_split_clause,[],[f44,f62,f154])).
fof(f44,axiom,(
  sk_c10 = multiply(sk_c3,sk_c11) | sk_c10 = inverse(sk_c2)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f146,plain,(
  spl0_14 | spl0_5),
  inference(avatar_split_clause,[],[f37,f77,f140])).
fof(f37,axiom,(
  sk_c10 = inverse(sk_c4) | sk_c11 = multiply(sk_c2,sk_c10)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f145,plain,(
  spl0_14 | spl0_4),
  inference(avatar_split_clause,[],[f36,f72,f140])).
fof(f36,axiom,(
  sk_c9 = multiply(sk_c4,sk_c10) | sk_c11 = multiply(sk_c2,sk_c10)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f144,plain,(
  spl0_14 | spl0_3),
  inference(avatar_split_clause,[],[f35,f67,f140])).
fof(f35,axiom,(
  sk_c11 = inverse(sk_c3) | sk_c11 = multiply(sk_c2,sk_c10)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f143,plain,(
  spl0_14 | spl0_2),
  inference(avatar_split_clause,[],[f34,f62,f140])).
fof(f34,axiom,(
  sk_c10 = multiply(sk_c3,sk_c11) | sk_c11 = multiply(sk_c2,sk_c10)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f132,plain,(
  spl0_13 | spl0_5),
  inference(avatar_split_clause,[],[f27,f77,f126])).
fof(f27,axiom,(
  sk_c10 = inverse(sk_c4) | sk_c11 = multiply(sk_c1,sk_c10)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f131,plain,(
  spl0_13 | spl0_4),
  inference(avatar_split_clause,[],[f26,f72,f126])).
fof(f26,axiom,(
  sk_c9 = multiply(sk_c4,sk_c10) | sk_c11 = multiply(sk_c1,sk_c10)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f130,plain,(
  spl0_13 | spl0_3),
  inference(avatar_split_clause,[],[f25,f67,f126])).
fof(f25,axiom,(
  sk_c11 = inverse(sk_c3) | sk_c11 = multiply(sk_c1,sk_c10)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f129,plain,(
  spl0_13 | spl0_2),
  inference(avatar_split_clause,[],[f24,f62,f126])).
fof(f24,axiom,(
  sk_c10 = multiply(sk_c3,sk_c11) | sk_c11 = multiply(sk_c1,sk_c10)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f118,plain,(
  spl0_12 | spl0_5),
  inference(avatar_split_clause,[],[f17,f77,f112])).
fof(f17,axiom,(
  sk_c10 = inverse(sk_c4) | sk_c11 = inverse(sk_c1)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f117,plain,(
  spl0_12 | spl0_4),
  inference(avatar_split_clause,[],[f16,f72,f112])).
fof(f16,axiom,(
  sk_c9 = multiply(sk_c4,sk_c10) | sk_c11 = inverse(sk_c1)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f116,plain,(
  spl0_12 | spl0_3),
  inference(avatar_split_clause,[],[f15,f67,f112])).
fof(f15,axiom,(
  sk_c11 = inverse(sk_c3) | sk_c11 = inverse(sk_c1)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f115,plain,(
  spl0_12 | spl0_2),
  inference(avatar_split_clause,[],[f14,f62,f112])).
fof(f14,axiom,(
  sk_c10 = multiply(sk_c3,sk_c11) | sk_c11 = inverse(sk_c1)),
  file('./p/GRP338-1.tptp',unknown)).
fof(f110,plain,(
  spl0_1 | spl0_11),
  inference(avatar_split_clause,[],[f13,f107,f58])).
fof(f13,axiom,(
  sk_c6 = multiply(sk_c7,sk_c8) | multiply(sk_c10,sk_c11) = sk_c9),
  file('./p/GRP338-1.tptp',unknown)).
fof(f105,plain,(
  spl0_1 | spl0_10),
  inference(avatar_split_clause,[],[f12,f102,f58])).
fof(f12,axiom,(
  sk_c8 = inverse(sk_c6) | multiply(sk_c10,sk_c11) = sk_c9),
  file('./p/GRP338-1.tptp',unknown)).
fof(f100,plain,(
  spl0_1 | spl0_9),
  inference(avatar_split_clause,[],[f11,f97,f58])).
fof(f11,axiom,(
  inverse(sk_c7) = sk_c6 | multiply(sk_c10,sk_c11) = sk_c9),
  file('./p/GRP338-1.tptp',unknown)).
fof(f95,plain,(
  spl0_1 | spl0_8),
  inference(avatar_split_clause,[],[f10,f92,f58])).
fof(f10,axiom,(
  sk_c11 = multiply(sk_c8,sk_c10) | multiply(sk_c10,sk_c11) = sk_c9),
  file('./p/GRP338-1.tptp',unknown)).
fof(f90,plain,(
  spl0_1 | spl0_7),
  inference(avatar_split_clause,[],[f9,f87,f58])).
fof(f9,axiom,(
  sk_c8 = inverse(sk_c5) | multiply(sk_c10,sk_c11) = sk_c9),
  file('./p/GRP338-1.tptp',unknown)).
fof(f85,plain,(
  spl0_1 | spl0_6),
  inference(avatar_split_clause,[],[f8,f82,f58])).
fof(f8,axiom,(
  sk_c11 = multiply(sk_c5,sk_c8) | multiply(sk_c10,sk_c11) = sk_c9),
  file('./p/GRP338-1.tptp',unknown)).
fof(f80,plain,(
  spl0_1 | spl0_5),
  inference(avatar_split_clause,[],[f7,f77,f58])).
fof(f7,axiom,(
  sk_c10 = inverse(sk_c4) | multiply(sk_c10,sk_c11) = sk_c9),
  file('./p/GRP338-1.tptp',unknown)).
fof(f75,plain,(
  spl0_1 | spl0_4),
  inference(avatar_split_clause,[],[f6,f72,f58])).
fof(f6,axiom,(
  sk_c9 = multiply(sk_c4,sk_c10) | multiply(sk_c10,sk_c11) = sk_c9),
  file('./p/GRP338-1.tptp',unknown)).
fof(f70,plain,(
  spl0_1 | spl0_3),
  inference(avatar_split_clause,[],[f5,f67,f58])).
fof(f5,axiom,(
  sk_c11 = inverse(sk_c3) | multiply(sk_c10,sk_c11) = sk_c9),
  file('./p/GRP338-1.tptp',unknown)).
fof(f65,plain,(
  spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f4,f62,f58])).
fof(f4,axiom,(
  sk_c10 = multiply(sk_c3,sk_c11) | multiply(sk_c10,sk_c11) = sk_c9),
  file('./p/GRP338-1.tptp',unknown)).
% SZS output end Proof for GRP338-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.009 s
% ------------------------------
% ------------------------------
