% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GRP389-1
% SZS output start Proof for GRP389-1
fof(f842,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f69,f74,f79,f84,f99,f104,f109,f114,f115,f118,f119,f127,f128,f131,f132,f140,f141,f142,f143,f144,f145,f146,f147,f148,f153,f154,f155,f156,f157,f158,f159,f160,f161,f166,f167,f168,f169,f172,f173,f174,f187,f302,f331,f348,f359,f482,f523,f592,f705,f711,f819,f840])).
fof(f840,plain,(
  ~spl0_11 | ~spl0_12 | ~spl0_18 | ~spl0_27),
  inference(avatar_contradiction_clause,[],[f839])).
fof(f839,plain,(
  $false | (~spl0_11 | ~spl0_12 | ~spl0_18 | ~spl0_27)),
  inference(subsumption_resolution,[],[f838,f834])).
fof(f834,plain,(
  sk_c11 != inverse(sk_c11) | (~spl0_11 | ~spl0_12 | ~spl0_18 | ~spl0_27)),
  inference(trivial_inequality_removal,[],[f831])).
fof(f831,plain,(
  identity != identity | sk_c11 != inverse(sk_c11) | (~spl0_11 | ~spl0_12 | ~spl0_18 | ~spl0_27)),
  inference(superposition,[],[f822,f803])).
fof(f803,plain,(
  identity = multiply(sk_c11,sk_c11) | (~spl0_11 | ~spl0_12 | ~spl0_27)),
  inference(forward_demodulation,[],[f444,f489])).
fof(f489,plain,(
  identity = sk_c10 | ~spl0_27),
  inference(avatar_component_clause,[],[f488])).
fof(f488,plain,(
  spl0_27 <=> identity = sk_c10),
  introduced(avatar_definition,[new_symbols(naming,[spl0_27])])).
fof(f444,plain,(
  sk_c10 = multiply(sk_c11,sk_c11) | (~spl0_11 | ~spl0_12)),
  inference(forward_demodulation,[],[f442,f113])).
fof(f113,plain,(
  sk_c11 = inverse(sk_c1) | ~spl0_11),
  inference(avatar_component_clause,[],[f111])).
fof(f111,plain,(
  spl0_11 <=> sk_c11 = inverse(sk_c1)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_11])])).
fof(f442,plain,(
  sk_c10 = multiply(inverse(sk_c1),sk_c11) | ~spl0_12),
  inference(superposition,[],[f203,f126])).
fof(f126,plain,(
  sk_c11 = multiply(sk_c1,sk_c10) | ~spl0_12),
  inference(avatar_component_clause,[],[f124])).
fof(f124,plain,(
  spl0_12 <=> sk_c11 = multiply(sk_c1,sk_c10)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_12])])).
fof(f203,plain,(
  ( ! [X6,X7] : (multiply(inverse(X6),multiply(X6,X7)) = X7) )),
  inference(forward_demodulation,[],[f194,f1])).
fof(f1,axiom,(
  ( ! [X0] : (multiply(identity,X0) = X0) )),
  file('./p/GRP389-1.tptp',unknown)).
fof(f194,plain,(
  ( ! [X6,X7] : (multiply(inverse(X6),multiply(X6,X7)) = multiply(identity,X7)) )),
  inference(superposition,[],[f3,f2])).
fof(f2,axiom,(
  ( ! [X0] : (identity = multiply(inverse(X0),X0)) )),
  file('./p/GRP389-1.tptp',unknown)).
fof(f3,axiom,(
  ( ! [X2,X0,X1] : (multiply(multiply(X0,X1),X2) = multiply(X0,multiply(X1,X2))) )),
  file('./p/GRP389-1.tptp',unknown)).
fof(f822,plain,(
  ( ! [X6] : (identity != multiply(X6,sk_c11) | sk_c11 != inverse(X6)) ) | (~spl0_18 | ~spl0_27)),
  inference(forward_demodulation,[],[f183,f489])).
fof(f183,plain,(
  ( ! [X6] : (sk_c11 != inverse(X6) | sk_c10 != multiply(X6,sk_c11)) ) | ~spl0_18),
  inference(avatar_component_clause,[],[f182])).
fof(f182,plain,(
  spl0_18 <=> ! [X6] : (sk_c11 != inverse(X6) | sk_c10 != multiply(X6,sk_c11))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_18])])).
fof(f838,plain,(
  sk_c11 = inverse(sk_c11) | (~spl0_11 | ~spl0_12 | ~spl0_27)),
  inference(backward_demodulation,[],[f835,f113])).
fof(f835,plain,(
  sk_c11 = sk_c1 | (~spl0_11 | ~spl0_12 | ~spl0_27)),
  inference(backward_demodulation,[],[f832,f548])).
fof(f548,plain,(
  sk_c1 = multiply(inverse(sk_c11),identity) | ~spl0_11),
  inference(superposition,[],[f203,f429])).
fof(f429,plain,(
  identity = multiply(sk_c11,sk_c1) | ~spl0_11),
  inference(superposition,[],[f2,f113])).
fof(f832,plain,(
  sk_c11 = multiply(inverse(sk_c11),identity) | (~spl0_11 | ~spl0_12 | ~spl0_27)),
  inference(superposition,[],[f203,f803])).
fof(f819,plain,(
  ~spl0_11 | ~spl0_12 | ~spl0_16 | ~spl0_27),
  inference(avatar_contradiction_clause,[],[f818])).
fof(f818,plain,(
  $false | (~spl0_11 | ~spl0_12 | ~spl0_16 | ~spl0_27)),
  inference(subsumption_resolution,[],[f817,f113])).
fof(f817,plain,(
  sk_c11 != inverse(sk_c1) | (~spl0_12 | ~spl0_16 | ~spl0_27)),
  inference(trivial_inequality_removal,[],[f814])).
fof(f814,plain,(
  sk_c11 != sk_c11 | sk_c11 != inverse(sk_c1) | (~spl0_12 | ~spl0_16 | ~spl0_27)),
  inference(superposition,[],[f736,f799])).
fof(f799,plain,(
  sk_c11 = multiply(sk_c1,identity) | (~spl0_12 | ~spl0_27)),
  inference(forward_demodulation,[],[f126,f489])).
fof(f736,plain,(
  ( ! [X3] : (sk_c11 != multiply(X3,identity) | sk_c11 != inverse(X3)) ) | (~spl0_16 | ~spl0_27)),
  inference(forward_demodulation,[],[f177,f489])).
fof(f177,plain,(
  ( ! [X3] : (sk_c11 != multiply(X3,sk_c10) | sk_c11 != inverse(X3)) ) | ~spl0_16),
  inference(avatar_component_clause,[],[f176])).
fof(f176,plain,(
  spl0_16 <=> ! [X3] : (sk_c11 != multiply(X3,sk_c10) | sk_c11 != inverse(X3))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_16])])).
fof(f711,plain,(
  ~spl0_30 | ~spl0_1 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_19 | ~spl0_27),
  inference(avatar_split_clause,[],[f710,f488,f185,f163,f150,f137,f62,f514])).
fof(f514,plain,(
  spl0_30 <=> identity = sk_c9),
  introduced(avatar_definition,[new_symbols(naming,[spl0_30])])).
fof(f62,plain,(
  spl0_1 <=> inverse(sk_c10) = sk_c9),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f137,plain,(
  spl0_13 <=> sk_c10 = multiply(sk_c2,sk_c3)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_13])])).
fof(f150,plain,(
  spl0_14 <=> sk_c3 = inverse(sk_c2)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_14])])).
fof(f163,plain,(
  spl0_15 <=> sk_c10 = multiply(sk_c3,sk_c9)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_15])])).
fof(f185,plain,(
  spl0_19 <=> ! [X7] : (sk_c10 != inverse(X7) | sk_c9 != multiply(X7,sk_c10))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_19])])).
fof(f710,plain,(
  identity != sk_c9 | (~spl0_1 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_19 | ~spl0_27)),
  inference(subsumption_resolution,[],[f709,f489])).
fof(f709,plain,(
  identity != sk_c10 | identity != sk_c9 | (~spl0_1 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_19 | ~spl0_27)),
  inference(forward_demodulation,[],[f544,f702])).
fof(f702,plain,(
  identity = inverse(sk_c9) | (~spl0_1 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_27)),
  inference(forward_demodulation,[],[f660,f663])).
fof(f663,plain,(
  identity = sk_c3 | (~spl0_1 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_27)),
  inference(backward_demodulation,[],[f661,f613])).
fof(f613,plain,(
  sk_c3 = multiply(sk_c3,identity) | (~spl0_13 | ~spl0_14 | ~spl0_27)),
  inference(backward_demodulation,[],[f489,f447])).
fof(f447,plain,(
  sk_c3 = multiply(sk_c3,sk_c10) | (~spl0_13 | ~spl0_14)),
  inference(forward_demodulation,[],[f445,f152])).
fof(f152,plain,(
  sk_c3 = inverse(sk_c2) | ~spl0_14),
  inference(avatar_component_clause,[],[f150])).
fof(f445,plain,(
  sk_c3 = multiply(inverse(sk_c2),sk_c10) | ~spl0_13),
  inference(superposition,[],[f203,f139])).
fof(f139,plain,(
  sk_c10 = multiply(sk_c2,sk_c3) | ~spl0_13),
  inference(avatar_component_clause,[],[f137])).
fof(f661,plain,(
  ( ! [X0] : (multiply(sk_c3,X0) = X0) ) | (~spl0_1 | ~spl0_14 | ~spl0_15 | ~spl0_27)),
  inference(forward_demodulation,[],[f659,f626])).
fof(f626,plain,(
  ( ! [X0] : (multiply(sk_c9,X0) = X0) ) | (~spl0_1 | ~spl0_27)),
  inference(backward_demodulation,[],[f593,f215])).
fof(f215,plain,(
  ( ! [X0] : (multiply(inverse(identity),X0) = X0) )),
  inference(superposition,[],[f203,f1])).
fof(f593,plain,(
  sk_c9 = inverse(identity) | (~spl0_1 | ~spl0_27)),
  inference(backward_demodulation,[],[f489,f64])).
fof(f64,plain,(
  inverse(sk_c10) = sk_c9 | ~spl0_1),
  inference(avatar_component_clause,[],[f62])).
fof(f659,plain,(
  ( ! [X0] : (multiply(sk_c3,multiply(sk_c9,X0)) = X0) ) | (~spl0_14 | ~spl0_15 | ~spl0_27)),
  inference(backward_demodulation,[],[f655,f430])).
fof(f430,plain,(
  ( ! [X0] : (multiply(sk_c3,multiply(sk_c2,X0)) = X0) ) | ~spl0_14),
  inference(superposition,[],[f203,f152])).
fof(f655,plain,(
  sk_c9 = sk_c2 | (~spl0_14 | ~spl0_15 | ~spl0_27)),
  inference(backward_demodulation,[],[f614,f558])).
fof(f558,plain,(
  sk_c2 = multiply(inverse(sk_c3),identity) | ~spl0_14),
  inference(superposition,[],[f203,f431])).
fof(f431,plain,(
  identity = multiply(sk_c3,sk_c2) | ~spl0_14),
  inference(superposition,[],[f2,f152])).
fof(f614,plain,(
  sk_c9 = multiply(inverse(sk_c3),identity) | (~spl0_15 | ~spl0_27)),
  inference(backward_demodulation,[],[f489,f448])).
fof(f448,plain,(
  sk_c9 = multiply(inverse(sk_c3),sk_c10) | ~spl0_15),
  inference(superposition,[],[f203,f165])).
fof(f165,plain,(
  sk_c10 = multiply(sk_c3,sk_c9) | ~spl0_15),
  inference(avatar_component_clause,[],[f163])).
fof(f660,plain,(
  sk_c3 = inverse(sk_c9) | (~spl0_14 | ~spl0_15 | ~spl0_27)),
  inference(backward_demodulation,[],[f655,f152])).
fof(f544,plain,(
  identity != sk_c9 | sk_c10 != inverse(sk_c9) | (~spl0_1 | ~spl0_19)),
  inference(superposition,[],[f186,f421])).
fof(f421,plain,(
  identity = multiply(sk_c9,sk_c10) | ~spl0_1),
  inference(superposition,[],[f2,f64])).
fof(f186,plain,(
  ( ! [X7] : (sk_c9 != multiply(X7,sk_c10) | sk_c10 != inverse(X7)) ) | ~spl0_19),
  inference(avatar_component_clause,[],[f185])).
fof(f705,plain,(
  spl0_30 | ~spl0_1 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_27),
  inference(avatar_split_clause,[],[f669,f488,f163,f150,f137,f62,f514])).
fof(f669,plain,(
  identity = sk_c9 | (~spl0_1 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_27)),
  inference(forward_demodulation,[],[f666,f2])).
fof(f666,plain,(
  sk_c9 = multiply(inverse(identity),identity) | (~spl0_1 | ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_27)),
  inference(backward_demodulation,[],[f663,f614])).
fof(f592,plain,(
  spl0_27 | ~spl0_13 | ~spl0_14),
  inference(avatar_split_clause,[],[f591,f150,f137,f488])).
fof(f591,plain,(
  identity = sk_c10 | (~spl0_13 | ~spl0_14)),
  inference(forward_demodulation,[],[f580,f2])).
fof(f580,plain,(
  sk_c10 = multiply(inverse(sk_c3),sk_c3) | (~spl0_13 | ~spl0_14)),
  inference(superposition,[],[f203,f447])).
fof(f523,plain,(
  ~spl0_4 | ~spl0_5 | ~spl0_19),
  inference(avatar_contradiction_clause,[],[f522])).
fof(f522,plain,(
  $false | (~spl0_4 | ~spl0_5 | ~spl0_19)),
  inference(subsumption_resolution,[],[f502,f83])).
fof(f83,plain,(
  sk_c10 = inverse(sk_c5) | ~spl0_5),
  inference(avatar_component_clause,[],[f81])).
fof(f81,plain,(
  spl0_5 <=> sk_c10 = inverse(sk_c5)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f502,plain,(
  sk_c10 != inverse(sk_c5) | (~spl0_4 | ~spl0_19)),
  inference(trivial_inequality_removal,[],[f499])).
fof(f499,plain,(
  sk_c9 != sk_c9 | sk_c10 != inverse(sk_c5) | (~spl0_4 | ~spl0_19)),
  inference(superposition,[],[f186,f78])).
fof(f78,plain,(
  sk_c9 = multiply(sk_c5,sk_c10) | ~spl0_4),
  inference(avatar_component_clause,[],[f76])).
fof(f76,plain,(
  spl0_4 <=> sk_c9 = multiply(sk_c5,sk_c10)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f482,plain,(
  ~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_17),
  inference(avatar_contradiction_clause,[],[f481])).
fof(f481,plain,(
  $false | (~spl0_13 | ~spl0_14 | ~spl0_15 | ~spl0_17)),
  inference(subsumption_resolution,[],[f480,f139])).
fof(f480,plain,(
  sk_c10 != multiply(sk_c2,sk_c3) | (~spl0_14 | ~spl0_15 | ~spl0_17)),
  inference(forward_demodulation,[],[f479,f152])).
fof(f479,plain,(
  sk_c10 != multiply(sk_c2,inverse(sk_c2)) | (~spl0_14 | ~spl0_15 | ~spl0_17)),
  inference(subsumption_resolution,[],[f455,f165])).
fof(f455,plain,(
  sk_c10 != multiply(sk_c3,sk_c9) | sk_c10 != multiply(sk_c2,inverse(sk_c2)) | (~spl0_14 | ~spl0_17)),
  inference(superposition,[],[f180,f152])).
fof(f180,plain,(
  ( ! [X4] : (sk_c10 != multiply(inverse(X4),sk_c9) | sk_c10 != multiply(X4,inverse(X4))) ) | ~spl0_17),
  inference(avatar_component_clause,[],[f179])).
fof(f179,plain,(
  spl0_17 <=> ! [X4] : (sk_c10 != multiply(inverse(X4),sk_c9) | sk_c10 != multiply(X4,inverse(X4)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_17])])).
fof(f359,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_18),
  inference(avatar_contradiction_clause,[],[f358])).
fof(f358,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_18)),
  inference(subsumption_resolution,[],[f354,f73])).
fof(f73,plain,(
  sk_c11 = inverse(sk_c4) | ~spl0_3),
  inference(avatar_component_clause,[],[f71])).
fof(f71,plain,(
  spl0_3 <=> sk_c11 = inverse(sk_c4)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f354,plain,(
  sk_c11 != inverse(sk_c4) | (~spl0_2 | ~spl0_3 | ~spl0_18)),
  inference(trivial_inequality_removal,[],[f353])).
fof(f353,plain,(
  identity != identity | sk_c11 != inverse(sk_c4) | (~spl0_2 | ~spl0_3 | ~spl0_18)),
  inference(superposition,[],[f349,f248])).
fof(f248,plain,(
  identity = multiply(sk_c4,sk_c11) | (~spl0_2 | ~spl0_3)),
  inference(backward_demodulation,[],[f246,f68])).
fof(f68,plain,(
  sk_c10 = multiply(sk_c4,sk_c11) | ~spl0_2),
  inference(avatar_component_clause,[],[f66])).
fof(f66,plain,(
  spl0_2 <=> sk_c10 = multiply(sk_c4,sk_c11)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f246,plain,(
  identity = sk_c10 | (~spl0_2 | ~spl0_3)),
  inference(forward_demodulation,[],[f244,f2])).
fof(f244,plain,(
  sk_c10 = multiply(inverse(sk_c11),sk_c11) | (~spl0_2 | ~spl0_3)),
  inference(superposition,[],[f203,f229])).
fof(f229,plain,(
  sk_c11 = multiply(sk_c11,sk_c10) | (~spl0_2 | ~spl0_3)),
  inference(forward_demodulation,[],[f220,f73])).
fof(f220,plain,(
  sk_c11 = multiply(inverse(sk_c4),sk_c10) | ~spl0_2),
  inference(superposition,[],[f203,f68])).
fof(f349,plain,(
  ( ! [X6] : (identity != multiply(X6,sk_c11) | sk_c11 != inverse(X6)) ) | (~spl0_2 | ~spl0_3 | ~spl0_18)),
  inference(forward_demodulation,[],[f183,f246])).
fof(f348,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_17),
  inference(avatar_contradiction_clause,[],[f347])).
fof(f347,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_17)),
  inference(subsumption_resolution,[],[f346,f1])).
fof(f346,plain,(
  identity != multiply(identity,identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_17)),
  inference(forward_demodulation,[],[f338,f275])).
fof(f275,plain,(
  identity = inverse(identity) | (~spl0_2 | ~spl0_3 | ~spl0_5)),
  inference(backward_demodulation,[],[f272,f250])).
fof(f250,plain,(
  identity = inverse(sk_c5) | (~spl0_2 | ~spl0_3 | ~spl0_5)),
  inference(backward_demodulation,[],[f246,f83])).
fof(f272,plain,(
  identity = sk_c5 | (~spl0_2 | ~spl0_3 | ~spl0_5)),
  inference(forward_demodulation,[],[f259,f2])).
fof(f259,plain,(
  sk_c5 = multiply(inverse(identity),identity) | (~spl0_2 | ~spl0_3 | ~spl0_5)),
  inference(backward_demodulation,[],[f246,f219])).
fof(f219,plain,(
  sk_c5 = multiply(inverse(sk_c10),identity) | ~spl0_5),
  inference(superposition,[],[f203,f189])).
fof(f189,plain,(
  identity = multiply(sk_c10,sk_c5) | ~spl0_5),
  inference(superposition,[],[f2,f83])).
fof(f338,plain,(
  identity != multiply(identity,inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_17)),
  inference(trivial_inequality_removal,[],[f337])).
fof(f337,plain,(
  identity != identity | identity != multiply(identity,inverse(identity)) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_17)),
  inference(superposition,[],[f334,f2])).
fof(f334,plain,(
  ( ! [X4] : (identity != multiply(inverse(X4),identity) | identity != multiply(X4,inverse(X4))) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_17)),
  inference(forward_demodulation,[],[f333,f246])).
fof(f333,plain,(
  ( ! [X4] : (identity != multiply(inverse(X4),identity) | sk_c10 != multiply(X4,inverse(X4))) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_17)),
  inference(forward_demodulation,[],[f332,f246])).
fof(f332,plain,(
  ( ! [X4] : (sk_c10 != multiply(inverse(X4),identity) | sk_c10 != multiply(X4,inverse(X4))) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10 | ~spl0_17)),
  inference(forward_demodulation,[],[f180,f300])).
fof(f300,plain,(
  identity = sk_c9 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f296,f2])).
fof(f296,plain,(
  sk_c9 = multiply(inverse(identity),identity) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(backward_demodulation,[],[f294,f260])).
fof(f260,plain,(
  sk_c9 = multiply(inverse(sk_c8),identity) | (~spl0_2 | ~spl0_3 | ~spl0_10)),
  inference(backward_demodulation,[],[f246,f226])).
fof(f226,plain,(
  sk_c9 = multiply(inverse(sk_c8),sk_c10) | ~spl0_10),
  inference(superposition,[],[f203,f108])).
fof(f108,plain,(
  sk_c10 = multiply(sk_c8,sk_c9) | ~spl0_10),
  inference(avatar_component_clause,[],[f106])).
fof(f106,plain,(
  spl0_10 <=> sk_c10 = multiply(sk_c8,sk_c9)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f294,plain,(
  identity = sk_c8 | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(forward_demodulation,[],[f266,f279])).
fof(f279,plain,(
  ( ! [X13] : (multiply(sk_c8,X13) = X13) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_10)),
  inference(backward_demodulation,[],[f277,f270])).
fof(f270,plain,(
  ( ! [X13] : (multiply(sk_c8,multiply(sk_c9,X13)) = X13) ) | (~spl0_2 | ~spl0_3 | ~spl0_10)),
  inference(forward_demodulation,[],[f257,f1])).
fof(f257,plain,(
  ( ! [X13] : (multiply(sk_c8,multiply(sk_c9,X13)) = multiply(identity,X13)) ) | (~spl0_2 | ~spl0_3 | ~spl0_10)),
  inference(backward_demodulation,[],[f246,f200])).
fof(f200,plain,(
  ( ! [X13] : (multiply(sk_c8,multiply(sk_c9,X13)) = multiply(sk_c10,X13)) ) | ~spl0_10),
  inference(superposition,[],[f3,f108])).
fof(f277,plain,(
  ( ! [X10] : (multiply(sk_c9,X10) = X10) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5)),
  inference(forward_demodulation,[],[f273,f1])).
fof(f273,plain,(
  ( ! [X10] : (multiply(sk_c9,X10) = multiply(identity,X10)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5)),
  inference(backward_demodulation,[],[f272,f268])).
fof(f268,plain,(
  ( ! [X10] : (multiply(sk_c9,X10) = multiply(sk_c5,X10)) ) | (~spl0_2 | ~spl0_3 | ~spl0_4)),
  inference(forward_demodulation,[],[f255,f1])).
fof(f255,plain,(
  ( ! [X10] : (multiply(sk_c9,X10) = multiply(sk_c5,multiply(identity,X10))) ) | (~spl0_2 | ~spl0_3 | ~spl0_4)),
  inference(backward_demodulation,[],[f246,f197])).
fof(f197,plain,(
  ( ! [X10] : (multiply(sk_c5,multiply(sk_c10,X10)) = multiply(sk_c9,X10)) ) | ~spl0_4),
  inference(superposition,[],[f3,f78])).
fof(f266,plain,(
  sk_c8 = multiply(sk_c8,identity) | (~spl0_2 | ~spl0_3 | ~spl0_8 | ~spl0_9)),
  inference(backward_demodulation,[],[f246,f239])).
fof(f239,plain,(
  sk_c8 = multiply(sk_c8,sk_c10) | (~spl0_8 | ~spl0_9)),
  inference(forward_demodulation,[],[f225,f103])).
fof(f103,plain,(
  sk_c8 = inverse(sk_c7) | ~spl0_9),
  inference(avatar_component_clause,[],[f101])).
fof(f101,plain,(
  spl0_9 <=> sk_c8 = inverse(sk_c7)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_9])])).
fof(f225,plain,(
  sk_c8 = multiply(inverse(sk_c7),sk_c10) | ~spl0_8),
  inference(superposition,[],[f203,f98])).
fof(f98,plain,(
  sk_c10 = multiply(sk_c7,sk_c8) | ~spl0_8),
  inference(avatar_component_clause,[],[f96])).
fof(f96,plain,(
  spl0_8 <=> sk_c10 = multiply(sk_c7,sk_c8)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_8])])).
fof(f331,plain,(
  ~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7 | ~spl0_16),
  inference(avatar_contradiction_clause,[],[f330])).
fof(f330,plain,(
  $false | (~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7 | ~spl0_16)),
  inference(subsumption_resolution,[],[f327,f73])).
fof(f327,plain,(
  sk_c11 != inverse(sk_c4) | (~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7 | ~spl0_16)),
  inference(trivial_inequality_removal,[],[f326])).
fof(f326,plain,(
  sk_c11 != sk_c11 | sk_c11 != inverse(sk_c4) | (~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7 | ~spl0_16)),
  inference(superposition,[],[f308,f263])).
fof(f263,plain,(
  sk_c11 = multiply(sk_c4,identity) | (~spl0_2 | ~spl0_3 | ~spl0_6 | ~spl0_7)),
  inference(backward_demodulation,[],[f246,f235])).
fof(f235,plain,(
  sk_c11 = multiply(sk_c4,sk_c10) | (~spl0_3 | ~spl0_6 | ~spl0_7)),
  inference(backward_demodulation,[],[f230,f93])).
fof(f93,plain,(
  sk_c11 = multiply(sk_c6,sk_c10) | ~spl0_7),
  inference(avatar_component_clause,[],[f91])).
fof(f91,plain,(
  spl0_7 <=> sk_c11 = multiply(sk_c6,sk_c10)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).
fof(f230,plain,(
  sk_c4 = sk_c6 | (~spl0_3 | ~spl0_6)),
  inference(forward_demodulation,[],[f222,f221])).
fof(f221,plain,(
  sk_c4 = multiply(inverse(sk_c11),identity) | ~spl0_3),
  inference(superposition,[],[f203,f188])).
fof(f188,plain,(
  identity = multiply(sk_c11,sk_c4) | ~spl0_3),
  inference(superposition,[],[f2,f73])).
fof(f222,plain,(
  sk_c6 = multiply(inverse(sk_c11),identity) | ~spl0_6),
  inference(superposition,[],[f203,f190])).
fof(f190,plain,(
  identity = multiply(sk_c11,sk_c6) | ~spl0_6),
  inference(superposition,[],[f2,f88])).
fof(f88,plain,(
  sk_c11 = inverse(sk_c6) | ~spl0_6),
  inference(avatar_component_clause,[],[f86])).
fof(f86,plain,(
  spl0_6 <=> sk_c11 = inverse(sk_c6)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).
fof(f308,plain,(
  ( ! [X3] : (sk_c11 != multiply(X3,identity) | sk_c11 != inverse(X3)) ) | (~spl0_2 | ~spl0_3 | ~spl0_16)),
  inference(forward_demodulation,[],[f177,f246])).
fof(f302,plain,(
  spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10),
  inference(avatar_contradiction_clause,[],[f301])).
fof(f301,plain,(
  $false | (spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_4 | ~spl0_5 | ~spl0_8 | ~spl0_9 | ~spl0_10)),
  inference(subsumption_resolution,[],[f300,f282])).
fof(f282,plain,(
  identity != sk_c9 | (spl0_1 | ~spl0_2 | ~spl0_3 | ~spl0_5)),
  inference(backward_demodulation,[],[f275,f247])).
fof(f247,plain,(
  sk_c9 != inverse(identity) | (spl0_1 | ~spl0_2 | ~spl0_3)),
  inference(backward_demodulation,[],[f246,f63])).
fof(f63,plain,(
  inverse(sk_c10) != sk_c9 | spl0_1),
  inference(avatar_component_clause,[],[f62])).
fof(f187,plain,(
  ~spl0_1 | spl0_16 | spl0_17 | spl0_18 | spl0_19 | spl0_16 | spl0_17),
  inference(avatar_split_clause,[],[f60,f179,f176,f185,f182,f179,f176,f62])).
fof(f60,plain,(
  ( ! [X6,X4,X8,X7,X3,X9] : (sk_c10 != multiply(inverse(X9),sk_c9) | sk_c10 != multiply(X9,inverse(X9)) | sk_c11 != multiply(X8,sk_c10) | sk_c11 != inverse(X8) | sk_c10 != inverse(X7) | sk_c9 != multiply(X7,sk_c10) | sk_c11 != inverse(X6) | sk_c10 != multiply(X6,sk_c11) | sk_c10 != multiply(inverse(X4),sk_c9) | sk_c10 != multiply(X4,inverse(X4)) | sk_c11 != multiply(X3,sk_c10) | sk_c11 != inverse(X3) | inverse(sk_c10) != sk_c9) )),
  inference(equality_resolution,[],[f59])).
fof(f59,plain,(
  ( ! [X6,X4,X8,X7,X5,X3,X9] : (sk_c10 != multiply(inverse(X9),sk_c9) | sk_c10 != multiply(X9,inverse(X9)) | sk_c11 != multiply(X8,sk_c10) | sk_c11 != inverse(X8) | sk_c10 != inverse(X7) | sk_c9 != multiply(X7,sk_c10) | sk_c11 != inverse(X6) | sk_c10 != multiply(X6,sk_c11) | sk_c10 != multiply(X5,sk_c9) | inverse(X4) != X5 | sk_c10 != multiply(X4,X5) | sk_c11 != multiply(X3,sk_c10) | sk_c11 != inverse(X3) | inverse(sk_c10) != sk_c9) )),
  inference(equality_resolution,[],[f58])).
fof(f58,axiom,(
  ( ! [X6,X4,X10,X8,X7,X5,X3,X9] : (sk_c10 != multiply(X10,sk_c9) | inverse(X9) != X10 | sk_c10 != multiply(X9,X10) | sk_c11 != multiply(X8,sk_c10) | sk_c11 != inverse(X8) | sk_c10 != inverse(X7) | sk_c9 != multiply(X7,sk_c10) | sk_c11 != inverse(X6) | sk_c10 != multiply(X6,sk_c11) | sk_c10 != multiply(X5,sk_c9) | inverse(X4) != X5 | sk_c10 != multiply(X4,X5) | sk_c11 != multiply(X3,sk_c10) | sk_c11 != inverse(X3) | inverse(sk_c10) != sk_c9) )),
  file('./p/GRP389-1.tptp',unknown)).
fof(f174,plain,(
  spl0_15 | spl0_10),
  inference(avatar_split_clause,[],[f57,f106,f163])).
fof(f57,axiom,(
  sk_c10 = multiply(sk_c8,sk_c9) | sk_c10 = multiply(sk_c3,sk_c9)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f173,plain,(
  spl0_15 | spl0_9),
  inference(avatar_split_clause,[],[f56,f101,f163])).
fof(f56,axiom,(
  sk_c8 = inverse(sk_c7) | sk_c10 = multiply(sk_c3,sk_c9)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f172,plain,(
  spl0_15 | spl0_8),
  inference(avatar_split_clause,[],[f55,f96,f163])).
fof(f55,axiom,(
  sk_c10 = multiply(sk_c7,sk_c8) | sk_c10 = multiply(sk_c3,sk_c9)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f169,plain,(
  spl0_15 | spl0_5),
  inference(avatar_split_clause,[],[f52,f81,f163])).
fof(f52,axiom,(
  sk_c10 = inverse(sk_c5) | sk_c10 = multiply(sk_c3,sk_c9)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f168,plain,(
  spl0_15 | spl0_4),
  inference(avatar_split_clause,[],[f51,f76,f163])).
fof(f51,axiom,(
  sk_c9 = multiply(sk_c5,sk_c10) | sk_c10 = multiply(sk_c3,sk_c9)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f167,plain,(
  spl0_15 | spl0_3),
  inference(avatar_split_clause,[],[f50,f71,f163])).
fof(f50,axiom,(
  sk_c11 = inverse(sk_c4) | sk_c10 = multiply(sk_c3,sk_c9)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f166,plain,(
  spl0_15 | spl0_2),
  inference(avatar_split_clause,[],[f49,f66,f163])).
fof(f49,axiom,(
  sk_c10 = multiply(sk_c4,sk_c11) | sk_c10 = multiply(sk_c3,sk_c9)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f161,plain,(
  spl0_14 | spl0_10),
  inference(avatar_split_clause,[],[f48,f106,f150])).
fof(f48,axiom,(
  sk_c10 = multiply(sk_c8,sk_c9) | sk_c3 = inverse(sk_c2)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f160,plain,(
  spl0_14 | spl0_9),
  inference(avatar_split_clause,[],[f47,f101,f150])).
fof(f47,axiom,(
  sk_c8 = inverse(sk_c7) | sk_c3 = inverse(sk_c2)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f159,plain,(
  spl0_14 | spl0_8),
  inference(avatar_split_clause,[],[f46,f96,f150])).
fof(f46,axiom,(
  sk_c10 = multiply(sk_c7,sk_c8) | sk_c3 = inverse(sk_c2)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f158,plain,(
  spl0_14 | spl0_7),
  inference(avatar_split_clause,[],[f45,f91,f150])).
fof(f45,axiom,(
  sk_c11 = multiply(sk_c6,sk_c10) | sk_c3 = inverse(sk_c2)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f157,plain,(
  spl0_14 | spl0_6),
  inference(avatar_split_clause,[],[f44,f86,f150])).
fof(f44,axiom,(
  sk_c11 = inverse(sk_c6) | sk_c3 = inverse(sk_c2)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f156,plain,(
  spl0_14 | spl0_5),
  inference(avatar_split_clause,[],[f43,f81,f150])).
fof(f43,axiom,(
  sk_c10 = inverse(sk_c5) | sk_c3 = inverse(sk_c2)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f155,plain,(
  spl0_14 | spl0_4),
  inference(avatar_split_clause,[],[f42,f76,f150])).
fof(f42,axiom,(
  sk_c9 = multiply(sk_c5,sk_c10) | sk_c3 = inverse(sk_c2)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f154,plain,(
  spl0_14 | spl0_3),
  inference(avatar_split_clause,[],[f41,f71,f150])).
fof(f41,axiom,(
  sk_c11 = inverse(sk_c4) | sk_c3 = inverse(sk_c2)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f153,plain,(
  spl0_14 | spl0_2),
  inference(avatar_split_clause,[],[f40,f66,f150])).
fof(f40,axiom,(
  sk_c10 = multiply(sk_c4,sk_c11) | sk_c3 = inverse(sk_c2)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f148,plain,(
  spl0_13 | spl0_10),
  inference(avatar_split_clause,[],[f39,f106,f137])).
fof(f39,axiom,(
  sk_c10 = multiply(sk_c8,sk_c9) | sk_c10 = multiply(sk_c2,sk_c3)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f147,plain,(
  spl0_13 | spl0_9),
  inference(avatar_split_clause,[],[f38,f101,f137])).
fof(f38,axiom,(
  sk_c8 = inverse(sk_c7) | sk_c10 = multiply(sk_c2,sk_c3)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f146,plain,(
  spl0_13 | spl0_8),
  inference(avatar_split_clause,[],[f37,f96,f137])).
fof(f37,axiom,(
  sk_c10 = multiply(sk_c7,sk_c8) | sk_c10 = multiply(sk_c2,sk_c3)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f145,plain,(
  spl0_13 | spl0_7),
  inference(avatar_split_clause,[],[f36,f91,f137])).
fof(f36,axiom,(
  sk_c11 = multiply(sk_c6,sk_c10) | sk_c10 = multiply(sk_c2,sk_c3)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f144,plain,(
  spl0_13 | spl0_6),
  inference(avatar_split_clause,[],[f35,f86,f137])).
fof(f35,axiom,(
  sk_c11 = inverse(sk_c6) | sk_c10 = multiply(sk_c2,sk_c3)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f143,plain,(
  spl0_13 | spl0_5),
  inference(avatar_split_clause,[],[f34,f81,f137])).
fof(f34,axiom,(
  sk_c10 = inverse(sk_c5) | sk_c10 = multiply(sk_c2,sk_c3)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f142,plain,(
  spl0_13 | spl0_4),
  inference(avatar_split_clause,[],[f33,f76,f137])).
fof(f33,axiom,(
  sk_c9 = multiply(sk_c5,sk_c10) | sk_c10 = multiply(sk_c2,sk_c3)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f141,plain,(
  spl0_13 | spl0_3),
  inference(avatar_split_clause,[],[f32,f71,f137])).
fof(f32,axiom,(
  sk_c11 = inverse(sk_c4) | sk_c10 = multiply(sk_c2,sk_c3)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f140,plain,(
  spl0_13 | spl0_2),
  inference(avatar_split_clause,[],[f31,f66,f137])).
fof(f31,axiom,(
  sk_c10 = multiply(sk_c4,sk_c11) | sk_c10 = multiply(sk_c2,sk_c3)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f132,plain,(
  spl0_12 | spl0_7),
  inference(avatar_split_clause,[],[f27,f91,f124])).
fof(f27,axiom,(
  sk_c11 = multiply(sk_c6,sk_c10) | sk_c11 = multiply(sk_c1,sk_c10)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f131,plain,(
  spl0_12 | spl0_6),
  inference(avatar_split_clause,[],[f26,f86,f124])).
fof(f26,axiom,(
  sk_c11 = inverse(sk_c6) | sk_c11 = multiply(sk_c1,sk_c10)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f128,plain,(
  spl0_12 | spl0_3),
  inference(avatar_split_clause,[],[f23,f71,f124])).
fof(f23,axiom,(
  sk_c11 = inverse(sk_c4) | sk_c11 = multiply(sk_c1,sk_c10)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f127,plain,(
  spl0_12 | spl0_2),
  inference(avatar_split_clause,[],[f22,f66,f124])).
fof(f22,axiom,(
  sk_c10 = multiply(sk_c4,sk_c11) | sk_c11 = multiply(sk_c1,sk_c10)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f119,plain,(
  spl0_11 | spl0_7),
  inference(avatar_split_clause,[],[f18,f91,f111])).
fof(f18,axiom,(
  sk_c11 = multiply(sk_c6,sk_c10) | sk_c11 = inverse(sk_c1)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f118,plain,(
  spl0_11 | spl0_6),
  inference(avatar_split_clause,[],[f17,f86,f111])).
fof(f17,axiom,(
  sk_c11 = inverse(sk_c6) | sk_c11 = inverse(sk_c1)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f115,plain,(
  spl0_11 | spl0_3),
  inference(avatar_split_clause,[],[f14,f71,f111])).
fof(f14,axiom,(
  sk_c11 = inverse(sk_c4) | sk_c11 = inverse(sk_c1)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f114,plain,(
  spl0_11 | spl0_2),
  inference(avatar_split_clause,[],[f13,f66,f111])).
fof(f13,axiom,(
  sk_c10 = multiply(sk_c4,sk_c11) | sk_c11 = inverse(sk_c1)),
  file('./p/GRP389-1.tptp',unknown)).
fof(f109,plain,(
  spl0_1 | spl0_10),
  inference(avatar_split_clause,[],[f12,f106,f62])).
fof(f12,axiom,(
  sk_c10 = multiply(sk_c8,sk_c9) | inverse(sk_c10) = sk_c9),
  file('./p/GRP389-1.tptp',unknown)).
fof(f104,plain,(
  spl0_1 | spl0_9),
  inference(avatar_split_clause,[],[f11,f101,f62])).
fof(f11,axiom,(
  sk_c8 = inverse(sk_c7) | inverse(sk_c10) = sk_c9),
  file('./p/GRP389-1.tptp',unknown)).
fof(f99,plain,(
  spl0_1 | spl0_8),
  inference(avatar_split_clause,[],[f10,f96,f62])).
fof(f10,axiom,(
  sk_c10 = multiply(sk_c7,sk_c8) | inverse(sk_c10) = sk_c9),
  file('./p/GRP389-1.tptp',unknown)).
fof(f84,plain,(
  spl0_1 | spl0_5),
  inference(avatar_split_clause,[],[f7,f81,f62])).
fof(f7,axiom,(
  sk_c10 = inverse(sk_c5) | inverse(sk_c10) = sk_c9),
  file('./p/GRP389-1.tptp',unknown)).
fof(f79,plain,(
  spl0_1 | spl0_4),
  inference(avatar_split_clause,[],[f6,f76,f62])).
fof(f6,axiom,(
  sk_c9 = multiply(sk_c5,sk_c10) | inverse(sk_c10) = sk_c9),
  file('./p/GRP389-1.tptp',unknown)).
fof(f74,plain,(
  spl0_1 | spl0_3),
  inference(avatar_split_clause,[],[f5,f71,f62])).
fof(f5,axiom,(
  sk_c11 = inverse(sk_c4) | inverse(sk_c10) = sk_c9),
  file('./p/GRP389-1.tptp',unknown)).
fof(f69,plain,(
  spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f4,f66,f62])).
fof(f4,axiom,(
  sk_c10 = multiply(sk_c4,sk_c11) | inverse(sk_c10) = sk_c9),
  file('./p/GRP389-1.tptp',unknown)).
% SZS output end Proof for GRP389-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5245
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
