% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GRP560-1
% SZS output start Proof for GRP560-1
fof(f2065,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f2064])).
fof(f2064,plain,(
  divide(a,inverse(b)) != divide(a,inverse(b))),
  inference(superposition,[],[f4,f1547])).
fof(f1547,plain,(
  ( ! [X0,X1] : (divide(X0,inverse(X1)) = divide(X1,inverse(X0))) )),
  inference(superposition,[],[f446,f713])).
fof(f713,plain,(
  ( ! [X3] : (inverse(inverse(X3)) = X3) )),
  inference(backward_demodulation,[],[f710,f259])).
fof(f259,plain,(
  ( ! [X2,X3] : (inverse(divide(divide(X2,X2),X3)) = X3) )),
  inference(superposition,[],[f63,f11])).
fof(f11,plain,(
  ( ! [X0,X1] : (divide(X0,inverse(divide(X1,X1))) = X0) )),
  inference(superposition,[],[f5,f1])).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (divide(X0,inverse(divide(divide(X1,X2),divide(X0,X2)))) = X1) )),
  file('./p/GRP560-1.tptp',unknown)).
fof(f5,plain,(
  ( ! [X2,X0,X3,X1] : (divide(X3,inverse(divide(X1,divide(X3,inverse(divide(divide(X1,X2),divide(X0,X2))))))) = X0) )),
  inference(superposition,[],[f1,f1])).
fof(f63,plain,(
  ( ! [X0,X1] : (divide(inverse(divide(X1,X0)),inverse(X1)) = X0) )),
  inference(superposition,[],[f26,f26])).
fof(f26,plain,(
  ( ! [X2,X3] : (divide(X3,inverse(divide(X2,X3))) = X2) )),
  inference(superposition,[],[f6,f11])).
fof(f6,plain,(
  ( ! [X2,X0,X3,X1] : (divide(X0,inverse(divide(divide(X3,inverse(divide(divide(X1,X2),divide(X0,X2)))),X1))) = X3) )),
  inference(superposition,[],[f1,f1])).
fof(f710,plain,(
  ( ! [X0,X1] : (inverse(X1) = divide(divide(X0,X0),X1)) )),
  inference(superposition,[],[f259,f363])).
fof(f363,plain,(
  ( ! [X38,X36] : (divide(X36,divide(X36,X38)) = X38) )),
  inference(backward_demodulation,[],[f346,f80])).
fof(f80,plain,(
  ( ! [X39,X37,X38,X36] : (divide(X36,inverse(divide(divide(X38,inverse(divide(divide(X39,inverse(divide(X37,X36))),X37))),X39))) = X38) )),
  inference(superposition,[],[f6,f26])).
fof(f346,plain,(
  ( ! [X2,X0,X3,X1] : (divide(X3,X0) = inverse(divide(divide(X0,inverse(divide(divide(X1,inverse(divide(X2,X3))),X2))),X1))) )),
  inference(superposition,[],[f196,f259])).
fof(f196,plain,(
  ( ! [X12,X15,X13,X11,X16] : (divide(X11,inverse(divide(divide(X15,inverse(divide(divide(X16,inverse(divide(divide(X12,inverse(divide(X13,X11))),X13))),X12))),X16))) = X15) )),
  inference(backward_demodulation,[],[f164,f19])).
fof(f19,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (divide(X11,inverse(divide(divide(X15,inverse(divide(divide(X16,inverse(divide(divide(X12,inverse(divide(divide(X13,X14),divide(X11,X14)))),X13))),X12))),X16))) = X15) )),
  inference(superposition,[],[f6,f6])).
fof(f164,plain,(
  ( ! [X6,X4,X7,X5] : (divide(X4,inverse(divide(divide(X5,X6),divide(X7,X6)))) = divide(X4,inverse(divide(X5,X7)))) )),
  inference(superposition,[],[f12,f26])).
fof(f12,plain,(
  ( ! [X4,X2,X0,X3,X1] : (divide(X0,inverse(divide(divide(X4,inverse(divide(X1,divide(X0,inverse(divide(divide(X1,X2),divide(X3,X2))))))),X3))) = X4) )),
  inference(superposition,[],[f1,f5])).
fof(f446,plain,(
  ( ! [X0,X1] : (divide(X0,X1) = divide(inverse(X1),inverse(X0))) )),
  inference(backward_demodulation,[],[f375,f254])).
fof(f254,plain,(
  ( ! [X0,X1] : (inverse(divide(X1,X0)) = divide(inverse(X1),inverse(X0))) )),
  inference(superposition,[],[f63,f26])).
fof(f375,plain,(
  ( ! [X4,X5] : (divide(X5,X4) = inverse(divide(X4,X5))) )),
  inference(backward_demodulation,[],[f374,f263])).
fof(f263,plain,(
  ( ! [X6,X4,X5] : (inverse(divide(X4,X5)) = divide(X6,inverse(divide(X5,divide(X6,inverse(X4)))))) )),
  inference(superposition,[],[f1,f63])).
fof(f374,plain,(
  ( ! [X4,X7,X5] : (divide(X4,X5) = divide(X7,inverse(divide(X4,divide(X7,inverse(X5)))))) )),
  inference(backward_demodulation,[],[f363,f71])).
fof(f71,plain,(
  ( ! [X6,X4,X7,X5] : (divide(X4,X5) = divide(X7,inverse(divide(X4,divide(X7,inverse(divide(X6,divide(X6,X5)))))))) )),
  inference(superposition,[],[f7,f26])).
fof(f7,plain,(
  ( ! [X4,X2,X0,X3,X1] : (divide(X3,inverse(divide(X0,divide(X3,inverse(divide(X1,divide(X4,inverse(divide(divide(X1,X2),divide(X0,X2)))))))))) = X4) )),
  inference(superposition,[],[f5,f1])).
fof(f4,plain,(
  divide(a,inverse(b)) != divide(b,inverse(a))),
  inference(definition_unfolding,[],[f3,f2,f2])).
fof(f2,axiom,(
  ( ! [X0,X1] : (multiply(X0,X1) = divide(X0,inverse(X1))) )),
  file('./p/GRP560-1.tptp',unknown)).
fof(f3,axiom,(
  multiply(a,b) != multiply(b,a)),
  file('./p/GRP560-1.tptp',unknown)).
% SZS output end Proof for GRP560-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 7036
% Time elapsed: 0.029 s
% ------------------------------
% ------------------------------
