% Refutation found. Thanks to Tanya!
% SZS status Theorem for GRP644+1
% SZS output start Proof for GRP644+1
fof(f463,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f441,f453,f462])).
fof(f462,plain,(
  spl17_12),
  inference(avatar_contradiction_clause,[],[f461])).
fof(f461,plain,(
  $false | spl17_12),
  inference(subsumption_resolution,[],[f460,f145])).
fof(f145,plain,(
  ~v3_struct_0(sK0)),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ((~r2_hidden(k3_group_1(sK0,sK2),k1_funct_1(k1_latsubgr(sK0),sK1)) & r2_hidden(sK2,k1_funct_1(k1_latsubgr(sK0),sK1)) & m1_subset_1(sK2,u1_struct_0(sK0))) & m1_group_2(sK1,sK0) & v1_group_1(sK1)) & l1_group_1(sK0) & v4_group_1(sK0) & v3_group_1(sK0) & ~v3_struct_0(sK0)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f63,f113,f112,f111])).
fof(f111,plain,(
  ? [X0] : (? [X1] : (? [X2] : (~r2_hidden(k3_group_1(X0,X2),k1_funct_1(k1_latsubgr(X0),X1)) & r2_hidden(X2,k1_funct_1(k1_latsubgr(X0),X1)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_group_2(X1,X0) & v1_group_1(X1)) & l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & ~v3_struct_0(X0)) => (? [X1] : (? [X2] : (~r2_hidden(k3_group_1(sK0,X2),k1_funct_1(k1_latsubgr(sK0),X1)) & r2_hidden(X2,k1_funct_1(k1_latsubgr(sK0),X1)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_group_2(X1,sK0) & v1_group_1(X1)) & l1_group_1(sK0) & v4_group_1(sK0) & v3_group_1(sK0) & ~v3_struct_0(sK0))),
  introduced(choice_axiom,[])).
fof(f112,plain,(
  ? [X1] : (? [X2] : (~r2_hidden(k3_group_1(sK0,X2),k1_funct_1(k1_latsubgr(sK0),X1)) & r2_hidden(X2,k1_funct_1(k1_latsubgr(sK0),X1)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_group_2(X1,sK0) & v1_group_1(X1)) => (? [X2] : (~r2_hidden(k3_group_1(sK0,X2),k1_funct_1(k1_latsubgr(sK0),sK1)) & r2_hidden(X2,k1_funct_1(k1_latsubgr(sK0),sK1)) & m1_subset_1(X2,u1_struct_0(sK0))) & m1_group_2(sK1,sK0) & v1_group_1(sK1))),
  introduced(choice_axiom,[])).
fof(f113,plain,(
  ? [X2] : (~r2_hidden(k3_group_1(sK0,X2),k1_funct_1(k1_latsubgr(sK0),sK1)) & r2_hidden(X2,k1_funct_1(k1_latsubgr(sK0),sK1)) & m1_subset_1(X2,u1_struct_0(sK0))) => (~r2_hidden(k3_group_1(sK0,sK2),k1_funct_1(k1_latsubgr(sK0),sK1)) & r2_hidden(sK2,k1_funct_1(k1_latsubgr(sK0),sK1)) & m1_subset_1(sK2,u1_struct_0(sK0)))),
  introduced(choice_axiom,[])).
fof(f63,plain,(
  ? [X0] : (? [X1] : (? [X2] : (~r2_hidden(k3_group_1(X0,X2),k1_funct_1(k1_latsubgr(X0),X1)) & r2_hidden(X2,k1_funct_1(k1_latsubgr(X0),X1)) & m1_subset_1(X2,u1_struct_0(X0))) & m1_group_2(X1,X0) & v1_group_1(X1)) & l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & ~v3_struct_0(X0))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ? [X0] : (? [X1] : (? [X2] : ((~r2_hidden(k3_group_1(X0,X2),k1_funct_1(k1_latsubgr(X0),X1)) & r2_hidden(X2,k1_funct_1(k1_latsubgr(X0),X1))) & m1_subset_1(X2,u1_struct_0(X0))) & (m1_group_2(X1,X0) & v1_group_1(X1))) & (l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & ~v3_struct_0(X0)))),
  inference(ennf_transformation,[],[f2])).
fof(f2,negated_conjecture,(
  ~! [X0] : ((l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & ~v3_struct_0(X0)) => ! [X1] : ((m1_group_2(X1,X0) & v1_group_1(X1)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r2_hidden(X2,k1_funct_1(k1_latsubgr(X0),X1)) => r2_hidden(k3_group_1(X0,X2),k1_funct_1(k1_latsubgr(X0),X1))))))),
  inference(negated_conjecture,[],[f1])).
fof(f1,conjecture,(
  ! [X0] : ((l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & ~v3_struct_0(X0)) => ! [X1] : ((m1_group_2(X1,X0) & v1_group_1(X1)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r2_hidden(X2,k1_funct_1(k1_latsubgr(X0),X1)) => r2_hidden(k3_group_1(X0,X2),k1_funct_1(k1_latsubgr(X0),X1))))))),
  file('./p/GRP644+1.tptp',unknown)).
fof(f460,plain,(
  v3_struct_0(sK0) | spl17_12),
  inference(subsumption_resolution,[],[f459,f146])).
fof(f146,plain,(
  v3_group_1(sK0)),
  inference(cnf_transformation,[],[f114])).
fof(f459,plain,(
  ~v3_group_1(sK0) | v3_struct_0(sK0) | spl17_12),
  inference(subsumption_resolution,[],[f458,f147])).
fof(f147,plain,(
  v4_group_1(sK0)),
  inference(cnf_transformation,[],[f114])).
fof(f458,plain,(
  ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0) | spl17_12),
  inference(subsumption_resolution,[],[f457,f148])).
fof(f148,plain,(
  l1_group_1(sK0)),
  inference(cnf_transformation,[],[f114])).
fof(f457,plain,(
  ~l1_group_1(sK0) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0) | spl17_12),
  inference(subsumption_resolution,[],[f456,f151])).
fof(f151,plain,(
  m1_subset_1(sK2,u1_struct_0(sK0))),
  inference(cnf_transformation,[],[f114])).
fof(f456,plain,(
  ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_group_1(sK0) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0) | spl17_12),
  inference(subsumption_resolution,[],[f455,f150])).
fof(f150,plain,(
  m1_group_2(sK1,sK0)),
  inference(cnf_transformation,[],[f114])).
fof(f455,plain,(
  ~m1_group_2(sK1,sK0) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_group_1(sK0) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0) | spl17_12),
  inference(subsumption_resolution,[],[f454,f419])).
fof(f419,plain,(
  r1_rlvect_1(sK1,sK2)),
  inference(subsumption_resolution,[],[f418,f145])).
fof(f418,plain,(
  r1_rlvect_1(sK1,sK2) | v3_struct_0(sK0)),
  inference(subsumption_resolution,[],[f417,f146])).
fof(f417,plain,(
  r1_rlvect_1(sK1,sK2) | ~v3_group_1(sK0) | v3_struct_0(sK0)),
  inference(subsumption_resolution,[],[f416,f147])).
fof(f416,plain,(
  r1_rlvect_1(sK1,sK2) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0)),
  inference(subsumption_resolution,[],[f415,f148])).
fof(f415,plain,(
  r1_rlvect_1(sK1,sK2) | ~l1_group_1(sK0) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0)),
  inference(subsumption_resolution,[],[f414,f149])).
fof(f149,plain,(
  v1_group_1(sK1)),
  inference(cnf_transformation,[],[f114])).
fof(f414,plain,(
  r1_rlvect_1(sK1,sK2) | ~v1_group_1(sK1) | ~l1_group_1(sK0) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0)),
  inference(subsumption_resolution,[],[f413,f150])).
fof(f413,plain,(
  r1_rlvect_1(sK1,sK2) | ~m1_group_2(sK1,sK0) | ~v1_group_1(sK1) | ~l1_group_1(sK0) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0)),
  inference(subsumption_resolution,[],[f411,f151])).
fof(f411,plain,(
  r1_rlvect_1(sK1,sK2) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~m1_group_2(sK1,sK0) | ~v1_group_1(sK1) | ~l1_group_1(sK0) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0)),
  inference(resolution,[],[f210,f152])).
fof(f152,plain,(
  r2_hidden(sK2,k1_funct_1(k1_latsubgr(sK0),sK1))),
  inference(cnf_transformation,[],[f114])).
fof(f210,plain,(
  ( ! [X2,X0,X1] : (~r2_hidden(X2,k1_funct_1(k1_latsubgr(X0),X1)) | r1_rlvect_1(X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_group_2(X1,X0) | ~v1_group_1(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : (! [X1] : (! [X2] : (((r2_hidden(X2,k1_funct_1(k1_latsubgr(X0),X1)) | ~r1_rlvect_1(X1,X2)) & (r1_rlvect_1(X1,X2) | ~r2_hidden(X2,k1_funct_1(k1_latsubgr(X0),X1)))) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_group_2(X1,X0) | ~v1_group_1(X1)) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(X2,k1_funct_1(k1_latsubgr(X0),X1)) <=> r1_rlvect_1(X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | ~m1_group_2(X1,X0) | ~v1_group_1(X1)) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : (! [X1] : (! [X2] : ((r2_hidden(X2,k1_funct_1(k1_latsubgr(X0),X1)) <=> r1_rlvect_1(X1,X2)) | ~m1_subset_1(X2,u1_struct_0(X0))) | (~m1_group_2(X1,X0) | ~v1_group_1(X1))) | (~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)))),
  inference(ennf_transformation,[],[f47])).
fof(f47,axiom,(
  ! [X0] : ((l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & ~v3_struct_0(X0)) => ! [X1] : ((m1_group_2(X1,X0) & v1_group_1(X1)) => ! [X2] : (m1_subset_1(X2,u1_struct_0(X0)) => (r2_hidden(X2,k1_funct_1(k1_latsubgr(X0),X1)) <=> r1_rlvect_1(X1,X2)))))),
  file('./p/GRP644+1.tptp',unknown)).
fof(f454,plain,(
  ~r1_rlvect_1(sK1,sK2) | ~m1_group_2(sK1,sK0) | ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_group_1(sK0) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0) | spl17_12),
  inference(resolution,[],[f440,f217])).
fof(f217,plain,(
  ( ! [X2,X0,X1] : (r1_rlvect_1(X2,k3_group_1(X0,X1)) | ~r1_rlvect_1(X2,X1) | ~m1_group_2(X2,X0) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : (! [X1] : (! [X2] : (r1_rlvect_1(X2,k3_group_1(X0,X1)) | ~r1_rlvect_1(X2,X1) | ~m1_group_2(X2,X0)) | ~m1_subset_1(X1,u1_struct_0(X0))) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : (! [X1] : (! [X2] : ((r1_rlvect_1(X2,k3_group_1(X0,X1)) | ~r1_rlvect_1(X2,X1)) | ~m1_group_2(X2,X0)) | ~m1_subset_1(X1,u1_struct_0(X0))) | (~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)))),
  inference(ennf_transformation,[],[f53])).
fof(f53,axiom,(
  ! [X0] : ((l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & ~v3_struct_0(X0)) => ! [X1] : (m1_subset_1(X1,u1_struct_0(X0)) => ! [X2] : (m1_group_2(X2,X0) => (r1_rlvect_1(X2,X1) => r1_rlvect_1(X2,k3_group_1(X0,X1))))))),
  file('./p/GRP644+1.tptp',unknown)).
fof(f440,plain,(
  ~r1_rlvect_1(sK1,k3_group_1(sK0,sK2)) | spl17_12),
  inference(avatar_component_clause,[],[f438])).
fof(f438,plain,(
  spl17_12 <=> r1_rlvect_1(sK1,k3_group_1(sK0,sK2))),
  introduced(avatar_definition,[new_symbols(naming,[spl17_12])])).
fof(f453,plain,(
  spl17_11),
  inference(avatar_contradiction_clause,[],[f452])).
fof(f452,plain,(
  $false | spl17_11),
  inference(subsumption_resolution,[],[f451,f145])).
fof(f451,plain,(
  v3_struct_0(sK0) | spl17_11),
  inference(subsumption_resolution,[],[f450,f146])).
fof(f450,plain,(
  ~v3_group_1(sK0) | v3_struct_0(sK0) | spl17_11),
  inference(subsumption_resolution,[],[f449,f147])).
fof(f449,plain,(
  ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0) | spl17_11),
  inference(subsumption_resolution,[],[f448,f148])).
fof(f448,plain,(
  ~l1_group_1(sK0) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0) | spl17_11),
  inference(subsumption_resolution,[],[f447,f151])).
fof(f447,plain,(
  ~m1_subset_1(sK2,u1_struct_0(sK0)) | ~l1_group_1(sK0) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0) | spl17_11),
  inference(resolution,[],[f436,f170])).
fof(f170,plain,(
  ( ! [X0,X1] : (m1_subset_1(k3_group_1(X0,X1),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : (m1_subset_1(k3_group_1(X0,X1),u1_struct_0(X0)) | ~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : (m1_subset_1(k3_group_1(X0,X1),u1_struct_0(X0)) | (~m1_subset_1(X1,u1_struct_0(X0)) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)))),
  inference(ennf_transformation,[],[f16])).
fof(f16,axiom,(
  ! [X0,X1] : ((m1_subset_1(X1,u1_struct_0(X0)) & l1_group_1(X0) & v4_group_1(X0) & v3_group_1(X0) & ~v3_struct_0(X0)) => m1_subset_1(k3_group_1(X0,X1),u1_struct_0(X0)))),
  file('./p/GRP644+1.tptp',unknown)).
fof(f436,plain,(
  ~m1_subset_1(k3_group_1(sK0,sK2),u1_struct_0(sK0)) | spl17_11),
  inference(avatar_component_clause,[],[f434])).
fof(f434,plain,(
  spl17_11 <=> m1_subset_1(k3_group_1(sK0,sK2),u1_struct_0(sK0))),
  introduced(avatar_definition,[new_symbols(naming,[spl17_11])])).
fof(f441,plain,(
  ~spl17_11 | ~spl17_12),
  inference(avatar_split_clause,[],[f432,f438,f434])).
fof(f432,plain,(
  ~r1_rlvect_1(sK1,k3_group_1(sK0,sK2)) | ~m1_subset_1(k3_group_1(sK0,sK2),u1_struct_0(sK0))),
  inference(subsumption_resolution,[],[f431,f145])).
fof(f431,plain,(
  ~r1_rlvect_1(sK1,k3_group_1(sK0,sK2)) | ~m1_subset_1(k3_group_1(sK0,sK2),u1_struct_0(sK0)) | v3_struct_0(sK0)),
  inference(subsumption_resolution,[],[f430,f146])).
fof(f430,plain,(
  ~r1_rlvect_1(sK1,k3_group_1(sK0,sK2)) | ~m1_subset_1(k3_group_1(sK0,sK2),u1_struct_0(sK0)) | ~v3_group_1(sK0) | v3_struct_0(sK0)),
  inference(subsumption_resolution,[],[f429,f147])).
fof(f429,plain,(
  ~r1_rlvect_1(sK1,k3_group_1(sK0,sK2)) | ~m1_subset_1(k3_group_1(sK0,sK2),u1_struct_0(sK0)) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0)),
  inference(subsumption_resolution,[],[f428,f148])).
fof(f428,plain,(
  ~r1_rlvect_1(sK1,k3_group_1(sK0,sK2)) | ~m1_subset_1(k3_group_1(sK0,sK2),u1_struct_0(sK0)) | ~l1_group_1(sK0) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0)),
  inference(subsumption_resolution,[],[f427,f149])).
fof(f427,plain,(
  ~r1_rlvect_1(sK1,k3_group_1(sK0,sK2)) | ~m1_subset_1(k3_group_1(sK0,sK2),u1_struct_0(sK0)) | ~v1_group_1(sK1) | ~l1_group_1(sK0) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0)),
  inference(subsumption_resolution,[],[f424,f150])).
fof(f424,plain,(
  ~r1_rlvect_1(sK1,k3_group_1(sK0,sK2)) | ~m1_subset_1(k3_group_1(sK0,sK2),u1_struct_0(sK0)) | ~m1_group_2(sK1,sK0) | ~v1_group_1(sK1) | ~l1_group_1(sK0) | ~v4_group_1(sK0) | ~v3_group_1(sK0) | v3_struct_0(sK0)),
  inference(resolution,[],[f211,f153])).
fof(f153,plain,(
  ~r2_hidden(k3_group_1(sK0,sK2),k1_funct_1(k1_latsubgr(sK0),sK1))),
  inference(cnf_transformation,[],[f114])).
fof(f211,plain,(
  ( ! [X2,X0,X1] : (r2_hidden(X2,k1_funct_1(k1_latsubgr(X0),X1)) | ~r1_rlvect_1(X1,X2) | ~m1_subset_1(X2,u1_struct_0(X0)) | ~m1_group_2(X1,X0) | ~v1_group_1(X1) | ~l1_group_1(X0) | ~v4_group_1(X0) | ~v3_group_1(X0) | v3_struct_0(X0)) )),
  inference(cnf_transformation,[],[f144])).
% SZS output end Proof for GRP644+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
