% Refutation found. Thanks to Tanya!
% SZS status Theorem for GRP667+1
% SZS output start Proof for GRP667+1
fof(f93219,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f92918])).
fof(f92918,plain,(
  mult(a,mult(b,mult(a,c))) != mult(a,mult(b,mult(a,c)))),
  inference(superposition,[],[f45,f68849])).
fof(f68849,plain,(
  ( ! [X4,X2,X3] : (mult(mult(X3,mult(X2,X3)),X4) = mult(X3,mult(X2,mult(X3,X4)))) )),
  inference(forward_demodulation,[],[f68848,f27])).
fof(f27,plain,(
  ( ! [X0,X1] : (ld(X1,mult(X1,X0)) = X0) )),
  inference(cnf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1] : ld(X1,mult(X1,X0)) = X0),
  file('./p/GRP667+1.tptp',unknown)).
fof(f68848,plain,(
  ( ! [X4,X2,X3] : (mult(X3,mult(X2,mult(X3,X4))) = mult(ld(X2,mult(X2,mult(X3,mult(X2,X3)))),X4)) )),
  inference(forward_demodulation,[],[f68633,f13271])).
fof(f13271,plain,(
  ( ! [X12,X11] : (mult(X11,mult(X12,mult(X11,X12))) = mult(mult(X11,X12),mult(X11,X12))) )),
  inference(forward_demodulation,[],[f13208,f100])).
fof(f100,plain,(
  ( ! [X6,X7,X5] : (mult(X5,mult(X6,mult(X7,X6))) = rd(mult(mult(X5,X6),mult(X7,mult(X6,X7))),X7)) )),
  inference(superposition,[],[f29,f40])).
fof(f40,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X2,X1),mult(X0,mult(X1,X0))) = mult(mult(X2,mult(X1,mult(X0,X1))),X0)) )),
  inference(forward_demodulation,[],[f39,f33])).
fof(f33,plain,(
  ( ! [X0,X1] : (mult(mult(X1,X0),X1) = mult(X1,mult(X0,X1))) )),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0,X1] : mult(mult(X1,X0),X1) = mult(X1,mult(X0,X1))),
  file('./p/GRP667+1.tptp',unknown)).
fof(f39,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X2,mult(mult(X1,X0),X1)),X0) = mult(mult(X2,X1),mult(X0,mult(X1,X0)))) )),
  inference(backward_demodulation,[],[f33,f32])).
fof(f32,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X2,X1),mult(mult(X0,X1),X0)) = mult(mult(X2,mult(mult(X1,X0),X1)),X0)) )),
  inference(cnf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1,X2] : mult(mult(X2,X1),mult(mult(X0,X1),X0)) = mult(mult(X2,mult(mult(X1,X0),X1)),X0)),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X2,X0,X1] : mult(mult(X1,X0),mult(mult(X2,X0),X2)) = mult(mult(X1,mult(mult(X0,X2),X0)),X2)),
  file('./p/GRP667+1.tptp',unknown)).
fof(f29,plain,(
  ( ! [X0,X1] : (rd(mult(X1,X0),X0) = X1) )),
  inference(cnf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1] : rd(mult(X1,X0),X0) = X1),
  file('./p/GRP667+1.tptp',unknown)).
fof(f13208,plain,(
  ( ! [X12,X11] : (mult(mult(X11,X12),mult(X11,X12)) = rd(mult(mult(X11,X12),mult(X11,mult(X12,X11))),X11)) )),
  inference(superposition,[],[f378,f33])).
fof(f378,plain,(
  ( ! [X12,X13] : (mult(X12,X12) = rd(mult(X12,mult(X12,X13)),X13)) )),
  inference(superposition,[],[f29,f152])).
fof(f152,plain,(
  ( ! [X15,X16] : (mult(mult(X15,X15),X16) = mult(X15,mult(X15,X16))) )),
  inference(forward_demodulation,[],[f151,f31])).
fof(f31,plain,(
  ( ! [X0] : (mult(unit,X0) = X0) )),
  inference(cnf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0] : mult(unit,X0) = X0),
  inference(rectify,[],[f6])).
fof(f6,axiom,(
  ! [X1] : mult(unit,X1) = X1),
  file('./p/GRP667+1.tptp',unknown)).
fof(f151,plain,(
  ( ! [X15,X16] : (mult(X15,mult(unit,mult(X15,X16))) = mult(mult(X15,X15),X16)) )),
  inference(forward_demodulation,[],[f150,f31])).
fof(f150,plain,(
  ( ! [X15,X16] : (mult(X15,mult(unit,mult(X15,X16))) = mult(mult(X15,mult(unit,X15)),X16)) )),
  inference(subsumption_resolution,[],[f149,f31])).
fof(f149,plain,(
  ( ! [X15,X16] : (mult(X15,mult(X16,X15)) != mult(unit,mult(X15,mult(X16,X15))) | mult(X15,mult(unit,mult(X15,X16))) = mult(mult(X15,mult(unit,X15)),X16)) )),
  inference(forward_demodulation,[],[f130,f33])).
fof(f130,plain,(
  ( ! [X15,X16] : (mult(unit,mult(X15,mult(X16,X15))) != mult(mult(X15,X16),X15) | mult(X15,mult(unit,mult(X15,X16))) = mult(mult(X15,mult(unit,X15)),X16)) )),
  inference(superposition,[],[f41,f31])).
fof(f41,plain,(
  ( ! [X2,X0,X1] : (mult(X0,mult(X1,mult(X2,X1))) != mult(mult(mult(X0,X1),X2),X1) | mult(X1,mult(X0,mult(X1,X2))) = mult(mult(X1,mult(X0,X1)),X2)) )),
  inference(forward_demodulation,[],[f35,f33])).
fof(f35,plain,(
  ( ! [X2,X0,X1] : (mult(X1,mult(X0,mult(X1,X2))) = mult(mult(mult(X1,X0),X1),X2) | mult(X0,mult(X1,mult(X2,X1))) != mult(mult(mult(X0,X1),X2),X1)) )),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2] : (mult(X1,mult(X0,mult(X1,X2))) = mult(mult(mult(X1,X0),X1),X2) | mult(X0,mult(X1,mult(X2,X1))) != mult(mult(mult(X0,X1),X2),X1))),
  inference(ennf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2] : (mult(X0,mult(X1,mult(X2,X1))) = mult(mult(mult(X0,X1),X2),X1) => mult(X1,mult(X0,mult(X1,X2))) = mult(mult(mult(X1,X0),X1),X2))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X3,X4,X5] : (mult(X3,mult(X4,mult(X5,X4))) = mult(mult(mult(X3,X4),X5),X4) => mult(X4,mult(X3,mult(X4,X5))) = mult(mult(mult(X4,X3),X4),X5))),
  file('./p/GRP667+1.tptp',unknown)).
fof(f68633,plain,(
  ( ! [X4,X2,X3] : (mult(X3,mult(X2,mult(X3,X4))) = mult(ld(X2,mult(mult(X2,X3),mult(X2,X3))),X4)) )),
  inference(superposition,[],[f67103,f27])).
fof(f67103,plain,(
  ( ! [X4,X2,X3] : (mult(ld(X2,X3),mult(X2,mult(ld(X2,X3),X4))) = mult(ld(X2,mult(X3,X3)),X4)) )),
  inference(trivial_inequality_removal,[],[f67081])).
fof(f67081,plain,(
  ( ! [X4,X2,X3] : (mult(mult(X3,X4),ld(X2,X3)) != mult(mult(X3,X4),ld(X2,X3)) | mult(ld(X2,X3),mult(X2,mult(ld(X2,X3),X4))) = mult(ld(X2,mult(X3,X3)),X4)) )),
  inference(backward_demodulation,[],[f66982,f28574])).
fof(f28574,plain,(
  ( ! [X4,X2,X3] : (mult(X2,mult(ld(X2,X3),mult(X4,ld(X2,X3)))) != mult(mult(X3,X4),ld(X2,X3)) | mult(ld(X2,X3),mult(X2,mult(ld(X2,X3),X4))) = mult(ld(X2,mult(X3,X3)),X4)) )),
  inference(backward_demodulation,[],[f28572,f145])).
fof(f145,plain,(
  ( ! [X4,X2,X3] : (mult(X2,mult(ld(X2,X3),mult(X4,ld(X2,X3)))) != mult(mult(X3,X4),ld(X2,X3)) | mult(ld(X2,X3),mult(X2,mult(ld(X2,X3),X4))) = mult(mult(ld(X2,X3),X3),X4)) )),
  inference(forward_demodulation,[],[f126,f26])).
fof(f26,plain,(
  ( ! [X0,X1] : (mult(X1,ld(X1,X0)) = X0) )),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1] : mult(X1,ld(X1,X0)) = X0),
  file('./p/GRP667+1.tptp',unknown)).
fof(f126,plain,(
  ( ! [X4,X2,X3] : (mult(X2,mult(ld(X2,X3),mult(X4,ld(X2,X3)))) != mult(mult(X3,X4),ld(X2,X3)) | mult(ld(X2,X3),mult(X2,mult(ld(X2,X3),X4))) = mult(mult(ld(X2,X3),mult(X2,ld(X2,X3))),X4)) )),
  inference(superposition,[],[f41,f26])).
fof(f28572,plain,(
  ( ! [X6,X7] : (mult(ld(X6,X7),X7) = ld(X6,mult(X7,X7))) )),
  inference(forward_demodulation,[],[f28477,f15503])).
fof(f15503,plain,(
  ( ! [X14,X13] : (ld(X14,mult(X13,X13)) = rd(ld(X14,mult(X13,mult(X13,X14))),X14)) )),
  inference(superposition,[],[f555,f152])).
fof(f555,plain,(
  ( ! [X2,X3] : (ld(X2,X3) = rd(ld(X2,mult(X3,X2)),X2)) )),
  inference(superposition,[],[f29,f237])).
fof(f237,plain,(
  ( ! [X4,X3] : (mult(ld(X3,X4),X3) = ld(X3,mult(X4,X3))) )),
  inference(superposition,[],[f27,f67])).
fof(f67,plain,(
  ( ! [X2,X1] : (mult(X2,X1) = mult(X1,mult(ld(X1,X2),X1))) )),
  inference(superposition,[],[f33,f26])).
fof(f28477,plain,(
  ( ! [X6,X7] : (mult(ld(X6,X7),X7) = rd(ld(X6,mult(X7,mult(X7,X6))),X6)) )),
  inference(backward_demodulation,[],[f28251,f25872])).
fof(f25872,plain,(
  ( ! [X6,X7] : (mult(ld(X6,X7),X7) = rd(mult(ld(X6,X7),mult(X7,X6)),X6)) )),
  inference(forward_demodulation,[],[f25871,f26])).
fof(f25871,plain,(
  ( ! [X6,X7] : (mult(ld(X6,X7),mult(X6,ld(X6,X7))) = rd(mult(ld(X6,X7),mult(X7,X6)),X6)) )),
  inference(forward_demodulation,[],[f25779,f26])).
fof(f25779,plain,(
  ( ! [X6,X7] : (mult(ld(X6,X7),mult(X6,ld(X6,X7))) = rd(mult(ld(X6,X7),mult(X6,ld(X6,mult(X7,X6)))),X6)) )),
  inference(superposition,[],[f697,f237])).
fof(f697,plain,(
  ( ! [X10,X11] : (mult(X10,mult(X11,X10)) = rd(mult(X10,mult(X11,mult(X10,X11))),X11)) )),
  inference(superposition,[],[f29,f114])).
fof(f114,plain,(
  ( ! [X10,X11] : (mult(mult(X10,mult(X11,X10)),X11) = mult(X10,mult(X11,mult(X10,X11)))) )),
  inference(forward_demodulation,[],[f95,f31])).
fof(f95,plain,(
  ( ! [X10,X11] : (mult(mult(unit,X10),mult(X11,mult(X10,X11))) = mult(mult(X10,mult(X11,X10)),X11)) )),
  inference(superposition,[],[f40,f31])).
fof(f28251,plain,(
  ( ! [X10,X11] : (mult(ld(X10,X11),mult(X11,X10)) = ld(X10,mult(X11,mult(X11,X10)))) )),
  inference(superposition,[],[f27,f1680])).
fof(f1680,plain,(
  ( ! [X0,X1] : (mult(X1,mult(X1,X0)) = mult(X0,mult(ld(X0,X1),mult(X1,X0)))) )),
  inference(forward_demodulation,[],[f1679,f714])).
fof(f714,plain,(
  ( ! [X2,X1] : (mult(mult(ld(X1,X2),X2),X1) = mult(ld(X1,X2),mult(X2,X1))) )),
  inference(forward_demodulation,[],[f713,f26])).
fof(f713,plain,(
  ( ! [X2,X1] : (mult(mult(ld(X1,X2),X2),X1) = mult(ld(X1,X2),mult(X1,ld(X1,mult(X2,X1))))) )),
  inference(forward_demodulation,[],[f669,f237])).
fof(f669,plain,(
  ( ! [X2,X1] : (mult(ld(X1,X2),mult(X1,mult(ld(X1,X2),X1))) = mult(mult(ld(X1,X2),X2),X1)) )),
  inference(superposition,[],[f114,f26])).
fof(f1679,plain,(
  ( ! [X0,X1] : (mult(X1,mult(X1,X0)) = mult(X0,mult(mult(ld(X0,X1),X1),X0))) )),
  inference(forward_demodulation,[],[f1622,f26])).
fof(f1622,plain,(
  ( ! [X0,X1] : (mult(X0,mult(mult(ld(X0,X1),X1),X0)) = mult(mult(X0,ld(X0,X1)),mult(X1,X0))) )),
  inference(superposition,[],[f105,f33])).
fof(f105,plain,(
  ( ! [X4,X2,X3] : (mult(mult(X4,mult(ld(X2,X3),X3)),X2) = mult(mult(X4,ld(X2,X3)),mult(X3,X2))) )),
  inference(forward_demodulation,[],[f85,f67])).
fof(f85,plain,(
  ( ! [X4,X2,X3] : (mult(mult(X4,ld(X2,X3)),mult(X2,mult(ld(X2,X3),X2))) = mult(mult(X4,mult(ld(X2,X3),X3)),X2)) )),
  inference(superposition,[],[f40,f26])).
fof(f66982,plain,(
  ( ! [X4,X2,X3] : (mult(X2,mult(ld(X2,X3),mult(X4,ld(X2,X3)))) = mult(mult(X3,X4),ld(X2,X3))) )),
  inference(backward_demodulation,[],[f66796,f1190])).
fof(f1190,plain,(
  ( ! [X4,X2,X3] : (mult(X2,mult(ld(X2,X3),mult(X4,ld(X2,X3)))) = rd(mult(X3,mult(X4,mult(ld(X2,X3),X4))),X4)) )),
  inference(superposition,[],[f100,f26])).
fof(f66796,plain,(
  ( ! [X37,X38,X36] : (rd(mult(X36,mult(X38,mult(X37,X38))),X38) = mult(mult(X36,X38),X37)) )),
  inference(backward_demodulation,[],[f66795,f1201])).
fof(f1201,plain,(
  ( ! [X37,X38,X36] : (mult(rd(X36,X37),mult(X37,mult(X38,X37))) = rd(mult(X36,mult(X38,mult(X37,X38))),X38)) )),
  inference(superposition,[],[f100,f28])).
fof(f28,plain,(
  ( ! [X0,X1] : (mult(rd(X1,X0),X0) = X1) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1] : mult(rd(X1,X0),X0) = X1),
  file('./p/GRP667+1.tptp',unknown)).
fof(f66795,plain,(
  ( ! [X4,X2,X3] : (mult(mult(X4,X3),X2) = mult(rd(X4,X2),mult(X2,mult(X3,X2)))) )),
  inference(forward_demodulation,[],[f66528,f33])).
fof(f66528,plain,(
  ( ! [X4,X2,X3] : (mult(mult(X4,X3),X2) = mult(rd(X4,X2),mult(mult(X2,X3),X2))) )),
  inference(superposition,[],[f60890,f27])).
fof(f60890,plain,(
  ( ! [X21,X19,X20] : (mult(mult(X20,ld(X21,X19)),X21) = mult(rd(X20,X21),mult(X19,X21))) )),
  inference(forward_demodulation,[],[f60632,f369])).
fof(f369,plain,(
  ( ! [X4,X5] : (mult(X4,X5) = mult(f(X4),mult(f(X4),X5))) )),
  inference(superposition,[],[f152,f34])).
fof(f34,plain,(
  ( ! [X0] : (mult(f(X0),f(X0)) = X0) )),
  inference(cnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0] : mult(f(X0),f(X0)) = X0),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X1] : mult(f(X1),f(X1)) = X1),
  file('./p/GRP667+1.tptp',unknown)).
fof(f60632,plain,(
  ( ! [X21,X19,X20] : (mult(rd(X20,X21),mult(f(X19),mult(f(X19),X21))) = mult(mult(X20,ld(X21,X19)),X21)) )),
  inference(superposition,[],[f60329,f34])).
fof(f60329,plain,(
  ( ! [X4,X2,X3] : (mult(mult(X4,ld(X2,mult(X3,X3))),X2) = mult(rd(X4,X2),mult(X3,mult(X3,X2)))) )),
  inference(backward_demodulation,[],[f60327,f28573])).
fof(f28573,plain,(
  ( ! [X4,X2,X3] : (mult(mult(X4,ld(X2,X3)),mult(X3,X2)) = mult(mult(X4,ld(X2,mult(X3,X3))),X2)) )),
  inference(backward_demodulation,[],[f28572,f105])).
fof(f60327,plain,(
  ( ! [X8,X7,X9] : (mult(mult(X9,ld(X8,X7)),mult(X7,X8)) = mult(rd(X9,X8),mult(X7,mult(X7,X8)))) )),
  inference(forward_demodulation,[],[f60292,f57981])).
fof(f57981,plain,(
  ( ! [X163,X162] : (rd(X163,X162) = mult(X163,ld(X162,unit))) )),
  inference(forward_demodulation,[],[f57980,f56141])).
fof(f56141,plain,(
  ( ! [X17,X16] : (mult(ld(X16,unit),mult(X16,X17)) = X17) )),
  inference(forward_demodulation,[],[f56140,f29])).
fof(f56140,plain,(
  ( ! [X17,X16] : (mult(ld(X16,unit),mult(X16,X17)) = rd(mult(X17,rd(X17,X16)),rd(X17,X16))) )),
  inference(forward_demodulation,[],[f55966,f31])).
fof(f55966,plain,(
  ( ! [X17,X16] : (mult(ld(X16,unit),mult(X16,X17)) = rd(mult(unit,mult(X17,rd(X17,X16))),rd(X17,X16))) )),
  inference(superposition,[],[f1285,f762])).
fof(f762,plain,(
  ( ! [X0] : (unit = mult(ld(X0,unit),X0)) )),
  inference(superposition,[],[f655,f58])).
fof(f58,plain,(
  ( ! [X2,X1] : (rd(X2,ld(X1,X2)) = X1) )),
  inference(superposition,[],[f29,f26])).
fof(f655,plain,(
  ( ! [X0] : (unit = mult(X0,rd(unit,X0))) )),
  inference(forward_demodulation,[],[f630,f59])).
fof(f59,plain,(
  ( ! [X3] : (unit = rd(X3,X3)) )),
  inference(superposition,[],[f29,f31])).
fof(f630,plain,(
  ( ! [X0] : (rd(X0,X0) = mult(X0,rd(unit,X0))) )),
  inference(superposition,[],[f282,f30])).
fof(f30,plain,(
  ( ! [X0] : (mult(X0,unit) = X0) )),
  inference(cnf_transformation,[],[f15])).
fof(f15,plain,(
  ! [X0] : mult(X0,unit) = X0),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X1] : mult(X1,unit) = X1),
  file('./p/GRP667+1.tptp',unknown)).
fof(f282,plain,(
  ( ! [X8,X9] : (mult(X9,rd(X8,X9)) = rd(mult(X9,X8),X9)) )),
  inference(superposition,[],[f74,f28])).
fof(f74,plain,(
  ( ! [X2,X1] : (mult(X1,X2) = rd(mult(X1,mult(X2,X1)),X1)) )),
  inference(superposition,[],[f29,f33])).
fof(f1285,plain,(
  ( ! [X6,X8,X7] : (mult(X8,mult(X7,X6)) = rd(mult(mult(X8,X7),mult(X6,rd(X6,X7))),rd(X6,X7))) )),
  inference(forward_demodulation,[],[f1227,f28])).
fof(f1227,plain,(
  ( ! [X6,X8,X7] : (mult(X8,mult(X7,mult(rd(X6,X7),X7))) = rd(mult(mult(X8,X7),mult(X6,rd(X6,X7))),rd(X6,X7))) )),
  inference(superposition,[],[f100,f69])).
fof(f69,plain,(
  ( ! [X6,X5] : (mult(rd(X5,X6),mult(X6,rd(X5,X6))) = mult(X5,rd(X5,X6))) )),
  inference(superposition,[],[f33,f28])).
fof(f57980,plain,(
  ( ! [X163,X162] : (mult(X163,ld(X162,unit)) = mult(ld(X162,unit),mult(X162,rd(X163,X162)))) )),
  inference(forward_demodulation,[],[f56977,f57885])).
fof(f57885,plain,(
  ( ! [X10,X11] : (mult(mult(X10,X11),ld(X10,unit)) = mult(X10,rd(X11,X10))) )),
  inference(backward_demodulation,[],[f57880,f35438])).
fof(f35438,plain,(
  ( ! [X10,X11] : (mult(X10,rd(X11,X10)) = mult(mult(X10,X11),ld(X11,rd(X11,X10)))) )),
  inference(backward_demodulation,[],[f35435,f34322])).
fof(f34322,plain,(
  ( ! [X10,X11] : (mult(mult(X10,X11),rd(X11,mult(X10,X11))) = mult(X10,rd(X11,X10))) )),
  inference(backward_demodulation,[],[f33805,f635])).
fof(f635,plain,(
  ( ! [X10,X11] : (mult(mult(X10,X11),rd(X11,mult(X10,X11))) = rd(mult(X10,mult(X11,X11)),mult(X10,X11))) )),
  inference(superposition,[],[f282,f104])).
fof(f104,plain,(
  ( ! [X0,X1] : (mult(mult(X1,X0),X0) = mult(X1,mult(X0,X0))) )),
  inference(forward_demodulation,[],[f103,f30])).
fof(f103,plain,(
  ( ! [X0,X1] : (mult(mult(X1,unit),mult(X0,X0)) = mult(mult(X1,X0),X0)) )),
  inference(forward_demodulation,[],[f84,f31])).
fof(f84,plain,(
  ( ! [X0,X1] : (mult(mult(X1,unit),mult(X0,mult(unit,X0))) = mult(mult(X1,mult(unit,X0)),X0)) )),
  inference(superposition,[],[f40,f30])).
fof(f33805,plain,(
  ( ! [X50,X51] : (mult(X50,rd(X51,X50)) = rd(mult(X50,mult(X51,X51)),mult(X50,X51))) )),
  inference(superposition,[],[f29,f27504])).
fof(f27504,plain,(
  ( ! [X35,X36] : (mult(mult(X36,rd(X35,X36)),mult(X36,X35)) = mult(X36,mult(X35,X35))) )),
  inference(backward_demodulation,[],[f27500,f2151])).
fof(f2151,plain,(
  ( ! [X35,X36] : (mult(mult(X36,rd(X35,X36)),mult(X36,X35)) = mult(X36,mult(rd(X35,X36),mult(X36,X35)))) )),
  inference(superposition,[],[f115,f28])).
fof(f115,plain,(
  ( ! [X0,X1] : (mult(mult(X0,X1),mult(X0,mult(X1,X0))) = mult(X0,mult(X1,mult(X0,mult(X1,X0))))) )),
  inference(forward_demodulation,[],[f96,f114])).
fof(f96,plain,(
  ( ! [X0,X1] : (mult(X0,mult(mult(X1,mult(X0,X1)),X0)) = mult(mult(X0,X1),mult(X0,mult(X1,X0)))) )),
  inference(superposition,[],[f40,f33])).
fof(f27500,plain,(
  ( ! [X35,X34] : (mult(X35,X35) = mult(rd(X35,X34),mult(X34,X35))) )),
  inference(forward_demodulation,[],[f27499,f378])).
fof(f27499,plain,(
  ( ! [X35,X34] : (mult(rd(X35,X34),mult(X34,X35)) = rd(mult(X35,mult(X35,rd(X35,X34))),rd(X35,X34))) )),
  inference(forward_demodulation,[],[f27316,f1820])).
fof(f1820,plain,(
  ( ! [X4,X3] : (mult(X3,mult(X3,rd(X3,X4))) = mult(rd(X3,X4),mult(X4,mult(X3,rd(X3,X4))))) )),
  inference(forward_demodulation,[],[f1819,f719])).
fof(f719,plain,(
  ( ! [X19,X18] : (mult(mult(X19,X18),rd(X18,X19)) = mult(X19,mult(X18,rd(X18,X19)))) )),
  inference(forward_demodulation,[],[f677,f69])).
fof(f677,plain,(
  ( ! [X19,X18] : (mult(X19,mult(rd(X18,X19),mult(X19,rd(X18,X19)))) = mult(mult(X19,X18),rd(X18,X19))) )),
  inference(superposition,[],[f114,f28])).
fof(f1819,plain,(
  ( ! [X4,X3] : (mult(rd(X3,X4),mult(mult(X4,X3),rd(X3,X4))) = mult(X3,mult(X3,rd(X3,X4)))) )),
  inference(forward_demodulation,[],[f1751,f28])).
fof(f1751,plain,(
  ( ! [X4,X3] : (mult(rd(X3,X4),mult(mult(X4,X3),rd(X3,X4))) = mult(mult(rd(X3,X4),X4),mult(X3,rd(X3,X4)))) )),
  inference(superposition,[],[f106,f33])).
fof(f106,plain,(
  ( ! [X14,X12,X13] : (mult(mult(X14,mult(X13,X12)),rd(X12,X13)) = mult(mult(X14,X13),mult(X12,rd(X12,X13)))) )),
  inference(forward_demodulation,[],[f88,f69])).
fof(f88,plain,(
  ( ! [X14,X12,X13] : (mult(mult(X14,X13),mult(rd(X12,X13),mult(X13,rd(X12,X13)))) = mult(mult(X14,mult(X13,X12)),rd(X12,X13))) )),
  inference(superposition,[],[f40,f28])).
fof(f27316,plain,(
  ( ! [X35,X34] : (mult(rd(X35,X34),mult(X34,X35)) = rd(mult(rd(X35,X34),mult(X34,mult(X35,rd(X35,X34)))),rd(X35,X34))) )),
  inference(superposition,[],[f74,f719])).
fof(f35435,plain,(
  ( ! [X239,X238] : (rd(X239,mult(X238,X239)) = ld(X239,rd(X239,X238))) )),
  inference(backward_demodulation,[],[f35434,f34444])).
fof(f34444,plain,(
  ( ! [X239,X238] : (ld(mult(X238,X239),mult(X238,rd(X239,X238))) = rd(X239,mult(X238,X239))) )),
  inference(forward_demodulation,[],[f33871,f346])).
fof(f346,plain,(
  ( ! [X8,X7] : (ld(mult(X7,X8),mult(X7,mult(X8,X8))) = X8) )),
  inference(superposition,[],[f27,f104])).
fof(f33871,plain,(
  ( ! [X239,X238] : (ld(mult(X238,X239),mult(X238,rd(X239,X238))) = rd(ld(mult(X238,X239),mult(X238,mult(X239,X239))),mult(X238,X239))) )),
  inference(superposition,[],[f555,f27504])).
fof(f35434,plain,(
  ( ! [X48,X49] : (ld(mult(X48,X49),mult(X48,rd(X49,X48))) = ld(X49,rd(X49,X48))) )),
  inference(forward_demodulation,[],[f35433,f346])).
fof(f35433,plain,(
  ( ! [X48,X49] : (ld(mult(X48,X49),mult(X48,rd(X49,X48))) = ld(ld(mult(X48,X49),mult(X48,mult(X49,X49))),rd(X49,X48))) )),
  inference(forward_demodulation,[],[f35432,f25438])).
fof(f25438,plain,(
  ( ! [X76,X77] : (rd(X76,X77) = ld(mult(X77,X76),mult(X77,mult(X76,rd(X76,X77))))) )),
  inference(forward_demodulation,[],[f25264,f69])).
fof(f25264,plain,(
  ( ! [X76,X77] : (rd(X76,X77) = ld(mult(X77,X76),mult(X77,mult(rd(X76,X77),mult(X77,rd(X76,X77)))))) )),
  inference(superposition,[],[f696,f28])).
fof(f696,plain,(
  ( ! [X8,X9] : (ld(mult(X8,mult(X9,X8)),mult(X8,mult(X9,mult(X8,X9)))) = X9) )),
  inference(superposition,[],[f27,f114])).
fof(f35432,plain,(
  ( ! [X48,X49] : (ld(mult(X48,X49),mult(X48,rd(X49,X48))) = ld(ld(mult(X48,X49),mult(X48,mult(X49,X49))),ld(mult(X48,X49),mult(X48,mult(X49,rd(X49,X48)))))) )),
  inference(forward_demodulation,[],[f35431,f69])).
fof(f35431,plain,(
  ( ! [X48,X49] : (ld(mult(X48,X49),mult(X48,rd(X49,X48))) = ld(ld(mult(X48,X49),mult(X48,mult(X49,X49))),ld(mult(X48,X49),mult(X48,mult(rd(X49,X48),mult(X48,rd(X49,X48))))))) )),
  inference(forward_demodulation,[],[f35277,f13271])).
fof(f35277,plain,(
  ( ! [X48,X49] : (ld(mult(X48,X49),mult(X48,rd(X49,X48))) = ld(ld(mult(X48,X49),mult(X48,mult(X49,X49))),ld(mult(X48,X49),mult(mult(X48,rd(X49,X48)),mult(X48,rd(X49,X48)))))) )),
  inference(superposition,[],[f28577,f27504])).
fof(f28577,plain,(
  ( ! [X2,X1] : (ld(X1,X2) = ld(ld(X1,mult(X2,X1)),ld(X1,mult(X2,X2)))) )),
  inference(backward_demodulation,[],[f28572,f328])).
fof(f328,plain,(
  ( ! [X2,X1] : (ld(X1,X2) = ld(ld(X1,mult(X2,X1)),mult(ld(X1,X2),X2))) )),
  inference(forward_demodulation,[],[f311,f237])).
fof(f311,plain,(
  ( ! [X2,X1] : (ld(X1,X2) = ld(mult(ld(X1,X2),X1),mult(ld(X1,X2),X2))) )),
  inference(superposition,[],[f75,f26])).
fof(f75,plain,(
  ( ! [X4,X3] : (ld(mult(X3,X4),mult(X3,mult(X4,X3))) = X3) )),
  inference(superposition,[],[f27,f33])).
fof(f57880,plain,(
  ( ! [X239,X238] : (ld(X239,rd(X239,X238)) = ld(X238,unit)) )),
  inference(backward_demodulation,[],[f56976,f35435])).
fof(f56976,plain,(
  ( ! [X161,X160] : (ld(X160,unit) = rd(X161,mult(X160,X161))) )),
  inference(superposition,[],[f29,f56141])).
fof(f56977,plain,(
  ( ! [X163,X162] : (mult(ld(X162,unit),mult(mult(X162,X163),ld(X162,unit))) = mult(X163,ld(X162,unit))) )),
  inference(superposition,[],[f33,f56141])).
fof(f60292,plain,(
  ( ! [X8,X7,X9] : (mult(mult(X9,ld(X8,X7)),mult(X7,X8)) = mult(mult(X9,ld(X8,unit)),mult(X7,mult(X7,X8)))) )),
  inference(backward_demodulation,[],[f60008,f34555])).
fof(f34555,plain,(
  ( ! [X8,X7,X9] : (mult(mult(X9,ld(mult(X7,X8),X7)),mult(X7,mult(X7,X8))) = mult(mult(X9,ld(X8,X7)),mult(X7,X8))) )),
  inference(superposition,[],[f28573,f26847])).
fof(f26847,plain,(
  ( ! [X6,X7] : (ld(X6,X7) = ld(mult(X7,X6),mult(X7,X7))) )),
  inference(backward_demodulation,[],[f26836,f25381])).
fof(f25381,plain,(
  ( ! [X6,X7] : (ld(X6,X7) = ld(mult(X7,X6),mult(X6,mult(ld(X6,X7),X7)))) )),
  inference(forward_demodulation,[],[f25380,f26])).
fof(f25380,plain,(
  ( ! [X6,X7] : (ld(X6,X7) = ld(mult(X6,ld(X6,mult(X7,X6))),mult(X6,mult(ld(X6,X7),X7)))) )),
  inference(forward_demodulation,[],[f25230,f26])).
fof(f25230,plain,(
  ( ! [X6,X7] : (ld(X6,X7) = ld(mult(X6,ld(X6,mult(X7,X6))),mult(X6,mult(ld(X6,X7),mult(X6,ld(X6,X7)))))) )),
  inference(superposition,[],[f696,f237])).
fof(f26836,plain,(
  ( ! [X66,X67] : (mult(X67,X67) = mult(X66,mult(ld(X66,X67),X67))) )),
  inference(forward_demodulation,[],[f26835,f378])).
fof(f26835,plain,(
  ( ! [X66,X67] : (mult(X66,mult(ld(X66,X67),X67)) = rd(mult(X67,mult(X67,X66)),X66)) )),
  inference(forward_demodulation,[],[f26629,f1680])).
fof(f26629,plain,(
  ( ! [X66,X67] : (mult(X66,mult(ld(X66,X67),X67)) = rd(mult(X66,mult(ld(X66,X67),mult(X67,X66))),X66)) )),
  inference(superposition,[],[f74,f714])).
fof(f60008,plain,(
  ( ! [X39,X38] : (ld(X39,unit) = ld(mult(X38,X39),X38)) )),
  inference(superposition,[],[f57880,f29])).
fof(f45,plain,(
  mult(a,mult(b,mult(a,c))) != mult(mult(a,mult(b,a)),c)),
  inference(forward_demodulation,[],[f38,f33])).
fof(f38,plain,(
  mult(a,mult(b,mult(a,c))) != mult(mult(mult(a,b),a),c)),
  inference(cnf_transformation,[],[f22])).
fof(f22,plain,(
  mult(a,mult(b,mult(a,c))) != mult(mult(mult(a,b),a),c)),
  inference(flattening,[],[f14])).
fof(f14,negated_conjecture,(
  ~mult(a,mult(b,mult(a,c))) = mult(mult(mult(a,b),a),c)),
  inference(negated_conjecture,[],[f13])).
fof(f13,conjecture,(
  mult(a,mult(b,mult(a,c))) = mult(mult(mult(a,b),a),c)),
  file('./p/GRP667+1.tptp',unknown)).
% SZS output end Proof for GRP667+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 104646
% Time elapsed: 4.522 s
% ------------------------------
% ------------------------------
