% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for GRP752-1
% SZS output start Proof for GRP752-1
fof(f143415,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f143381])).
fof(f143381,plain,(
  mult(mult(a,b),c) != mult(mult(a,b),c)),
  inference(superposition,[],[f8,f98551])).
fof(f98551,plain,(
  ( ! [X116,X117,X115] : (mult(mult(X117,X116),mult(X115,ld(X116,X116))) = mult(mult(X117,X115),X116)) )),
  inference(forward_demodulation,[],[f98525,f10])).
fof(f10,plain,(
  ( ! [X0,X1] : (rd(X1,ld(X0,X1)) = X0) )),
  inference(superposition,[],[f4,f1])).
fof(f1,axiom,(
  ( ! [X0,X1] : (mult(X0,ld(X0,X1)) = X1) )),
  file('./p/GRP752-1.tptp',unknown)).
fof(f4,axiom,(
  ( ! [X0,X1] : (rd(mult(X0,X1),X1) = X0) )),
  file('./p/GRP752-1.tptp',unknown)).
fof(f98525,plain,(
  ( ! [X116,X117,X115] : (mult(mult(X117,X116),mult(X115,ld(X116,X116))) = mult(mult(X117,rd(X116,ld(X115,X116))),X116)) )),
  inference(superposition,[],[f98449,f91344])).
fof(f91344,plain,(
  ( ! [X2,X0,X1] : (rd(X2,ld(X1,X0)) = rd(mult(X1,ld(X2,X2)),ld(X2,X0))) )),
  inference(superposition,[],[f89136,f10])).
fof(f89136,plain,(
  ( ! [X17,X15,X16] : (rd(X16,X17) = rd(mult(rd(X15,X17),ld(X16,X16)),ld(X16,X15))) )),
  inference(forward_demodulation,[],[f89135,f4])).
fof(f89135,plain,(
  ( ! [X17,X15,X16] : (rd(X16,X17) = rd(mult(rd(X15,X17),rd(mult(ld(X16,X16),X17),X17)),ld(X16,X15))) )),
  inference(forward_demodulation,[],[f89096,f6589])).
fof(f6589,plain,(
  ( ! [X4,X5,X3] : (mult(rd(X5,X4),rd(X3,X4)) = rd(mult(X5,X3),mult(X4,X4))) )),
  inference(superposition,[],[f526,f3])).
fof(f3,axiom,(
  ( ! [X0,X1] : (mult(rd(X0,X1),X1) = X0) )),
  file('./p/GRP752-1.tptp',unknown)).
fof(f526,plain,(
  ( ! [X23,X21,X22] : (mult(rd(X21,X22),X23) = rd(mult(X21,mult(X23,X22)),mult(X22,X22))) )),
  inference(superposition,[],[f4,f40])).
fof(f40,plain,(
  ( ! [X4,X5,X3] : (mult(mult(rd(X3,X4),X5),mult(X4,X4)) = mult(X3,mult(X5,X4))) )),
  inference(superposition,[],[f7,f3])).
fof(f7,axiom,(
  ( ! [X2,X0,X1] : (mult(mult(X0,X1),mult(X2,X2)) = mult(mult(X0,X2),mult(X1,X2))) )),
  file('./p/GRP752-1.tptp',unknown)).
fof(f89096,plain,(
  ( ! [X17,X15,X16] : (rd(X16,X17) = rd(rd(mult(X15,mult(ld(X16,X16),X17)),mult(X17,X17)),ld(X16,X15))) )),
  inference(superposition,[],[f30345,f85992])).
fof(f85992,plain,(
  ( ! [X17,X15,X16] : (rd(mult(X17,mult(ld(X15,X15),X16)),mult(ld(X15,X17),X16)) = X15) )),
  inference(superposition,[],[f48853,f2])).
fof(f2,axiom,(
  ( ! [X0,X1] : (ld(X0,mult(X0,X1)) = X1) )),
  file('./p/GRP752-1.tptp',unknown)).
fof(f48853,plain,(
  ( ! [X14,X12,X13] : (rd(mult(X13,X14),mult(ld(X12,X13),ld(ld(X12,X12),X14))) = X12) )),
  inference(forward_demodulation,[],[f48852,f10252])).
fof(f10252,plain,(
  ( ! [X304,X306,X303,X305] : (mult(ld(X303,X306),ld(ld(X303,X304),X305)) = ld(mult(X303,X303),mult(X306,ld(X304,mult(mult(X303,X303),X305))))) )),
  inference(superposition,[],[f170,f933])).
fof(f933,plain,(
  ( ! [X2,X0,X1] : (mult(X1,ld(ld(X1,X0),X2)) = ld(X0,mult(mult(X1,X1),X2))) )),
  inference(superposition,[],[f2,f104])).
fof(f104,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X0,X0),X2) = mult(X1,mult(X0,ld(ld(X0,X1),X2)))) )),
  inference(superposition,[],[f16,f1])).
fof(f16,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X2,X0),mult(X2,ld(X0,X1))) = mult(mult(X2,X2),X1)) )),
  inference(superposition,[],[f6,f1])).
fof(f6,axiom,(
  ( ! [X2,X0,X1] : (mult(mult(X0,X0),mult(X1,X2)) = mult(mult(X0,X1),mult(X0,X2))) )),
  file('./p/GRP752-1.tptp',unknown)).
fof(f170,plain,(
  ( ! [X12,X13,X11] : (mult(ld(X11,X12),X13) = ld(mult(X11,X11),mult(X12,mult(X11,X13)))) )),
  inference(superposition,[],[f2,f20])).
fof(f20,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X0,X0),mult(ld(X0,X1),X2)) = mult(X1,mult(X0,X2))) )),
  inference(superposition,[],[f6,f1])).
fof(f48852,plain,(
  ( ! [X14,X12,X13] : (rd(mult(X13,X14),ld(mult(X12,X12),mult(X13,ld(X12,mult(mult(X12,X12),X14))))) = X12) )),
  inference(forward_demodulation,[],[f48851,f46585])).
fof(f46585,plain,(
  ( ! [X21,X19,X22,X20] : (ld(X22,mult(X20,ld(X19,X21))) = ld(X19,ld(mult(X19,X22),mult(mult(X19,X20),X21)))) )),
  inference(superposition,[],[f11258,f193])).
fof(f193,plain,(
  ( ! [X19,X20,X18] : (mult(X19,ld(X18,X20)) = ld(mult(X18,X18),mult(mult(X18,X19),X20))) )),
  inference(superposition,[],[f2,f23])).
fof(f23,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X0,X0),mult(X2,ld(X0,X1))) = mult(mult(X0,X2),X1)) )),
  inference(superposition,[],[f6,f1])).
fof(f11258,plain,(
  ( ! [X21,X22,X20] : (ld(X21,ld(mult(X20,X20),X22)) = ld(X20,ld(mult(X20,X21),X22))) )),
  inference(superposition,[],[f2,f1157])).
fof(f1157,plain,(
  ( ! [X47,X48,X46] : (mult(X46,ld(X48,ld(mult(X46,X46),X47))) = ld(mult(X46,X48),X47)) )),
  inference(superposition,[],[f121,f1])).
fof(f121,plain,(
  ( ! [X14,X12,X13] : (mult(X12,ld(X13,X14)) = ld(mult(X12,X13),mult(mult(X12,X12),X14))) )),
  inference(superposition,[],[f2,f16])).
fof(f48851,plain,(
  ( ! [X14,X12,X13] : (rd(mult(X13,X14),ld(X12,ld(mult(X12,mult(X12,X12)),mult(mult(X12,X13),mult(mult(X12,X12),X14))))) = X12) )),
  inference(forward_demodulation,[],[f48789,f11258])).
fof(f48789,plain,(
  ( ! [X14,X12,X13] : (rd(mult(X13,X14),ld(mult(X12,X12),ld(mult(X12,X12),mult(mult(X12,X13),mult(mult(X12,X12),X14))))) = X12) )),
  inference(superposition,[],[f11259,f85])).
fof(f85,plain,(
  ( ! [X14,X12,X13] : (mult(X13,X14) = ld(mult(X12,mult(X12,X12)),mult(mult(X12,X13),mult(mult(X12,X12),X14)))) )),
  inference(superposition,[],[f2,f5])).
fof(f5,axiom,(
  ( ! [X2,X0,X1] : (mult(mult(X0,mult(X0,X0)),mult(X1,X2)) = mult(mult(X0,X1),mult(mult(X0,X0),X2))) )),
  file('./p/GRP752-1.tptp',unknown)).
fof(f11259,plain,(
  ( ! [X24,X23,X25] : (rd(ld(mult(X23,X24),X25),ld(X24,ld(mult(X23,X23),X25))) = X23) )),
  inference(superposition,[],[f4,f1157])).
fof(f30345,plain,(
  ( ! [X28,X26,X27] : (rd(rd(X26,mult(X27,X27)),X28) = rd(rd(X26,mult(X28,X27)),X27)) )),
  inference(superposition,[],[f4,f5182])).
fof(f5182,plain,(
  ( ! [X2,X0,X1] : (mult(rd(rd(X0,mult(X1,X1)),X2),X1) = rd(X0,mult(X2,X1))) )),
  inference(superposition,[],[f482,f3])).
fof(f482,plain,(
  ( ! [X21,X19,X20] : (mult(rd(X19,X20),X21) = rd(mult(X19,mult(X21,X21)),mult(X20,X21))) )),
  inference(superposition,[],[f4,f32])).
fof(f32,plain,(
  ( ! [X4,X5,X3] : (mult(mult(rd(X3,X4),X5),mult(X4,X5)) = mult(X3,mult(X5,X5))) )),
  inference(superposition,[],[f7,f3])).
fof(f98449,plain,(
  ( ! [X121,X120,X119] : (mult(mult(X119,X120),X121) = mult(mult(X119,rd(X121,ld(X120,X120))),X120)) )),
  inference(superposition,[],[f98406,f4])).
fof(f98406,plain,(
  ( ! [X17,X15,X16] : (mult(X15,X17) = mult(mult(rd(X15,X16),rd(X17,ld(X16,X16))),X16)) )),
  inference(forward_demodulation,[],[f98405,f2])).
fof(f98405,plain,(
  ( ! [X17,X15,X16] : (mult(X15,X17) = mult(mult(rd(X15,X16),ld(X15,mult(X15,rd(X17,ld(X16,X16))))),X16)) )),
  inference(backward_demodulation,[],[f98404,f98038])).
fof(f98038,plain,(
  ( ! [X17,X15,X16] : (mult(X15,X17) = mult(mult(rd(X15,X16),ld(X15,rd(mult(mult(X15,X16),X17),X16))),X16)) )),
  inference(backward_demodulation,[],[f98037,f94397])).
fof(f94397,plain,(
  ( ! [X17,X15,X16] : (mult(X15,X17) = mult(mult(rd(X15,X16),rd(X15,ld(rd(X17,ld(X16,X15)),X16))),X16)) )),
  inference(forward_demodulation,[],[f94391,f35770])).
fof(f35770,plain,(
  ( ! [X70,X68,X69] : (rd(mult(X68,X68),mult(X70,X69)) = mult(rd(X68,X70),rd(X68,X69))) )),
  inference(backward_demodulation,[],[f35730,f5278])).
fof(f5278,plain,(
  ( ! [X70,X68,X69] : (mult(rd(mult(rd(X68,X69),rd(X68,X69)),X70),X69) = rd(mult(X68,X68),mult(X70,X69))) )),
  inference(forward_demodulation,[],[f5202,f3])).
fof(f5202,plain,(
  ( ! [X70,X68,X69] : (mult(rd(mult(rd(X68,X69),rd(X68,X69)),X70),X69) = rd(mult(mult(rd(X68,X69),X69),X68),mult(X70,X69))) )),
  inference(superposition,[],[f482,f24])).
fof(f24,plain,(
  ( ! [X4,X5,X3] : (mult(mult(rd(X3,X4),rd(X3,X4)),mult(X5,X4)) = mult(mult(rd(X3,X4),X5),X3)) )),
  inference(superposition,[],[f6,f3])).
fof(f35730,plain,(
  ( ! [X4,X5,X3] : (mult(rd(X3,X5),rd(X3,X4)) = mult(rd(mult(rd(X3,X4),rd(X3,X4)),X5),X4)) )),
  inference(superposition,[],[f6704,f9])).
fof(f9,plain,(
  ( ! [X0,X1] : (ld(rd(X0,X1),X0) = X1) )),
  inference(superposition,[],[f2,f3])).
fof(f6704,plain,(
  ( ! [X14,X12,X13] : (mult(rd(mult(X12,X12),X14),ld(X12,X13)) = mult(rd(X13,X14),X12)) )),
  inference(forward_demodulation,[],[f6626,f526])).
fof(f6626,plain,(
  ( ! [X14,X12,X13] : (mult(rd(mult(X12,X12),X14),ld(X12,X13)) = rd(mult(X13,mult(X12,X14)),mult(X14,X14))) )),
  inference(superposition,[],[f526,f20])).
fof(f94391,plain,(
  ( ! [X17,X15,X16] : (mult(X15,X17) = mult(rd(mult(X15,X15),mult(X16,ld(rd(X17,ld(X16,X15)),X16))),X16)) )),
  inference(backward_demodulation,[],[f94367,f29086])).
fof(f29086,plain,(
  ( ! [X17,X15,X16] : (mult(X15,X17) = mult(rd(mult(X15,X15),rd(mult(X15,ld(X17,mult(X16,X16))),X16)),X16)) )),
  inference(superposition,[],[f4021,f122])).
fof(f122,plain,(
  ( ! [X17,X15,X16] : (mult(X15,X16) = rd(mult(mult(X15,X15),X17),mult(X15,ld(X16,X17)))) )),
  inference(superposition,[],[f4,f16])).
fof(f4021,plain,(
  ( ! [X39,X38,X40] : (mult(rd(X38,rd(X39,X40)),X40) = rd(mult(X38,mult(X40,X40)),X39)) )),
  inference(superposition,[],[f4,f459])).
fof(f459,plain,(
  ( ! [X4,X5,X3] : (mult(X5,mult(X4,X4)) = mult(mult(rd(X5,rd(X3,X4)),X4),X3)) )),
  inference(superposition,[],[f32,f3])).
fof(f94367,plain,(
  ( ! [X17,X15,X16] : (rd(mult(X17,ld(X15,mult(X16,X16))),X16) = mult(X16,ld(rd(X15,ld(X16,X17)),X16))) )),
  inference(superposition,[],[f91533,f84528])).
fof(f84528,plain,(
  ( ! [X658,X657,X659] : (ld(rd(X657,X658),X659) = rd(ld(X657,mult(X659,X659)),ld(X658,X659))) )),
  inference(superposition,[],[f26658,f3])).
fof(f26658,plain,(
  ( ! [X366,X368,X367] : (ld(X366,X367) = rd(ld(mult(X366,X368),mult(X367,X367)),ld(X368,X367))) )),
  inference(superposition,[],[f4,f3462])).
fof(f3462,plain,(
  ( ! [X35,X36,X34] : (mult(ld(X35,X36),ld(X34,X36)) = ld(mult(X35,X34),mult(X36,X36))) )),
  inference(forward_demodulation,[],[f3461,f1157])).
fof(f3461,plain,(
  ( ! [X35,X36,X34] : (mult(X35,ld(X34,ld(mult(X35,X35),mult(X36,X36)))) = mult(ld(X35,X36),ld(X34,X36))) )),
  inference(forward_demodulation,[],[f3377,f193])).
fof(f3377,plain,(
  ( ! [X35,X36,X34] : (mult(X35,ld(X34,ld(mult(X35,X35),mult(X36,X36)))) = ld(mult(X34,X34),mult(mult(X34,ld(X35,X36)),X36))) )),
  inference(superposition,[],[f193,f2018])).
fof(f2018,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X2,ld(X0,X1)),X1) = mult(mult(X2,X0),ld(mult(X0,X0),mult(X1,X1)))) )),
  inference(backward_demodulation,[],[f1947,f45])).
fof(f45,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X2,X0),mult(ld(X0,X1),ld(X0,X1))) = mult(mult(X2,ld(X0,X1)),X1)) )),
  inference(superposition,[],[f7,f1])).
fof(f1947,plain,(
  ( ! [X2,X0,X1] : (mult(ld(X0,X2),ld(X0,X1)) = ld(mult(X0,X0),mult(X2,X1))) )),
  inference(superposition,[],[f170,f1])).
fof(f91533,plain,(
  ( ! [X10,X11,X9] : (rd(mult(X10,X11),X9) = mult(X9,rd(X11,ld(ld(X9,X10),X9)))) )),
  inference(forward_demodulation,[],[f91478,f54044])).
fof(f54044,plain,(
  ( ! [X2,X0,X1] : (rd(mult(X0,X1),ld(X2,mult(X0,X0))) = rd(mult(X2,X1),X0)) )),
  inference(superposition,[],[f14493,f2])).
fof(f14493,plain,(
  ( ! [X17,X15,X16] : (rd(X15,ld(X16,mult(X17,X17))) = rd(mult(X16,ld(X17,X15)),X17)) )),
  inference(superposition,[],[f4,f1481])).
fof(f1481,plain,(
  ( ! [X14,X12,X13] : (mult(X12,ld(X14,X13)) = mult(rd(X13,ld(X12,mult(X14,X14))),X14)) )),
  inference(forward_demodulation,[],[f1430,f121])).
fof(f1430,plain,(
  ( ! [X14,X12,X13] : (mult(rd(X13,ld(X12,mult(X14,X14))),X14) = ld(mult(X12,X14),mult(mult(X12,X12),X13))) )),
  inference(superposition,[],[f55,f136])).
fof(f136,plain,(
  ( ! [X2,X0,X1] : (mult(mult(X0,X0),X2) = mult(mult(X0,rd(X2,ld(X0,X1))),X1)) )),
  inference(superposition,[],[f17,f1])).
fof(f17,plain,(
  ( ! [X4,X5,X3] : (mult(mult(X5,rd(X3,X4)),mult(X5,X4)) = mult(mult(X5,X5),X3)) )),
  inference(superposition,[],[f6,f3])).
fof(f55,plain,(
  ( ! [X6,X4,X5] : (mult(X6,X5) = ld(mult(X4,X5),mult(mult(X4,X6),mult(X5,X5)))) )),
  inference(superposition,[],[f2,f7])).
fof(f91478,plain,(
  ( ! [X10,X11,X9] : (mult(X9,rd(X11,ld(ld(X9,X10),X9))) = rd(mult(X9,X11),ld(X10,mult(X9,X9)))) )),
  inference(superposition,[],[f89152,f85409])).
fof(f85409,plain,(
  ( ! [X2,X0,X1] : (ld(X2,mult(X0,X0)) = mult(ld(X1,X0),ld(ld(X1,X2),X0))) )),
  inference(superposition,[],[f38450,f10])).
fof(f38450,plain,(
  ( ! [X14,X15,X16] : (ld(X16,mult(X14,X14)) = mult(X15,ld(ld(rd(X14,X15),X16),X14))) )),
  inference(superposition,[],[f6707,f1])).
fof(f6707,plain,(
  ( ! [X14,X12,X13] : (mult(X14,ld(X12,X13)) = ld(mult(rd(X13,X14),X12),mult(X13,X13))) )),
  inference(backward_demodulation,[],[f6704,f4658])).
fof(f4658,plain,(
  ( ! [X14,X12,X13] : (mult(X14,ld(X12,X13)) = ld(mult(rd(mult(X12,X12),X14),ld(X12,X13)),mult(X13,X13))) )),
  inference(forward_demodulation,[],[f4616,f1])).
fof(f4616,plain,(
  ( ! [X14,X12,X13] : (mult(X14,ld(X12,X13)) = ld(mult(rd(mult(X12,X12),X14),ld(X12,X13)),mult(X13,mult(X12,ld(X12,X13))))) )),
  inference(superposition,[],[f481,f20])).
fof(f481,plain,(
  ( ! [X17,X18,X16] : (mult(X17,X18) = ld(mult(rd(X16,X17),X18),mult(X16,mult(X18,X18)))) )),
  inference(superposition,[],[f2,f32])).
fof(f89152,plain,(
  ( ! [X408,X407,X409] : (rd(mult(X408,X407),mult(ld(X408,X408),X409)) = mult(X408,rd(X407,X409))) )),
  inference(forward_demodulation,[],[f89105,f84419])).
fof(f84419,plain,(
  ( ! [X90,X91,X89] : (mult(X89,rd(X90,X91)) = rd(mult(X90,X90),mult(ld(X89,X90),X91))) )),
  inference(superposition,[],[f10,f26607])).
fof(f26607,plain,(
  ( ! [X4,X5,X3] : (ld(mult(X5,rd(X3,X4)),mult(X3,X3)) = mult(ld(X5,X3),X4)) )),
  inference(superposition,[],[f3462,f9])).
fof(f89105,plain,(
  ( ! [X408,X407,X409] : (rd(mult(X407,X407),mult(ld(X408,X407),X409)) = rd(mult(X408,X407),mult(ld(X408,X408),X409))) )),
  inference(superposition,[],[f84792,f85992])).
fof(f84792,plain,(
  ( ! [X14,X12,X13] : (rd(mult(X12,X12),X14) = rd(mult(rd(mult(X12,X13),X14),X12),X13)) )),
  inference(forward_demodulation,[],[f84791,f4])).
fof(f84791,plain,(
  ( ! [X14,X12,X13] : (rd(mult(X12,X12),X14) = rd(mult(rd(mult(X12,X13),X14),rd(mult(X12,X14),X14)),X13)) )),
  inference(forward_demodulation,[],[f84759,f6589])).
fof(f84759,plain,(
  ( ! [X14,X12,X13] : (rd(mult(X12,X12),X14) = rd(rd(mult(mult(X12,X13),mult(X12,X14)),mult(X14,X14)),X13)) )),
  inference(superposition,[],[f30345,f27])).
fof(f27,plain,(
  ( ! [X2,X0,X1] : (mult(X0,X0) = rd(mult(mult(X0,X1),mult(X0,X2)),mult(X1,X2))) )),
  inference(superposition,[],[f4,f6])).
fof(f98037,plain,(
  ( ! [X134,X136,X135,X133] : (rd(X135,ld(rd(X133,ld(X134,X135)),X136)) = ld(X135,rd(mult(mult(X135,X134),X133),X136))) )),
  inference(forward_demodulation,[],[f97990,f86348])).
fof(f86348,plain,(
  ( ! [X28,X26,X27,X25] : (ld(X25,rd(mult(mult(X25,X26),X27),X28)) = rd(mult(X26,ld(X25,X27)),ld(X25,X28))) )),
  inference(superposition,[],[f60680,f193])).
fof(f60680,plain,(
  ( ! [X549,X547,X548] : (ld(X549,rd(X547,X548)) = rd(ld(mult(X549,X549),X547),ld(X549,X548))) )),
  inference(superposition,[],[f17306,f3])).
fof(f17306,plain,(
  ( ! [X325,X327,X326] : (ld(X325,X326) = rd(ld(mult(X325,X325),mult(X326,X327)),ld(X325,X327))) )),
  inference(superposition,[],[f4,f1947])).
fof(f97990,plain,(
  ( ! [X134,X136,X135,X133] : (rd(X135,ld(rd(X133,ld(X134,X135)),X136)) = rd(mult(X134,ld(X135,X133)),ld(X135,X136))) )),
  inference(superposition,[],[f91344,f93359])).
fof(f93359,plain,(
  ( ! [X33,X34,X32] : (mult(rd(X33,ld(X32,X34)),ld(X34,X34)) = mult(X32,ld(X34,X33))) )),
  inference(superposition,[],[f1,f91298])).
fof(f91298,plain,(
  ( ! [X2,X0,X1] : (ld(X0,X2) = ld(X1,mult(rd(X2,ld(X1,X0)),ld(X0,X0)))) )),
  inference(superposition,[],[f89134,f10])).
fof(f89134,plain,(
  ( ! [X14,X12,X13] : (ld(X13,X12) = ld(rd(X13,X14),mult(rd(X12,X14),ld(X13,X13)))) )),
  inference(forward_demodulation,[],[f89133,f4])).
fof(f89133,plain,(
  ( ! [X14,X12,X13] : (ld(X13,X12) = ld(rd(X13,X14),mult(rd(X12,X14),rd(mult(ld(X13,X13),X14),X14)))) )),
  inference(forward_demodulation,[],[f89095,f6589])).
fof(f89095,plain,(
  ( ! [X14,X12,X13] : (ld(X13,X12) = ld(rd(X13,X14),rd(mult(X12,mult(ld(X13,X13),X14)),mult(X14,X14)))) )),
  inference(superposition,[],[f34429,f85992])).
fof(f34429,plain,(
  ( ! [X30,X28,X29] : (ld(rd(rd(X28,mult(X29,X30)),X30),rd(X28,mult(X30,X30))) = X29) )),
  inference(superposition,[],[f2,f6622])).
fof(f6622,plain,(
  ( ! [X2,X0,X1] : (rd(X0,mult(X2,X2)) = mult(rd(rd(X0,mult(X1,X2)),X2),X1)) )),
  inference(superposition,[],[f526,f3])).
fof(f98404,plain,(
  ( ! [X121,X120,X119] : (rd(mult(mult(X119,X120),X121),X120) = mult(X119,rd(X121,ld(X120,X120)))) )),
  inference(forward_demodulation,[],[f98354,f93916])).
fof(f93916,plain,(
  ( ! [X109,X107,X110,X108] : (mult(X107,rd(X110,ld(X108,X109))) = rd(mult(X107,X110),ld(mult(X107,X108),mult(X107,X109)))) )),
  inference(superposition,[],[f89152,f91399])).
fof(f91399,plain,(
  ( ! [X2,X0,X1] : (mult(ld(X2,X2),ld(X1,X0)) = ld(mult(X2,X1),mult(X2,X0))) )),
  inference(superposition,[],[f89151,f10])).
fof(f89151,plain,(
  ( ! [X405,X404,X406] : (mult(ld(X405,X405),X406) = ld(mult(X405,rd(X404,X406)),mult(X405,X404))) )),
  inference(forward_demodulation,[],[f89104,f84419])).
fof(f89104,plain,(
  ( ! [X405,X404,X406] : (mult(ld(X405,X405),X406) = ld(rd(mult(X404,X404),mult(ld(X405,X404),X406)),mult(X405,X404))) )),
  inference(superposition,[],[f85053,f85992])).
fof(f85053,plain,(
  ( ! [X14,X12,X13] : (ld(rd(mult(X12,X12),X14),mult(rd(mult(X12,X13),X14),X12)) = X13) )),
  inference(forward_demodulation,[],[f85052,f4])).
fof(f85052,plain,(
  ( ! [X14,X12,X13] : (ld(rd(mult(X12,X12),X14),mult(rd(mult(X12,X13),X14),rd(mult(X12,X14),X14))) = X13) )),
  inference(forward_demodulation,[],[f85019,f6589])).
fof(f85019,plain,(
  ( ! [X14,X12,X13] : (ld(rd(mult(X12,X12),X14),rd(mult(mult(X12,X13),mult(X12,X14)),mult(X14,X14))) = X13) )),
  inference(superposition,[],[f34429,f27])).
fof(f98354,plain,(
  ( ! [X121,X120,X119] : (rd(mult(mult(X119,X120),X121),X120) = rd(mult(X119,X121),ld(mult(X119,X120),mult(X119,X120)))) )),
  inference(superposition,[],[f93517,f4])).
fof(f93517,plain,(
  ( ! [X144,X145,X143] : (rd(mult(rd(X143,X144),X145),ld(X143,X143)) = rd(mult(X143,X145),X144)) )),
  inference(forward_demodulation,[],[f93475,f87035])).
fof(f87035,plain,(
  ( ! [X121,X122,X120] : (rd(mult(X120,X120),ld(X122,mult(X120,X121))) = rd(mult(X122,X120),X121)) )),
  inference(superposition,[],[f84792,f10])).
fof(f93475,plain,(
  ( ! [X144,X145,X143] : (rd(mult(rd(X143,X144),X145),ld(X143,X143)) = rd(mult(X145,X145),ld(X143,mult(X145,X144)))) )),
  inference(superposition,[],[f87035,f91301])).
fof(f91301,plain,(
  ( ! [X72,X71,X73] : (ld(X73,mult(X71,X72)) = ld(rd(X73,X72),mult(X71,ld(X73,X73)))) )),
  inference(superposition,[],[f89134,f4])).
fof(f8,axiom,(
  mult(mult(a,b),c) != mult(mult(a,c),mult(b,ld(c,c)))),
  file('./p/GRP752-1.tptp',unknown)).
% SZS output end Proof for GRP752-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 357051
% Time elapsed: 35.592 s
% ------------------------------
% ------------------------------
