% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for HWV068-1
% SZS output start Proof for HWV068-1
fof(f1056045,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f50711,f53894,f61141,f70096,f71532,f73199,f120144,f1056044])).
fof(f1056044,plain,(
  spl0_7114 | spl0_7114 | spl0_7114 | spl0_7114 | ~spl0_199 | ~spl0_2783 | ~spl0_4943 | ~spl0_5282),
  inference(avatar_split_clause,[],[f1056043,f71530,f70094,f61139,f50709,f115267,f115267,f115267,f115267])).
fof(f115267,plain,(
  spl0_7114 <=> ! [X5] : ~p(X5)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7114])])).
fof(f50709,plain,(
  spl0_199 <=> ! [X1,X3,X0,X2] : (~p(X0) | ~p(X1) | ~p(X3) | ~p(X2) | esk16687_4(X2,X0,X3,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_199])])).
fof(f61139,plain,(
  spl0_2783 <=> ! [X1,X3,X0,X2] : (~p(X0) | ~p(X3) | ~p(X1) | ~p(X2) | esk16679_4(X2,X0,X3,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2783])])).
fof(f70094,plain,(
  spl0_4943 <=> ! [X1,X3,X0,X2] : (~p(X1) | ~p(X3) | ~p(X2) | ~p(X0) | ~esk16644_4(X2,X0,X3,X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4943])])).
fof(f71530,plain,(
  spl0_5282 <=> ! [X1,X3,X0,X2] : (~esk16753_4(X2,X0,X3,X1) | ~p(X0) | ~p(X1) | ~p(X2) | ~p(X3))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5282])])).
fof(f1056043,plain,(
  ( ! [X6,X4,X7,X5] : (~p(X4) | ~p(X5) | ~p(X6) | ~p(X7)) ) | (~spl0_199 | ~spl0_2783 | ~spl0_4943 | ~spl0_5282)),
  inference(subsumption_resolution,[],[f1056042,f70095])).
fof(f70095,plain,(
  ( ! [X2,X0,X3,X1] : (~esk16644_4(X2,X0,X3,X1) | ~p(X3) | ~p(X2) | ~p(X0) | ~p(X1)) ) | ~spl0_4943),
  inference(avatar_component_clause,[],[f70094])).
fof(f1056042,plain,(
  ( ! [X6,X4,X7,X5] : (~p(X4) | ~p(X5) | ~p(X6) | ~p(X7) | esk16644_4(X6,X4,X7,X5)) ) | (~spl0_199 | ~spl0_2783 | ~spl0_5282)),
  inference(resolution,[],[f1052851,f761])).
fof(f761,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk17351_4(X2,X0,X3,X1) | esk16644_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f1052851,plain,(
  ( ! [X26,X24,X27,X25] : (esk17351_4(X27,X25,X24,X26) | ~p(X25) | ~p(X26) | ~p(X27) | ~p(X24)) ) | (~spl0_199 | ~spl0_2783 | ~spl0_5282)),
  inference(subsumption_resolution,[],[f1052850,f1052839])).
fof(f1052839,plain,(
  ( ! [X59,X57,X58,X56] : (~p(X56) | ~p(X57) | ~p(X58) | ~p(X59) | esk17351_4(X59,X57,X56,X58) | esk16786_4(X59,X57,X56,X58)) ) | (~spl0_199 | ~spl0_2783 | ~spl0_5282)),
  inference(duplicate_literal_removal,[],[f1052824])).
fof(f1052824,plain,(
  ( ! [X59,X57,X58,X56] : (~p(X56) | ~p(X57) | ~p(X58) | ~p(X59) | esk17351_4(X59,X57,X56,X58) | ~p(X58) | ~p(X59) | esk16786_4(X59,X57,X56,X58) | ~p(X57) | ~p(X56)) ) | (~spl0_199 | ~spl0_2783 | ~spl0_5282)),
  inference(resolution,[],[f1052741,f516129])).
fof(f516129,plain,(
  ( ! [X6,X4,X7,X5] : (~esk16571_4(X6,X7,X4,X5) | ~p(X5) | ~p(X6) | esk16786_4(X6,X7,X4,X5) | ~p(X7) | ~p(X4)) ) | ~spl0_2783),
  inference(resolution,[],[f515819,f25321])).
fof(f25321,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk17260_4(X2,X0,X3,X1) | ~esk16571_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f515819,plain,(
  ( ! [X6,X4,X7,X5] : (esk17260_4(X7,X4,X5,X6) | ~p(X5) | ~p(X6) | ~p(X7) | esk16786_4(X7,X4,X5,X6) | ~p(X4)) ) | ~spl0_2783),
  inference(resolution,[],[f515375,f31506])).
fof(f31506,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk17261_4(X2,X0,X3,X1) | esk17260_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f515375,plain,(
  ( ! [X2,X0,X3,X1] : (esk17261_4(X0,X1,X2,X3) | ~p(X1) | ~p(X2) | ~p(X3) | ~p(X0) | esk16786_4(X0,X1,X2,X3)) ) | ~spl0_2783),
  inference(resolution,[],[f515242,f33465])).
fof(f33465,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk17262_4(X2,X0,X3,X1) | esk16786_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f515242,plain,(
  ( ! [X2,X0,X3,X1] : (esk17262_4(X1,X2,X3,X0) | ~p(X1) | ~p(X2) | ~p(X3) | ~p(X0) | esk17261_4(X1,X2,X3,X0)) ) | ~spl0_2783),
  inference(resolution,[],[f514845,f23173])).
fof(f23173,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk17263_4(X2,X0,X3,X1) | esk17262_4(X2,X0,X3,X1) | esk17261_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f514845,plain,(
  ( ! [X2,X0,X3,X1] : (esk17263_4(X2,X3,X0,X1) | ~p(X1) | ~p(X2) | ~p(X3) | ~p(X0)) ) | ~spl0_2783),
  inference(resolution,[],[f61140,f36992])).
fof(f36992,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk16679_4(X2,X0,X3,X1) | esk17263_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f61140,plain,(
  ( ! [X2,X0,X3,X1] : (esk16679_4(X2,X0,X3,X1) | ~p(X3) | ~p(X1) | ~p(X2) | ~p(X0)) ) | ~spl0_2783),
  inference(avatar_component_clause,[],[f61139])).
fof(f1052741,plain,(
  ( ! [X2,X0,X3,X1] : (esk16571_4(X0,X2,X1,X3) | ~p(X1) | ~p(X2) | ~p(X3) | ~p(X0) | esk17351_4(X0,X2,X1,X3)) ) | (~spl0_199 | ~spl0_5282)),
  inference(resolution,[],[f1052453,f38962])).
fof(f38962,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk17349_4(X2,X0,X3,X1) | esk17351_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f1052453,plain,(
  ( ! [X2,X0,X3,X1] : (esk17349_4(X1,X3,X2,X0) | ~p(X1) | ~p(X2) | ~p(X3) | ~p(X0) | esk16571_4(X1,X3,X2,X0)) ) | (~spl0_199 | ~spl0_5282)),
  inference(resolution,[],[f1052401,f120757])).
fof(f120757,plain,(
  ( ! [X2,X0,X3,X1] : (~esk17197_4(X0,X1,X2,X3) | esk17349_4(X0,X1,X2,X3) | esk16571_4(X0,X1,X2,X3)) )),
  inference(resolution,[],[f11720,f27801])).
fof(f27801,axiom,(
  ( ! [X2,X0,X3,X1] : (esk17348_4(X2,X0,X3,X1) | esk16571_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f11720,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk17348_4(X2,X0,X3,X1) | esk17349_4(X2,X0,X3,X1) | ~esk17197_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f1052401,plain,(
  ( ! [X10,X8,X11,X9] : (esk17197_4(X10,X8,X11,X9) | ~p(X9) | ~p(X10) | ~p(X11) | ~p(X8)) ) | (~spl0_199 | ~spl0_5282)),
  inference(resolution,[],[f1052389,f2121])).
fof(f2121,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk17195_4(X2,X0,X3,X1) | esk17197_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f1052389,plain,(
  ( ! [X6,X4,X7,X5] : (esk17195_4(X4,X5,X7,X6) | ~p(X5) | ~p(X6) | ~p(X4) | ~p(X7)) ) | (~spl0_199 | ~spl0_5282)),
  inference(subsumption_resolution,[],[f1052357,f71531])).
fof(f71531,plain,(
  ( ! [X2,X0,X3,X1] : (~esk16753_4(X2,X0,X3,X1) | ~p(X0) | ~p(X1) | ~p(X2) | ~p(X3)) ) | ~spl0_5282),
  inference(avatar_component_clause,[],[f71530])).
fof(f1052357,plain,(
  ( ! [X6,X4,X7,X5] : (~p(X4) | ~p(X5) | ~p(X6) | esk17195_4(X4,X5,X7,X6) | ~p(X7) | esk16753_4(X4,X5,X7,X6)) ) | ~spl0_199),
  inference(resolution,[],[f543027,f20417])).
fof(f20417,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk16372_4(X2,X0,X3,X1) | esk16753_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f543027,plain,(
  ( ! [X23,X21,X22,X20] : (esk16372_4(X21,X22,X20,X23) | ~p(X21) | ~p(X22) | ~p(X23) | esk17195_4(X21,X22,X20,X23) | ~p(X20)) ) | ~spl0_199),
  inference(resolution,[],[f542922,f117967])).
fof(f117967,plain,(
  ( ! [X2,X0,X3,X1] : (~esk16645_4(X0,X1,X2,X3) | esk17195_4(X0,X1,X2,X3) | esk16372_4(X0,X1,X2,X3)) )),
  inference(resolution,[],[f8378,f2681])).
fof(f2681,axiom,(
  ( ! [X2,X0,X3,X1] : (esk17194_4(X2,X0,X3,X1) | esk16372_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f8378,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk17194_4(X2,X0,X3,X1) | ~esk16645_4(X2,X0,X3,X1) | esk17195_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f542922,plain,(
  ( ! [X2,X0,X3,X1] : (esk16645_4(X2,X3,X1,X0) | ~p(X1) | ~p(X2) | ~p(X3) | ~p(X0)) ) | ~spl0_199),
  inference(resolution,[],[f50710,f43887])).
fof(f43887,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk16687_4(X2,X0,X3,X1) | esk16645_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f50710,plain,(
  ( ! [X2,X0,X3,X1] : (esk16687_4(X2,X0,X3,X1) | ~p(X1) | ~p(X3) | ~p(X2) | ~p(X0)) ) | ~spl0_199),
  inference(avatar_component_clause,[],[f50709])).
fof(f1052850,plain,(
  ( ! [X26,X24,X27,X25] : (~p(X24) | ~p(X25) | ~p(X26) | ~p(X27) | esk17351_4(X27,X25,X24,X26) | ~esk16786_4(X27,X25,X24,X26)) ) | (~spl0_199 | ~spl0_5282)),
  inference(subsumption_resolution,[],[f1052816,f2795])).
fof(f2795,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk17350_4(X2,X0,X3,X1) | esk17351_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f1052816,plain,(
  ( ! [X26,X24,X27,X25] : (~p(X24) | ~p(X25) | ~p(X26) | ~p(X27) | esk17351_4(X27,X25,X24,X26) | esk17350_4(X27,X25,X24,X26) | ~esk16786_4(X27,X25,X24,X26)) ) | (~spl0_199 | ~spl0_5282)),
  inference(resolution,[],[f1052741,f48003])).
fof(f48003,axiom,(
  ( ! [X2,X0,X3,X1] : (~esk16571_4(X2,X0,X3,X1) | esk17350_4(X2,X0,X3,X1) | ~esk16786_4(X2,X0,X3,X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f120144,plain,(
  ~spl0_7114),
  inference(avatar_contradiction_clause,[],[f120143])).
fof(f120143,plain,(
  $false | ~spl0_7114),
  inference(resolution,[],[f115268,f49911])).
fof(f49911,axiom,(
  p(true)),
  file('./p/HWV068-1.tptp',unknown)).
fof(f115268,plain,(
  ( ! [X5] : (~p(X5)) ) | ~spl0_7114),
  inference(avatar_component_clause,[],[f115267])).
fof(f73199,plain,(
  ~spl0_13),
  inference(avatar_split_clause,[],[f37238,f49963])).
fof(f49963,plain,(
  spl0_13 <=> esk1_0),
  introduced(avatar_definition,[new_symbols(naming,[spl0_13])])).
fof(f37238,axiom,(
  ~esk1_0),
  file('./p/HWV068-1.tptp',unknown)).
fof(f71532,plain,(
  spl0_5282 | spl0_13),
  inference(avatar_split_clause,[],[f34002,f49963,f71530])).
fof(f34002,axiom,(
  ( ! [X2,X0,X3,X1] : (esk1_0 | ~esk16753_4(X2,X0,X3,X1) | ~p(X3) | ~p(X2) | ~p(X1) | ~p(X0)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f70096,plain,(
  spl0_13 | spl0_4943),
  inference(avatar_split_clause,[],[f31426,f70094,f49963])).
fof(f31426,axiom,(
  ( ! [X2,X0,X3,X1] : (~p(X1) | ~esk16644_4(X2,X0,X3,X1) | ~p(X0) | ~p(X2) | ~p(X3) | esk1_0) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f61141,plain,(
  ~spl0_31 | spl0_2783),
  inference(avatar_split_clause,[],[f15112,f61139,f50035])).
fof(f50035,plain,(
  spl0_31 <=> esk2_0),
  introduced(avatar_definition,[new_symbols(naming,[spl0_31])])).
fof(f15112,axiom,(
  ( ! [X2,X0,X3,X1] : (~p(X0) | esk16679_4(X2,X0,X3,X1) | ~p(X2) | ~p(X1) | ~esk2_0 | ~p(X3)) )),
  file('./p/HWV068-1.tptp',unknown)).
fof(f53894,plain,(
  spl0_31),
  inference(avatar_split_clause,[],[f4556,f50035])).
fof(f4556,axiom,(
  esk2_0),
  file('./p/HWV068-1.tptp',unknown)).
fof(f50711,plain,(
  ~spl0_31 | spl0_199),
  inference(avatar_split_clause,[],[f814,f50709,f50035])).
fof(f814,axiom,(
  ( ! [X2,X0,X3,X1] : (~p(X0) | esk16687_4(X2,X0,X3,X1) | ~p(X2) | ~p(X3) | ~esk2_0 | ~p(X1)) )),
  file('./p/HWV068-1.tptp',unknown)).
% SZS output end Proof for HWV068-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 777812
% Time elapsed: 44.184 s
% ------------------------------
% ------------------------------
