% Refutation found. Thanks to Tanya!
% SZS status Theorem for KLE150+2
% SZS output start Proof for KLE150+2
fof(f7114,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f58,f7097,f7111])).
fof(f7111,plain,(
  spl1_2),
  inference(avatar_contradiction_clause,[],[f7110])).
fof(f7110,plain,(
  $false | spl1_2),
  inference(subsumption_resolution,[],[f7109,f73])).
fof(f73,plain,(
  ( ! [X1] : (leq(X1,X1)) )),
  inference(trivial_inequality_removal,[],[f68])).
fof(f68,plain,(
  ( ! [X1] : (X1 != X1 | leq(X1,X1)) )),
  inference(superposition,[],[f48,f33])).
fof(f33,plain,(
  ( ! [X0] : (addition(X0,X0) = X0) )),
  inference(cnf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0] : addition(X0,X0) = X0),
  file('./p/KLE150+2.tptp',unknown)).
fof(f48,plain,(
  ( ! [X0,X1] : (addition(X0,X1) != X1 | leq(X0,X1)) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : ((leq(X0,X1) | addition(X0,X1) != X1) & (addition(X0,X1) = X1 | ~leq(X0,X1)))),
  inference(nnf_transformation,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1] : (leq(X0,X1) <=> addition(X0,X1) = X1)),
  file('./p/KLE150+2.tptp',unknown)).
fof(f7109,plain,(
  ~leq(star(multiplication(sK0,zero)),star(multiplication(sK0,zero))) | spl1_2),
  inference(forward_demodulation,[],[f7108,f2355])).
fof(f2355,plain,(
  ( ! [X4] : (star(multiplication(X4,zero)) = addition(one,multiplication(X4,zero))) )),
  inference(superposition,[],[f134,f39])).
fof(f39,plain,(
  ( ! [X0] : (zero = multiplication(zero,X0)) )),
  inference(cnf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0] : zero = multiplication(zero,X0)),
  file('./p/KLE150+2.tptp',unknown)).
fof(f134,plain,(
  ( ! [X2,X3] : (star(multiplication(X2,X3)) = addition(one,multiplication(X2,multiplication(X3,star(multiplication(X2,X3)))))) )),
  inference(superposition,[],[f40,f34])).
fof(f34,plain,(
  ( ! [X2,X0,X1] : (multiplication(X0,multiplication(X1,X2)) = multiplication(multiplication(X0,X1),X2)) )),
  inference(cnf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2] : multiplication(X0,multiplication(X1,X2)) = multiplication(multiplication(X0,X1),X2)),
  file('./p/KLE150+2.tptp',unknown)).
fof(f40,plain,(
  ( ! [X0] : (star(X0) = addition(one,multiplication(X0,star(X0)))) )),
  inference(cnf_transformation,[],[f11])).
fof(f11,axiom,(
  ! [X0] : star(X0) = addition(one,multiplication(X0,star(X0)))),
  file('./p/KLE150+2.tptp',unknown)).
fof(f7108,plain,(
  ~leq(addition(one,multiplication(sK0,zero)),star(multiplication(sK0,zero))) | spl1_2),
  inference(forward_demodulation,[],[f57,f7092])).
fof(f7092,plain,(
  ( ! [X4] : (star(multiplication(X4,zero)) = strong_iteration(multiplication(X4,zero))) )),
  inference(backward_demodulation,[],[f7091,f2426])).
fof(f2426,plain,(
  ( ! [X4] : (addition(multiplication(X4,zero),one) = strong_iteration(multiplication(X4,zero))) )),
  inference(superposition,[],[f135,f39])).
fof(f135,plain,(
  ( ! [X4,X5] : (strong_iteration(multiplication(X4,X5)) = addition(multiplication(X4,multiplication(X5,strong_iteration(multiplication(X4,X5)))),one)) )),
  inference(superposition,[],[f44,f34])).
fof(f44,plain,(
  ( ! [X0] : (strong_iteration(X0) = addition(multiplication(X0,strong_iteration(X0)),one)) )),
  inference(cnf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0] : strong_iteration(X0) = addition(multiplication(X0,strong_iteration(X0)),one)),
  file('./p/KLE150+2.tptp',unknown)).
fof(f7091,plain,(
  ( ! [X51] : (star(multiplication(X51,zero)) = addition(multiplication(X51,zero),one)) )),
  inference(forward_demodulation,[],[f7048,f320])).
fof(f320,plain,(
  ( ! [X19] : (star(X19) = addition(one,star(X19))) )),
  inference(superposition,[],[f105,f40])).
fof(f105,plain,(
  ( ! [X2,X3] : (addition(X2,X3) = addition(X2,addition(X2,X3))) )),
  inference(superposition,[],[f31,f33])).
fof(f31,plain,(
  ( ! [X2,X0,X1] : (addition(X2,addition(X1,X0)) = addition(addition(X2,X1),X0)) )),
  inference(cnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2] : addition(X2,addition(X1,X0)) = addition(addition(X2,X1),X0)),
  inference(rectify,[],[f2])).
fof(f2,axiom,(
  ! [X2,X1,X0] : addition(X0,addition(X1,X2)) = addition(addition(X0,X1),X2)),
  file('./p/KLE150+2.tptp',unknown)).
fof(f7048,plain,(
  ( ! [X51] : (addition(multiplication(X51,zero),one) = addition(one,star(multiplication(X51,zero)))) )),
  inference(superposition,[],[f488,f2355])).
fof(f488,plain,(
  ( ! [X0,X1] : (addition(X1,X0) = addition(X0,addition(X0,X1))) )),
  inference(superposition,[],[f116,f33])).
fof(f116,plain,(
  ( ! [X8,X7,X9] : (addition(X7,addition(X8,X9)) = addition(X9,addition(X7,X8))) )),
  inference(superposition,[],[f31,f30])).
fof(f30,plain,(
  ( ! [X0,X1] : (addition(X0,X1) = addition(X1,X0)) )),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1] : addition(X0,X1) = addition(X1,X0)),
  file('./p/KLE150+2.tptp',unknown)).
fof(f57,plain,(
  ~leq(addition(one,multiplication(sK0,zero)),strong_iteration(multiplication(sK0,zero))) | spl1_2),
  inference(avatar_component_clause,[],[f55])).
fof(f55,plain,(
  spl1_2 <=> leq(addition(one,multiplication(sK0,zero)),strong_iteration(multiplication(sK0,zero)))),
  introduced(avatar_definition,[new_symbols(naming,[spl1_2])])).
fof(f7097,plain,(
  spl1_1),
  inference(avatar_contradiction_clause,[],[f7096])).
fof(f7096,plain,(
  $false | spl1_1),
  inference(subsumption_resolution,[],[f7094,f73])).
fof(f7094,plain,(
  ~leq(star(multiplication(sK0,zero)),star(multiplication(sK0,zero))) | spl1_1),
  inference(backward_demodulation,[],[f7092,f2387])).
fof(f2387,plain,(
  ~leq(strong_iteration(multiplication(sK0,zero)),star(multiplication(sK0,zero))) | spl1_1),
  inference(backward_demodulation,[],[f2355,f53])).
fof(f53,plain,(
  ~leq(strong_iteration(multiplication(sK0,zero)),addition(one,multiplication(sK0,zero))) | spl1_1),
  inference(avatar_component_clause,[],[f51])).
fof(f51,plain,(
  spl1_1 <=> leq(strong_iteration(multiplication(sK0,zero)),addition(one,multiplication(sK0,zero)))),
  introduced(avatar_definition,[new_symbols(naming,[spl1_1])])).
fof(f58,plain,(
  ~spl1_1 | ~spl1_2),
  inference(avatar_split_clause,[],[f49,f55,f51])).
fof(f49,plain,(
  ~leq(addition(one,multiplication(sK0,zero)),strong_iteration(multiplication(sK0,zero))) | ~leq(strong_iteration(multiplication(sK0,zero)),addition(one,multiplication(sK0,zero)))),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ~leq(addition(one,multiplication(sK0,zero)),strong_iteration(multiplication(sK0,zero))) | ~leq(strong_iteration(multiplication(sK0,zero)),addition(one,multiplication(sK0,zero)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f26,f28])).
fof(f28,plain,(
  ? [X0] : (~leq(addition(one,multiplication(X0,zero)),strong_iteration(multiplication(X0,zero))) | ~leq(strong_iteration(multiplication(X0,zero)),addition(one,multiplication(X0,zero)))) => (~leq(addition(one,multiplication(sK0,zero)),strong_iteration(multiplication(sK0,zero))) | ~leq(strong_iteration(multiplication(sK0,zero)),addition(one,multiplication(sK0,zero))))),
  introduced(choice_axiom,[])).
fof(f26,plain,(
  ? [X0] : (~leq(addition(one,multiplication(X0,zero)),strong_iteration(multiplication(X0,zero))) | ~leq(strong_iteration(multiplication(X0,zero)),addition(one,multiplication(X0,zero))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0] : (leq(addition(one,multiplication(X0,zero)),strong_iteration(multiplication(X0,zero))) & leq(strong_iteration(multiplication(X0,zero)),addition(one,multiplication(X0,zero))))),
  inference(rectify,[],[f20])).
fof(f20,negated_conjecture,(
  ~! [X3] : (leq(addition(one,multiplication(X3,zero)),strong_iteration(multiplication(X3,zero))) & leq(strong_iteration(multiplication(X3,zero)),addition(one,multiplication(X3,zero))))),
  inference(negated_conjecture,[],[f19])).
fof(f19,conjecture,(
  ! [X3] : (leq(addition(one,multiplication(X3,zero)),strong_iteration(multiplication(X3,zero))) & leq(strong_iteration(multiplication(X3,zero)),addition(one,multiplication(X3,zero))))),
  file('./p/KLE150+2.tptp',unknown)).
% SZS output end Proof for KLE150+2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 11257
% Time elapsed: 0.097 s
% ------------------------------
% ------------------------------
