% Refutation found. Thanks to Tanya!
% SZS status Theorem for KRS146+1
% SZS output start Proof for KRS146+1
fof(f536,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f345,f351,f403,f409,f480,f483,f498,f506,f535])).
fof(f535,plain,(
  spl8_2),
  inference(avatar_contradiction_clause,[],[f534])).
fof(f534,plain,(
  $false | spl8_2),
  inference(subsumption_resolution,[],[f533,f292])).
fof(f292,plain,(
  ~cC86(iV822576)),
  inference(cnf_transformation,[],[f78])).
fof(f78,axiom,(
  ~cC86(iV822576)),
  file('./p/KRS146+1.tptp',unknown)).
fof(f533,plain,(
  cC86(iV822576) | spl8_2),
  inference(subsumption_resolution,[],[f529,f293])).
fof(f293,plain,(
  ~cC92(iV822576)),
  inference(cnf_transformation,[],[f79])).
fof(f79,axiom,(
  ~cC92(iV822576)),
  file('./p/KRS146+1.tptp',unknown)).
fof(f529,plain,(
  cC92(iV822576) | cC86(iV822576) | spl8_2),
  inference(resolution,[],[f276,f320])).
fof(f320,plain,(
  ~cC94(iV822576) | spl8_2),
  inference(avatar_component_clause,[],[f318])).
fof(f318,plain,(
  spl8_2 <=> cC94(iV822576)),
  introduced(avatar_definition,[new_symbols(naming,[spl8_2])])).
fof(f276,plain,(
  ( ! [X0] : (cC94(X0) | cC92(X0) | cC86(X0)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0] : ((cC94(X0) | cC92(X0) | cC86(X0)) & ((~cC92(X0) & ~cC86(X0)) | ~cC94(X0)))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0] : ((cC94(X0) | (cC92(X0) | cC86(X0))) & ((~cC92(X0) & ~cC86(X0)) | ~cC94(X0)))),
  inference(nnf_transformation,[],[f65])).
fof(f65,axiom,(
  ! [X0] : (cC94(X0) <=> (~cC92(X0) & ~cC86(X0)))),
  file('./p/KRS146+1.tptp',unknown)).
fof(f506,plain,(
  spl8_7 | ~spl8_6 | ~spl8_8),
  inference(avatar_split_clause,[],[f505,f342,f334,f338])).
fof(f338,plain,(
  spl8_7 <=> cC136(iV822576)),
  introduced(avatar_definition,[new_symbols(naming,[spl8_7])])).
fof(f334,plain,(
  spl8_6 <=> cC114(iV822576)),
  introduced(avatar_definition,[new_symbols(naming,[spl8_6])])).
fof(f342,plain,(
  spl8_8 <=> cC80(iV822576)),
  introduced(avatar_definition,[new_symbols(naming,[spl8_8])])).
fof(f505,plain,(
  cC136(iV822576) | (~spl8_6 | ~spl8_8)),
  inference(subsumption_resolution,[],[f504,f286])).
fof(f286,plain,(
  ~cC134(iV822576)),
  inference(cnf_transformation,[],[f72])).
fof(f72,axiom,(
  ~cC134(iV822576)),
  file('./p/KRS146+1.tptp',unknown)).
fof(f504,plain,(
  cC134(iV822576) | cC136(iV822576) | (~spl8_6 | ~spl8_8)),
  inference(resolution,[],[f500,f245])).
fof(f245,plain,(
  ( ! [X0] : (~cC116(X0) | cC134(X0) | cC136(X0)) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : (cC136(X0) | cC134(X0) | ~cC116(X0))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0] : (cC136(X0) | (cC134(X0) | ~cC116(X0)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : ((~cC134(X0) & cC116(X0)) => cC136(X0))),
  inference(unused_predicate_definition_removal,[],[f22])).
fof(f22,axiom,(
  ! [X0] : (cC136(X0) <=> (~cC134(X0) & cC116(X0)))),
  file('./p/KRS146+1.tptp',unknown)).
fof(f500,plain,(
  cC116(iV822576) | (~spl8_6 | ~spl8_8)),
  inference(subsumption_resolution,[],[f499,f335])).
fof(f335,plain,(
  cC114(iV822576) | ~spl8_6),
  inference(avatar_component_clause,[],[f334])).
fof(f499,plain,(
  cC116(iV822576) | ~cC114(iV822576) | ~spl8_8),
  inference(resolution,[],[f343,f217])).
fof(f217,plain,(
  ( ! [X0] : (~cC80(X0) | cC116(X0) | ~cC114(X0)) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : (cC116(X0) | ~cC80(X0) | ~cC114(X0))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : (cC116(X0) | (~cC80(X0) | ~cC114(X0)))),
  inference(ennf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : ((cC80(X0) & cC114(X0)) => cC116(X0))),
  inference(unused_predicate_definition_removal,[],[f11])).
fof(f11,axiom,(
  ! [X0] : (cC116(X0) <=> (cC80(X0) & cC114(X0)))),
  file('./p/KRS146+1.tptp',unknown)).
fof(f343,plain,(
  cC80(iV822576) | ~spl8_8),
  inference(avatar_component_clause,[],[f342])).
fof(f498,plain,(
  spl8_8),
  inference(avatar_split_clause,[],[f497,f342])).
fof(f497,plain,(
  cC80(iV822576)),
  inference(subsumption_resolution,[],[f494,f301])).
fof(f301,plain,(
  ~cC76(iV822576)),
  inference(cnf_transformation,[],[f87])).
fof(f87,axiom,(
  ~cC76(iV822576)),
  file('./p/KRS146+1.tptp',unknown)).
fof(f494,plain,(
  cC76(iV822576) | cC80(iV822576)),
  inference(resolution,[],[f255,f303])).
fof(f303,plain,(
  ~cC78(iV822576)),
  inference(cnf_transformation,[],[f89])).
fof(f89,axiom,(
  ~cC78(iV822576)),
  file('./p/KRS146+1.tptp',unknown)).
fof(f255,plain,(
  ( ! [X0] : (cC78(X0) | cC76(X0) | cC80(X0)) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : (cC80(X0) | cC76(X0) | cC78(X0))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : (cC80(X0) | (cC76(X0) | cC78(X0)))),
  inference(ennf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : ((~cC76(X0) & ~cC78(X0)) => cC80(X0))),
  inference(unused_predicate_definition_removal,[],[f58])).
fof(f58,axiom,(
  ! [X0] : (cC80(X0) <=> (~cC76(X0) & ~cC78(X0)))),
  file('./p/KRS146+1.tptp',unknown)).
fof(f483,plain,(
  spl8_3 | ~spl8_4),
  inference(avatar_split_clause,[],[f482,f326,f322])).
fof(f322,plain,(
  spl8_3 <=> cC58(iV822576)),
  introduced(avatar_definition,[new_symbols(naming,[spl8_3])])).
fof(f326,plain,(
  spl8_4 <=> cC56(iV822576)),
  introduced(avatar_definition,[new_symbols(naming,[spl8_4])])).
fof(f482,plain,(
  cC58(iV822576) | ~spl8_4),
  inference(subsumption_resolution,[],[f481,f289])).
fof(f289,plain,(
  cC34(iV822576)),
  inference(cnf_transformation,[],[f75])).
fof(f75,axiom,(
  cC34(iV822576)),
  file('./p/KRS146+1.tptp',unknown)).
fof(f481,plain,(
  cC58(iV822576) | ~cC34(iV822576) | ~spl8_4),
  inference(resolution,[],[f248,f327])).
fof(f327,plain,(
  cC56(iV822576) | ~spl8_4),
  inference(avatar_component_clause,[],[f326])).
fof(f248,plain,(
  ( ! [X0] : (~cC56(X0) | cC58(X0) | ~cC34(X0)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : (cC58(X0) | ~cC56(X0) | ~cC34(X0))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : (cC58(X0) | (~cC56(X0) | ~cC34(X0)))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((cC56(X0) & cC34(X0)) => cC58(X0))),
  inference(unused_predicate_definition_removal,[],[f45])).
fof(f45,axiom,(
  ! [X0] : (cC58(X0) <=> (cC56(X0) & cC34(X0)))),
  file('./p/KRS146+1.tptp',unknown)).
fof(f480,plain,(
  spl8_4),
  inference(avatar_split_clause,[],[f479,f326])).
fof(f479,plain,(
  cC56(iV822576)),
  inference(subsumption_resolution,[],[f476,f299])).
fof(f299,plain,(
  ~cC10(iV822576)),
  inference(cnf_transformation,[],[f85])).
fof(f85,axiom,(
  ~cC10(iV822576)),
  file('./p/KRS146+1.tptp',unknown)).
fof(f476,plain,(
  cC10(iV822576) | cC56(iV822576)),
  inference(resolution,[],[f247,f290])).
fof(f290,plain,(
  cC4(iV822576)),
  inference(cnf_transformation,[],[f76])).
fof(f76,axiom,(
  cC4(iV822576)),
  file('./p/KRS146+1.tptp',unknown)).
fof(f247,plain,(
  ( ! [X0] : (~cC4(X0) | cC10(X0) | cC56(X0)) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : (cC56(X0) | cC10(X0) | ~cC4(X0))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : (cC56(X0) | (cC10(X0) | ~cC4(X0)))),
  inference(ennf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : ((~cC10(X0) & cC4(X0)) => cC56(X0))),
  inference(unused_predicate_definition_removal,[],[f44])).
fof(f44,axiom,(
  ! [X0] : (cC56(X0) <=> (~cC10(X0) & cC4(X0)))),
  file('./p/KRS146+1.tptp',unknown)).
fof(f409,plain,(
  spl8_6),
  inference(avatar_split_clause,[],[f408,f334])).
fof(f408,plain,(
  cC114(iV822576)),
  inference(subsumption_resolution,[],[f406,f302])).
fof(f302,plain,(
  ~cC112(iV822576)),
  inference(cnf_transformation,[],[f88])).
fof(f88,axiom,(
  ~cC112(iV822576)),
  file('./p/KRS146+1.tptp',unknown)).
fof(f406,plain,(
  cC112(iV822576) | cC114(iV822576)),
  inference(resolution,[],[f216,f291])).
fof(f291,plain,(
  ~cC96(iV822576)),
  inference(cnf_transformation,[],[f77])).
fof(f77,axiom,(
  ~cC96(iV822576)),
  file('./p/KRS146+1.tptp',unknown)).
fof(f216,plain,(
  ( ! [X0] : (cC96(X0) | cC112(X0) | cC114(X0)) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : (cC114(X0) | cC112(X0) | cC96(X0))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : (cC114(X0) | (cC112(X0) | cC96(X0)))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : ((~cC112(X0) & ~cC96(X0)) => cC114(X0))),
  inference(unused_predicate_definition_removal,[],[f10])).
fof(f10,axiom,(
  ! [X0] : (cC114(X0) <=> (~cC112(X0) & ~cC96(X0)))),
  file('./p/KRS146+1.tptp',unknown)).
fof(f403,plain,(
  spl8_5),
  inference(avatar_split_clause,[],[f402,f330])).
fof(f330,plain,(
  spl8_5 <=> cC110(iV822576)),
  introduced(avatar_definition,[new_symbols(naming,[spl8_5])])).
fof(f402,plain,(
  cC110(iV822576)),
  inference(subsumption_resolution,[],[f399,f302])).
fof(f399,plain,(
  cC112(iV822576) | cC110(iV822576)),
  inference(resolution,[],[f215,f290])).
fof(f215,plain,(
  ( ! [X0] : (~cC4(X0) | cC112(X0) | cC110(X0)) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : ((cC112(X0) | ~cC4(X0) | cC110(X0)) & ((cC4(X0) & ~cC110(X0)) | ~cC112(X0)))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0] : ((cC112(X0) | (~cC4(X0) | cC110(X0))) & ((cC4(X0) & ~cC110(X0)) | ~cC112(X0)))),
  inference(nnf_transformation,[],[f9])).
fof(f9,axiom,(
  ! [X0] : (cC112(X0) <=> (cC4(X0) & ~cC110(X0)))),
  file('./p/KRS146+1.tptp',unknown)).
fof(f351,plain,(
  spl8_1 | ~spl8_2 | ~spl8_3 | ~spl8_4 | ~spl8_5 | ~spl8_6 | ~spl8_7 | ~spl8_8),
  inference(avatar_split_clause,[],[f350,f342,f338,f334,f330,f326,f322,f318,f314])).
fof(f314,plain,(
  spl8_1 <=> xsd_string(sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl8_1])])).
fof(f350,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC58(iV822576) | ~cC94(iV822576) | xsd_string(sK6)),
  inference(subsumption_resolution,[],[f349,f194])).
fof(f194,plain,(
  ( ! [X0] : (xsd_integer(X0) | xsd_string(X0)) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : ((xsd_string(X0) | xsd_integer(X0)) & (~xsd_integer(X0) | ~xsd_string(X0)))),
  inference(nnf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (xsd_string(X0) <=> ~xsd_integer(X0))),
  file('./p/KRS146+1.tptp',unknown)).
fof(f349,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC58(iV822576) | ~cC94(iV822576) | ~xsd_integer(sK6) | xsd_string(sK6)),
  inference(subsumption_resolution,[],[f348,f217])).
fof(f348,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC116(iV822576) | ~cC58(iV822576) | ~cC94(iV822576) | ~xsd_integer(sK6) | xsd_string(sK6)),
  inference(subsumption_resolution,[],[f347,f191])).
fof(f191,plain,(
  ( ! [X0] : (cowlThing(X0)) )),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0] : (~cowlNothing(X0) & cowlThing(X0))),
  file('./p/KRS146+1.tptp',unknown)).
fof(f347,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC116(iV822576) | ~cC58(iV822576) | ~cC94(iV822576) | ~xsd_integer(sK6) | xsd_string(sK6) | ~cowlThing(sK7)),
  inference(subsumption_resolution,[],[f346,f192])).
fof(f192,plain,(
  ( ! [X0] : (~cowlNothing(X0)) )),
  inference(cnf_transformation,[],[f1])).
fof(f346,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC116(iV822576) | ~cC58(iV822576) | ~cC94(iV822576) | ~xsd_integer(sK6) | xsd_string(sK6) | cowlNothing(sK7) | ~cowlThing(sK7)),
  inference(subsumption_resolution,[],[f306,f191])).
fof(f306,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC116(iV822576) | ~cC58(iV822576) | ~cowlThing(iV822576) | ~cC94(iV822576) | ~xsd_integer(sK6) | xsd_string(sK6) | cowlNothing(sK7) | ~cowlThing(sK7)),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC116(iV822576) | ~cC58(iV822576) | ~cowlThing(iV822576) | ~cC94(iV822576) | ((xsd_integer(sK6) | ~xsd_string(sK6)) & (~xsd_integer(sK6) | xsd_string(sK6))) | (cowlNothing(sK7) | ~cowlThing(sK7))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6,sK7])],[f187,f189,f188])).
fof(f188,plain,(
  ? [X0] : ((xsd_integer(X0) | ~xsd_string(X0)) & (~xsd_integer(X0) | xsd_string(X0))) => ((xsd_integer(sK6) | ~xsd_string(sK6)) & (~xsd_integer(sK6) | xsd_string(sK6)))),
  introduced(choice_axiom,[])).
fof(f189,plain,(
  ? [X1] : (cowlNothing(X1) | ~cowlThing(X1)) => (cowlNothing(sK7) | ~cowlThing(sK7))),
  introduced(choice_axiom,[])).
fof(f187,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC116(iV822576) | ~cC58(iV822576) | ~cowlThing(iV822576) | ~cC94(iV822576) | ? [X0] : ((xsd_integer(X0) | ~xsd_string(X0)) & (~xsd_integer(X0) | xsd_string(X0))) | ? [X1] : (cowlNothing(X1) | ~cowlThing(X1))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC116(iV822576) | ~cC58(iV822576) | ~cowlThing(iV822576) | ~cC94(iV822576) | ? [X0] : (xsd_string(X0) <~> ~xsd_integer(X0)) | ? [X1] : (cowlNothing(X1) | ~cowlThing(X1))),
  inference(ennf_transformation,[],[f94])).
fof(f94,plain,(
  ~(cC80(iV822576) & cC136(iV822576) & cC114(iV822576) & cC110(iV822576) & cC56(iV822576) & cC116(iV822576) & cC58(iV822576) & cowlThing(iV822576) & cC94(iV822576) & ! [X0] : (xsd_string(X0) <=> ~xsd_integer(X0)) & ! [X1] : (~cowlNothing(X1) & cowlThing(X1)))),
  inference(rectify,[],[f93])).
fof(f93,negated_conjecture,(
  ~(cC80(iV822576) & cC136(iV822576) & cC114(iV822576) & cC110(iV822576) & cC56(iV822576) & cC116(iV822576) & cC58(iV822576) & cowlThing(iV822576) & cC94(iV822576) & ! [X0] : (xsd_string(X0) <=> ~xsd_integer(X0)) & ! [X0] : (~cowlNothing(X0) & cowlThing(X0)))),
  inference(negated_conjecture,[],[f92])).
fof(f92,conjecture,(
  cC80(iV822576) & cC136(iV822576) & cC114(iV822576) & cC110(iV822576) & cC56(iV822576) & cC116(iV822576) & cC58(iV822576) & cowlThing(iV822576) & cC94(iV822576) & ! [X0] : (xsd_string(X0) <=> ~xsd_integer(X0)) & ! [X0] : (~cowlNothing(X0) & cowlThing(X0))),
  file('./p/KRS146+1.tptp',unknown)).
fof(f345,plain,(
  ~spl8_1 | ~spl8_2 | ~spl8_3 | ~spl8_4 | ~spl8_5 | ~spl8_6 | ~spl8_7 | ~spl8_8),
  inference(avatar_split_clause,[],[f312,f342,f338,f334,f330,f326,f322,f318,f314])).
fof(f312,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC58(iV822576) | ~cC94(iV822576) | ~xsd_string(sK6)),
  inference(subsumption_resolution,[],[f311,f217])).
fof(f311,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC116(iV822576) | ~cC58(iV822576) | ~cC94(iV822576) | ~xsd_string(sK6)),
  inference(subsumption_resolution,[],[f310,f193])).
fof(f193,plain,(
  ( ! [X0] : (~xsd_integer(X0) | ~xsd_string(X0)) )),
  inference(cnf_transformation,[],[f120])).
fof(f310,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC116(iV822576) | ~cC58(iV822576) | ~cC94(iV822576) | xsd_integer(sK6) | ~xsd_string(sK6)),
  inference(subsumption_resolution,[],[f309,f191])).
fof(f309,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC116(iV822576) | ~cC58(iV822576) | ~cC94(iV822576) | xsd_integer(sK6) | ~xsd_string(sK6) | ~cowlThing(sK7)),
  inference(subsumption_resolution,[],[f308,f192])).
fof(f308,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC116(iV822576) | ~cC58(iV822576) | ~cC94(iV822576) | xsd_integer(sK6) | ~xsd_string(sK6) | cowlNothing(sK7) | ~cowlThing(sK7)),
  inference(subsumption_resolution,[],[f307,f191])).
fof(f307,plain,(
  ~cC80(iV822576) | ~cC136(iV822576) | ~cC114(iV822576) | ~cC110(iV822576) | ~cC56(iV822576) | ~cC116(iV822576) | ~cC58(iV822576) | ~cowlThing(iV822576) | ~cC94(iV822576) | xsd_integer(sK6) | ~xsd_string(sK6) | cowlNothing(sK7) | ~cowlThing(sK7)),
  inference(cnf_transformation,[],[f190])).
% SZS output end Proof for KRS146+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
