% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for KRS121+1
% SZS output start Proof for KRS121+1
fof(f237,plain,(
  $false),
  inference(subsumption_resolution,[],[f236,f177])).
fof(f177,plain,(
  cp1(i2003_11_14_17_21_5199)),
  inference(resolution,[],[f137,f156])).
fof(f156,plain,(
  cUnsatisfiable(i2003_11_14_17_21_5199)),
  inference(cnf_transformation,[],[f31])).
fof(f31,axiom,(
  cUnsatisfiable(i2003_11_14_17_21_5199)),
  file('./p/KRS121+1.tptp',unknown)).
fof(f137,plain,(
  ( ! [X0] : (~cUnsatisfiable(X0) | cp1(X0)) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((cUnsatisfiable(X0) | ~cp1(X0) | ! [X1] : (~ca_Vx3(X1) | ~rr(X0,X1))) & ((cp1(X0) & (ca_Vx3(sK0(X0)) & rr(X0,sK0(X0)))) | ~cUnsatisfiable(X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f90,f91])).
fof(f91,plain,(
  ! [X0] : (? [X2] : (ca_Vx3(X2) & rr(X0,X2)) => (ca_Vx3(sK0(X0)) & rr(X0,sK0(X0))))),
  introduced(choice_axiom,[])).
fof(f90,plain,(
  ! [X0] : ((cUnsatisfiable(X0) | ~cp1(X0) | ! [X1] : (~ca_Vx3(X1) | ~rr(X0,X1))) & ((cp1(X0) & ? [X2] : (ca_Vx3(X2) & rr(X0,X2))) | ~cUnsatisfiable(X0)))),
  inference(rectify,[],[f89])).
fof(f89,plain,(
  ! [X0] : ((cUnsatisfiable(X0) | ~cp1(X0) | ! [X1] : (~ca_Vx3(X1) | ~rr(X0,X1))) & ((cp1(X0) & ? [X1] : (ca_Vx3(X1) & rr(X0,X1))) | ~cUnsatisfiable(X0)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ((cUnsatisfiable(X0) | (~cp1(X0) | ! [X1] : (~ca_Vx3(X1) | ~rr(X0,X1)))) & ((cp1(X0) & ? [X1] : (ca_Vx3(X1) & rr(X0,X1))) | ~cUnsatisfiable(X0)))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : (cUnsatisfiable(X0) <=> (cp1(X0) & ? [X1] : (ca_Vx3(X1) & rr(X0,X1))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X3] : (cUnsatisfiable(X3) <=> (cp1(X3) & ? [X4] : (ca_Vx3(X4) & rr(X3,X4))))),
  file('./p/KRS121+1.tptp',unknown)).
fof(f236,plain,(
  ~cp1(i2003_11_14_17_21_5199)),
  inference(resolution,[],[f229,f182])).
fof(f182,plain,(
  ( ! [X1] : (~cp1xcomp(X1) | ~cp1(X1)) )),
  inference(resolution,[],[f141,f139])).
fof(f139,plain,(
  ( ! [X2,X0] : (~ra_Px1(X0,X2) | ~cp1(X0)) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : ((cp1(X0) | ra_Px1(X0,sK1(X0))) & (! [X2] : ~ra_Px1(X0,X2) | ~cp1(X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f94,f95])).
fof(f95,plain,(
  ! [X0] : (? [X1] : ra_Px1(X0,X1) => ra_Px1(X0,sK1(X0)))),
  introduced(choice_axiom,[])).
fof(f94,plain,(
  ! [X0] : ((cp1(X0) | ? [X1] : ra_Px1(X0,X1)) & (! [X2] : ~ra_Px1(X0,X2) | ~cp1(X0)))),
  inference(rectify,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((cp1(X0) | ? [X1] : ra_Px1(X0,X1)) & (! [X1] : ~ra_Px1(X0,X1) | ~cp1(X0)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (cp1(X0) <=> ! [X1] : ~ra_Px1(X0,X1))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (cp1(X0) <=> ~? [X1] : ra_Px1(X0,X1))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X3] : (cp1(X3) <=> ~? [X4] : ra_Px1(X3,X4))),
  file('./p/KRS121+1.tptp',unknown)).
fof(f141,plain,(
  ( ! [X0] : (ra_Px1(X0,sK2(X0)) | ~cp1xcomp(X0)) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : ((cp1xcomp(X0) | ! [X1] : ~ra_Px1(X0,X1)) & (ra_Px1(X0,sK2(X0)) | ~cp1xcomp(X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f98,f99])).
fof(f99,plain,(
  ! [X0] : (? [X2] : ra_Px1(X0,X2) => ra_Px1(X0,sK2(X0)))),
  introduced(choice_axiom,[])).
fof(f98,plain,(
  ! [X0] : ((cp1xcomp(X0) | ! [X1] : ~ra_Px1(X0,X1)) & (? [X2] : ra_Px1(X0,X2) | ~cp1xcomp(X0)))),
  inference(rectify,[],[f97])).
fof(f97,plain,(
  ! [X0] : ((cp1xcomp(X0) | ! [X1] : ~ra_Px1(X0,X1)) & (? [X1] : ra_Px1(X0,X1) | ~cp1xcomp(X0)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (cp1xcomp(X0) <=> ? [X1] : ra_Px1(X0,X1))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X3] : (cp1xcomp(X3) <=> ? [X5] : ra_Px1(X3,X5))),
  file('./p/KRS121+1.tptp',unknown)).
fof(f229,plain,(
  cp1xcomp(i2003_11_14_17_21_5199)),
  inference(resolution,[],[f225,f191])).
fof(f191,plain,(
  ( ! [X0] : (~ca_Vx3(X0) | cp1xcomp(X0)) )),
  inference(subsumption_resolution,[],[f189,f148])).
fof(f148,plain,(
  ( ! [X0] : (ca_Ax2(sK4(X0)) | ~ca_Vx3(X0)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : ((ca_Vx3(X0) | ! [X1] : (~ca_Ax2(X1) | ~rr(X0,X1))) & ((ca_Ax2(sK4(X0)) & rr(X0,sK4(X0))) | ~ca_Vx3(X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f107,f108])).
fof(f108,plain,(
  ! [X0] : (? [X2] : (ca_Ax2(X2) & rr(X0,X2)) => (ca_Ax2(sK4(X0)) & rr(X0,sK4(X0))))),
  introduced(choice_axiom,[])).
fof(f107,plain,(
  ! [X0] : ((ca_Vx3(X0) | ! [X1] : (~ca_Ax2(X1) | ~rr(X0,X1))) & (? [X2] : (ca_Ax2(X2) & rr(X0,X2)) | ~ca_Vx3(X0)))),
  inference(rectify,[],[f106])).
fof(f106,plain,(
  ! [X0] : ((ca_Vx3(X0) | ! [X1] : (~ca_Ax2(X1) | ~rr(X0,X1))) & (? [X1] : (ca_Ax2(X1) & rr(X0,X1)) | ~ca_Vx3(X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (ca_Vx3(X0) <=> ? [X1] : (ca_Ax2(X1) & rr(X0,X1)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X3] : (ca_Vx3(X3) <=> ? [X4] : (ca_Ax2(X4) & rr(X3,X4)))),
  file('./p/KRS121+1.tptp',unknown)).
fof(f189,plain,(
  ( ! [X0] : (~ca_Vx3(X0) | cp1xcomp(X0) | ~ca_Ax2(sK4(X0))) )),
  inference(resolution,[],[f185,f144])).
fof(f144,plain,(
  ( ! [X2,X0] : (~rinvR(X0,X2) | cp1xcomp(X2) | ~ca_Ax2(X0)) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : ((ca_Ax2(X0) | (~cp1xcomp(sK3(X0)) & rinvR(X0,sK3(X0))) | ~cp1(X0)) & ((! [X2] : (cp1xcomp(X2) | ~rinvR(X0,X2)) & cp1(X0)) | ~ca_Ax2(X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f103,f104])).
fof(f104,plain,(
  ! [X0] : (? [X1] : (~cp1xcomp(X1) & rinvR(X0,X1)) => (~cp1xcomp(sK3(X0)) & rinvR(X0,sK3(X0))))),
  introduced(choice_axiom,[])).
fof(f103,plain,(
  ! [X0] : ((ca_Ax2(X0) | ? [X1] : (~cp1xcomp(X1) & rinvR(X0,X1)) | ~cp1(X0)) & ((! [X2] : (cp1xcomp(X2) | ~rinvR(X0,X2)) & cp1(X0)) | ~ca_Ax2(X0)))),
  inference(rectify,[],[f102])).
fof(f102,plain,(
  ! [X0] : ((ca_Ax2(X0) | ? [X1] : (~cp1xcomp(X1) & rinvR(X0,X1)) | ~cp1(X0)) & ((! [X1] : (cp1xcomp(X1) | ~rinvR(X0,X1)) & cp1(X0)) | ~ca_Ax2(X0)))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ((ca_Ax2(X0) | (? [X1] : (~cp1xcomp(X1) & rinvR(X0,X1)) | ~cp1(X0))) & ((! [X1] : (cp1xcomp(X1) | ~rinvR(X0,X1)) & cp1(X0)) | ~ca_Ax2(X0)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (ca_Ax2(X0) <=> (! [X1] : (cp1xcomp(X1) | ~rinvR(X0,X1)) & cp1(X0)))),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (ca_Ax2(X0) <=> (! [X1] : (rinvR(X0,X1) => cp1xcomp(X1)) & cp1(X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X3] : (ca_Ax2(X3) <=> (! [X4] : (rinvR(X3,X4) => cp1xcomp(X4)) & cp1(X3)))),
  file('./p/KRS121+1.tptp',unknown)).
fof(f185,plain,(
  ( ! [X1] : (rinvR(sK4(X1),X1) | ~ca_Vx3(X1)) )),
  inference(resolution,[],[f154,f147])).
fof(f147,plain,(
  ( ! [X0] : (rr(X0,sK4(X0)) | ~ca_Vx3(X0)) )),
  inference(cnf_transformation,[],[f109])).
fof(f154,plain,(
  ( ! [X0,X1] : (~rr(X1,X0) | rinvR(X0,X1)) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : ((rinvR(X0,X1) | ~rr(X1,X0)) & (rr(X1,X0) | ~rinvR(X0,X1)))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : (rinvR(X0,X1) <=> rr(X1,X0))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X3,X4] : (rinvR(X3,X4) <=> rr(X4,X3))),
  file('./p/KRS121+1.tptp',unknown)).
fof(f225,plain,(
  ca_Vx3(i2003_11_14_17_21_5199)),
  inference(resolution,[],[f222,f156])).
fof(f222,plain,(
  ( ! [X0] : (~cUnsatisfiable(X0) | ca_Vx3(X0)) )),
  inference(subsumption_resolution,[],[f217,f136])).
fof(f136,plain,(
  ( ! [X0] : (ca_Vx3(sK0(X0)) | ~cUnsatisfiable(X0)) )),
  inference(cnf_transformation,[],[f92])).
fof(f217,plain,(
  ( ! [X0] : (~ca_Vx3(sK0(X0)) | ca_Vx3(X0) | ~cUnsatisfiable(X0)) )),
  inference(resolution,[],[f216,f135])).
fof(f135,plain,(
  ( ! [X0] : (rr(X0,sK0(X0)) | ~cUnsatisfiable(X0)) )),
  inference(cnf_transformation,[],[f92])).
fof(f216,plain,(
  ( ! [X6,X5] : (~rr(X5,X6) | ~ca_Vx3(X6) | ca_Vx3(X5)) )),
  inference(subsumption_resolution,[],[f214,f148])).
fof(f214,plain,(
  ( ! [X6,X5] : (~rr(X5,X6) | ~ca_Vx3(X6) | ~ca_Ax2(sK4(X6)) | ca_Vx3(X5)) )),
  inference(resolution,[],[f201,f149])).
fof(f149,plain,(
  ( ! [X0,X1] : (~rr(X0,X1) | ~ca_Ax2(X1) | ca_Vx3(X0)) )),
  inference(cnf_transformation,[],[f109])).
fof(f201,plain,(
  ( ! [X2,X3] : (rr(X2,sK4(X3)) | ~rr(X2,X3) | ~ca_Vx3(X3)) )),
  inference(resolution,[],[f155,f147])).
fof(f155,plain,(
  ( ! [X2,X0,X1] : (~rr(X1,X2) | rr(X0,X2) | ~rr(X0,X1)) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : (rr(X0,X2) | ~rr(X1,X2) | ~rr(X0,X1))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : (rr(X0,X2) | (~rr(X1,X2) | ~rr(X0,X1)))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2] : ((rr(X1,X2) & rr(X0,X1)) => rr(X0,X2))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X3,X4,X6] : ((rr(X4,X6) & rr(X3,X4)) => rr(X3,X6))),
  file('./p/KRS121+1.tptp',unknown)).
% SZS output end Proof for KRS121+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
