% Refutation found. Thanks to Tanya!
% SZS status Theorem for KRS203+1
% SZS output start Proof for KRS203+1
fof(f165599,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f342,f5290,f5577,f5583,f5806,f165598])).
fof(f165598,plain,(
  ~spl59_160),
  inference(avatar_contradiction_clause,[],[f165597])).
fof(f165597,plain,(
  $false | ~spl59_160),
  inference(subsumption_resolution,[],[f165521,f292])).
fof(f292,plain,(
  ~model(sK53,sK50)),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ~model(sK52,sK51) & (model(sK53,sK51) & ~model(sK53,sK50)) & ! [X4] : (model(X4,sK51) | ~model(X4,sK50)) & model(sK54,sK50)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK50,sK51,sK52,sK53,sK54])],[f71,f200,f199,f198,f197])).
fof(f197,plain,(
  ? [X0,X1] : (? [X2] : ~model(X2,X1) & ? [X3] : (model(X3,X1) & ~model(X3,X0)) & ! [X4] : (model(X4,X1) | ~model(X4,X0)) & ? [X5] : model(X5,X0)) => (? [X2] : ~model(X2,sK51) & ? [X3] : (model(X3,sK51) & ~model(X3,sK50)) & ! [X4] : (model(X4,sK51) | ~model(X4,sK50)) & ? [X5] : model(X5,sK50))),
  introduced(choice_axiom,[])).
fof(f198,plain,(
  ? [X2] : ~model(X2,sK51) => ~model(sK52,sK51)),
  introduced(choice_axiom,[])).
fof(f199,plain,(
  ? [X3] : (model(X3,sK51) & ~model(X3,sK50)) => (model(sK53,sK51) & ~model(sK53,sK50))),
  introduced(choice_axiom,[])).
fof(f200,plain,(
  ? [X5] : model(X5,sK50) => model(sK54,sK50)),
  introduced(choice_axiom,[])).
fof(f71,plain,(
  ? [X0,X1] : (? [X2] : ~model(X2,X1) & ? [X3] : (model(X3,X1) & ~model(X3,X0)) & ! [X4] : (model(X4,X1) | ~model(X4,X0)) & ? [X5] : model(X5,X0))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ? [X0,X1] : (? [X2] : ~model(X2,X1) & ? [X3] : (model(X3,X1) & ~model(X3,X0)) & ! [X4] : (model(X4,X0) => model(X4,X1)) & ? [X5] : model(X5,X0))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ? [X0,X1] : (? [X5] : ~model(X5,X1) & ? [X4] : (model(X4,X1) & ~model(X4,X0)) & ! [X3] : (model(X3,X0) => model(X3,X1)) & ? [X2] : model(X2,X0))),
  file('./p/KRS203+1.tptp',unknown)).
fof(f165521,plain,(
  model(sK53,sK50) | ~spl59_160),
  inference(resolution,[],[f165177,f293])).
fof(f293,plain,(
  model(sK53,sK51)),
  inference(cnf_transformation,[],[f201])).
fof(f165177,plain,(
  ( ! [X40] : (~model(X40,sK51) | model(X40,sK50)) ) | ~spl59_160),
  inference(subsumption_resolution,[],[f164959,f298])).
fof(f298,plain,(
  ~nota(sat,thm)),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ~nota(sat,thm)),
  inference(flattening,[],[f34])).
fof(f34,negated_conjecture,(
  ~nota(sat,thm)),
  inference(negated_conjecture,[],[f33])).
fof(f33,conjecture,(
  nota(sat,thm)),
  file('./p/KRS203+1.tptp',unknown)).
fof(f164959,plain,(
  ( ! [X40] : (nota(sat,thm) | ~model(X40,sK51) | model(X40,sK50)) ) | ~spl59_160),
  inference(resolution,[],[f7257,f220])).
fof(f220,plain,(
  ( ! [X2,X0,X1] : (~status(X0,X1,thm) | ~model(X2,X0) | model(X2,X1)) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : ((! [X2] : (model(X2,X1) | ~model(X2,X0)) | ~status(X0,X1,thm)) & (status(X0,X1,thm) | (~model(sK9(X0,X1),X1) & model(sK9(X0,X1),X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f96,f97])).
fof(f97,plain,(
  ! [X1,X0] : (? [X3] : (~model(X3,X1) & model(X3,X0)) => (~model(sK9(X0,X1),X1) & model(sK9(X0,X1),X0)))),
  introduced(choice_axiom,[])).
fof(f96,plain,(
  ! [X0,X1] : ((! [X2] : (model(X2,X1) | ~model(X2,X0)) | ~status(X0,X1,thm)) & (status(X0,X1,thm) | ? [X3] : (~model(X3,X1) & model(X3,X0))))),
  inference(rectify,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : ((! [X2] : (model(X2,X1) | ~model(X2,X0)) | ~status(X0,X1,thm)) & (status(X0,X1,thm) | ? [X2] : (~model(X2,X1) & model(X2,X0))))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : (! [X2] : (model(X2,X1) | ~model(X2,X0)) <=> status(X0,X1,thm))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1] : (! [X2] : (model(X2,X0) => model(X2,X1)) <=> status(X0,X1,thm))),
  file('./p/KRS203+1.tptp',unknown)).
fof(f7257,plain,(
  ( ! [X10] : (status(sK51,sK50,X10) | nota(sat,X10)) ) | ~spl59_160),
  inference(resolution,[],[f5576,f278])).
fof(f278,plain,(
  ( ! [X2,X0,X3,X1] : (~status(X2,X3,X0) | status(X2,X3,X1) | nota(X0,X1)) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : (nota(X0,X1) | ! [X2,X3] : (status(X2,X3,X1) | ~status(X2,X3,X0)))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : (? [X2,X3] : (~status(X2,X3,X1) & status(X2,X3,X0)) => nota(X0,X1))),
  inference(unused_predicate_definition_removal,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : (? [X2,X3] : (~status(X2,X3,X1) & status(X2,X3,X0)) <=> nota(X0,X1))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6,X7] : (? [X0,X1] : (~status(X0,X1,X7) & status(X0,X1,X6)) <=> nota(X6,X7))),
  file('./p/KRS203+1.tptp',unknown)).
fof(f5576,plain,(
  status(sK51,sK50,sat) | ~spl59_160),
  inference(avatar_component_clause,[],[f5574])).
fof(f5574,plain,(
  spl59_160 <=> status(sK51,sK50,sat)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_160])])).
fof(f5806,plain,(
  ~spl59_161),
  inference(avatar_contradiction_clause,[],[f5781])).
fof(f5781,plain,(
  $false | ~spl59_161),
  inference(resolution,[],[f5582,f1657])).
fof(f1657,plain,(
  model(sK4(sK50),sK50)),
  inference(subsumption_resolution,[],[f1654,f283])).
fof(f283,plain,(
  ( ! [X1] : (model(X1,sK41)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X1] : model(X1,sK41)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK41])],[f53,f185])).
fof(f185,plain,(
  ? [X0] : ! [X1] : model(X1,X0) => ! [X1] : model(X1,sK41)),
  introduced(choice_axiom,[])).
fof(f53,plain,(
  ? [X0] : ! [X1] : model(X1,X0)),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ? [X9] : ! [X8] : model(X8,X9)),
  file('./p/KRS203+1.tptp',unknown)).
fof(f1654,plain,(
  ( ! [X0] : (~model(X0,sK41) | model(sK4(sK50),sK50)) )),
  inference(resolution,[],[f1602,f214])).
fof(f214,plain,(
  ( ! [X0,X3,X1] : (~status(X0,X1,esa) | ~model(X3,X1) | model(sK4(X0),X0)) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : ((((model(sK4(X0),X0) | ! [X3] : ~model(X3,X1)) & (model(sK5(X1),X1) | ! [X5] : ~model(X5,X0))) | ~status(X0,X1,esa)) & (status(X0,X1,esa) | ((! [X6] : ~model(X6,X1) | ! [X7] : ~model(X7,X0)) & (model(sK6(X1),X1) | model(sK7(X0),X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5,sK6,sK7])],[f85,f89,f88,f87,f86])).
fof(f86,plain,(
  ! [X0] : (? [X2] : model(X2,X0) => model(sK4(X0),X0))),
  introduced(choice_axiom,[])).
fof(f87,plain,(
  ! [X1] : (? [X4] : model(X4,X1) => model(sK5(X1),X1))),
  introduced(choice_axiom,[])).
fof(f88,plain,(
  ! [X1] : (? [X8] : model(X8,X1) => model(sK6(X1),X1))),
  introduced(choice_axiom,[])).
fof(f89,plain,(
  ! [X0] : (? [X9] : model(X9,X0) => model(sK7(X0),X0))),
  introduced(choice_axiom,[])).
fof(f85,plain,(
  ! [X0,X1] : ((((? [X2] : model(X2,X0) | ! [X3] : ~model(X3,X1)) & (? [X4] : model(X4,X1) | ! [X5] : ~model(X5,X0))) | ~status(X0,X1,esa)) & (status(X0,X1,esa) | ((! [X6] : ~model(X6,X1) | ! [X7] : ~model(X7,X0)) & (? [X8] : model(X8,X1) | ? [X9] : model(X9,X0)))))),
  inference(rectify,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : ((((? [X2] : model(X2,X0) | ! [X3] : ~model(X3,X1)) & (? [X3] : model(X3,X1) | ! [X2] : ~model(X2,X0))) | ~status(X0,X1,esa)) & (status(X0,X1,esa) | ((! [X3] : ~model(X3,X1) | ! [X2] : ~model(X2,X0)) & (? [X3] : model(X3,X1) | ? [X2] : model(X2,X0)))))),
  inference(nnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1] : ((? [X2] : model(X2,X0) <=> ? [X3] : model(X3,X1)) <=> status(X0,X1,esa))),
  file('./p/KRS203+1.tptp',unknown)).
fof(f1602,plain,(
  status(sK50,sK41,esa)),
  inference(resolution,[],[f484,f283])).
fof(f484,plain,(
  ( ! [X70,X71] : (~model(X70,X71) | status(sK50,X71,esa)) )),
  inference(resolution,[],[f212,f290])).
fof(f290,plain,(
  model(sK54,sK50)),
  inference(cnf_transformation,[],[f201])).
fof(f212,plain,(
  ( ! [X6,X0,X7,X1] : (~model(X7,X0) | ~model(X6,X1) | status(X0,X1,esa)) )),
  inference(cnf_transformation,[],[f90])).
fof(f5582,plain,(
  ( ! [X1] : (~model(X1,sK50)) ) | ~spl59_161),
  inference(avatar_component_clause,[],[f5581])).
fof(f5581,plain,(
  spl59_161 <=> ! [X1] : ~model(X1,sK50)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_161])])).
fof(f5583,plain,(
  spl59_161 | spl59_18 | ~spl59_4),
  inference(avatar_split_clause,[],[f5578,f340,f646,f5581])).
fof(f646,plain,(
  spl59_18 <=> ! [X16] : model(sK0(X16),X16)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_18])])).
fof(f340,plain,(
  spl59_4 <=> ! [X9] : status(X9,sK50,unp)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_4])])).
fof(f5578,plain,(
  ( ! [X0,X1] : (model(sK0(X0),X0) | ~model(X1,sK50)) ) | ~spl59_4),
  inference(resolution,[],[f341,f207])).
fof(f207,plain,(
  ( ! [X2,X0,X1] : (~status(X0,X1,unp) | model(sK0(X0),X0) | ~model(X2,X1)) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : ((! [X2] : ~model(X2,X1) | model(sK0(X0),X0) | ~status(X0,X1,unp)) & (status(X0,X1,unp) | (model(sK1(X1),X1) & ! [X5] : ~model(X5,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f74,f76,f75])).
fof(f75,plain,(
  ! [X0] : (? [X3] : model(X3,X0) => model(sK0(X0),X0))),
  introduced(choice_axiom,[])).
fof(f76,plain,(
  ! [X1] : (? [X4] : model(X4,X1) => model(sK1(X1),X1))),
  introduced(choice_axiom,[])).
fof(f74,plain,(
  ! [X0,X1] : ((! [X2] : ~model(X2,X1) | ? [X3] : model(X3,X0) | ~status(X0,X1,unp)) & (status(X0,X1,unp) | (? [X4] : model(X4,X1) & ! [X5] : ~model(X5,X0))))),
  inference(rectify,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ((! [X3] : ~model(X3,X1) | ? [X2] : model(X2,X0) | ~status(X0,X1,unp)) & (status(X0,X1,unp) | (? [X3] : model(X3,X1) & ! [X2] : ~model(X2,X0))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : (((! [X3] : ~model(X3,X1) | ? [X2] : model(X2,X0)) | ~status(X0,X1,unp)) & (status(X0,X1,unp) | (? [X3] : model(X3,X1) & ! [X2] : ~model(X2,X0))))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ((! [X3] : ~model(X3,X1) | ? [X2] : model(X2,X0)) <=> status(X0,X1,unp))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1] : ((~? [X2] : model(X2,X0) => ~? [X3] : model(X3,X1)) <=> status(X0,X1,unp))),
  file('./p/KRS203+1.tptp',unknown)).
fof(f341,plain,(
  ( ! [X9] : (status(X9,sK50,unp)) ) | ~spl59_4),
  inference(avatar_component_clause,[],[f340])).
fof(f5577,plain,(
  spl59_4 | spl59_160 | ~spl59_3),
  inference(avatar_split_clause,[],[f5566,f336,f5574,f340])).
fof(f336,plain,(
  spl59_3 <=> model(sK1(sK50),sK51)),
  introduced(avatar_definition,[new_symbols(naming,[spl59_3])])).
fof(f5566,plain,(
  ( ! [X7] : (status(sK51,sK50,sat) | status(X7,sK50,unp)) ) | ~spl59_3),
  inference(resolution,[],[f490,f338])).
fof(f338,plain,(
  model(sK1(sK50),sK51) | ~spl59_3),
  inference(avatar_component_clause,[],[f336])).
fof(f490,plain,(
  ( ! [X6,X7,X5] : (~model(sK1(X6),X5) | status(X5,X6,sat) | status(X7,X6,unp)) )),
  inference(resolution,[],[f215,f206])).
fof(f206,plain,(
  ( ! [X0,X1] : (status(X0,X1,unp) | model(sK1(X1),X1)) )),
  inference(cnf_transformation,[],[f77])).
fof(f215,plain,(
  ( ! [X0,X3,X1] : (~model(X3,X1) | status(X0,X1,sat) | ~model(X3,X0)) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : (((model(sK8(X0,X1),X1) & model(sK8(X0,X1),X0)) | ~status(X0,X1,sat)) & (status(X0,X1,sat) | ! [X3] : (~model(X3,X1) | ~model(X3,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f92,f93])).
fof(f93,plain,(
  ! [X1,X0] : (? [X2] : (model(X2,X1) & model(X2,X0)) => (model(sK8(X0,X1),X1) & model(sK8(X0,X1),X0)))),
  introduced(choice_axiom,[])).
fof(f92,plain,(
  ! [X0,X1] : ((? [X2] : (model(X2,X1) & model(X2,X0)) | ~status(X0,X1,sat)) & (status(X0,X1,sat) | ! [X3] : (~model(X3,X1) | ~model(X3,X0))))),
  inference(rectify,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : ((? [X2] : (model(X2,X1) & model(X2,X0)) | ~status(X0,X1,sat)) & (status(X0,X1,sat) | ! [X2] : (~model(X2,X1) | ~model(X2,X0))))),
  inference(nnf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1] : (? [X2] : (model(X2,X1) & model(X2,X0)) <=> status(X0,X1,sat))),
  file('./p/KRS203+1.tptp',unknown)).
fof(f5290,plain,(
  ~spl59_18),
  inference(avatar_contradiction_clause,[],[f5288])).
fof(f5288,plain,(
  $false | ~spl59_18),
  inference(resolution,[],[f647,f286])).
fof(f286,plain,(
  ( ! [X1] : (~model(X1,sK45)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X1] : ~model(X1,sK45)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK45])],[f55,f191])).
fof(f191,plain,(
  ? [X0] : ! [X1] : ~model(X1,X0) => ! [X1] : ~model(X1,sK45)),
  introduced(choice_axiom,[])).
fof(f55,plain,(
  ? [X0] : ! [X1] : ~model(X1,X0)),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ? [X9] : ! [X8] : ~model(X8,X9)),
  file('./p/KRS203+1.tptp',unknown)).
fof(f647,plain,(
  ( ! [X16] : (model(sK0(X16),X16)) ) | ~spl59_18),
  inference(avatar_component_clause,[],[f646])).
fof(f342,plain,(
  spl59_3 | spl59_4),
  inference(avatar_split_clause,[],[f334,f340,f336])).
fof(f334,plain,(
  ( ! [X9] : (status(X9,sK50,unp) | model(sK1(sK50),sK51)) )),
  inference(resolution,[],[f206,f291])).
fof(f291,plain,(
  ( ! [X4] : (~model(X4,sK50) | model(X4,sK51)) )),
  inference(cnf_transformation,[],[f201])).
% SZS output end Proof for KRS203+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 67418
% Time elapsed: 10.904 s
% ------------------------------
% ------------------------------
