% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for LCL090-1
% SZS output start Proof for LCL090-1
fof(f2309,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f18,f25,f59,f872,f898,f1399,f1718,f1767,f2196,f2246,f2308])).
fof(f2308,plain,(
  ~spl0_41),
  inference(avatar_contradiction_clause,[],[f2286])).
fof(f2286,plain,(
  $false | ~spl0_41),
  inference(resolution,[],[f2245,f3])).
fof(f3,axiom,(
  ~is_a_theorem(implies(a,a))),
  file('./p/LCL090-1.tptp',unknown)).
fof(f2245,plain,(
  ( ! [X19] : (is_a_theorem(implies(X19,X19))) ) | ~spl0_41),
  inference(avatar_component_clause,[],[f2244])).
fof(f2244,plain,(
  spl0_41 <=> ! [X19] : is_a_theorem(implies(X19,X19))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_41])])).
fof(f2246,plain,(
  spl0_41 | spl0_2 | ~spl0_40),
  inference(avatar_split_clause,[],[f2226,f2194,f16,f2244])).
fof(f16,plain,(
  spl0_2 <=> ! [X0] : ~is_a_theorem(X0)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f2194,plain,(
  spl0_40 <=> ! [X25,X26] : is_a_theorem(implies(X25,implies(X26,X26)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_40])])).
fof(f2226,plain,(
  ( ! [X19,X18] : (~is_a_theorem(X18) | is_a_theorem(implies(X19,X19))) ) | ~spl0_40),
  inference(resolution,[],[f2195,f1])).
fof(f1,axiom,(
  ( ! [X0,X1] : (~is_a_theorem(implies(X0,X1)) | ~is_a_theorem(X0) | is_a_theorem(X1)) )),
  file('./p/LCL090-1.tptp',unknown)).
fof(f2195,plain,(
  ( ! [X26,X25] : (is_a_theorem(implies(X25,implies(X26,X26)))) ) | ~spl0_40),
  inference(avatar_component_clause,[],[f2194])).
fof(f2196,plain,(
  spl0_40 | spl0_2 | ~spl0_26 | ~spl0_35),
  inference(avatar_split_clause,[],[f2176,f1765,f1397,f16,f2194])).
fof(f1397,plain,(
  spl0_26 <=> ! [X29,X28,X30] : is_a_theorem(implies(implies(X28,implies(X29,implies(X30,X30))),implies(X29,implies(X30,X30))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_26])])).
fof(f1765,plain,(
  spl0_35 <=> ! [X23,X24] : is_a_theorem(implies(implies(X23,X24),implies(X24,X24)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_35])])).
fof(f2176,plain,(
  ( ! [X26,X24,X25] : (~is_a_theorem(X24) | is_a_theorem(implies(X25,implies(X26,X26)))) ) | (~spl0_26 | ~spl0_35)),
  inference(resolution,[],[f2049,f1])).
fof(f2049,plain,(
  ( ! [X4,X5,X3] : (is_a_theorem(implies(X3,implies(X4,implies(X5,X5))))) ) | (~spl0_26 | ~spl0_35)),
  inference(resolution,[],[f1654,f1766])).
fof(f1766,plain,(
  ( ! [X24,X23] : (is_a_theorem(implies(implies(X23,X24),implies(X24,X24)))) ) | ~spl0_35),
  inference(avatar_component_clause,[],[f1765])).
fof(f1654,plain,(
  ( ! [X94,X92,X93,X91] : (~is_a_theorem(implies(implies(X91,X91),X92)) | is_a_theorem(implies(X93,implies(X94,X92)))) ) | ~spl0_26),
  inference(resolution,[],[f1419,f1])).
fof(f1419,plain,(
  ( ! [X6,X4,X5,X3] : (is_a_theorem(implies(implies(implies(X3,X3),X4),implies(X5,implies(X6,X4))))) ) | ~spl0_26),
  inference(resolution,[],[f1398,f4])).
fof(f4,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~is_a_theorem(implies(implies(X0,X1),implies(X2,X3))) | is_a_theorem(implies(implies(X3,X0),implies(X4,implies(X2,X0))))) )),
  inference(resolution,[],[f2,f1])).
fof(f2,axiom,(
  ( ! [X6,X4,X2,X5,X3] : (is_a_theorem(implies(implies(implies(X2,X3),implies(X4,X5)),implies(implies(X5,X2),implies(X6,implies(X4,X2)))))) )),
  file('./p/LCL090-1.tptp',unknown)).
fof(f1398,plain,(
  ( ! [X30,X28,X29] : (is_a_theorem(implies(implies(X28,implies(X29,implies(X30,X30))),implies(X29,implies(X30,X30))))) ) | ~spl0_26),
  inference(avatar_component_clause,[],[f1397])).
fof(f1767,plain,(
  spl0_35 | spl0_2 | ~spl0_34),
  inference(avatar_split_clause,[],[f1748,f1716,f16,f1765])).
fof(f1716,plain,(
  spl0_34 <=> ! [X29,X28,X30] : is_a_theorem(implies(X28,implies(implies(X29,X30),implies(X30,X30))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_34])])).
fof(f1748,plain,(
  ( ! [X24,X23,X22] : (~is_a_theorem(X22) | is_a_theorem(implies(implies(X23,X24),implies(X24,X24)))) ) | ~spl0_34),
  inference(resolution,[],[f1717,f1])).
fof(f1717,plain,(
  ( ! [X30,X28,X29] : (is_a_theorem(implies(X28,implies(implies(X29,X30),implies(X30,X30))))) ) | ~spl0_34),
  inference(avatar_component_clause,[],[f1716])).
fof(f1718,plain,(
  spl0_34 | spl0_2 | ~spl0_26),
  inference(avatar_split_clause,[],[f1695,f1397,f16,f1716])).
fof(f1695,plain,(
  ( ! [X30,X28,X29,X27] : (~is_a_theorem(X27) | is_a_theorem(implies(X28,implies(implies(X29,X30),implies(X30,X30))))) ) | ~spl0_26),
  inference(resolution,[],[f1420,f1])).
fof(f1420,plain,(
  ( ! [X10,X8,X7,X9] : (is_a_theorem(implies(X7,implies(X8,implies(implies(X9,X10),implies(X10,X10)))))) ) | ~spl0_26),
  inference(resolution,[],[f1398,f50])).
fof(f50,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25,X22] : (~is_a_theorem(implies(implies(X22,implies(X23,implies(X24,X25))),implies(implies(X26,X25),implies(X25,X27)))) | is_a_theorem(implies(X28,implies(X29,implies(implies(X26,X25),implies(X25,X27)))))) )),
  inference(resolution,[],[f20,f1])).
fof(f20,plain,(
  ( ! [X6,X12,X10,X8,X7,X13,X11,X9] : (is_a_theorem(implies(implies(implies(X6,implies(X7,implies(X8,X9))),implies(implies(X10,X9),implies(X9,X11))),implies(X12,implies(X13,implies(implies(X10,X9),implies(X9,X11))))))) )),
  inference(resolution,[],[f6,f4])).
fof(f6,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (is_a_theorem(implies(implies(implies(implies(X0,X1),implies(X1,X2)),implies(X3,implies(X4,X1))),implies(X5,implies(X6,implies(X3,implies(X4,X1))))))) )),
  inference(resolution,[],[f5,f4])).
fof(f5,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (is_a_theorem(implies(implies(implies(X0,implies(X1,X2)),implies(X2,X3)),implies(X4,implies(implies(X5,X2),implies(X2,X3)))))) )),
  inference(resolution,[],[f4,f2])).
fof(f1399,plain,(
  spl0_26 | spl0_2 | ~spl0_4 | ~spl0_19),
  inference(avatar_split_clause,[],[f1378,f896,f57,f16,f1397])).
fof(f57,plain,(
  spl0_4 <=> ! [X9,X7,X8,X10] : is_a_theorem(implies(implies(X7,X8),implies(X8,implies(X9,implies(X10,X8)))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f896,plain,(
  spl0_19 <=> ! [X34,X36,X33,X35] : is_a_theorem(implies(implies(implies(X33,X34),X33),implies(X35,implies(implies(X36,X33),X33))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_19])])).
fof(f1378,plain,(
  ( ! [X30,X28,X29,X27] : (~is_a_theorem(X27) | is_a_theorem(implies(implies(X28,implies(X29,implies(X30,X30))),implies(X29,implies(X30,X30))))) ) | (~spl0_4 | ~spl0_19)),
  inference(resolution,[],[f946,f1])).
fof(f946,plain,(
  ( ! [X26,X24,X23,X25] : (is_a_theorem(implies(X23,implies(implies(X24,implies(X25,implies(X26,X26))),implies(X25,implies(X26,X26)))))) ) | (~spl0_4 | ~spl0_19)),
  inference(resolution,[],[f923,f63])).
fof(f63,plain,(
  ( ! [X4,X2,X0,X3,X1] : (is_a_theorem(implies(implies(implies(X0,implies(X1,X2)),X3),implies(X4,implies(X2,X3))))) ) | ~spl0_4),
  inference(resolution,[],[f58,f4])).
fof(f58,plain,(
  ( ! [X10,X8,X7,X9] : (is_a_theorem(implies(implies(X7,X8),implies(X8,implies(X9,implies(X10,X8)))))) ) | ~spl0_4),
  inference(avatar_component_clause,[],[f57])).
fof(f923,plain,(
  ( ! [X54,X52,X55,X53] : (~is_a_theorem(implies(implies(X52,X53),X52)) | is_a_theorem(implies(X54,implies(implies(X55,X52),X52)))) ) | ~spl0_19),
  inference(resolution,[],[f897,f1])).
fof(f897,plain,(
  ( ! [X35,X33,X36,X34] : (is_a_theorem(implies(implies(implies(X33,X34),X33),implies(X35,implies(implies(X36,X33),X33))))) ) | ~spl0_19),
  inference(avatar_component_clause,[],[f896])).
fof(f898,plain,(
  spl0_19 | spl0_2 | ~spl0_18),
  inference(avatar_split_clause,[],[f883,f870,f16,f896])).
fof(f870,plain,(
  spl0_18 <=> ! [X40,X42,X38,X41,X39] : is_a_theorem(implies(X38,implies(implies(implies(X39,X40),X39),implies(X41,implies(implies(X42,X39),X39)))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_18])])).
fof(f883,plain,(
  ( ! [X35,X33,X36,X34,X32] : (~is_a_theorem(X32) | is_a_theorem(implies(implies(implies(X33,X34),X33),implies(X35,implies(implies(X36,X33),X33))))) ) | ~spl0_18),
  inference(resolution,[],[f871,f1])).
fof(f871,plain,(
  ( ! [X39,X41,X38,X42,X40] : (is_a_theorem(implies(X38,implies(implies(implies(X39,X40),X39),implies(X41,implies(implies(X42,X39),X39)))))) ) | ~spl0_18),
  inference(avatar_component_clause,[],[f870])).
fof(f872,plain,(
  spl0_18 | spl0_2),
  inference(avatar_split_clause,[],[f853,f16,f870])).
fof(f853,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (~is_a_theorem(X37) | is_a_theorem(implies(X38,implies(implies(implies(X39,X40),X39),implies(X41,implies(implies(X42,X39),X39)))))) )),
  inference(resolution,[],[f553,f1])).
fof(f553,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (is_a_theorem(implies(X0,implies(X1,implies(implies(implies(X2,X3),X2),implies(X4,implies(implies(X5,X2),X2))))))) )),
  inference(resolution,[],[f198,f2])).
fof(f198,plain,(
  ( ! [X47,X54,X52,X50,X48,X55,X53,X51,X49] : (~is_a_theorem(implies(implies(X47,implies(implies(X48,X49),implies(X49,X50))),implies(X51,implies(X52,implies(X53,X49))))) | is_a_theorem(implies(X54,implies(X55,implies(X51,implies(X52,implies(X53,X49))))))) )),
  inference(resolution,[],[f48,f1])).
fof(f48,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X13,X11,X9] : (is_a_theorem(implies(implies(implies(X6,implies(implies(X7,X8),implies(X8,X9))),implies(X10,implies(X11,implies(X12,X8)))),implies(X13,implies(X14,implies(X10,implies(X11,implies(X12,X8)))))))) )),
  inference(resolution,[],[f20,f4])).
fof(f59,plain,(
  spl0_4 | spl0_2 | ~spl0_1),
  inference(avatar_split_clause,[],[f52,f13,f16,f57])).
fof(f13,plain,(
  spl0_1 <=> ! [X1,X3,X2,X4] : is_a_theorem(implies(implies(X1,X2),implies(X2,implies(implies(X3,X4),implies(X4,X2)))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f52,plain,(
  ( ! [X6,X10,X8,X7,X9] : (~is_a_theorem(X6) | is_a_theorem(implies(implies(X7,X8),implies(X8,implies(X9,implies(X10,X8)))))) ) | ~spl0_1),
  inference(resolution,[],[f44,f1])).
fof(f44,plain,(
  ( ! [X6,X10,X8,X7,X9] : (is_a_theorem(implies(X6,implies(implies(X7,X8),implies(X8,implies(X9,implies(X10,X8))))))) ) | ~spl0_1),
  inference(resolution,[],[f26,f7])).
fof(f7,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (~is_a_theorem(implies(implies(X7,implies(X8,X9)),implies(X9,X10))) | is_a_theorem(implies(X11,implies(implies(X12,X9),implies(X9,X10))))) )),
  inference(resolution,[],[f5,f1])).
fof(f26,plain,(
  ( ! [X4,X2,X0,X3,X1] : (is_a_theorem(implies(implies(implies(implies(X0,X1),implies(X1,X2)),X3),implies(X4,implies(X2,X3))))) ) | ~spl0_1),
  inference(resolution,[],[f14,f4])).
fof(f14,plain,(
  ( ! [X4,X2,X3,X1] : (is_a_theorem(implies(implies(X1,X2),implies(X2,implies(implies(X3,X4),implies(X4,X2)))))) ) | ~spl0_1),
  inference(avatar_component_clause,[],[f13])).
fof(f25,plain,(
  ~spl0_2),
  inference(avatar_contradiction_clause,[],[f24])).
fof(f24,plain,(
  $false | ~spl0_2),
  inference(subsumption_resolution,[],[f20,f17])).
fof(f17,plain,(
  ( ! [X0] : (~is_a_theorem(X0)) ) | ~spl0_2),
  inference(avatar_component_clause,[],[f16])).
fof(f18,plain,(
  spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f9,f16,f13])).
fof(f9,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~is_a_theorem(X0) | is_a_theorem(implies(implies(X1,X2),implies(X2,implies(implies(X3,X4),implies(X4,X2)))))) )),
  inference(resolution,[],[f8,f1])).
fof(f8,plain,(
  ( ! [X4,X2,X0,X3,X1] : (is_a_theorem(implies(X0,implies(implies(X1,X2),implies(X2,implies(implies(X3,X4),implies(X4,X2))))))) )),
  inference(resolution,[],[f7,f5])).
% SZS output end Proof for LCL090-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 7419
% Time elapsed: 0.032 s
% ------------------------------
% ------------------------------
