% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for LCL021-1
% SZS output start Proof for LCL021-1
fof(f718922,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f125,f135,f311,f1967,f2020,f2142,f5401,f6017,f6162,f6925,f9327,f10077,f13774,f171626,f718918])).
fof(f718918,plain,(
  ~spl0_3 | ~spl0_100 | ~spl0_276 | ~spl0_289 | ~spl0_297 | ~spl0_325 | ~spl0_330 | ~spl0_374),
  inference(avatar_contradiction_clause,[],[f718917])).
fof(f718917,plain,(
  $false | (~spl0_3 | ~spl0_100 | ~spl0_276 | ~spl0_289 | ~spl0_297 | ~spl0_325 | ~spl0_330 | ~spl0_374)),
  inference(resolution,[],[f593856,f3])).
fof(f3,axiom,(
  ~is_a_theorem(equivalent(a,equivalent(equivalent(b,c),equivalent(equivalent(a,c),b))))),
  file('./p/LCL021-1.tptp',unknown)).
fof(f593856,plain,(
  ( ! [X389,X391,X390] : (is_a_theorem(equivalent(X389,equivalent(equivalent(X390,X391),equivalent(equivalent(X389,X391),X390))))) ) | (~spl0_3 | ~spl0_100 | ~spl0_276 | ~spl0_289 | ~spl0_297 | ~spl0_325 | ~spl0_330 | ~spl0_374)),
  inference(resolution,[],[f357346,f120393])).
fof(f120393,plain,(
  ( ! [X151,X149,X150] : (~is_a_theorem(equivalent(equivalent(X149,equivalent(X150,X149)),X151)) | is_a_theorem(equivalent(X150,X151))) ) | (~spl0_100 | ~spl0_276 | ~spl0_289 | ~spl0_325 | ~spl0_330)),
  inference(resolution,[],[f50406,f50368])).
fof(f50368,plain,(
  ( ! [X61,X62,X60] : (~is_a_theorem(equivalent(X60,equivalent(X61,X62))) | ~is_a_theorem(X62) | is_a_theorem(equivalent(X60,X61))) ) | (~spl0_100 | ~spl0_276 | ~spl0_289 | ~spl0_325 | ~spl0_330)),
  inference(resolution,[],[f13326,f2207])).
fof(f2207,plain,(
  ( ! [X47,X48,X49] : (is_a_theorem(equivalent(X49,equivalent(equivalent(X49,X48),X47))) | ~is_a_theorem(X48) | ~is_a_theorem(X47)) ) | ~spl0_100),
  inference(resolution,[],[f2141,f82])).
fof(f82,plain,(
  ( ! [X2,X0,X3,X1] : (~is_a_theorem(equivalent(equivalent(X2,equivalent(X0,X3)),X1)) | ~is_a_theorem(X3) | is_a_theorem(equivalent(X0,equivalent(X1,X2)))) )),
  inference(resolution,[],[f13,f4])).
fof(f4,plain,(
  ( ! [X2,X0,X1] : (is_a_theorem(equivalent(equivalent(X1,X2),equivalent(equivalent(X2,X0),X1))) | ~is_a_theorem(X0)) )),
  inference(resolution,[],[f2,f1])).
fof(f1,axiom,(
  ( ! [X0,X1] : (~is_a_theorem(equivalent(X0,X1)) | ~is_a_theorem(X0) | is_a_theorem(X1)) )),
  file('./p/LCL021-1.tptp',unknown)).
fof(f2,axiom,(
  ( ! [X2,X0,X1] : (is_a_theorem(equivalent(X0,equivalent(equivalent(X1,X2),equivalent(equivalent(X2,X0),X1))))) )),
  file('./p/LCL021-1.tptp',unknown)).
fof(f13,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~is_a_theorem(equivalent(X4,equivalent(equivalent(equivalent(X5,X6),equivalent(equivalent(X6,equivalent(X7,X8)),X5)),X7))) | is_a_theorem(X4) | ~is_a_theorem(X8)) )),
  inference(resolution,[],[f8,f2])).
fof(f8,plain,(
  ( ! [X6,X4,X7,X5] : (~is_a_theorem(equivalent(equivalent(X4,X5),X6)) | ~is_a_theorem(equivalent(X7,equivalent(X6,X4))) | is_a_theorem(X7) | ~is_a_theorem(X5)) )),
  inference(resolution,[],[f6,f4])).
fof(f6,plain,(
  ( ! [X2,X0,X1] : (~is_a_theorem(equivalent(X1,X2)) | ~is_a_theorem(X2) | ~is_a_theorem(equivalent(X0,X1)) | is_a_theorem(X0)) )),
  inference(resolution,[],[f5,f1])).
fof(f5,plain,(
  ( ! [X2,X0,X1] : (is_a_theorem(equivalent(equivalent(X2,X0),X1)) | ~is_a_theorem(equivalent(X1,X2)) | ~is_a_theorem(X0)) )),
  inference(resolution,[],[f4,f1])).
fof(f2141,plain,(
  ( ! [X21,X19] : (is_a_theorem(equivalent(equivalent(X19,X21),X21)) | ~is_a_theorem(X19)) ) | ~spl0_100),
  inference(avatar_component_clause,[],[f2140])).
fof(f2140,plain,(
  spl0_100 <=> ! [X19,X21] : (~is_a_theorem(X19) | is_a_theorem(equivalent(equivalent(X19,X21),X21)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_100])])).
fof(f13326,plain,(
  ( ! [X26,X27,X25] : (~is_a_theorem(equivalent(X25,equivalent(X26,equivalent(X27,X26)))) | is_a_theorem(equivalent(X27,X25))) ) | (~spl0_276 | ~spl0_289 | ~spl0_325 | ~spl0_330)),
  inference(resolution,[],[f10645,f6715])).
fof(f6715,plain,(
  ( ! [X2,X1] : (is_a_theorem(equivalent(equivalent(X1,equivalent(X2,X1)),X2))) ) | (~spl0_276 | ~spl0_289)),
  inference(resolution,[],[f6243,f2])).
fof(f6243,plain,(
  ( ! [X114,X115] : (~is_a_theorem(equivalent(X115,equivalent(X115,X114))) | is_a_theorem(X114)) ) | (~spl0_276 | ~spl0_289)),
  inference(resolution,[],[f6161,f5444])).
fof(f5444,plain,(
  ( ! [X2,X0,X1] : (~is_a_theorem(equivalent(X0,equivalent(X1,X2))) | is_a_theorem(X2) | ~is_a_theorem(equivalent(X1,X0))) ) | ~spl0_276),
  inference(resolution,[],[f5400,f35])).
fof(f35,plain,(
  ( ! [X6,X4,X8,X7,X5] : (~is_a_theorem(equivalent(X5,equivalent(equivalent(equivalent(X6,X7),equivalent(equivalent(X7,equivalent(X8,X4)),X6)),X8))) | is_a_theorem(X4) | ~is_a_theorem(X5)) )),
  inference(resolution,[],[f19,f2])).
fof(f19,plain,(
  ( ! [X2,X0,X3,X1] : (~is_a_theorem(equivalent(equivalent(X3,X1),X2)) | is_a_theorem(X1) | ~is_a_theorem(equivalent(X0,equivalent(X2,X3))) | ~is_a_theorem(X0)) )),
  inference(resolution,[],[f9,f2])).
fof(f9,plain,(
  ( ! [X10,X8,X11,X9] : (~is_a_theorem(equivalent(X9,equivalent(X10,X11))) | ~is_a_theorem(X8) | is_a_theorem(X9) | ~is_a_theorem(equivalent(X8,X10)) | ~is_a_theorem(X11)) )),
  inference(resolution,[],[f6,f5])).
fof(f5400,plain,(
  ( ! [X6,X4,X7] : (is_a_theorem(equivalent(X6,equivalent(equivalent(X6,equivalent(X4,X7)),X7))) | ~is_a_theorem(X4)) ) | ~spl0_276),
  inference(avatar_component_clause,[],[f5399])).
fof(f5399,plain,(
  spl0_276 <=> ! [X7,X4,X6] : (~is_a_theorem(X4) | is_a_theorem(equivalent(X6,equivalent(equivalent(X6,equivalent(X4,X7)),X7))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_276])])).
fof(f6161,plain,(
  ( ! [X14] : (is_a_theorem(equivalent(X14,X14))) ) | ~spl0_289),
  inference(avatar_component_clause,[],[f6160])).
fof(f6160,plain,(
  spl0_289 <=> ! [X14] : is_a_theorem(equivalent(X14,X14))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_289])])).
fof(f10645,plain,(
  ( ! [X43,X41,X42] : (~is_a_theorem(equivalent(X43,X41)) | ~is_a_theorem(equivalent(X42,X43)) | is_a_theorem(equivalent(X41,X42))) ) | (~spl0_325 | ~spl0_330)),
  inference(resolution,[],[f10099,f9326])).
fof(f9326,plain,(
  ( ! [X37,X38,X36] : (is_a_theorem(equivalent(X36,equivalent(X38,X37))) | ~is_a_theorem(X38) | ~is_a_theorem(equivalent(X36,X37))) ) | ~spl0_325),
  inference(avatar_component_clause,[],[f9325])).
fof(f9325,plain,(
  spl0_325 <=> ! [X36,X38,X37] : (~is_a_theorem(equivalent(X36,X37)) | ~is_a_theorem(X38) | is_a_theorem(equivalent(X36,equivalent(X38,X37))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_325])])).
fof(f10099,plain,(
  ( ! [X39,X41,X40] : (~is_a_theorem(equivalent(X39,equivalent(equivalent(X40,X39),X41))) | is_a_theorem(equivalent(X41,X40))) ) | ~spl0_330),
  inference(resolution,[],[f10076,f2])).
fof(f10076,plain,(
  ( ! [X14,X15,X16] : (~is_a_theorem(equivalent(X14,equivalent(X15,X16))) | ~is_a_theorem(equivalent(X14,X16)) | is_a_theorem(X15)) ) | ~spl0_330),
  inference(avatar_component_clause,[],[f10075])).
fof(f10075,plain,(
  spl0_330 <=> ! [X16,X15,X14] : (~is_a_theorem(equivalent(X14,equivalent(X15,X16))) | ~is_a_theorem(equivalent(X14,X16)) | is_a_theorem(X15))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_330])])).
fof(f50406,plain,(
  ( ! [X210,X208,X209] : (is_a_theorem(equivalent(X208,equivalent(X209,equivalent(equivalent(X210,equivalent(X208,X210)),X209))))) ) | (~spl0_276 | ~spl0_289 | ~spl0_325 | ~spl0_330)),
  inference(resolution,[],[f13326,f6715])).
fof(f357346,plain,(
  ( ! [X401,X400,X402] : (is_a_theorem(equivalent(equivalent(X400,X401),equivalent(equivalent(X402,X400),equivalent(X401,X402))))) ) | (~spl0_3 | ~spl0_276 | ~spl0_297 | ~spl0_325 | ~spl0_330 | ~spl0_374)),
  inference(resolution,[],[f205308,f59504])).
fof(f59504,plain,(
  ( ! [X140,X141,X139] : (~is_a_theorem(equivalent(X139,equivalent(X140,X141))) | is_a_theorem(equivalent(equivalent(X141,X139),X140))) ) | (~spl0_3 | ~spl0_325 | ~spl0_330 | ~spl0_374)),
  inference(resolution,[],[f31732,f2])).
fof(f31732,plain,(
  ( ! [X165,X163,X164] : (~is_a_theorem(equivalent(X163,equivalent(X164,X165))) | ~is_a_theorem(equivalent(X163,X164)) | is_a_theorem(X165)) ) | (~spl0_3 | ~spl0_325 | ~spl0_330 | ~spl0_374)),
  inference(resolution,[],[f13967,f168])).
fof(f168,plain,(
  ( ! [X37,X35,X36] : (~is_a_theorem(equivalent(equivalent(X36,X37),X35)) | ~is_a_theorem(equivalent(X35,X36)) | is_a_theorem(X37)) ) | ~spl0_3),
  inference(resolution,[],[f62,f16])).
fof(f16,plain,(
  ( ! [X2,X0,X1] : (~is_a_theorem(equivalent(equivalent(X1,X2),equivalent(equivalent(X2,X0),X1))) | is_a_theorem(X0)) )),
  inference(resolution,[],[f10,f2])).
fof(f10,plain,(
  ( ! [X0,X1] : (~is_a_theorem(equivalent(X0,X1)) | is_a_theorem(X0) | ~is_a_theorem(X1)) )),
  inference(resolution,[],[f7,f4])).
fof(f7,plain,(
  ( ! [X2,X0,X3,X1] : (~is_a_theorem(equivalent(equivalent(X0,X1),equivalent(equivalent(X1,X2),X0))) | ~is_a_theorem(equivalent(X3,X2)) | is_a_theorem(X3)) )),
  inference(resolution,[],[f6,f2])).
fof(f62,plain,(
  ( ! [X19,X20] : (is_a_theorem(equivalent(X19,X20)) | ~is_a_theorem(X19) | ~is_a_theorem(X20)) ) | ~spl0_3),
  inference(avatar_component_clause,[],[f61])).
fof(f61,plain,(
  spl0_3 <=> ! [X20,X19] : (~is_a_theorem(X19) | is_a_theorem(equivalent(X19,X20)) | ~is_a_theorem(X20))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f13967,plain,(
  ( ! [X229,X230,X228] : (is_a_theorem(equivalent(equivalent(equivalent(X229,X230),X230),X228)) | ~is_a_theorem(equivalent(X228,X229))) ) | (~spl0_325 | ~spl0_330 | ~spl0_374)),
  inference(resolution,[],[f13773,f10645])).
fof(f13773,plain,(
  ( ! [X64,X63] : (is_a_theorem(equivalent(X63,equivalent(equivalent(X63,X64),X64)))) ) | ~spl0_374),
  inference(avatar_component_clause,[],[f13772])).
fof(f13772,plain,(
  spl0_374 <=> ! [X63,X64] : is_a_theorem(equivalent(X63,equivalent(equivalent(X63,X64),X64)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_374])])).
fof(f205308,plain,(
  ( ! [X802,X804,X803] : (is_a_theorem(equivalent(X802,equivalent(equivalent(equivalent(X803,X804),equivalent(X802,X803)),X804)))) ) | (~spl0_276 | ~spl0_297 | ~spl0_325 | ~spl0_330 | ~spl0_374)),
  inference(resolution,[],[f79440,f14183])).
fof(f14183,plain,(
  ( ! [X294,X293,X295] : (~is_a_theorem(equivalent(X293,equivalent(equivalent(X294,X295),X295))) | is_a_theorem(equivalent(X294,X293))) ) | (~spl0_297 | ~spl0_325 | ~spl0_330 | ~spl0_374)),
  inference(resolution,[],[f13966,f10645])).
fof(f13966,plain,(
  ( ! [X227,X226] : (is_a_theorem(equivalent(equivalent(equivalent(X226,X227),X227),X226))) ) | (~spl0_297 | ~spl0_330 | ~spl0_374)),
  inference(resolution,[],[f13773,f10643])).
fof(f10643,plain,(
  ( ! [X37,X38] : (~is_a_theorem(equivalent(X38,X37)) | is_a_theorem(equivalent(X37,X38))) ) | (~spl0_297 | ~spl0_330)),
  inference(resolution,[],[f10099,f6924])).
fof(f6924,plain,(
  ( ! [X24,X25] : (is_a_theorem(equivalent(X24,equivalent(X25,X24))) | ~is_a_theorem(X25)) ) | ~spl0_297),
  inference(avatar_component_clause,[],[f6923])).
fof(f6923,plain,(
  spl0_297 <=> ! [X25,X24] : (is_a_theorem(equivalent(X24,equivalent(X25,X24))) | ~is_a_theorem(X25))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_297])])).
fof(f79440,plain,(
  ( ! [X344,X343,X345] : (is_a_theorem(equivalent(equivalent(equivalent(equivalent(X343,X344),X345),X344),equivalent(X345,X343)))) ) | (~spl0_276 | ~spl0_297 | ~spl0_330)),
  inference(resolution,[],[f10831,f5806])).
fof(f5806,plain,(
  ( ! [X2,X0,X1] : (~is_a_theorem(equivalent(equivalent(X2,X0),X1)) | is_a_theorem(equivalent(equivalent(X0,X1),X2))) ) | ~spl0_276),
  inference(resolution,[],[f5444,f2])).
fof(f10831,plain,(
  ( ! [X167,X166,X168] : (is_a_theorem(equivalent(equivalent(equivalent(X166,X167),equivalent(equivalent(X167,X168),X166)),X168))) ) | (~spl0_297 | ~spl0_330)),
  inference(resolution,[],[f10643,f2])).
fof(f171626,plain,(
  ~spl0_98),
  inference(avatar_contradiction_clause,[],[f171520])).
fof(f171520,plain,(
  $false | ~spl0_98),
  inference(resolution,[],[f2019,f2])).
fof(f2019,plain,(
  ( ! [X17,X16] : (~is_a_theorem(equivalent(X16,X17))) ) | ~spl0_98),
  inference(avatar_component_clause,[],[f2018])).
fof(f2018,plain,(
  spl0_98 <=> ! [X16,X17] : ~is_a_theorem(equivalent(X16,X17))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_98])])).
fof(f13774,plain,(
  spl0_374 | spl0_1 | ~spl0_330),
  inference(avatar_split_clause,[],[f13758,f10075,f54,f13772])).
fof(f54,plain,(
  spl0_1 <=> ! [X10] : ~is_a_theorem(X10)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f13758,plain,(
  ( ! [X64,X62,X63] : (~is_a_theorem(X62) | is_a_theorem(equivalent(X63,equivalent(equivalent(X63,X64),X64)))) ) | ~spl0_330),
  inference(duplicate_literal_removal,[],[f13623])).
fof(f13623,plain,(
  ( ! [X64,X62,X63] : (~is_a_theorem(X62) | is_a_theorem(equivalent(X63,equivalent(equivalent(X63,X64),X64))) | ~is_a_theorem(X62)) ) | ~spl0_330),
  inference(resolution,[],[f10649,f17])).
fof(f17,plain,(
  ( ! [X4,X5,X3] : (~is_a_theorem(equivalent(equivalent(X4,X5),X3)) | is_a_theorem(equivalent(X3,X4)) | ~is_a_theorem(X5)) )),
  inference(resolution,[],[f10,f4])).
fof(f10649,plain,(
  ( ! [X54,X56,X55] : (is_a_theorem(equivalent(equivalent(equivalent(equivalent(X54,X55),X55),X56),X54)) | ~is_a_theorem(X56)) ) | ~spl0_330),
  inference(resolution,[],[f10099,f8822])).
fof(f8822,plain,(
  ( ! [X24,X23,X22] : (is_a_theorem(equivalent(X22,equivalent(X23,equivalent(equivalent(X23,X22),X24)))) | ~is_a_theorem(X24)) )),
  inference(resolution,[],[f943,f2])).
fof(f943,plain,(
  ( ! [X14,X12,X15,X13] : (~is_a_theorem(equivalent(equivalent(X12,equivalent(X13,X14)),equivalent(X13,equivalent(X15,X12)))) | is_a_theorem(X15) | ~is_a_theorem(X14)) )),
  inference(resolution,[],[f300,f26])).
fof(f26,plain,(
  ( ! [X6,X4,X5,X3] : (is_a_theorem(equivalent(equivalent(equivalent(X3,X4),equivalent(equivalent(X4,equivalent(X5,X6)),X3)),X5)) | ~is_a_theorem(X6)) )),
  inference(resolution,[],[f17,f2])).
fof(f300,plain,(
  ( ! [X2,X0,X3,X1] : (~is_a_theorem(equivalent(equivalent(X3,equivalent(X1,X0)),X2)) | ~is_a_theorem(equivalent(X1,equivalent(X2,X3))) | is_a_theorem(X0)) )),
  inference(resolution,[],[f35,f4])).
fof(f10077,plain,(
  spl0_330 | spl0_1 | ~spl0_19),
  inference(avatar_split_clause,[],[f10072,f309,f54,f10075])).
fof(f309,plain,(
  spl0_19 <=> ! [X7,X8,X4,X6] : (is_a_theorem(X4) | ~is_a_theorem(X8) | ~is_a_theorem(equivalent(equivalent(X6,X7),equivalent(equivalent(X7,equivalent(X8,X4)),X6))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_19])])).
fof(f10072,plain,(
  ( ! [X14,X15,X13,X16] : (~is_a_theorem(X13) | ~is_a_theorem(equivalent(X14,equivalent(X15,X16))) | is_a_theorem(X15) | ~is_a_theorem(equivalent(X14,X16))) ) | ~spl0_19),
  inference(duplicate_literal_removal,[],[f10012])).
fof(f10012,plain,(
  ( ! [X14,X15,X13,X16] : (~is_a_theorem(X13) | ~is_a_theorem(equivalent(X14,equivalent(X15,X16))) | is_a_theorem(X15) | ~is_a_theorem(X13) | ~is_a_theorem(equivalent(X14,X16))) ) | ~spl0_19),
  inference(resolution,[],[f1175,f9128])).
fof(f9128,plain,(
  ( ! [X33,X31,X32] : (is_a_theorem(equivalent(X31,equivalent(X32,X33))) | ~is_a_theorem(X33) | ~is_a_theorem(equivalent(X32,X31))) )),
  inference(resolution,[],[f945,f2])).
fof(f945,plain,(
  ( ! [X23,X21,X22,X20] : (~is_a_theorem(equivalent(equivalent(X20,X21),equivalent(X22,equivalent(X23,X20)))) | is_a_theorem(X23) | ~is_a_theorem(X21) | ~is_a_theorem(X22)) )),
  inference(resolution,[],[f300,f31])).
fof(f31,plain,(
  ( ! [X2,X0,X3,X1] : (is_a_theorem(equivalent(equivalent(equivalent(X1,X2),equivalent(equivalent(X2,X0),X1)),X3)) | ~is_a_theorem(X0) | ~is_a_theorem(X3)) )),
  inference(resolution,[],[f18,f2])).
fof(f18,plain,(
  ( ! [X6,X8,X7] : (~is_a_theorem(equivalent(X8,X6)) | ~is_a_theorem(X8) | is_a_theorem(equivalent(X6,X7)) | ~is_a_theorem(X7)) )),
  inference(resolution,[],[f10,f5])).
fof(f1175,plain,(
  ( ! [X206,X204,X207,X205] : (~is_a_theorem(equivalent(X206,equivalent(X204,X207))) | ~is_a_theorem(X207) | ~is_a_theorem(equivalent(X204,equivalent(X205,X206))) | is_a_theorem(X205)) ) | ~spl0_19),
  inference(resolution,[],[f154,f310])).
fof(f310,plain,(
  ( ! [X6,X4,X8,X7] : (~is_a_theorem(equivalent(equivalent(X6,X7),equivalent(equivalent(X7,equivalent(X8,X4)),X6))) | ~is_a_theorem(X8) | is_a_theorem(X4)) ) | ~spl0_19),
  inference(avatar_component_clause,[],[f309])).
fof(f154,plain,(
  ( ! [X6,X4,X8,X7,X5] : (is_a_theorem(equivalent(X5,equivalent(equivalent(equivalent(X6,X7),equivalent(equivalent(X7,equivalent(X8,X4)),X6)),X8))) | ~is_a_theorem(X5) | ~is_a_theorem(X4)) )),
  inference(resolution,[],[f77,f2])).
fof(f77,plain,(
  ( ! [X2,X0,X3,X1] : (~is_a_theorem(equivalent(equivalent(X2,X3),X1)) | ~is_a_theorem(X3) | ~is_a_theorem(X0) | is_a_theorem(equivalent(X0,equivalent(X1,X2)))) )),
  inference(resolution,[],[f20,f2])).
fof(f20,plain,(
  ( ! [X6,X4,X7,X5] : (~is_a_theorem(equivalent(X4,equivalent(X6,X7))) | is_a_theorem(equivalent(X5,X6)) | ~is_a_theorem(X4) | ~is_a_theorem(X5) | ~is_a_theorem(X7)) )),
  inference(resolution,[],[f9,f4])).
fof(f9327,plain,(
  spl0_325 | spl0_1 | ~spl0_3),
  inference(avatar_split_clause,[],[f9315,f61,f54,f9325])).
fof(f9315,plain,(
  ( ! [X37,X35,X38,X36] : (~is_a_theorem(X35) | ~is_a_theorem(equivalent(X36,X37)) | is_a_theorem(equivalent(X36,equivalent(X38,X37))) | ~is_a_theorem(X38)) ) | ~spl0_3),
  inference(duplicate_literal_removal,[],[f9186])).
fof(f9186,plain,(
  ( ! [X37,X35,X38,X36] : (~is_a_theorem(X35) | ~is_a_theorem(equivalent(X36,X37)) | is_a_theorem(equivalent(X36,equivalent(X38,X37))) | ~is_a_theorem(X35) | ~is_a_theorem(X38)) ) | ~spl0_3),
  inference(resolution,[],[f9128,f211])).
fof(f211,plain,(
  ( ! [X21,X19,X20,X18] : (~is_a_theorem(equivalent(X21,equivalent(X19,X18))) | is_a_theorem(equivalent(X19,equivalent(X20,X21))) | ~is_a_theorem(X18) | ~is_a_theorem(X20)) ) | ~spl0_3),
  inference(resolution,[],[f82,f62])).
fof(f6925,plain,(
  spl0_297 | spl0_1 | ~spl0_3 | ~spl0_276 | ~spl0_289),
  inference(avatar_split_clause,[],[f6910,f6160,f5399,f61,f54,f6923])).
fof(f6910,plain,(
  ( ! [X24,X23,X25] : (~is_a_theorem(X23) | is_a_theorem(equivalent(X24,equivalent(X25,X24))) | ~is_a_theorem(X25)) ) | (~spl0_3 | ~spl0_276 | ~spl0_289)),
  inference(duplicate_literal_removal,[],[f6809])).
fof(f6809,plain,(
  ( ! [X24,X23,X25] : (~is_a_theorem(X23) | is_a_theorem(equivalent(X24,equivalent(X25,X24))) | ~is_a_theorem(X23) | ~is_a_theorem(X25)) ) | (~spl0_3 | ~spl0_276 | ~spl0_289)),
  inference(resolution,[],[f6244,f211])).
fof(f6244,plain,(
  ( ! [X116,X117] : (is_a_theorem(equivalent(X116,equivalent(X116,X117))) | ~is_a_theorem(X117)) ) | (~spl0_276 | ~spl0_289)),
  inference(resolution,[],[f6161,f5445])).
fof(f5445,plain,(
  ( ! [X4,X5,X3] : (~is_a_theorem(equivalent(X3,equivalent(X4,X5))) | is_a_theorem(equivalent(X4,X3)) | ~is_a_theorem(X5)) ) | ~spl0_276),
  inference(resolution,[],[f5400,f13])).
fof(f6162,plain,(
  spl0_289 | spl0_1 | ~spl0_287),
  inference(avatar_split_clause,[],[f6152,f6015,f54,f6160])).
fof(f6015,plain,(
  spl0_287 <=> ! [X46,X45] : (is_a_theorem(equivalent(equivalent(X45,X46),X45)) | ~is_a_theorem(X46))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_287])])).
fof(f6152,plain,(
  ( ! [X14,X13] : (~is_a_theorem(X13) | is_a_theorem(equivalent(X14,X14))) ) | ~spl0_287),
  inference(duplicate_literal_removal,[],[f6052])).
fof(f6052,plain,(
  ( ! [X14,X13] : (~is_a_theorem(X13) | is_a_theorem(equivalent(X14,X14)) | ~is_a_theorem(X13)) ) | ~spl0_287),
  inference(resolution,[],[f6016,f17])).
fof(f6016,plain,(
  ( ! [X45,X46] : (is_a_theorem(equivalent(equivalent(X45,X46),X45)) | ~is_a_theorem(X46)) ) | ~spl0_287),
  inference(avatar_component_clause,[],[f6015])).
fof(f6017,plain,(
  spl0_20 | spl0_287 | ~spl0_276),
  inference(avatar_split_clause,[],[f6013,f5399,f6015,f314])).
fof(f314,plain,(
  spl0_20 <=> ! [X11,X10] : (~is_a_theorem(equivalent(X10,X11)) | ~is_a_theorem(X11))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_20])])).
fof(f6013,plain,(
  ( ! [X47,X45,X48,X46] : (is_a_theorem(equivalent(equivalent(X45,X46),X45)) | ~is_a_theorem(equivalent(X47,X48)) | ~is_a_theorem(X48) | ~is_a_theorem(X46)) ) | ~spl0_276),
  inference(subsumption_resolution,[],[f5979,f10])).
fof(f5979,plain,(
  ( ! [X47,X45,X48,X46] : (is_a_theorem(equivalent(equivalent(X45,X46),X45)) | ~is_a_theorem(equivalent(X47,X48)) | ~is_a_theorem(X48) | ~is_a_theorem(X46) | ~is_a_theorem(X47)) ) | ~spl0_276),
  inference(resolution,[],[f5445,f219])).
fof(f219,plain,(
  ( ! [X28,X26,X27,X25] : (is_a_theorem(equivalent(X28,equivalent(equivalent(X28,X27),equivalent(X25,X26)))) | ~is_a_theorem(X26) | ~is_a_theorem(X27) | ~is_a_theorem(X25)) )),
  inference(resolution,[],[f25,f82])).
fof(f25,plain,(
  ( ! [X2,X0,X1] : (is_a_theorem(equivalent(equivalent(equivalent(X0,X1),X2),X2)) | ~is_a_theorem(X0) | ~is_a_theorem(X1)) )),
  inference(resolution,[],[f17,f4])).
fof(f5401,plain,(
  spl0_1 | spl0_276),
  inference(avatar_split_clause,[],[f5397,f5399,f54])).
fof(f5397,plain,(
  ( ! [X6,X4,X7,X5] : (~is_a_theorem(X4) | ~is_a_theorem(X5) | is_a_theorem(equivalent(X6,equivalent(equivalent(X6,equivalent(X4,X7)),X7)))) )),
  inference(duplicate_literal_removal,[],[f5308])).
fof(f5308,plain,(
  ( ! [X6,X4,X7,X5] : (~is_a_theorem(X4) | ~is_a_theorem(X5) | ~is_a_theorem(X5) | is_a_theorem(equivalent(X6,equivalent(equivalent(X6,equivalent(X4,X7)),X7)))) )),
  inference(resolution,[],[f423,f82])).
fof(f423,plain,(
  ( ! [X2,X0,X3,X1] : (is_a_theorem(equivalent(equivalent(X2,equivalent(X3,X0)),equivalent(X3,equivalent(X1,X2)))) | ~is_a_theorem(X1) | ~is_a_theorem(X0)) )),
  inference(resolution,[],[f26,f82])).
fof(f2142,plain,(
  spl0_1 | spl0_100 | ~spl0_96),
  inference(avatar_split_clause,[],[f2131,f1965,f2140,f54])).
fof(f1965,plain,(
  spl0_96 <=> ! [X82,X83,X80] : (~is_a_theorem(X80) | ~is_a_theorem(X83) | is_a_theorem(equivalent(equivalent(X82,X80),equivalent(X83,X82))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_96])])).
fof(f2131,plain,(
  ( ! [X21,X19,X20] : (~is_a_theorem(X19) | ~is_a_theorem(X20) | is_a_theorem(equivalent(equivalent(X19,X21),X21))) ) | ~spl0_96),
  inference(duplicate_literal_removal,[],[f2055])).
fof(f2055,plain,(
  ( ! [X21,X19,X20] : (~is_a_theorem(X19) | ~is_a_theorem(X20) | is_a_theorem(equivalent(equivalent(X19,X21),X21)) | ~is_a_theorem(X20)) ) | ~spl0_96),
  inference(resolution,[],[f1966,f17])).
fof(f1966,plain,(
  ( ! [X80,X83,X82] : (is_a_theorem(equivalent(equivalent(X82,X80),equivalent(X83,X82))) | ~is_a_theorem(X83) | ~is_a_theorem(X80)) ) | ~spl0_96),
  inference(avatar_component_clause,[],[f1965])).
fof(f2020,plain,(
  spl0_1 | spl0_98 | spl0_1 | ~spl0_20),
  inference(avatar_split_clause,[],[f2004,f314,f54,f2018,f54])).
fof(f2004,plain,(
  ( ! [X17,X15,X18,X16] : (~is_a_theorem(X15) | ~is_a_theorem(equivalent(X16,X17)) | ~is_a_theorem(X18)) ) | ~spl0_20),
  inference(duplicate_literal_removal,[],[f1991])).
fof(f1991,plain,(
  ( ! [X17,X15,X18,X16] : (~is_a_theorem(X15) | ~is_a_theorem(equivalent(X16,X17)) | ~is_a_theorem(X15) | ~is_a_theorem(X18)) ) | ~spl0_20),
  inference(resolution,[],[f315,f32])).
fof(f32,plain,(
  ( ! [X6,X4,X7,X5] : (is_a_theorem(equivalent(equivalent(equivalent(X5,X6),X4),X7)) | ~is_a_theorem(equivalent(X4,X5)) | ~is_a_theorem(X7) | ~is_a_theorem(X6)) )),
  inference(resolution,[],[f18,f4])).
fof(f315,plain,(
  ( ! [X10,X11] : (~is_a_theorem(equivalent(X10,X11)) | ~is_a_theorem(X11)) ) | ~spl0_20),
  inference(avatar_component_clause,[],[f314])).
fof(f1967,plain,(
  spl0_96 | spl0_20 | ~spl0_3),
  inference(avatar_split_clause,[],[f1963,f61,f314,f1965])).
fof(f1963,plain,(
  ( ! [X80,X83,X81,X79,X82] : (~is_a_theorem(X79) | ~is_a_theorem(X80) | is_a_theorem(equivalent(equivalent(X82,X80),equivalent(X83,X82))) | ~is_a_theorem(equivalent(X81,X79)) | ~is_a_theorem(X83)) ) | ~spl0_3),
  inference(subsumption_resolution,[],[f1908,f10])).
fof(f1908,plain,(
  ( ! [X80,X83,X81,X79,X82] : (~is_a_theorem(X79) | ~is_a_theorem(X80) | ~is_a_theorem(X81) | is_a_theorem(equivalent(equivalent(X82,X80),equivalent(X83,X82))) | ~is_a_theorem(equivalent(X81,X79)) | ~is_a_theorem(X83)) ) | ~spl0_3),
  inference(resolution,[],[f219,f211])).
fof(f311,plain,(
  spl0_1 | spl0_19),
  inference(avatar_split_clause,[],[f307,f309,f54])).
fof(f307,plain,(
  ( ! [X6,X4,X8,X7,X5] : (is_a_theorem(X4) | ~is_a_theorem(X5) | ~is_a_theorem(equivalent(equivalent(X6,X7),equivalent(equivalent(X7,equivalent(X8,X4)),X6))) | ~is_a_theorem(X8)) )),
  inference(duplicate_literal_removal,[],[f301])).
fof(f301,plain,(
  ( ! [X6,X4,X8,X7,X5] : (is_a_theorem(X4) | ~is_a_theorem(X5) | ~is_a_theorem(equivalent(equivalent(X6,X7),equivalent(equivalent(X7,equivalent(X8,X4)),X6))) | ~is_a_theorem(X8) | ~is_a_theorem(X5)) )),
  inference(resolution,[],[f35,f29])).
fof(f29,plain,(
  ( ! [X2,X0,X1] : (is_a_theorem(equivalent(X0,equivalent(X1,X2))) | ~is_a_theorem(X1) | ~is_a_theorem(X2) | ~is_a_theorem(X0)) )),
  inference(resolution,[],[f12,f4])).
fof(f12,plain,(
  ( ! [X2,X0,X3,X1] : (~is_a_theorem(equivalent(X0,equivalent(equivalent(equivalent(X1,X2),X3),X3))) | is_a_theorem(X0) | ~is_a_theorem(X1) | ~is_a_theorem(X2)) )),
  inference(resolution,[],[f8,f4])).
fof(f135,plain,(
  ~spl0_1),
  inference(avatar_contradiction_clause,[],[f131])).
fof(f131,plain,(
  $false | ~spl0_1),
  inference(resolution,[],[f55,f2])).
fof(f55,plain,(
  ( ! [X10] : (~is_a_theorem(X10)) ) | ~spl0_1),
  inference(avatar_component_clause,[],[f54])).
fof(f125,plain,(
  spl0_1 | spl0_1 | spl0_3),
  inference(avatar_split_clause,[],[f81,f61,f54,f54])).
fof(f81,plain,(
  ( ! [X10,X8,X11,X9] : (is_a_theorem(equivalent(X8,X9)) | ~is_a_theorem(X10) | ~is_a_theorem(X8) | ~is_a_theorem(X11) | ~is_a_theorem(X9)) )),
  inference(duplicate_literal_removal,[],[f79])).
fof(f79,plain,(
  ( ! [X10,X8,X11,X9] : (is_a_theorem(equivalent(X8,X9)) | ~is_a_theorem(X10) | ~is_a_theorem(X8) | ~is_a_theorem(X11) | ~is_a_theorem(X9) | ~is_a_theorem(X11) | ~is_a_theorem(X10)) )),
  inference(resolution,[],[f20,f29])).
% SZS output end Proof for LCL021-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 422679
% Time elapsed: 58.881 s
% ------------------------------
% ------------------------------
