% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for LCL072-1
% SZS output start Proof for LCL072-1
fof(f4094,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f58,f81,f4093])).
fof(f4093,plain,(
  ~spl0_1),
  inference(avatar_contradiction_clause,[],[f4092])).
fof(f4092,plain,(
  $false | ~spl0_1),
  inference(resolution,[],[f4007,f5])).
fof(f5,axiom,(
  ~is_a_theorem(implies(a,implies(not(a),b)))),
  file('./p/LCL072-1.tptp',unknown)).
fof(f4007,plain,(
  ( ! [X8,X7] : (is_a_theorem(implies(X7,implies(not(X7),X8)))) ) | ~spl0_1),
  inference(resolution,[],[f2279,f15])).
fof(f15,plain,(
  ( ! [X4,X5,X3] : (is_a_theorem(implies(not(implies(X3,X4)),implies(not(X3),X5)))) )),
  inference(resolution,[],[f7,f3])).
fof(f3,axiom,(
  ( ! [X2,X0,X1] : (is_a_theorem(implies(implies(implies(X0,X1),X2),implies(not(X0),X2)))) )),
  file('./p/LCL072-1.tptp',unknown)).
fof(f7,plain,(
  ( ! [X2,X0,X1] : (~is_a_theorem(implies(implies(X0,X1),X2)) | is_a_theorem(implies(not(X0),X2))) )),
  inference(resolution,[],[f3,f1])).
fof(f1,axiom,(
  ( ! [X0,X1] : (~is_a_theorem(implies(X0,X1)) | ~is_a_theorem(X0) | is_a_theorem(X1)) )),
  file('./p/LCL072-1.tptp',unknown)).
fof(f2279,plain,(
  ( ! [X97,X96] : (~is_a_theorem(implies(not(implies(X96,X97)),X97)) | is_a_theorem(implies(X96,X97))) ) | ~spl0_1),
  inference(resolution,[],[f2177,f1019])).
fof(f1019,plain,(
  ( ! [X20] : (~is_a_theorem(implies(not(X20),X20)) | is_a_theorem(X20)) ) | ~spl0_1),
  inference(subsumption_resolution,[],[f1016,f828])).
fof(f828,plain,(
  ( ! [X0] : (is_a_theorem(implies(X0,X0))) ) | ~spl0_1),
  inference(resolution,[],[f760,f2])).
fof(f2,axiom,(
  ( ! [X2,X0,X1] : (is_a_theorem(implies(implies(implies(X0,X1),X2),implies(X1,X2)))) )),
  file('./p/LCL072-1.tptp',unknown)).
fof(f760,plain,(
  ( ! [X21,X19,X20] : (~is_a_theorem(implies(implies(X19,X20),implies(X20,X21))) | is_a_theorem(implies(X20,X21))) ) | ~spl0_1),
  inference(resolution,[],[f647,f2])).
fof(f647,plain,(
  ( ! [X6,X8,X7,X9] : (~is_a_theorem(implies(implies(X9,X7),X6)) | ~is_a_theorem(implies(X6,implies(X7,X8))) | is_a_theorem(implies(X7,X8))) ) | ~spl0_1),
  inference(resolution,[],[f623,f14])).
fof(f14,plain,(
  ( ! [X2,X0,X1] : (is_a_theorem(implies(not(implies(X0,X1)),implies(X1,X2)))) )),
  inference(resolution,[],[f7,f2])).
fof(f623,plain,(
  ( ! [X41,X42,X40] : (~is_a_theorem(implies(not(X42),X41)) | ~is_a_theorem(implies(X40,X41)) | ~is_a_theorem(implies(X42,X40)) | is_a_theorem(X41)) ) | ~spl0_1),
  inference(resolution,[],[f301,f1])).
fof(f301,plain,(
  ( ! [X14,X15,X16] : (is_a_theorem(implies(implies(X14,X16),X15)) | ~is_a_theorem(implies(X16,X15)) | ~is_a_theorem(implies(not(X14),X15))) ) | ~spl0_1),
  inference(resolution,[],[f54,f1])).
fof(f54,plain,(
  ( ! [X47,X48,X49] : (is_a_theorem(implies(implies(X47,X48),implies(implies(X49,X47),X48))) | ~is_a_theorem(implies(not(X49),X48))) ) | ~spl0_1),
  inference(avatar_component_clause,[],[f53])).
fof(f53,plain,(
  spl0_1 <=> ! [X49,X48,X47] : (is_a_theorem(implies(implies(X47,X48),implies(implies(X49,X47),X48))) | ~is_a_theorem(implies(not(X49),X48)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f1016,plain,(
  ( ! [X20] : (~is_a_theorem(implies(not(X20),X20)) | ~is_a_theorem(implies(X20,X20)) | is_a_theorem(X20)) ) | ~spl0_1),
  inference(resolution,[],[f955,f1])).
fof(f955,plain,(
  ( ! [X23] : (is_a_theorem(implies(implies(X23,X23),X23)) | ~is_a_theorem(implies(not(X23),X23))) ) | ~spl0_1),
  inference(resolution,[],[f830,f54])).
fof(f830,plain,(
  ( ! [X2,X3] : (~is_a_theorem(implies(X2,implies(X2,X3))) | is_a_theorem(implies(X2,X3))) ) | ~spl0_1),
  inference(resolution,[],[f760,f332])).
fof(f332,plain,(
  ( ! [X14,X15,X13,X16] : (is_a_theorem(implies(implies(implies(X16,X14),X13),implies(X14,X15))) | ~is_a_theorem(implies(X13,implies(X14,X15)))) )),
  inference(resolution,[],[f140,f1])).
fof(f140,plain,(
  ( ! [X2,X0,X3,X1] : (is_a_theorem(implies(implies(X0,implies(X1,X2)),implies(implies(implies(X3,X1),X0),implies(X1,X2))))) )),
  inference(resolution,[],[f40,f2])).
fof(f40,plain,(
  ( ! [X2,X0,X3,X1] : (~is_a_theorem(implies(implies(X0,X1),X2)) | is_a_theorem(implies(implies(X3,X2),implies(implies(X0,X3),X2)))) )),
  inference(resolution,[],[f22,f3])).
fof(f22,plain,(
  ( ! [X2,X0,X3,X1] : (~is_a_theorem(implies(X0,implies(not(X1),X2))) | ~is_a_theorem(X0) | is_a_theorem(implies(implies(X3,X2),implies(implies(X1,X3),X2)))) )),
  inference(resolution,[],[f11,f1])).
fof(f11,plain,(
  ( ! [X6,X4,X7,X5] : (is_a_theorem(implies(X4,implies(implies(X7,X6),implies(implies(X5,X7),X6)))) | ~is_a_theorem(implies(X4,implies(not(X5),X6)))) )),
  inference(resolution,[],[f4,f1])).
fof(f4,axiom,(
  ( ! [X2,X0,X3,X1] : (is_a_theorem(implies(implies(X0,implies(not(X1),X2)),implies(X0,implies(implies(X3,X2),implies(implies(X1,X3),X2)))))) )),
  file('./p/LCL072-1.tptp',unknown)).
fof(f2177,plain,(
  ( ! [X316,X318,X317] : (is_a_theorem(implies(X316,implies(X318,X317))) | ~is_a_theorem(implies(X316,X317))) ) | ~spl0_1),
  inference(resolution,[],[f904,f13])).
fof(f13,plain,(
  ( ! [X4,X2,X3] : (is_a_theorem(implies(X2,implies(X3,implies(X4,X2))))) )),
  inference(resolution,[],[f8,f6])).
fof(f6,plain,(
  ( ! [X2,X0,X1] : (~is_a_theorem(implies(implies(X0,X1),X2)) | is_a_theorem(implies(X1,X2))) )),
  inference(resolution,[],[f2,f1])).
fof(f8,plain,(
  ( ! [X0,X1] : (is_a_theorem(implies(X0,implies(X1,X0)))) )),
  inference(resolution,[],[f6,f2])).
fof(f904,plain,(
  ( ! [X2,X0,X1] : (~is_a_theorem(implies(X0,implies(X1,X2))) | ~is_a_theorem(implies(X1,X0)) | is_a_theorem(implies(X1,X2))) ) | ~spl0_1),
  inference(resolution,[],[f891,f623])).
fof(f891,plain,(
  ( ! [X24,X23] : (is_a_theorem(implies(not(X23),implies(X23,X24)))) ) | ~spl0_1),
  inference(resolution,[],[f828,f7])).
fof(f81,plain,(
  ~spl0_2),
  inference(avatar_contradiction_clause,[],[f59])).
fof(f59,plain,(
  $false | ~spl0_2),
  inference(resolution,[],[f57,f2])).
fof(f57,plain,(
  ( ! [X46] : (~is_a_theorem(X46)) ) | ~spl0_2),
  inference(avatar_component_clause,[],[f56])).
fof(f56,plain,(
  spl0_2 <=> ! [X46] : ~is_a_theorem(X46)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f58,plain,(
  spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f51,f56,f53])).
fof(f51,plain,(
  ( ! [X47,X48,X46,X49] : (~is_a_theorem(X46) | is_a_theorem(implies(implies(X47,X48),implies(implies(X49,X47),X48))) | ~is_a_theorem(implies(not(X49),X48))) )),
  inference(resolution,[],[f22,f12])).
fof(f12,plain,(
  ( ! [X0,X1] : (is_a_theorem(implies(X1,X0)) | ~is_a_theorem(X0)) )),
  inference(resolution,[],[f8,f1])).
% SZS output end Proof for LCL072-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 8059
% Time elapsed: 0.104 s
% ------------------------------
% ------------------------------
