% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for LCL099-1
% SZS output start Proof for LCL099-1
fof(f899588,plain,(
  $false),
  inference(subsumption_resolution,[],[f899024,f300681])).
fof(f300681,plain,(
  ( ! [X144,X142,X143,X141] : (is_a_theorem(equivalent(equivalent(equivalent(X141,equivalent(equivalent(X142,X143),equivalent(equivalent(X143,X142),X141))),X144),X144))) )),
  inference(resolution,[],[f300642,f7135])).
fof(f7135,plain,(
  ( ! [X78,X76,X77,X75] : (~is_a_theorem(equivalent(X75,equivalent(equivalent(X76,X77),equivalent(equivalent(X77,X76),X78)))) | is_a_theorem(equivalent(X75,X78))) )),
  inference(resolution,[],[f294,f3829])).
fof(f3829,plain,(
  ( ! [X47,X48,X46] : (is_a_theorem(equivalent(equivalent(equivalent(X46,X47),equivalent(equivalent(X47,X46),X48)),X48))) )),
  inference(resolution,[],[f1986,f398])).
fof(f398,plain,(
  ( ! [X68,X69,X67] : (~is_a_theorem(equivalent(X67,equivalent(equivalent(X68,X68),X69))) | is_a_theorem(equivalent(X67,X69))) )),
  inference(resolution,[],[f67,f289])).
fof(f289,plain,(
  ( ! [X12,X11] : (is_a_theorem(equivalent(equivalent(equivalent(X11,X11),X12),X12))) )),
  inference(resolution,[],[f58,f42])).
fof(f42,plain,(
  ( ! [X2,X0,X3,X1] : (is_a_theorem(equivalent(equivalent(equivalent(X0,X1),X2),equivalent(equivalent(equivalent(X3,X0),equivalent(X3,X1)),X2)))) )),
  inference(resolution,[],[f18,f2])).
fof(f2,axiom,(
  ( ! [X2,X0,X3,X1] : (is_a_theorem(equivalent(equivalent(equivalent(equivalent(equivalent(X0,X1),equivalent(X0,X2)),equivalent(X1,X2)),X3),X3))) )),
  file('./p/LCL099-1.tptp',unknown)).
fof(f18,plain,(
  ( ! [X6,X8,X7,X9] : (~is_a_theorem(equivalent(equivalent(equivalent(X6,X7),equivalent(X6,X8)),X9)) | is_a_theorem(equivalent(equivalent(X7,X8),X9))) )),
  inference(resolution,[],[f10,f1])).
fof(f1,axiom,(
  ( ! [X0,X1] : (~is_a_theorem(equivalent(X0,X1)) | ~is_a_theorem(X0) | is_a_theorem(X1)) )),
  file('./p/LCL099-1.tptp',unknown)).
fof(f10,plain,(
  ( ! [X6,X4,X5,X3] : (is_a_theorem(equivalent(equivalent(equivalent(equivalent(X3,X4),equivalent(X3,X5)),X6),equivalent(equivalent(X4,X5),X6)))) )),
  inference(resolution,[],[f7,f6])).
fof(f6,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~is_a_theorem(equivalent(equivalent(equivalent(equivalent(equivalent(X0,X1),equivalent(X0,X2)),X3),equivalent(equivalent(X1,X2),X3)),X4)) | is_a_theorem(X4)) )),
  inference(resolution,[],[f3,f1])).
fof(f3,axiom,(
  ( ! [X4,X2,X0,X3,X1] : (is_a_theorem(equivalent(equivalent(equivalent(equivalent(equivalent(equivalent(X0,X1),equivalent(X0,X2)),X3),equivalent(equivalent(X1,X2),X3)),X4),X4))) )),
  file('./p/LCL099-1.tptp',unknown)).
fof(f7,plain,(
  ( ! [X0,X1] : (is_a_theorem(equivalent(equivalent(X0,X1),equivalent(X0,X1)))) )),
  inference(resolution,[],[f6,f2])).
fof(f58,plain,(
  ( ! [X2,X0,X3,X1] : (~is_a_theorem(equivalent(equivalent(X0,X1),equivalent(equivalent(X2,X0),equivalent(X2,X3)))) | is_a_theorem(equivalent(X1,X3))) )),
  inference(resolution,[],[f23,f1])).
fof(f23,plain,(
  ( ! [X2,X0,X3,X1] : (is_a_theorem(equivalent(equivalent(equivalent(X0,X1),equivalent(equivalent(X2,X0),equivalent(X2,X3))),equivalent(X1,X3)))) )),
  inference(resolution,[],[f12,f5])).
fof(f5,plain,(
  ( ! [X2,X0,X3,X1] : (~is_a_theorem(equivalent(equivalent(equivalent(equivalent(X0,X1),equivalent(X0,X2)),equivalent(X1,X2)),X3)) | is_a_theorem(X3)) )),
  inference(resolution,[],[f2,f1])).
fof(f12,plain,(
  ( ! [X2,X0,X3,X1] : (is_a_theorem(equivalent(equivalent(X0,equivalent(X1,X2)),equivalent(equivalent(equivalent(X1,X3),X0),equivalent(X3,X2))))) )),
  inference(resolution,[],[f8,f5])).
fof(f8,plain,(
  ( ! [X4,X2,X5,X3] : (is_a_theorem(equivalent(equivalent(equivalent(X2,X3),X4),equivalent(equivalent(X5,X3),equivalent(equivalent(X2,X5),X4))))) )),
  inference(resolution,[],[f6,f3])).
fof(f67,plain,(
  ( ! [X14,X15,X13,X16] : (~is_a_theorem(equivalent(equivalent(X14,X16),X13)) | ~is_a_theorem(equivalent(X13,equivalent(X14,X15))) | is_a_theorem(equivalent(X16,X15))) )),
  inference(resolution,[],[f25,f1])).
fof(f25,plain,(
  ( ! [X12,X10,X11,X9] : (is_a_theorem(equivalent(equivalent(equivalent(X10,X12),X9),equivalent(X12,X11))) | ~is_a_theorem(equivalent(X9,equivalent(X10,X11)))) )),
  inference(resolution,[],[f12,f1])).
fof(f1986,plain,(
  ( ! [X2,X0,X3,X1] : (is_a_theorem(equivalent(equivalent(equivalent(X0,X1),equivalent(equivalent(X2,X0),X3)),equivalent(equivalent(X2,X1),X3)))) )),
  inference(resolution,[],[f196,f2])).
fof(f196,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (~is_a_theorem(equivalent(equivalent(equivalent(X32,X33),equivalent(X32,X34)),equivalent(equivalent(X35,X36),equivalent(X35,X37)))) | is_a_theorem(equivalent(equivalent(X33,X34),equivalent(X36,X37)))) )),
  inference(resolution,[],[f37,f1])).
fof(f37,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (is_a_theorem(equivalent(equivalent(equivalent(equivalent(X0,X1),equivalent(X0,X2)),equivalent(equivalent(X3,X4),equivalent(X3,X5))),equivalent(equivalent(X1,X2),equivalent(X4,X5))))) )),
  inference(resolution,[],[f13,f5])).
fof(f13,plain,(
  ( ! [X6,X4,X8,X7,X5] : (is_a_theorem(equivalent(equivalent(X4,X5),equivalent(equivalent(equivalent(equivalent(X6,X7),equivalent(X6,X8)),X4),equivalent(equivalent(X7,X8),X5))))) )),
  inference(resolution,[],[f8,f6])).
fof(f294,plain,(
  ( ! [X24,X23,X25,X22] : (~is_a_theorem(equivalent(equivalent(X25,X24),equivalent(X25,X23))) | ~is_a_theorem(equivalent(X22,X24)) | is_a_theorem(equivalent(X22,X23))) )),
  inference(resolution,[],[f58,f97])).
fof(f97,plain,(
  ( ! [X64,X62,X63] : (is_a_theorem(equivalent(equivalent(X63,X62),X64)) | ~is_a_theorem(equivalent(X62,X63)) | ~is_a_theorem(X64)) )),
  inference(resolution,[],[f32,f33])).
fof(f33,plain,(
  ( ! [X0,X1] : (is_a_theorem(equivalent(equivalent(X1,X1),X0)) | ~is_a_theorem(X0)) )),
  inference(resolution,[],[f21,f1])).
fof(f21,plain,(
  ( ! [X10,X9] : (is_a_theorem(equivalent(X9,equivalent(equivalent(X10,X10),X9)))) )),
  inference(resolution,[],[f15,f8])).
fof(f15,plain,(
  ( ! [X2,X0,X1] : (~is_a_theorem(equivalent(equivalent(X0,X1),equivalent(X0,X2))) | is_a_theorem(equivalent(X1,X2))) )),
  inference(resolution,[],[f9,f1])).
fof(f9,plain,(
  ( ! [X2,X0,X1] : (is_a_theorem(equivalent(equivalent(equivalent(X0,X1),equivalent(X0,X2)),equivalent(X1,X2)))) )),
  inference(resolution,[],[f7,f5])).
fof(f32,plain,(
  ( ! [X14,X12,X13,X11] : (~is_a_theorem(equivalent(equivalent(X11,X12),X13)) | ~is_a_theorem(equivalent(X14,X12)) | is_a_theorem(equivalent(equivalent(X11,X14),X13))) )),
  inference(resolution,[],[f14,f1])).
fof(f14,plain,(
  ( ! [X12,X10,X11,X9] : (is_a_theorem(equivalent(equivalent(X12,X10),equivalent(equivalent(X9,X12),X11))) | ~is_a_theorem(equivalent(equivalent(X9,X10),X11))) )),
  inference(resolution,[],[f8,f1])).
fof(f300642,plain,(
  ( ! [X1531,X1533,X1532,X1534] : (is_a_theorem(equivalent(X1531,equivalent(equivalent(equivalent(equivalent(X1532,X1533),equivalent(equivalent(X1533,X1532),X1534)),X1534),X1531)))) )),
  inference(resolution,[],[f7758,f13129])).
fof(f13129,plain,(
  ( ! [X21,X19,X22,X20] : (is_a_theorem(equivalent(X19,equivalent(X20,equivalent(equivalent(equivalent(equivalent(X21,X22),X20),equivalent(X22,X21)),X19))))) )),
  inference(resolution,[],[f4724,f2702])).
fof(f2702,plain,(
  ( ! [X191,X189,X192,X190] : (~is_a_theorem(equivalent(X189,equivalent(equivalent(X190,X191),equivalent(X190,X192)))) | is_a_theorem(equivalent(X189,equivalent(X191,X192)))) )),
  inference(resolution,[],[f54,f1])).
fof(f54,plain,(
  ( ! [X2,X0,X3,X1] : (is_a_theorem(equivalent(equivalent(X0,equivalent(equivalent(X1,X2),equivalent(X1,X3))),equivalent(X0,equivalent(X2,X3))))) )),
  inference(resolution,[],[f48,f5])).
fof(f48,plain,(
  ( ! [X24,X23,X25] : (is_a_theorem(equivalent(equivalent(X23,X24),equivalent(equivalent(X25,X23),equivalent(X25,X24))))) )),
  inference(resolution,[],[f18,f7])).
fof(f4724,plain,(
  ( ! [X87,X85,X86,X84] : (is_a_theorem(equivalent(X84,equivalent(X85,equivalent(equivalent(X86,X87),equivalent(equivalent(X85,equivalent(X87,X86)),X84)))))) )),
  inference(resolution,[],[f3308,f205])).
fof(f205,plain,(
  ( ! [X6,X4,X8,X7,X5] : (is_a_theorem(equivalent(equivalent(equivalent(X4,equivalent(X5,X6)),X7),equivalent(equivalent(X8,X6),equivalent(equivalent(X4,equivalent(X5,X8)),X7))))) )),
  inference(resolution,[],[f38,f6])).
fof(f38,plain,(
  ( ! [X6,X12,X10,X8,X7,X11,X9] : (is_a_theorem(equivalent(equivalent(equivalent(equivalent(X6,X7),equivalent(X6,X8)),equivalent(equivalent(equivalent(X9,X10),equivalent(X9,X11)),X12)),equivalent(equivalent(X7,X8),equivalent(equivalent(X10,X11),X12))))) )),
  inference(resolution,[],[f13,f6])).
fof(f3308,plain,(
  ( ! [X94,X92,X93,X91] : (~is_a_theorem(equivalent(equivalent(equivalent(X92,equivalent(X94,X94)),X91),X93)) | is_a_theorem(equivalent(X91,equivalent(X92,X93)))) )),
  inference(resolution,[],[f414,f95])).
fof(f95,plain,(
  ( ! [X54,X56,X55,X53] : (is_a_theorem(equivalent(equivalent(X55,X53),equivalent(equivalent(X56,X55),equivalent(X56,X54)))) | ~is_a_theorem(equivalent(X53,X54))) )),
  inference(resolution,[],[f32,f48])).
fof(f414,plain,(
  ( ! [X140,X138,X139,X137] : (~is_a_theorem(equivalent(equivalent(equivalent(X137,X137),equivalent(X138,X139)),equivalent(X138,X140))) | is_a_theorem(equivalent(X139,X140))) )),
  inference(resolution,[],[f67,f21])).
fof(f7758,plain,(
  ( ! [X39,X37,X38,X36,X40] : (~is_a_theorem(equivalent(equivalent(equivalent(X38,X39),X36),equivalent(equivalent(X40,X39),equivalent(equivalent(X38,X40),X37)))) | is_a_theorem(equivalent(X36,X37))) )),
  inference(resolution,[],[f684,f405])).
fof(f405,plain,(
  ( ! [X101,X99,X97,X100,X98] : (~is_a_theorem(equivalent(equivalent(equivalent(X97,X98),X99),equivalent(X98,X100))) | is_a_theorem(equivalent(X101,X100)) | ~is_a_theorem(equivalent(equivalent(X97,X101),X99))) )),
  inference(resolution,[],[f67,f14])).
fof(f684,plain,(
  ( ! [X54,X52,X56,X55,X53] : (is_a_theorem(equivalent(equivalent(equivalent(equivalent(X52,X53),X54),equivalent(equivalent(X55,X53),equivalent(equivalent(X52,X55),X56))),equivalent(X54,X56)))) )),
  inference(resolution,[],[f90,f42])).
fof(f90,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~is_a_theorem(equivalent(X30,equivalent(equivalent(X31,X32),equivalent(X31,X33)))) | is_a_theorem(equivalent(equivalent(equivalent(X32,X34),X30),equivalent(X34,X33)))) )),
  inference(resolution,[],[f32,f23])).
fof(f899024,plain,(
  ~is_a_theorem(equivalent(equivalent(equivalent(e,equivalent(equivalent(b,a),equivalent(equivalent(a,b),e))),falsehood),falsehood))),
  inference(resolution,[],[f3843,f57305])).
fof(f57305,plain,(
  ~is_a_theorem(equivalent(equivalent(equivalent(a,equivalent(equivalent(b,a),equivalent(equivalent(a,b),e))),equivalent(equivalent(e,a),falsehood)),falsehood))),
  inference(resolution,[],[f2304,f4])).
fof(f4,axiom,(
  ~is_a_theorem(equivalent(equivalent(equivalent(a,equivalent(equivalent(b,a),c)),equivalent(equivalent(e,a),falsehood)),equivalent(equivalent(equivalent(equivalent(a,b),e),c),falsehood)))),
  file('./p/LCL099-1.tptp',unknown)).
fof(f2304,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (is_a_theorem(equivalent(equivalent(equivalent(X23,equivalent(X24,X25)),X26),equivalent(equivalent(X27,X25),X28))) | ~is_a_theorem(equivalent(equivalent(equivalent(X23,equivalent(X24,X27)),X26),X28))) )),
  inference(resolution,[],[f211,f39])).
fof(f39,plain,(
  ( ! [X14,X17,X15,X13,X16] : (is_a_theorem(equivalent(equivalent(equivalent(equivalent(X15,X16),equivalent(X15,X17)),X13),equivalent(equivalent(X16,X17),X14))) | ~is_a_theorem(equivalent(X13,X14))) )),
  inference(resolution,[],[f13,f1])).
fof(f211,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (~is_a_theorem(equivalent(equivalent(equivalent(X43,X44),equivalent(X43,X45)),equivalent(equivalent(equivalent(X46,X47),equivalent(X46,X48)),X49))) | is_a_theorem(equivalent(equivalent(X44,X45),equivalent(equivalent(X47,X48),X49)))) )),
  inference(resolution,[],[f38,f1])).
fof(f3843,plain,(
  ( ! [X127,X130,X128,X126,X129] : (is_a_theorem(equivalent(equivalent(equivalent(X130,X127),equivalent(equivalent(X126,X130),X128)),X129)) | ~is_a_theorem(equivalent(equivalent(equivalent(X126,X127),X128),X129))) )),
  inference(resolution,[],[f1986,f63])).
fof(f63,plain,(
  ( ! [X17,X15,X16] : (~is_a_theorem(equivalent(X17,X15)) | ~is_a_theorem(equivalent(X15,X16)) | is_a_theorem(equivalent(X17,X16))) )),
  inference(resolution,[],[f57,f1])).
fof(f57,plain,(
  ( ! [X14,X15,X13] : (is_a_theorem(equivalent(equivalent(X15,X13),equivalent(X15,X14))) | ~is_a_theorem(equivalent(X13,X14))) )),
  inference(resolution,[],[f48,f1])).
% SZS output end Proof for LCL099-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 346945
% Time elapsed: 29.263 s
% ------------------------------
% ------------------------------
