% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for LCL144-1
% SZS output start Proof for LCL144-1
fof(f11689,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f20,f21,f9361,f11682])).
fof(f11682,plain,(
  ~spl0_1 | spl0_2),
  inference(avatar_contradiction_clause,[],[f11681])).
fof(f11681,plain,(
  $false | (~spl0_1 | spl0_2)),
  inference(subsumption_resolution,[],[f11680,f18])).
fof(f18,plain,(
  ~ordered(y,implies(x,z)) | spl0_2),
  inference(avatar_component_clause,[],[f17])).
fof(f17,plain,(
  spl0_2 <=> ordered(y,implies(x,z))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f11680,plain,(
  ordered(y,implies(x,z)) | ~spl0_1),
  inference(forward_demodulation,[],[f11629,f1])).
fof(f1,axiom,(
  ( ! [X0] : (implies(truth,X0) = X0) )),
  file('./p/LCL144-1.tptp',unknown)).
fof(f11629,plain,(
  ordered(y,implies(truth,implies(x,z))) | ~spl0_1),
  inference(superposition,[],[f167,f11316])).
fof(f11316,plain,(
  truth = implies(y,implies(x,z)) | ~spl0_1),
  inference(superposition,[],[f10421,f1])).
fof(f10421,plain,(
  truth = implies(truth,implies(y,implies(x,z))) | ~spl0_1),
  inference(superposition,[],[f2851,f9462])).
fof(f9462,plain,(
  ( ! [X26] : (truth = implies(implies(implies(y,z),X26),implies(x,X26))) ) | ~spl0_1),
  inference(forward_demodulation,[],[f9407,f188])).
fof(f188,plain,(
  ( ! [X4] : (truth = implies(X4,truth)) )),
  inference(backward_demodulation,[],[f174,f155])).
fof(f155,plain,(
  ( ! [X4] : (truth = implies(X4,implies(X4,X4))) )),
  inference(superposition,[],[f91,f27])).
fof(f27,plain,(
  ( ! [X0] : (implies(implies(X0,truth),truth) = implies(X0,X0)) )),
  inference(superposition,[],[f3,f1])).
fof(f3,axiom,(
  ( ! [X0,X1] : (implies(implies(X0,X1),X1) = implies(implies(X1,X0),X0)) )),
  file('./p/LCL144-1.tptp',unknown)).
fof(f91,plain,(
  ( ! [X2,X3] : (truth = implies(X2,implies(implies(X2,X3),X3))) )),
  inference(forward_demodulation,[],[f60,f1])).
fof(f60,plain,(
  ( ! [X2,X3] : (truth = implies(X2,implies(implies(X2,X3),implies(truth,X3)))) )),
  inference(superposition,[],[f2,f1])).
fof(f2,axiom,(
  ( ! [X2,X0,X1] : (truth = implies(implies(X0,X1),implies(implies(X1,X2),implies(X0,X2)))) )),
  file('./p/LCL144-1.tptp',unknown)).
fof(f174,plain,(
  ( ! [X2] : (truth = implies(X2,X2)) )),
  inference(forward_demodulation,[],[f156,f1])).
fof(f156,plain,(
  ( ! [X2] : (truth = implies(implies(truth,X2),X2)) )),
  inference(superposition,[],[f91,f1])).
fof(f9407,plain,(
  ( ! [X26] : (implies(implies(implies(y,z),X26),implies(x,X26)) = implies(implies(implies(implies(implies(y,z),X26),implies(x,X26)),truth),truth)) ) | ~spl0_1),
  inference(superposition,[],[f97,f9368])).
fof(f9368,plain,(
  truth = implies(x,implies(y,z)) | ~spl0_1),
  inference(resolution,[],[f15,f7])).
fof(f7,axiom,(
  ( ! [X0,X1] : (~ordered(X0,X1) | truth = implies(X0,X1)) )),
  file('./p/LCL144-1.tptp',unknown)).
fof(f15,plain,(
  ordered(x,implies(y,z)) | ~spl0_1),
  inference(avatar_component_clause,[],[f13])).
fof(f13,plain,(
  spl0_1 <=> ordered(x,implies(y,z))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f97,plain,(
  ( ! [X2,X0,X1] : (implies(implies(X1,X2),implies(X0,X2)) = implies(implies(implies(implies(X1,X2),implies(X0,X2)),implies(X0,X1)),implies(X0,X1))) )),
  inference(forward_demodulation,[],[f86,f1])).
fof(f86,plain,(
  ( ! [X2,X0,X1] : (implies(implies(implies(implies(X1,X2),implies(X0,X2)),implies(X0,X1)),implies(X0,X1)) = implies(truth,implies(implies(X1,X2),implies(X0,X2)))) )),
  inference(superposition,[],[f3,f2])).
fof(f2851,plain,(
  ( ! [X2,X0,X1] : (truth = implies(implies(implies(implies(X0,X1),X1),X2),implies(X0,X2))) )),
  inference(forward_demodulation,[],[f2685,f188])).
fof(f2685,plain,(
  ( ! [X2,X0,X1] : (implies(implies(implies(implies(X0,X1),X1),X2),implies(X0,X2)) = implies(implies(implies(implies(implies(implies(X0,X1),X1),X2),implies(X0,X2)),truth),truth)) )),
  inference(superposition,[],[f97,f91])).
fof(f167,plain,(
  ( ! [X19,X18] : (ordered(X18,implies(implies(X18,X19),X19))) )),
  inference(trivial_inequality_removal,[],[f165])).
fof(f165,plain,(
  ( ! [X19,X18] : (truth != truth | ordered(X18,implies(implies(X18,X19),X19))) )),
  inference(superposition,[],[f8,f91])).
fof(f8,axiom,(
  ( ! [X0,X1] : (truth != implies(X0,X1) | ordered(X0,X1)) )),
  file('./p/LCL144-1.tptp',unknown)).
fof(f9361,plain,(
  spl0_1 | ~spl0_2),
  inference(avatar_contradiction_clause,[],[f9360])).
fof(f9360,plain,(
  $false | (spl0_1 | ~spl0_2)),
  inference(subsumption_resolution,[],[f9359,f14])).
fof(f14,plain,(
  ~ordered(x,implies(y,z)) | spl0_1),
  inference(avatar_component_clause,[],[f13])).
fof(f9359,plain,(
  ordered(x,implies(y,z)) | ~spl0_2),
  inference(forward_demodulation,[],[f9293,f1])).
fof(f9293,plain,(
  ordered(x,implies(truth,implies(y,z))) | ~spl0_2),
  inference(superposition,[],[f167,f8763])).
fof(f8763,plain,(
  truth = implies(x,implies(y,z)) | ~spl0_2),
  inference(superposition,[],[f7476,f1])).
fof(f7476,plain,(
  truth = implies(truth,implies(x,implies(y,z))) | ~spl0_2),
  inference(superposition,[],[f2851,f2916])).
fof(f2916,plain,(
  ( ! [X71] : (truth = implies(implies(implies(x,z),X71),implies(y,X71))) ) | ~spl0_2),
  inference(forward_demodulation,[],[f2708,f188])).
fof(f2708,plain,(
  ( ! [X71] : (implies(implies(implies(x,z),X71),implies(y,X71)) = implies(implies(implies(implies(implies(x,z),X71),implies(y,X71)),truth),truth)) ) | ~spl0_2),
  inference(superposition,[],[f97,f22])).
fof(f22,plain,(
  truth = implies(y,implies(x,z)) | ~spl0_2),
  inference(resolution,[],[f7,f19])).
fof(f19,plain,(
  ordered(y,implies(x,z)) | ~spl0_2),
  inference(avatar_component_clause,[],[f17])).
fof(f21,plain,(
  ~spl0_1 | ~spl0_2),
  inference(avatar_split_clause,[],[f10,f17,f13])).
fof(f10,axiom,(
  ~ordered(y,implies(x,z)) | ~ordered(x,implies(y,z))),
  file('./p/LCL144-1.tptp',unknown)).
fof(f20,plain,(
  spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f9,f17,f13])).
fof(f9,axiom,(
  ordered(y,implies(x,z)) | ordered(x,implies(y,z))),
  file('./p/LCL144-1.tptp',unknown)).
% SZS output end Proof for LCL144-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 11385
% Time elapsed: 0.114 s
% ------------------------------
% ------------------------------
