% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for LCL240-3
% SZS output start Proof for LCL240-3
fof(f13385,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f72,f13384])).
fof(f13384,plain,(
  ~spl0_3),
  inference(avatar_contradiction_clause,[],[f13383])).
fof(f13383,plain,(
  $false | ~spl0_3),
  inference(subsumption_resolution,[],[f13346,f361])).
fof(f361,plain,(
  ( ! [X10,X11] : (theorem(or(X10,or(not(X10),X11)))) )),
  inference(resolution,[],[f323,f33])).
fof(f33,plain,(
  ( ! [X2,X0,X1] : (~theorem(or(X0,or(X1,X2))) | theorem(or(X1,or(X0,X2)))) )),
  inference(resolution,[],[f23,f16])).
fof(f16,plain,(
  ( ! [X4,X3] : (~theorem(or(not(X4),X3)) | ~theorem(X4) | theorem(X3)) )),
  inference(definition_unfolding,[],[f8,f6])).
fof(f6,axiom,(
  ( ! [X4,X3] : (implies(X3,X4) = or(not(X3),X4)) )),
  file('./p/LCL240-3.tptp',unknown)).
fof(f8,axiom,(
  ( ! [X4,X3] : (~theorem(X4) | ~theorem(implies(X4,X3)) | theorem(X3)) )),
  file('./p/LCL240-3.tptp',unknown)).
fof(f23,plain,(
  ( ! [X2,X0,X1] : (theorem(or(not(or(X0,or(X1,X2))),or(X1,or(X0,X2))))) )),
  inference(resolution,[],[f14,f7])).
fof(f7,axiom,(
  ( ! [X3] : (~axiom(X3) | theorem(X3)) )),
  file('./p/LCL240-3.tptp',unknown)).
fof(f14,plain,(
  ( ! [X2,X0,X1] : (axiom(or(not(or(X0,or(X1,X2))),or(X1,or(X0,X2))))) )),
  inference(definition_unfolding,[],[f4,f6])).
fof(f4,axiom,(
  ( ! [X2,X0,X1] : (axiom(implies(or(X0,or(X1,X2)),or(X1,or(X0,X2))))) )),
  file('./p/LCL240-3.tptp',unknown)).
fof(f323,plain,(
  ( ! [X19,X20] : (theorem(or(not(X19),or(X19,X20)))) )),
  inference(resolution,[],[f101,f19])).
fof(f19,plain,(
  ( ! [X0,X1] : (theorem(or(not(X0),or(X1,X0)))) )),
  inference(resolution,[],[f12,f7])).
fof(f12,plain,(
  ( ! [X0,X1] : (axiom(or(not(X0),or(X1,X0)))) )),
  inference(definition_unfolding,[],[f2,f6])).
fof(f2,axiom,(
  ( ! [X0,X1] : (axiom(implies(X0,or(X1,X0)))) )),
  file('./p/LCL240-3.tptp',unknown)).
fof(f101,plain,(
  ( ! [X4,X2,X3] : (~theorem(or(X2,or(X3,X4))) | theorem(or(X2,or(X4,X3)))) )),
  inference(resolution,[],[f57,f20])).
fof(f20,plain,(
  ( ! [X0,X1] : (theorem(or(not(or(X0,X1)),or(X1,X0)))) )),
  inference(resolution,[],[f13,f7])).
fof(f13,plain,(
  ( ! [X0,X1] : (axiom(or(not(or(X0,X1)),or(X1,X0)))) )),
  inference(definition_unfolding,[],[f3,f6])).
fof(f3,axiom,(
  ( ! [X0,X1] : (axiom(implies(or(X0,X1),or(X1,X0)))) )),
  file('./p/LCL240-3.tptp',unknown)).
fof(f57,plain,(
  ( ! [X2,X0,X1] : (~theorem(or(not(X0),X1)) | ~theorem(or(X2,X0)) | theorem(or(X2,X1))) )),
  inference(resolution,[],[f46,f16])).
fof(f46,plain,(
  ( ! [X2,X0,X1] : (theorem(or(not(or(X2,X0)),or(X2,X1))) | ~theorem(or(not(X0),X1))) )),
  inference(resolution,[],[f27,f16])).
fof(f27,plain,(
  ( ! [X2,X0,X1] : (theorem(or(not(or(not(X0),X1)),or(not(or(X2,X0)),or(X2,X1))))) )),
  inference(resolution,[],[f15,f7])).
fof(f15,plain,(
  ( ! [X2,X0,X1] : (axiom(or(not(or(not(X0),X1)),or(not(or(X2,X0)),or(X2,X1))))) )),
  inference(definition_unfolding,[],[f5,f6,f6,f6])).
fof(f5,axiom,(
  ( ! [X2,X0,X1] : (axiom(implies(implies(X0,X1),implies(or(X2,X0),or(X2,X1))))) )),
  file('./p/LCL240-3.tptp',unknown)).
fof(f13346,plain,(
  ~theorem(or(p,or(not(p),not(q)))) | ~spl0_3),
  inference(resolution,[],[f5475,f17])).
fof(f17,plain,(
  ~theorem(or(not(not(or(not(p),not(q)))),p))),
  inference(definition_unfolding,[],[f10,f6,f9])).
fof(f9,axiom,(
  ( ! [X6,X5] : (and(X5,X6) = not(or(not(X5),not(X6)))) )),
  file('./p/LCL240-3.tptp',unknown)).
fof(f10,axiom,(
  ~theorem(implies(and(p,q),p))),
  file('./p/LCL240-3.tptp',unknown)).
fof(f5475,plain,(
  ( ! [X59,X58] : (theorem(or(not(not(X59)),X58)) | ~theorem(or(X58,X59))) ) | ~spl0_3),
  inference(resolution,[],[f110,f28])).
fof(f28,plain,(
  ( ! [X0,X1] : (~theorem(or(X0,X1)) | theorem(or(X1,X0))) )),
  inference(resolution,[],[f20,f16])).
fof(f110,plain,(
  ( ! [X30,X31] : (theorem(or(X30,not(not(X31)))) | ~theorem(or(X30,X31))) ) | ~spl0_3),
  inference(resolution,[],[f57,f74])).
fof(f74,plain,(
  ( ! [X1] : (theorem(or(X1,not(X1)))) ) | ~spl0_3),
  inference(resolution,[],[f67,f28])).
fof(f67,plain,(
  ( ! [X1] : (theorem(or(not(X1),X1))) ) | ~spl0_3),
  inference(avatar_component_clause,[],[f66])).
fof(f66,plain,(
  spl0_3 <=> ! [X1] : theorem(or(not(X1),X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f72,plain,(
  spl0_3),
  inference(avatar_split_clause,[],[f63,f66])).
fof(f63,plain,(
  ( ! [X4] : (theorem(or(not(X4),X4))) )),
  inference(resolution,[],[f52,f21])).
fof(f21,plain,(
  ( ! [X0] : (~theorem(or(X0,X0)) | theorem(X0)) )),
  inference(resolution,[],[f16,f18])).
fof(f18,plain,(
  ( ! [X0] : (theorem(or(not(or(X0,X0)),X0))) )),
  inference(resolution,[],[f11,f7])).
fof(f11,plain,(
  ( ! [X0] : (axiom(or(not(or(X0,X0)),X0))) )),
  inference(definition_unfolding,[],[f1,f6])).
fof(f1,axiom,(
  ( ! [X0] : (axiom(implies(or(X0,X0),X0))) )),
  file('./p/LCL240-3.tptp',unknown)).
fof(f52,plain,(
  ( ! [X10,X11] : (theorem(or(X10,or(not(X11),X11)))) )),
  inference(resolution,[],[f33,f19])).
% SZS output end Proof for LCL240-3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 14328
% Time elapsed: 0.735 s
% ------------------------------
% ------------------------------
