% Refutation found. Thanks to Tanya!
% SZS status Theorem for MGT013+1
% SZS output start Proof for MGT013+1
fof(f148,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f108,f118,f143])).
fof(f143,plain,(
  ~spl7_1),
  inference(avatar_contradiction_clause,[],[f142])).
fof(f142,plain,(
  $false | ~spl7_1),
  inference(subsumption_resolution,[],[f141,f51])).
fof(f51,plain,(
  ( ! [X1] : (~greater(X1,X1)) )),
  inference(equality_resolution,[],[f34])).
fof(f34,plain,(
  ( ! [X0,X1] : (X0 != X1 | ~greater(X0,X1)) )),
  inference(cnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : (X0 != X1 | ~greater(X0,X1))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1] : ~(X0 = X1 & greater(X0,X1))),
  file('./p/MGT013+1.tptp',unknown)).
fof(f141,plain,(
  greater(sK1,sK1) | ~spl7_1),
  inference(backward_demodulation,[],[f133,f49])).
fof(f49,plain,(
  greater(sK2,sK1)),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  greater(sK3,sK4) & greater(sK2,sK1) & size(sK0,sK4,sK6) & size(sK0,sK3,sK5) & complexity(sK0,sK2,sK6) & complexity(sK0,sK1,sK5) & reorganization_free(sK0,sK5,sK6) & organization(sK0,sK6) & organization(sK0,sK5)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6])],[f31,f32])).
fof(f32,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (greater(X3,X4) & greater(X2,X1) & size(X0,X4,X6) & size(X0,X3,X5) & complexity(X0,X2,X6) & complexity(X0,X1,X5) & reorganization_free(X0,X5,X6) & organization(X0,X6) & organization(X0,X5)) => (greater(sK3,sK4) & greater(sK2,sK1) & size(sK0,sK4,sK6) & size(sK0,sK3,sK5) & complexity(sK0,sK2,sK6) & complexity(sK0,sK1,sK5) & reorganization_free(sK0,sK5,sK6) & organization(sK0,sK6) & organization(sK0,sK5))),
  introduced(choice_axiom,[])).
fof(f31,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (greater(X3,X4) & greater(X2,X1) & size(X0,X4,X6) & size(X0,X3,X5) & complexity(X0,X2,X6) & complexity(X0,X1,X5) & reorganization_free(X0,X5,X6) & organization(X0,X6) & organization(X0,X5))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6] : (greater(X3,X4) & (greater(X2,X1) & size(X0,X4,X6) & size(X0,X3,X5) & complexity(X0,X2,X6) & complexity(X0,X1,X5) & reorganization_free(X0,X5,X6) & organization(X0,X6) & organization(X0,X5)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6] : ((greater(X2,X1) & size(X0,X4,X6) & size(X0,X3,X5) & complexity(X0,X2,X6) & complexity(X0,X1,X5) & reorganization_free(X0,X5,X6) & organization(X0,X6) & organization(X0,X5)) => ~greater(X3,X4))),
  inference(rectify,[],[f10])).
fof(f10,negated_conjecture,(
  ~! [X0,X5,X6,X7,X8,X3,X4] : ((greater(X6,X5) & size(X0,X8,X4) & size(X0,X7,X3) & complexity(X0,X6,X4) & complexity(X0,X5,X3) & reorganization_free(X0,X3,X4) & organization(X0,X4) & organization(X0,X3)) => ~greater(X7,X8))),
  inference(negated_conjecture,[],[f9])).
fof(f9,conjecture,(
  ! [X0,X5,X6,X7,X8,X3,X4] : ((greater(X6,X5) & size(X0,X8,X4) & size(X0,X7,X3) & complexity(X0,X6,X4) & complexity(X0,X5,X3) & reorganization_free(X0,X3,X4) & organization(X0,X4) & organization(X0,X3)) => ~greater(X7,X8))),
  file('./p/MGT013+1.tptp',unknown)).
fof(f133,plain,(
  sK1 = sK2 | ~spl7_1),
  inference(resolution,[],[f121,f71])).
fof(f71,plain,(
  ( ! [X0] : (~complexity(sK0,X0,sK5) | sK1 = X0) )),
  inference(subsumption_resolution,[],[f69,f42])).
fof(f42,plain,(
  organization(sK0,sK5)),
  inference(cnf_transformation,[],[f33])).
fof(f69,plain,(
  ( ! [X0] : (sK1 = X0 | ~complexity(sK0,X0,sK5) | ~organization(sK0,sK5)) )),
  inference(resolution,[],[f53,f45])).
fof(f45,plain,(
  complexity(sK0,sK1,sK5)),
  inference(cnf_transformation,[],[f33])).
fof(f53,plain,(
  ( ! [X4,X2,X0,X1] : (~complexity(X0,X2,X4) | X1 = X2 | ~complexity(X0,X1,X4) | ~organization(X0,X4)) )),
  inference(duplicate_literal_removal,[],[f52])).
fof(f52,plain,(
  ( ! [X4,X2,X0,X1] : (X1 = X2 | ~complexity(X0,X2,X4) | ~complexity(X0,X1,X4) | ~organization(X0,X4) | ~organization(X0,X4)) )),
  inference(equality_resolution,[],[f39])).
fof(f39,plain,(
  ( ! [X4,X2,X0,X3,X1] : (X1 = X2 | X3 != X4 | ~complexity(X0,X2,X4) | ~complexity(X0,X1,X3) | ~organization(X0,X4) | ~organization(X0,X3)) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2,X3,X4] : (X1 = X2 | X3 != X4 | ~complexity(X0,X2,X4) | ~complexity(X0,X1,X3) | ~organization(X0,X4) | ~organization(X0,X3))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2,X3,X4] : (X1 = X2 | (X3 != X4 | ~complexity(X0,X2,X4) | ~complexity(X0,X1,X3) | ~organization(X0,X4) | ~organization(X0,X3)))),
  inference(ennf_transformation,[],[f14])).
fof(f14,plain,(
  ! [X0,X1,X2,X3,X4] : ((X3 = X4 & complexity(X0,X2,X4) & complexity(X0,X1,X3) & organization(X0,X4) & organization(X0,X3)) => X1 = X2)),
  inference(rectify,[],[f6])).
fof(f6,axiom,(
  ! [X0,X5,X6,X3,X4] : ((X3 = X4 & complexity(X0,X6,X4) & complexity(X0,X5,X3) & organization(X0,X4) & organization(X0,X3)) => X5 = X6)),
  file('./p/MGT013+1.tptp',unknown)).
fof(f121,plain,(
  complexity(sK0,sK2,sK5) | ~spl7_1),
  inference(backward_demodulation,[],[f100,f46])).
fof(f46,plain,(
  complexity(sK0,sK2,sK6)),
  inference(cnf_transformation,[],[f33])).
fof(f100,plain,(
  sK5 = sK6 | ~spl7_1),
  inference(avatar_component_clause,[],[f98])).
fof(f98,plain,(
  spl7_1 <=> sK5 = sK6),
  introduced(avatar_definition,[new_symbols(naming,[spl7_1])])).
fof(f118,plain,(
  ~spl7_2),
  inference(avatar_split_clause,[],[f117,f102])).
fof(f102,plain,(
  spl7_2 <=> greater(sK5,sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl7_2])])).
fof(f117,plain,(
  ~greater(sK5,sK6)),
  inference(subsumption_resolution,[],[f116,f43])).
fof(f43,plain,(
  organization(sK0,sK6)),
  inference(cnf_transformation,[],[f33])).
fof(f116,plain,(
  ~greater(sK5,sK6) | ~organization(sK0,sK6)),
  inference(subsumption_resolution,[],[f111,f58])).
fof(f58,plain,(
  reorganization_free(sK0,sK6,sK5)),
  inference(resolution,[],[f38,f44])).
fof(f44,plain,(
  reorganization_free(sK0,sK5,sK6)),
  inference(cnf_transformation,[],[f33])).
fof(f38,plain,(
  ( ! [X2,X0,X1] : (~reorganization_free(X0,X1,X2) | reorganization_free(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2] : (reorganization_free(X0,X2,X1) | ~reorganization_free(X0,X1,X2))),
  inference(ennf_transformation,[],[f13])).
fof(f13,plain,(
  ! [X0,X1,X2] : (reorganization_free(X0,X1,X2) => reorganization_free(X0,X2,X1))),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X0,X3,X4] : (reorganization_free(X0,X3,X4) => reorganization_free(X0,X4,X3))),
  file('./p/MGT013+1.tptp',unknown)).
fof(f111,plain,(
  ~greater(sK5,sK6) | ~reorganization_free(sK0,sK6,sK5) | ~organization(sK0,sK6)),
  inference(subsumption_resolution,[],[f110,f49])).
fof(f110,plain,(
  ~greater(sK2,sK1) | ~greater(sK5,sK6) | ~reorganization_free(sK0,sK6,sK5) | ~organization(sK0,sK6)),
  inference(resolution,[],[f82,f46])).
fof(f82,plain,(
  ( ! [X0,X1] : (~complexity(sK0,X1,X0) | ~greater(X1,sK1) | ~greater(sK5,X0) | ~reorganization_free(sK0,X0,sK5) | ~organization(sK0,X0)) )),
  inference(subsumption_resolution,[],[f80,f42])).
fof(f80,plain,(
  ( ! [X0,X1] : (~greater(sK5,X0) | ~greater(X1,sK1) | ~complexity(sK0,X1,X0) | ~reorganization_free(sK0,X0,sK5) | ~organization(sK0,sK5) | ~organization(sK0,X0)) )),
  inference(resolution,[],[f41,f45])).
fof(f41,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~complexity(X0,X2,X4) | ~greater(X4,X3) | ~greater(X1,X2) | ~complexity(X0,X1,X3) | ~reorganization_free(X0,X3,X4) | ~organization(X0,X4) | ~organization(X0,X3)) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2,X3,X4] : (~greater(X1,X2) | ~greater(X4,X3) | ~complexity(X0,X2,X4) | ~complexity(X0,X1,X3) | ~reorganization_free(X0,X3,X4) | ~organization(X0,X4) | ~organization(X0,X3))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2,X3,X4] : (~greater(X1,X2) | (~greater(X4,X3) | ~complexity(X0,X2,X4) | ~complexity(X0,X1,X3) | ~reorganization_free(X0,X3,X4) | ~organization(X0,X4) | ~organization(X0,X3)))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0,X1,X2,X3,X4] : ((greater(X4,X3) & complexity(X0,X2,X4) & complexity(X0,X1,X3) & reorganization_free(X0,X3,X4) & organization(X0,X4) & organization(X0,X3)) => ~greater(X1,X2))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X0,X5,X6,X3,X4] : ((greater(X4,X3) & complexity(X0,X6,X4) & complexity(X0,X5,X3) & reorganization_free(X0,X3,X4) & organization(X0,X4) & organization(X0,X3)) => ~greater(X5,X6))),
  file('./p/MGT013+1.tptp',unknown)).
fof(f108,plain,(
  spl7_1 | spl7_2),
  inference(avatar_split_clause,[],[f107,f102,f98])).
fof(f107,plain,(
  greater(sK5,sK6) | sK5 = sK6),
  inference(subsumption_resolution,[],[f106,f54])).
fof(f54,plain,(
  time(sK5)),
  inference(resolution,[],[f36,f42])).
fof(f36,plain,(
  ( ! [X0,X1] : (~organization(X0,X1) | time(X1)) )),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1] : (time(X1) | ~organization(X0,X1))),
  inference(ennf_transformation,[],[f11])).
fof(f11,plain,(
  ! [X0,X1] : (organization(X0,X1) => time(X1))),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X0,X2] : (organization(X0,X2) => time(X2))),
  file('./p/MGT013+1.tptp',unknown)).
fof(f106,plain,(
  greater(sK5,sK6) | sK5 = sK6 | ~time(sK5)),
  inference(subsumption_resolution,[],[f94,f55])).
fof(f55,plain,(
  time(sK6)),
  inference(resolution,[],[f36,f43])).
fof(f94,plain,(
  greater(sK5,sK6) | sK5 = sK6 | ~time(sK6) | ~time(sK5)),
  inference(resolution,[],[f92,f37])).
fof(f37,plain,(
  ( ! [X0,X1] : (greater(X1,X0) | greater(X0,X1) | X0 = X1 | ~time(X1) | ~time(X0)) )),
  inference(cnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1] : (greater(X1,X0) | X0 = X1 | greater(X0,X1) | ~time(X1) | ~time(X0))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : ((greater(X1,X0) | X0 = X1 | greater(X0,X1)) | (~time(X1) | ~time(X0)))),
  inference(ennf_transformation,[],[f12])).
fof(f12,plain,(
  ! [X0,X1] : ((time(X1) & time(X0)) => (greater(X1,X0) | X0 = X1 | greater(X0,X1)))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X3,X4] : ((time(X4) & time(X3)) => (greater(X4,X3) | X3 = X4 | greater(X3,X4)))),
  file('./p/MGT013+1.tptp',unknown)).
fof(f92,plain,(
  ~greater(sK6,sK5)),
  inference(subsumption_resolution,[],[f91,f42])).
fof(f91,plain,(
  ~greater(sK6,sK5) | ~organization(sK0,sK5)),
  inference(subsumption_resolution,[],[f90,f44])).
fof(f90,plain,(
  ~greater(sK6,sK5) | ~reorganization_free(sK0,sK5,sK6) | ~organization(sK0,sK5)),
  inference(subsumption_resolution,[],[f88,f50])).
fof(f50,plain,(
  greater(sK3,sK4)),
  inference(cnf_transformation,[],[f33])).
fof(f88,plain,(
  ~greater(sK3,sK4) | ~greater(sK6,sK5) | ~reorganization_free(sK0,sK5,sK6) | ~organization(sK0,sK5)),
  inference(resolution,[],[f78,f47])).
fof(f47,plain,(
  size(sK0,sK3,sK5)),
  inference(cnf_transformation,[],[f33])).
fof(f78,plain,(
  ( ! [X2,X3] : (~size(sK0,X3,X2) | ~greater(X3,sK4) | ~greater(sK6,X2) | ~reorganization_free(sK0,X2,sK6) | ~organization(sK0,X2)) )),
  inference(subsumption_resolution,[],[f76,f43])).
fof(f76,plain,(
  ( ! [X2,X3] : (~greater(sK6,X2) | ~greater(X3,sK4) | ~size(sK0,X3,X2) | ~reorganization_free(sK0,X2,sK6) | ~organization(sK0,sK6) | ~organization(sK0,X2)) )),
  inference(resolution,[],[f40,f48])).
fof(f48,plain,(
  size(sK0,sK4,sK6)),
  inference(cnf_transformation,[],[f33])).
fof(f40,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~size(X0,X2,X4) | ~greater(X4,X3) | ~greater(X1,X2) | ~size(X0,X1,X3) | ~reorganization_free(X0,X3,X4) | ~organization(X0,X4) | ~organization(X0,X3)) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3,X4] : (~greater(X1,X2) | ~greater(X4,X3) | ~size(X0,X2,X4) | ~size(X0,X1,X3) | ~reorganization_free(X0,X3,X4) | ~organization(X0,X4) | ~organization(X0,X3))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2,X3,X4] : (~greater(X1,X2) | (~greater(X4,X3) | ~size(X0,X2,X4) | ~size(X0,X1,X3) | ~reorganization_free(X0,X3,X4) | ~organization(X0,X4) | ~organization(X0,X3)))),
  inference(ennf_transformation,[],[f15])).
fof(f15,plain,(
  ! [X0,X1,X2,X3,X4] : ((greater(X4,X3) & size(X0,X2,X4) & size(X0,X1,X3) & reorganization_free(X0,X3,X4) & organization(X0,X4) & organization(X0,X3)) => ~greater(X1,X2))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X0,X7,X8,X3,X4] : ((greater(X4,X3) & size(X0,X8,X4) & size(X0,X7,X3) & reorganization_free(X0,X3,X4) & organization(X0,X4) & organization(X0,X3)) => ~greater(X7,X8))),
  file('./p/MGT013+1.tptp',unknown)).
% SZS output end Proof for MGT013+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
