% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for MGT002-1
% SZS output start Proof for MGT002-1
fof(f121,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f96,f109,f120])).
fof(f120,plain,(
  ~spl0_10),
  inference(avatar_contradiction_clause,[],[f119])).
fof(f119,plain,(
  $false | ~spl0_10),
  inference(subsumption_resolution,[],[f118,f8])).
fof(f8,axiom,(
  organization(sk2,sk6)),
  file('./p/MGT002-1.tptp',unknown)).
fof(f118,plain,(
  ~organization(sk2,sk6) | ~spl0_10),
  inference(resolution,[],[f95,f1])).
fof(f1,axiom,(
  ( ! [X0,X1] : (reproducibility(X0,sk1(X1,X0),X1) | ~organization(X0,X1)) )),
  file('./p/MGT002-1.tptp',unknown)).
fof(f95,plain,(
  ( ! [X0] : (~reproducibility(sk2,X0,sk6)) ) | ~spl0_10),
  inference(avatar_component_clause,[],[f94])).
fof(f94,plain,(
  spl0_10 <=> ! [X0] : ~reproducibility(sk2,X0,sk6)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f109,plain,(
  ~spl0_9),
  inference(avatar_contradiction_clause,[],[f108])).
fof(f108,plain,(
  $false | ~spl0_9),
  inference(subsumption_resolution,[],[f107,f7])).
fof(f7,axiom,(
  organization(sk2,sk5)),
  file('./p/MGT002-1.tptp',unknown)).
fof(f107,plain,(
  ~organization(sk2,sk5) | ~spl0_9),
  inference(resolution,[],[f92,f1])).
fof(f92,plain,(
  ( ! [X1] : (~reproducibility(sk2,X1,sk5)) ) | ~spl0_9),
  inference(avatar_component_clause,[],[f91])).
fof(f91,plain,(
  spl0_9 <=> ! [X1] : ~reproducibility(sk2,X1,sk5)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_9])])).
fof(f96,plain,(
  spl0_9 | spl0_10),
  inference(avatar_split_clause,[],[f89,f94,f91])).
fof(f89,plain,(
  ( ! [X0,X1] : (~reproducibility(sk2,X0,sk6) | ~reproducibility(sk2,X1,sk5)) )),
  inference(subsumption_resolution,[],[f88,f27])).
fof(f27,plain,(
  ( ! [X0,X1] : (~reproducibility(sk2,X1,sk5) | ~reproducibility(sk2,X0,sk6) | greater(X0,X1)) )),
  inference(subsumption_resolution,[],[f26,f7])).
fof(f26,plain,(
  ( ! [X0,X1] : (~reproducibility(sk2,X0,sk6) | ~reproducibility(sk2,X1,sk5) | greater(X0,X1) | ~organization(sk2,sk5)) )),
  inference(subsumption_resolution,[],[f25,f8])).
fof(f25,plain,(
  ( ! [X0,X1] : (~reproducibility(sk2,X0,sk6) | ~reproducibility(sk2,X1,sk5) | greater(X0,X1) | ~organization(sk2,sk6) | ~organization(sk2,sk5)) )),
  inference(subsumption_resolution,[],[f20,f12])).
fof(f12,axiom,(
  greater(sk6,sk5)),
  file('./p/MGT002-1.tptp',unknown)).
fof(f20,plain,(
  ( ! [X0,X1] : (~greater(sk6,sk5) | ~reproducibility(sk2,X0,sk6) | ~reproducibility(sk2,X1,sk5) | greater(X0,X1) | ~organization(sk2,sk6) | ~organization(sk2,sk5)) )),
  inference(resolution,[],[f6,f9])).
fof(f9,axiom,(
  reorganization_free(sk2,sk5,sk6)),
  file('./p/MGT002-1.tptp',unknown)).
fof(f6,axiom,(
  ( ! [X4,X2,X0,X3,X1] : (~reorganization_free(X0,X1,X2) | ~greater(X2,X1) | ~reproducibility(X0,X4,X2) | ~reproducibility(X0,X3,X1) | greater(X4,X3) | ~organization(X0,X2) | ~organization(X0,X1)) )),
  file('./p/MGT002-1.tptp',unknown)).
fof(f88,plain,(
  ( ! [X0,X1] : (~greater(X0,X1) | ~reproducibility(sk2,X0,sk6) | ~reproducibility(sk2,X1,sk5)) )),
  inference(subsumption_resolution,[],[f87,f7])).
fof(f87,plain,(
  ( ! [X0,X1] : (~greater(X0,X1) | ~reproducibility(sk2,X0,sk6) | ~reproducibility(sk2,X1,sk5) | ~organization(sk2,sk5)) )),
  inference(subsumption_resolution,[],[f86,f14])).
fof(f14,plain,(
  reorganization_free(sk2,sk5,sk5)),
  inference(resolution,[],[f2,f9])).
fof(f2,axiom,(
  ( ! [X2,X0,X1] : (~reorganization_free(X0,X1,X2) | reorganization_free(X0,X1,X1)) )),
  file('./p/MGT002-1.tptp',unknown)).
fof(f86,plain,(
  ( ! [X0,X1] : (~greater(X0,X1) | ~reproducibility(sk2,X0,sk6) | ~reproducibility(sk2,X1,sk5) | ~reorganization_free(sk2,sk5,sk5) | ~organization(sk2,sk5)) )),
  inference(subsumption_resolution,[],[f84,f13])).
fof(f13,axiom,(
  ~greater(sk4,sk3)),
  file('./p/MGT002-1.tptp',unknown)).
fof(f84,plain,(
  ( ! [X0,X1] : (greater(sk4,sk3) | ~greater(X0,X1) | ~reproducibility(sk2,X0,sk6) | ~reproducibility(sk2,X1,sk5) | ~reorganization_free(sk2,sk5,sk5) | ~organization(sk2,sk5)) )),
  inference(resolution,[],[f51,f10])).
fof(f10,axiom,(
  inertia(sk2,sk3,sk5)),
  file('./p/MGT002-1.tptp',unknown)).
fof(f51,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~inertia(X8,X7,X9) | greater(sk4,X7) | ~greater(X5,X6) | ~reproducibility(sk2,X5,sk6) | ~reproducibility(X8,X6,X9) | ~reorganization_free(X8,X9,X9) | ~organization(X8,X9)) )),
  inference(subsumption_resolution,[],[f50,f8])).
fof(f50,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~greater(X5,X6) | greater(sk4,X7) | ~inertia(X8,X7,X9) | ~reproducibility(sk2,X5,sk6) | ~reproducibility(X8,X6,X9) | ~reorganization_free(X8,X9,X9) | ~organization(sk2,sk6) | ~organization(X8,X9)) )),
  inference(subsumption_resolution,[],[f47,f16])).
fof(f16,plain,(
  reorganization_free(sk2,sk6,sk6)),
  inference(resolution,[],[f3,f9])).
fof(f3,axiom,(
  ( ! [X2,X0,X1] : (~reorganization_free(X0,X1,X2) | reorganization_free(X0,X2,X2)) )),
  file('./p/MGT002-1.tptp',unknown)).
fof(f47,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~greater(X5,X6) | greater(sk4,X7) | ~inertia(X8,X7,X9) | ~reproducibility(sk2,X5,sk6) | ~reproducibility(X8,X6,X9) | ~reorganization_free(sk2,sk6,sk6) | ~reorganization_free(X8,X9,X9) | ~organization(sk2,sk6) | ~organization(X8,X9)) )),
  inference(resolution,[],[f4,f11])).
fof(f11,axiom,(
  inertia(sk2,sk4,sk6)),
  file('./p/MGT002-1.tptp',unknown)).
fof(f4,axiom,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~inertia(X2,X7,X3) | ~greater(X5,X4) | greater(X7,X6) | ~inertia(X0,X6,X1) | ~reproducibility(X2,X5,X3) | ~reproducibility(X0,X4,X1) | ~reorganization_free(X2,X3,X3) | ~reorganization_free(X0,X1,X1) | ~organization(X2,X3) | ~organization(X0,X1)) )),
  file('./p/MGT002-1.tptp',unknown)).
% SZS output end Proof for MGT002-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
