% Refutation found. Thanks to Tanya!
% SZS status Theorem for MGT026+1
% SZS output start Proof for MGT026+1
fof(f300,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f166,f178,f209,f236,f243,f263,f281,f282,f296,f299])).
fof(f299,plain,(
  ~spl2_16),
  inference(avatar_contradiction_clause,[],[f298])).
fof(f298,plain,(
  $false | ~spl2_16),
  inference(subsumption_resolution,[],[f297,f58])).
fof(f58,plain,(
  environment(sK0)),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ~selection_favors(efficient_producers,first_movers,sK1) & greater(sK1,critical_point(sK0)) & in_environment(sK0,sK1) & environment(sK0)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f39,f42])).
fof(f42,plain,(
  ? [X0,X1] : (~selection_favors(efficient_producers,first_movers,X1) & greater(X1,critical_point(X0)) & in_environment(X0,X1) & environment(X0)) => (~selection_favors(efficient_producers,first_movers,sK1) & greater(sK1,critical_point(sK0)) & in_environment(sK0,sK1) & environment(sK0))),
  introduced(choice_axiom,[])).
fof(f39,plain,(
  ? [X0,X1] : (~selection_favors(efficient_producers,first_movers,X1) & greater(X1,critical_point(X0)) & in_environment(X0,X1) & environment(X0))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ? [X0,X1] : (~selection_favors(efficient_producers,first_movers,X1) & (greater(X1,critical_point(X0)) & in_environment(X0,X1) & environment(X0)))),
  inference(ennf_transformation,[],[f20])).
fof(f20,plain,(
  ~! [X0,X1] : ((greater(X1,critical_point(X0)) & in_environment(X0,X1) & environment(X0)) => selection_favors(efficient_producers,first_movers,X1))),
  inference(rectify,[],[f12])).
fof(f12,negated_conjecture,(
  ~! [X0,X3] : ((greater(X3,critical_point(X0)) & in_environment(X0,X3) & environment(X0)) => selection_favors(efficient_producers,first_movers,X3))),
  inference(negated_conjecture,[],[f11])).
fof(f11,conjecture,(
  ! [X0,X3] : ((greater(X3,critical_point(X0)) & in_environment(X0,X3) & environment(X0)) => selection_favors(efficient_producers,first_movers,X3))),
  file('./p/MGT026+1.tptp',unknown)).
fof(f297,plain,(
  ~environment(sK0) | ~spl2_16),
  inference(resolution,[],[f280,f59])).
fof(f59,plain,(
  in_environment(sK0,sK1)),
  inference(cnf_transformation,[],[f43])).
fof(f280,plain,(
  ( ! [X0] : (~in_environment(X0,sK1) | ~environment(X0)) ) | ~spl2_16),
  inference(avatar_component_clause,[],[f279])).
fof(f279,plain,(
  spl2_16 <=> ! [X0] : (~in_environment(X0,sK1) | ~environment(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl2_16])])).
fof(f296,plain,(
  spl2_16 | ~spl2_1 | ~spl2_13),
  inference(avatar_split_clause,[],[f295,f175,f76,f279])).
fof(f76,plain,(
  spl2_1 <=> zero = cardinality_at_time(first_movers,sK1)),
  introduced(avatar_definition,[new_symbols(naming,[spl2_1])])).
fof(f175,plain,(
  spl2_13 <=> greater(cardinality_at_time(efficient_producers,sK1),zero)),
  introduced(avatar_definition,[new_symbols(naming,[spl2_13])])).
fof(f295,plain,(
  ( ! [X0] : (~environment(X0) | ~in_environment(X0,sK1)) ) | (~spl2_1 | ~spl2_13)),
  inference(subsumption_resolution,[],[f294,f48])).
fof(f48,plain,(
  ( ! [X0,X1] : (subpopulation(first_movers,X0,X1) | ~in_environment(X0,X1) | ~environment(X0)) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ((subpopulation(efficient_producers,X0,X1) & subpopulation(first_movers,X0,X1)) | ~in_environment(X0,X1) | ~environment(X0))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : ((subpopulation(efficient_producers,X0,X1) & subpopulation(first_movers,X0,X1)) | (~in_environment(X0,X1) | ~environment(X0)))),
  inference(ennf_transformation,[],[f15])).
fof(f15,plain,(
  ! [X0,X1] : ((in_environment(X0,X1) & environment(X0)) => (subpopulation(efficient_producers,X0,X1) & subpopulation(first_movers,X0,X1)))),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X0,X3] : ((in_environment(X0,X3) & environment(X0)) => (subpopulation(efficient_producers,X0,X3) & subpopulation(first_movers,X0,X3)))),
  file('./p/MGT026+1.tptp',unknown)).
fof(f294,plain,(
  ( ! [X0] : (~subpopulation(first_movers,X0,sK1) | ~environment(X0) | ~in_environment(X0,sK1)) ) | (~spl2_1 | ~spl2_13)),
  inference(duplicate_literal_removal,[],[f293])).
fof(f293,plain,(
  ( ! [X0] : (~subpopulation(first_movers,X0,sK1) | ~environment(X0) | ~in_environment(X0,sK1) | ~environment(X0)) ) | (~spl2_1 | ~spl2_13)),
  inference(resolution,[],[f292,f49])).
fof(f49,plain,(
  ( ! [X0,X1] : (subpopulation(efficient_producers,X0,X1) | ~in_environment(X0,X1) | ~environment(X0)) )),
  inference(cnf_transformation,[],[f30])).
fof(f292,plain,(
  ( ! [X0] : (~subpopulation(efficient_producers,X0,sK1) | ~subpopulation(first_movers,X0,sK1) | ~environment(X0)) ) | (~spl2_1 | ~spl2_13)),
  inference(subsumption_resolution,[],[f288,f61])).
fof(f61,plain,(
  ~selection_favors(efficient_producers,first_movers,sK1)),
  inference(cnf_transformation,[],[f43])).
fof(f288,plain,(
  ( ! [X0] : (selection_favors(efficient_producers,first_movers,sK1) | ~subpopulation(first_movers,X0,sK1) | ~subpopulation(efficient_producers,X0,sK1) | ~environment(X0)) ) | (~spl2_1 | ~spl2_13)),
  inference(resolution,[],[f287,f177])).
fof(f177,plain,(
  greater(cardinality_at_time(efficient_producers,sK1),zero) | ~spl2_13),
  inference(avatar_component_clause,[],[f175])).
fof(f287,plain,(
  ( ! [X2,X1] : (~greater(cardinality_at_time(X1,sK1),zero) | selection_favors(X1,first_movers,sK1) | ~subpopulation(first_movers,X2,sK1) | ~subpopulation(X1,X2,sK1) | ~environment(X2)) ) | ~spl2_1),
  inference(trivial_inequality_removal,[],[f286])).
fof(f286,plain,(
  ( ! [X2,X1] : (zero != zero | selection_favors(X1,first_movers,sK1) | ~greater(cardinality_at_time(X1,sK1),zero) | ~subpopulation(first_movers,X2,sK1) | ~subpopulation(X1,X2,sK1) | ~environment(X2)) ) | ~spl2_1),
  inference(superposition,[],[f45,f78])).
fof(f78,plain,(
  zero = cardinality_at_time(first_movers,sK1) | ~spl2_1),
  inference(avatar_component_clause,[],[f76])).
fof(f45,plain,(
  ( ! [X2,X0,X3,X1] : (zero != cardinality_at_time(X2,X3) | selection_favors(X1,X2,X3) | ~greater(cardinality_at_time(X1,X3),zero) | ~subpopulation(X2,X0,X3) | ~subpopulation(X1,X0,X3) | ~environment(X0)) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2,X3] : (selection_favors(X1,X2,X3) | zero != cardinality_at_time(X2,X3) | ~greater(cardinality_at_time(X1,X3),zero) | ~subpopulation(X2,X0,X3) | ~subpopulation(X1,X0,X3) | ~environment(X0))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2,X3] : (selection_favors(X1,X2,X3) | (zero != cardinality_at_time(X2,X3) | ~greater(cardinality_at_time(X1,X3),zero) | ~subpopulation(X2,X0,X3) | ~subpopulation(X1,X0,X3) | ~environment(X0)))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1,X2,X3] : ((zero = cardinality_at_time(X2,X3) & greater(cardinality_at_time(X1,X3),zero) & subpopulation(X2,X0,X3) & subpopulation(X1,X0,X3) & environment(X0)) => selection_favors(X1,X2,X3))),
  file('./p/MGT026+1.tptp',unknown)).
fof(f282,plain,(
  spl2_1 | spl2_2),
  inference(avatar_split_clause,[],[f210,f80,f76])).
fof(f80,plain,(
  spl2_2 <=> greater(cardinality_at_time(first_movers,sK1),zero)),
  introduced(avatar_definition,[new_symbols(naming,[spl2_2])])).
fof(f210,plain,(
  greater(cardinality_at_time(first_movers,sK1),zero) | zero = cardinality_at_time(first_movers,sK1)),
  inference(resolution,[],[f73,f52])).
fof(f52,plain,(
  ( ! [X0,X1] : (~greater_or_equal(X0,X1) | greater(X0,X1) | X0 = X1) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ((greater_or_equal(X0,X1) | (X0 != X1 & ~greater(X0,X1))) & (X0 = X1 | greater(X0,X1) | ~greater_or_equal(X0,X1)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : ((greater_or_equal(X0,X1) | (X0 != X1 & ~greater(X0,X1))) & ((X0 = X1 | greater(X0,X1)) | ~greater_or_equal(X0,X1)))),
  inference(nnf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1] : (greater_or_equal(X0,X1) <=> (X0 = X1 | greater(X0,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X4,X5] : (greater_or_equal(X4,X5) <=> (X4 = X5 | greater(X4,X5)))),
  file('./p/MGT026+1.tptp',unknown)).
fof(f73,plain,(
  greater_or_equal(cardinality_at_time(first_movers,sK1),zero)),
  inference(subsumption_resolution,[],[f72,f58])).
fof(f72,plain,(
  greater_or_equal(cardinality_at_time(first_movers,sK1),zero) | ~environment(sK0)),
  inference(resolution,[],[f47,f59])).
fof(f47,plain,(
  ( ! [X0,X1] : (~in_environment(X0,X1) | greater_or_equal(cardinality_at_time(first_movers,X1),zero) | ~environment(X0)) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1] : (greater_or_equal(cardinality_at_time(first_movers,X1),zero) | ~in_environment(X0,X1) | ~environment(X0))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : (greater_or_equal(cardinality_at_time(first_movers,X1),zero) | (~in_environment(X0,X1) | ~environment(X0)))),
  inference(ennf_transformation,[],[f14])).
fof(f14,plain,(
  ! [X0,X1] : ((in_environment(X0,X1) & environment(X0)) => greater_or_equal(cardinality_at_time(first_movers,X1),zero))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X0,X3] : ((in_environment(X0,X3) & environment(X0)) => greater_or_equal(cardinality_at_time(first_movers,X3),zero))),
  file('./p/MGT026+1.tptp',unknown)).
fof(f281,plain,(
  spl2_16 | ~spl2_2 | ~spl2_10 | ~spl2_13),
  inference(avatar_split_clause,[],[f277,f175,f160,f80,f279])).
fof(f160,plain,(
  spl2_10 <=> greater(growth_rate(efficient_producers,sK1),growth_rate(first_movers,sK1))),
  introduced(avatar_definition,[new_symbols(naming,[spl2_10])])).
fof(f277,plain,(
  ( ! [X0] : (~greater(cardinality_at_time(first_movers,sK1),zero) | ~in_environment(X0,sK1) | ~environment(X0)) ) | (~spl2_10 | ~spl2_13)),
  inference(subsumption_resolution,[],[f246,f269])).
fof(f269,plain,(
  ( ! [X0] : (~subpopulations(first_movers,efficient_producers,X0,sK1) | ~environment(X0)) ) | ~spl2_10),
  inference(subsumption_resolution,[],[f265,f61])).
fof(f265,plain,(
  ( ! [X0] : (selection_favors(efficient_producers,first_movers,sK1) | ~subpopulations(first_movers,efficient_producers,X0,sK1) | ~environment(X0)) ) | ~spl2_10),
  inference(resolution,[],[f162,f44])).
fof(f44,plain,(
  ( ! [X2,X0,X3,X1] : (~greater(growth_rate(X2,X3),growth_rate(X1,X3)) | selection_favors(X2,X1,X3) | ~subpopulations(X1,X2,X0,X3) | ~environment(X0)) )),
  inference(cnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3] : (selection_favors(X2,X1,X3) | ~greater(growth_rate(X2,X3),growth_rate(X1,X3)) | ~subpopulations(X1,X2,X0,X3) | ~environment(X0))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2,X3] : (selection_favors(X2,X1,X3) | (~greater(growth_rate(X2,X3),growth_rate(X1,X3)) | ~subpopulations(X1,X2,X0,X3) | ~environment(X0)))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1,X2,X3] : ((greater(growth_rate(X2,X3),growth_rate(X1,X3)) & subpopulations(X1,X2,X0,X3) & environment(X0)) => selection_favors(X2,X1,X3))),
  file('./p/MGT026+1.tptp',unknown)).
fof(f162,plain,(
  greater(growth_rate(efficient_producers,sK1),growth_rate(first_movers,sK1)) | ~spl2_10),
  inference(avatar_component_clause,[],[f160])).
fof(f246,plain,(
  ( ! [X0] : (subpopulations(first_movers,efficient_producers,X0,sK1) | ~greater(cardinality_at_time(first_movers,sK1),zero) | ~in_environment(X0,sK1) | ~environment(X0)) ) | ~spl2_13),
  inference(resolution,[],[f177,f46])).
fof(f46,plain,(
  ( ! [X0,X1] : (~greater(cardinality_at_time(efficient_producers,X1),zero) | subpopulations(first_movers,efficient_producers,X0,X1) | ~greater(cardinality_at_time(first_movers,X1),zero) | ~in_environment(X0,X1) | ~environment(X0)) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : (subpopulations(first_movers,efficient_producers,X0,X1) | ~greater(cardinality_at_time(efficient_producers,X1),zero) | ~greater(cardinality_at_time(first_movers,X1),zero) | ~in_environment(X0,X1) | ~environment(X0))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : (subpopulations(first_movers,efficient_producers,X0,X1) | (~greater(cardinality_at_time(efficient_producers,X1),zero) | ~greater(cardinality_at_time(first_movers,X1),zero) | ~in_environment(X0,X1) | ~environment(X0)))),
  inference(ennf_transformation,[],[f13])).
fof(f13,plain,(
  ! [X0,X1] : ((greater(cardinality_at_time(efficient_producers,X1),zero) & greater(cardinality_at_time(first_movers,X1),zero) & in_environment(X0,X1) & environment(X0)) => subpopulations(first_movers,efficient_producers,X0,X1))),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X0,X3] : ((greater(cardinality_at_time(efficient_producers,X3),zero) & greater(cardinality_at_time(first_movers,X3),zero) & in_environment(X0,X3) & environment(X0)) => subpopulations(first_movers,efficient_producers,X0,X3))),
  file('./p/MGT026+1.tptp',unknown)).
fof(f263,plain,(
  spl2_10 | spl2_11 | ~spl2_2 | ~spl2_13),
  inference(avatar_split_clause,[],[f262,f175,f80,f164,f160])).
fof(f164,plain,(
  spl2_11 <=> ! [X0] : (~in_environment(X0,sK1) | ~greater(sK1,critical_point(X0)) | ~environment(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl2_11])])).
fof(f262,plain,(
  ( ! [X0] : (~in_environment(X0,sK1) | ~environment(X0) | ~greater(sK1,critical_point(X0)) | greater(growth_rate(efficient_producers,sK1),growth_rate(first_movers,sK1))) ) | (~spl2_2 | ~spl2_13)),
  inference(duplicate_literal_removal,[],[f261])).
fof(f261,plain,(
  ( ! [X0] : (~in_environment(X0,sK1) | ~environment(X0) | ~greater(sK1,critical_point(X0)) | greater(growth_rate(efficient_producers,sK1),growth_rate(first_movers,sK1)) | ~environment(X0)) ) | (~spl2_2 | ~spl2_13)),
  inference(resolution,[],[f250,f63])).
fof(f63,plain,(
  ( ! [X2,X0] : (~subpopulations(first_movers,efficient_producers,X0,X2) | ~greater(X2,critical_point(X0)) | greater(growth_rate(efficient_producers,X2),growth_rate(first_movers,X2)) | ~environment(X0)) )),
  inference(equality_resolution,[],[f56])).
fof(f56,plain,(
  ( ! [X2,X0,X1] : (greater(growth_rate(efficient_producers,X2),growth_rate(first_movers,X2)) | ~greater(X2,X1) | ~subpopulations(first_movers,efficient_producers,X0,X2) | critical_point(X0) != X1 | ~environment(X0)) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : ((! [X2] : (greater(growth_rate(efficient_producers,X2),growth_rate(first_movers,X2)) | ~greater(X2,X1) | ~subpopulations(first_movers,efficient_producers,X0,X2)) & ~greater(growth_rate(efficient_producers,X1),growth_rate(first_movers,X1))) | critical_point(X0) != X1 | ~environment(X0))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : ((! [X2] : (greater(growth_rate(efficient_producers,X2),growth_rate(first_movers,X2)) | (~greater(X2,X1) | ~subpopulations(first_movers,efficient_producers,X0,X2))) & ~greater(growth_rate(efficient_producers,X1),growth_rate(first_movers,X1))) | (critical_point(X0) != X1 | ~environment(X0)))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : ((critical_point(X0) = X1 & environment(X0)) => (! [X2] : ((greater(X2,X1) & subpopulations(first_movers,efficient_producers,X0,X2)) => greater(growth_rate(efficient_producers,X2),growth_rate(first_movers,X2))) & ~greater(growth_rate(efficient_producers,X1),growth_rate(first_movers,X1))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X0,X7] : ((critical_point(X0) = X7 & environment(X0)) => (! [X3] : ((greater(X3,X7) & subpopulations(first_movers,efficient_producers,X0,X3)) => greater(growth_rate(efficient_producers,X3),growth_rate(first_movers,X3))) & ~greater(growth_rate(efficient_producers,X7),growth_rate(first_movers,X7))))),
  file('./p/MGT026+1.tptp',unknown)).
fof(f250,plain,(
  ( ! [X0] : (subpopulations(first_movers,efficient_producers,X0,sK1) | ~in_environment(X0,sK1) | ~environment(X0)) ) | (~spl2_2 | ~spl2_13)),
  inference(subsumption_resolution,[],[f246,f82])).
fof(f82,plain,(
  greater(cardinality_at_time(first_movers,sK1),zero) | ~spl2_2),
  inference(avatar_component_clause,[],[f80])).
fof(f243,plain,(
  ~spl2_11),
  inference(avatar_contradiction_clause,[],[f242])).
fof(f242,plain,(
  $false | ~spl2_11),
  inference(subsumption_resolution,[],[f241,f58])).
fof(f241,plain,(
  ~environment(sK0) | ~spl2_11),
  inference(subsumption_resolution,[],[f239,f59])).
fof(f239,plain,(
  ~in_environment(sK0,sK1) | ~environment(sK0) | ~spl2_11),
  inference(resolution,[],[f165,f60])).
fof(f60,plain,(
  greater(sK1,critical_point(sK0))),
  inference(cnf_transformation,[],[f43])).
fof(f165,plain,(
  ( ! [X0] : (~greater(sK1,critical_point(X0)) | ~in_environment(X0,sK1) | ~environment(X0)) ) | ~spl2_11),
  inference(avatar_component_clause,[],[f164])).
fof(f236,plain,(
  ~spl2_2 | ~spl2_6 | ~spl2_10),
  inference(avatar_contradiction_clause,[],[f235])).
fof(f235,plain,(
  $false | (~spl2_2 | ~spl2_6 | ~spl2_10)),
  inference(subsumption_resolution,[],[f234,f58])).
fof(f234,plain,(
  ~environment(sK0) | (~spl2_2 | ~spl2_6 | ~spl2_10)),
  inference(resolution,[],[f233,f59])).
fof(f233,plain,(
  ( ! [X0] : (~in_environment(X0,sK1) | ~environment(X0)) ) | (~spl2_2 | ~spl2_6 | ~spl2_10)),
  inference(subsumption_resolution,[],[f232,f218])).
fof(f218,plain,(
  ( ! [X0] : (~subpopulations(first_movers,efficient_producers,X0,sK1) | ~environment(X0)) ) | ~spl2_10),
  inference(subsumption_resolution,[],[f214,f61])).
fof(f214,plain,(
  ( ! [X0] : (selection_favors(efficient_producers,first_movers,sK1) | ~subpopulations(first_movers,efficient_producers,X0,sK1) | ~environment(X0)) ) | ~spl2_10),
  inference(resolution,[],[f162,f44])).
fof(f232,plain,(
  ( ! [X0] : (subpopulations(first_movers,efficient_producers,X0,sK1) | ~in_environment(X0,sK1) | ~environment(X0)) ) | (~spl2_2 | ~spl2_6)),
  inference(subsumption_resolution,[],[f231,f65])).
fof(f65,plain,(
  greater_or_equal(sK1,critical_point(sK0))),
  inference(resolution,[],[f53,f60])).
fof(f53,plain,(
  ( ! [X0,X1] : (~greater(X0,X1) | greater_or_equal(X0,X1)) )),
  inference(cnf_transformation,[],[f41])).
fof(f231,plain,(
  ( ! [X0] : (subpopulations(first_movers,efficient_producers,X0,sK1) | ~greater_or_equal(sK1,critical_point(sK0)) | ~in_environment(X0,sK1) | ~environment(X0)) ) | (~spl2_2 | ~spl2_6)),
  inference(subsumption_resolution,[],[f230,f59])).
fof(f230,plain,(
  ( ! [X0] : (~in_environment(sK0,sK1) | subpopulations(first_movers,efficient_producers,X0,sK1) | ~greater_or_equal(sK1,critical_point(sK0)) | ~in_environment(X0,sK1) | ~environment(X0)) ) | (~spl2_2 | ~spl2_6)),
  inference(resolution,[],[f196,f82])).
fof(f196,plain,(
  ( ! [X0,X1] : (~greater(cardinality_at_time(first_movers,X0),zero) | ~in_environment(sK0,X0) | subpopulations(first_movers,efficient_producers,X1,X0) | ~greater_or_equal(X0,critical_point(sK0)) | ~in_environment(X1,X0) | ~environment(X1)) ) | ~spl2_6),
  inference(resolution,[],[f193,f46])).
fof(f193,plain,(
  ( ! [X0] : (greater(cardinality_at_time(efficient_producers,X0),zero) | ~greater_or_equal(X0,critical_point(sK0)) | ~in_environment(sK0,X0)) ) | ~spl2_6),
  inference(subsumption_resolution,[],[f186,f58])).
fof(f186,plain,(
  ( ! [X0] : (~greater_or_equal(X0,critical_point(sK0)) | greater(cardinality_at_time(efficient_producers,X0),zero) | ~in_environment(sK0,X0) | ~environment(sK0)) ) | ~spl2_6),
  inference(superposition,[],[f57,f120])).
fof(f120,plain,(
  critical_point(sK0) = appear(efficient_producers,sK0) | ~spl2_6),
  inference(avatar_component_clause,[],[f118])).
fof(f118,plain,(
  spl2_6 <=> critical_point(sK0) = appear(efficient_producers,sK0)),
  introduced(avatar_definition,[new_symbols(naming,[spl2_6])])).
fof(f57,plain,(
  ( ! [X0,X1] : (~greater_or_equal(X1,appear(efficient_producers,X0)) | greater(cardinality_at_time(efficient_producers,X1),zero) | ~in_environment(X0,X1) | ~environment(X0)) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : (greater(cardinality_at_time(efficient_producers,X1),zero) | ~greater_or_equal(X1,appear(efficient_producers,X0)) | ~in_environment(X0,X1) | ~environment(X0))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : (greater(cardinality_at_time(efficient_producers,X1),zero) | (~greater_or_equal(X1,appear(efficient_producers,X0)) | ~in_environment(X0,X1) | ~environment(X0)))),
  inference(ennf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1] : ((greater_or_equal(X1,appear(efficient_producers,X0)) & in_environment(X0,X1) & environment(X0)) => greater(cardinality_at_time(efficient_producers,X1),zero))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X0,X3] : ((greater_or_equal(X3,appear(efficient_producers,X0)) & in_environment(X0,X3) & environment(X0)) => greater(cardinality_at_time(efficient_producers,X3),zero))),
  file('./p/MGT026+1.tptp',unknown)).
fof(f209,plain,(
  ~spl2_1 | ~spl2_6),
  inference(avatar_contradiction_clause,[],[f208])).
fof(f208,plain,(
  $false | (~spl2_1 | ~spl2_6)),
  inference(subsumption_resolution,[],[f207,f58])).
fof(f207,plain,(
  ~environment(sK0) | (~spl2_1 | ~spl2_6)),
  inference(resolution,[],[f206,f59])).
fof(f206,plain,(
  ( ! [X0] : (~in_environment(X0,sK1) | ~environment(X0)) ) | (~spl2_1 | ~spl2_6)),
  inference(subsumption_resolution,[],[f205,f48])).
fof(f205,plain,(
  ( ! [X0] : (~subpopulation(first_movers,X0,sK1) | ~environment(X0) | ~in_environment(X0,sK1)) ) | (~spl2_1 | ~spl2_6)),
  inference(duplicate_literal_removal,[],[f204])).
fof(f204,plain,(
  ( ! [X0] : (~subpopulation(first_movers,X0,sK1) | ~environment(X0) | ~in_environment(X0,sK1) | ~environment(X0)) ) | (~spl2_1 | ~spl2_6)),
  inference(resolution,[],[f203,f49])).
fof(f203,plain,(
  ( ! [X2] : (~subpopulation(efficient_producers,X2,sK1) | ~subpopulation(first_movers,X2,sK1) | ~environment(X2)) ) | (~spl2_1 | ~spl2_6)),
  inference(subsumption_resolution,[],[f202,f61])).
fof(f202,plain,(
  ( ! [X2] : (selection_favors(efficient_producers,first_movers,sK1) | ~subpopulation(first_movers,X2,sK1) | ~subpopulation(efficient_producers,X2,sK1) | ~environment(X2)) ) | (~spl2_1 | ~spl2_6)),
  inference(subsumption_resolution,[],[f201,f59])).
fof(f201,plain,(
  ( ! [X2] : (~in_environment(sK0,sK1) | selection_favors(efficient_producers,first_movers,sK1) | ~subpopulation(first_movers,X2,sK1) | ~subpopulation(efficient_producers,X2,sK1) | ~environment(X2)) ) | (~spl2_1 | ~spl2_6)),
  inference(subsumption_resolution,[],[f197,f65])).
fof(f197,plain,(
  ( ! [X2] : (~greater_or_equal(sK1,critical_point(sK0)) | ~in_environment(sK0,sK1) | selection_favors(efficient_producers,first_movers,sK1) | ~subpopulation(first_movers,X2,sK1) | ~subpopulation(efficient_producers,X2,sK1) | ~environment(X2)) ) | (~spl2_1 | ~spl2_6)),
  inference(resolution,[],[f193,f184])).
fof(f184,plain,(
  ( ! [X0,X1] : (~greater(cardinality_at_time(X0,sK1),zero) | selection_favors(X0,first_movers,sK1) | ~subpopulation(first_movers,X1,sK1) | ~subpopulation(X0,X1,sK1) | ~environment(X1)) ) | ~spl2_1),
  inference(trivial_inequality_removal,[],[f183])).
fof(f183,plain,(
  ( ! [X0,X1] : (zero != zero | selection_favors(X0,first_movers,sK1) | ~greater(cardinality_at_time(X0,sK1),zero) | ~subpopulation(first_movers,X1,sK1) | ~subpopulation(X0,X1,sK1) | ~environment(X1)) ) | ~spl2_1),
  inference(superposition,[],[f45,f78])).
fof(f178,plain,(
  spl2_13 | spl2_6),
  inference(avatar_split_clause,[],[f173,f118,f175])).
fof(f173,plain,(
  critical_point(sK0) = appear(efficient_producers,sK0) | greater(cardinality_at_time(efficient_producers,sK1),zero)),
  inference(subsumption_resolution,[],[f172,f59])).
fof(f172,plain,(
  critical_point(sK0) = appear(efficient_producers,sK0) | greater(cardinality_at_time(efficient_producers,sK1),zero) | ~in_environment(sK0,sK1)),
  inference(subsumption_resolution,[],[f147,f58])).
fof(f147,plain,(
  critical_point(sK0) = appear(efficient_producers,sK0) | ~environment(sK0) | greater(cardinality_at_time(efficient_producers,sK1),zero) | ~in_environment(sK0,sK1)),
  inference(resolution,[],[f110,f60])).
fof(f110,plain,(
  ( ! [X0,X1] : (~greater(X1,critical_point(X0)) | critical_point(X0) = appear(efficient_producers,X0) | ~environment(X0) | greater(cardinality_at_time(efficient_producers,X1),zero) | ~in_environment(X0,X1)) )),
  inference(duplicate_literal_removal,[],[f108])).
fof(f108,plain,(
  ( ! [X0,X1] : (critical_point(X0) = appear(efficient_producers,X0) | ~greater(X1,critical_point(X0)) | ~environment(X0) | greater(cardinality_at_time(efficient_producers,X1),zero) | ~in_environment(X0,X1) | ~environment(X0)) )),
  inference(resolution,[],[f107,f57])).
fof(f107,plain,(
  ( ! [X6,X5] : (greater_or_equal(X6,appear(efficient_producers,X5)) | critical_point(X5) = appear(efficient_producers,X5) | ~greater(X6,critical_point(X5)) | ~environment(X5)) )),
  inference(resolution,[],[f103,f53])).
fof(f103,plain,(
  ( ! [X2,X1] : (greater(X2,appear(efficient_producers,X1)) | ~environment(X1) | critical_point(X1) = appear(efficient_producers,X1) | ~greater(X2,critical_point(X1))) )),
  inference(resolution,[],[f69,f51])).
fof(f51,plain,(
  ( ! [X2,X0,X1] : (~greater(X1,X2) | greater(X0,X2) | ~greater(X0,X1)) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2] : (greater(X0,X2) | ~greater(X1,X2) | ~greater(X0,X1))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : (greater(X0,X2) | (~greater(X1,X2) | ~greater(X0,X1)))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0,X1,X2] : ((greater(X1,X2) & greater(X0,X1)) => greater(X0,X2))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X4,X5,X6] : ((greater(X5,X6) & greater(X4,X5)) => greater(X4,X6))),
  file('./p/MGT026+1.tptp',unknown)).
fof(f69,plain,(
  ( ! [X1] : (greater(critical_point(X1),appear(efficient_producers,X1)) | critical_point(X1) = appear(efficient_producers,X1) | ~environment(X1)) )),
  inference(resolution,[],[f52,f50])).
fof(f50,plain,(
  ( ! [X0] : (greater_or_equal(critical_point(X0),appear(efficient_producers,X0)) | ~environment(X0)) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (greater_or_equal(critical_point(X0),appear(efficient_producers,X0)) | ~environment(X0))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0] : (environment(X0) => greater_or_equal(critical_point(X0),appear(efficient_producers,X0)))),
  file('./p/MGT026+1.tptp',unknown)).
fof(f166,plain,(
  spl2_10 | spl2_11 | ~spl2_2 | ~spl2_6),
  inference(avatar_split_clause,[],[f158,f118,f80,f164,f160])).
fof(f158,plain,(
  ( ! [X0] : (~in_environment(X0,sK1) | ~environment(X0) | ~greater(sK1,critical_point(X0)) | greater(growth_rate(efficient_producers,sK1),growth_rate(first_movers,sK1))) ) | (~spl2_2 | ~spl2_6)),
  inference(duplicate_literal_removal,[],[f157])).
fof(f157,plain,(
  ( ! [X0] : (~in_environment(X0,sK1) | ~environment(X0) | ~greater(sK1,critical_point(X0)) | greater(growth_rate(efficient_producers,sK1),growth_rate(first_movers,sK1)) | ~environment(X0)) ) | (~spl2_2 | ~spl2_6)),
  inference(resolution,[],[f156,f63])).
fof(f156,plain,(
  ( ! [X0] : (subpopulations(first_movers,efficient_producers,X0,sK1) | ~in_environment(X0,sK1) | ~environment(X0)) ) | (~spl2_2 | ~spl2_6)),
  inference(subsumption_resolution,[],[f155,f65])).
fof(f155,plain,(
  ( ! [X0] : (subpopulations(first_movers,efficient_producers,X0,sK1) | ~greater_or_equal(sK1,critical_point(sK0)) | ~in_environment(X0,sK1) | ~environment(X0)) ) | (~spl2_2 | ~spl2_6)),
  inference(subsumption_resolution,[],[f154,f59])).
fof(f154,plain,(
  ( ! [X0] : (~in_environment(sK0,sK1) | subpopulations(first_movers,efficient_producers,X0,sK1) | ~greater_or_equal(sK1,critical_point(sK0)) | ~in_environment(X0,sK1) | ~environment(X0)) ) | (~spl2_2 | ~spl2_6)),
  inference(resolution,[],[f149,f82])).
fof(f149,plain,(
  ( ! [X0,X1] : (~greater(cardinality_at_time(first_movers,X0),zero) | ~in_environment(sK0,X0) | subpopulations(first_movers,efficient_producers,X1,X0) | ~greater_or_equal(X0,critical_point(sK0)) | ~in_environment(X1,X0) | ~environment(X1)) ) | ~spl2_6),
  inference(resolution,[],[f135,f46])).
fof(f135,plain,(
  ( ! [X0] : (greater(cardinality_at_time(efficient_producers,X0),zero) | ~greater_or_equal(X0,critical_point(sK0)) | ~in_environment(sK0,X0)) ) | ~spl2_6),
  inference(subsumption_resolution,[],[f128,f58])).
fof(f128,plain,(
  ( ! [X0] : (~greater_or_equal(X0,critical_point(sK0)) | greater(cardinality_at_time(efficient_producers,X0),zero) | ~in_environment(sK0,X0) | ~environment(sK0)) ) | ~spl2_6),
  inference(superposition,[],[f57,f120])).
% SZS output end Proof for MGT026+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
