% Refutation found. Thanks to Tanya!
% SZS status Theorem for MGT008+1
% SZS output start Proof for MGT008+1
fof(f143,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f123,f136,f139,f142])).
fof(f142,plain,(
  ~spl10_14),
  inference(avatar_contradiction_clause,[],[f141])).
fof(f141,plain,(
  $false | ~spl10_14),
  inference(subsumption_resolution,[],[f140,f24])).
fof(f24,plain,(
  organization(sK2,sK9)),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ~greater(sK5,sK4) & greater(sK7,sK6) & size(sK2,sK7,sK9) & size(sK1,sK6,sK8) & survival_chance(sK2,sK5,sK9) & survival_chance(sK1,sK4,sK8) & class(sK2,sK3,sK9) & class(sK1,sK3,sK8) & reorganization_free(sK2,sK9,sK9) & reorganization_free(sK1,sK8,sK8) & organization(sK2,sK9) & organization(sK1,sK8)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8,sK9])],[f15,f18])).
fof(f18,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (~greater(X4,X3) & greater(X6,X5) & size(X1,X6,X8) & size(X0,X5,X7) & survival_chance(X1,X4,X8) & survival_chance(X0,X3,X7) & class(X1,X2,X8) & class(X0,X2,X7) & reorganization_free(X1,X8,X8) & reorganization_free(X0,X7,X7) & organization(X1,X8) & organization(X0,X7)) => (~greater(sK5,sK4) & greater(sK7,sK6) & size(sK2,sK7,sK9) & size(sK1,sK6,sK8) & survival_chance(sK2,sK5,sK9) & survival_chance(sK1,sK4,sK8) & class(sK2,sK3,sK9) & class(sK1,sK3,sK8) & reorganization_free(sK2,sK9,sK9) & reorganization_free(sK1,sK8,sK8) & organization(sK2,sK9) & organization(sK1,sK8))),
  introduced(choice_axiom,[])).
fof(f15,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (~greater(X4,X3) & greater(X6,X5) & size(X1,X6,X8) & size(X0,X5,X7) & survival_chance(X1,X4,X8) & survival_chance(X0,X3,X7) & class(X1,X2,X8) & class(X0,X2,X7) & reorganization_free(X1,X8,X8) & reorganization_free(X0,X7,X7) & organization(X1,X8) & organization(X0,X7))),
  inference(flattening,[],[f14])).
fof(f14,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (~greater(X4,X3) & (greater(X6,X5) & size(X1,X6,X8) & size(X0,X5,X7) & survival_chance(X1,X4,X8) & survival_chance(X0,X3,X7) & class(X1,X2,X8) & class(X0,X2,X7) & reorganization_free(X1,X8,X8) & reorganization_free(X0,X7,X7) & organization(X1,X8) & organization(X0,X7)))),
  inference(ennf_transformation,[],[f8])).
fof(f8,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((greater(X6,X5) & size(X1,X6,X8) & size(X0,X5,X7) & survival_chance(X1,X4,X8) & survival_chance(X0,X3,X7) & class(X1,X2,X8) & class(X0,X2,X7) & reorganization_free(X1,X8,X8) & reorganization_free(X0,X7,X7) & organization(X1,X8) & organization(X0,X7)) => greater(X4,X3))),
  inference(rectify,[],[f5])).
fof(f5,negated_conjecture,(
  ~! [X0,X3,X4,X11,X12,X5,X6,X9,X10] : ((greater(X6,X5) & size(X3,X6,X10) & size(X0,X5,X9) & survival_chance(X3,X12,X10) & survival_chance(X0,X11,X9) & class(X3,X4,X10) & class(X0,X4,X9) & reorganization_free(X3,X10,X10) & reorganization_free(X0,X9,X9) & organization(X3,X10) & organization(X0,X9)) => greater(X12,X11))),
  inference(negated_conjecture,[],[f4])).
fof(f4,conjecture,(
  ! [X0,X3,X4,X11,X12,X5,X6,X9,X10] : ((greater(X6,X5) & size(X3,X6,X10) & size(X0,X5,X9) & survival_chance(X3,X12,X10) & survival_chance(X0,X11,X9) & class(X3,X4,X10) & class(X0,X4,X9) & reorganization_free(X3,X10,X10) & reorganization_free(X0,X9,X9) & organization(X3,X10) & organization(X0,X9)) => greater(X12,X11))),
  file('./p/MGT008+1.tptp',unknown)).
fof(f140,plain,(
  ~organization(sK2,sK9) | ~spl10_14),
  inference(resolution,[],[f122,f20])).
fof(f20,plain,(
  ( ! [X0,X1] : (inertia(X0,sK0(X0,X1),X1) | ~organization(X0,X1)) )),
  inference(cnf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1] : (inertia(X0,sK0(X0,X1),X1) | ~organization(X0,X1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f9,f16])).
fof(f16,plain,(
  ! [X1,X0] : (? [X2] : inertia(X0,X2,X1) => inertia(X0,sK0(X0,X1),X1))),
  introduced(choice_axiom,[])).
fof(f9,plain,(
  ! [X0,X1] : (? [X2] : inertia(X0,X2,X1) | ~organization(X0,X1))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1] : (organization(X0,X1) => ? [X2] : inertia(X0,X2,X1))),
  file('./p/MGT008+1.tptp',unknown)).
fof(f122,plain,(
  ( ! [X0] : (~inertia(sK2,X0,sK9)) ) | ~spl10_14),
  inference(avatar_component_clause,[],[f121])).
fof(f121,plain,(
  spl10_14 <=> ! [X0] : ~inertia(sK2,X0,sK9)),
  introduced(avatar_definition,[new_symbols(naming,[spl10_14])])).
fof(f139,plain,(
  ~spl10_4),
  inference(avatar_contradiction_clause,[],[f138])).
fof(f138,plain,(
  $false | ~spl10_4),
  inference(subsumption_resolution,[],[f137,f27])).
fof(f27,plain,(
  class(sK1,sK3,sK8)),
  inference(cnf_transformation,[],[f19])).
fof(f137,plain,(
  ~class(sK1,sK3,sK8) | ~spl10_4),
  inference(resolution,[],[f57,f28])).
fof(f28,plain,(
  class(sK2,sK3,sK9)),
  inference(cnf_transformation,[],[f19])).
fof(f57,plain,(
  ( ! [X5] : (~class(sK2,X5,sK9) | ~class(sK1,X5,sK8)) ) | ~spl10_4),
  inference(avatar_component_clause,[],[f56])).
fof(f56,plain,(
  spl10_4 <=> ! [X5] : (~class(sK1,X5,sK8) | ~class(sK2,X5,sK9))),
  introduced(avatar_definition,[new_symbols(naming,[spl10_4])])).
fof(f136,plain,(
  ~spl10_13),
  inference(avatar_contradiction_clause,[],[f135])).
fof(f135,plain,(
  $false | ~spl10_13),
  inference(subsumption_resolution,[],[f134,f23])).
fof(f23,plain,(
  organization(sK1,sK8)),
  inference(cnf_transformation,[],[f19])).
fof(f134,plain,(
  ~organization(sK1,sK8) | ~spl10_13),
  inference(resolution,[],[f119,f20])).
fof(f119,plain,(
  ( ! [X1] : (~inertia(sK1,X1,sK8)) ) | ~spl10_13),
  inference(avatar_component_clause,[],[f118])).
fof(f118,plain,(
  spl10_13 <=> ! [X1] : ~inertia(sK1,X1,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl10_13])])).
fof(f123,plain,(
  spl10_4 | spl10_13 | spl10_14),
  inference(avatar_split_clause,[],[f116,f121,f118,f56])).
fof(f116,plain,(
  ( ! [X2,X0,X1] : (~inertia(sK2,X0,sK9) | ~inertia(sK1,X1,sK8) | ~class(sK2,X2,sK9) | ~class(sK1,X2,sK8)) )),
  inference(subsumption_resolution,[],[f115,f98])).
fof(f98,plain,(
  ( ! [X0,X1] : (~greater(X0,X1) | ~inertia(sK2,X0,sK9) | ~inertia(sK1,X1,sK8)) )),
  inference(subsumption_resolution,[],[f97,f23])).
fof(f97,plain,(
  ( ! [X0,X1] : (~greater(X0,X1) | ~inertia(sK2,X0,sK9) | ~inertia(sK1,X1,sK8) | ~organization(sK1,sK8)) )),
  inference(subsumption_resolution,[],[f96,f25])).
fof(f25,plain,(
  reorganization_free(sK1,sK8,sK8)),
  inference(cnf_transformation,[],[f19])).
fof(f96,plain,(
  ( ! [X0,X1] : (~greater(X0,X1) | ~inertia(sK2,X0,sK9) | ~inertia(sK1,X1,sK8) | ~reorganization_free(sK1,sK8,sK8) | ~organization(sK1,sK8)) )),
  inference(subsumption_resolution,[],[f94,f34])).
fof(f34,plain,(
  ~greater(sK5,sK4)),
  inference(cnf_transformation,[],[f19])).
fof(f94,plain,(
  ( ! [X0,X1] : (greater(sK5,sK4) | ~greater(X0,X1) | ~inertia(sK2,X0,sK9) | ~inertia(sK1,X1,sK8) | ~reorganization_free(sK1,sK8,sK8) | ~organization(sK1,sK8)) )),
  inference(resolution,[],[f71,f29])).
fof(f29,plain,(
  survival_chance(sK1,sK4,sK8)),
  inference(cnf_transformation,[],[f19])).
fof(f71,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~survival_chance(X8,X7,X9) | greater(sK5,X7) | ~greater(X5,X6) | ~inertia(sK2,X5,sK9) | ~inertia(X8,X6,X9) | ~reorganization_free(X8,X9,X9) | ~organization(X8,X9)) )),
  inference(subsumption_resolution,[],[f70,f24])).
fof(f70,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~greater(X5,X6) | greater(sK5,X7) | ~survival_chance(X8,X7,X9) | ~inertia(sK2,X5,sK9) | ~inertia(X8,X6,X9) | ~reorganization_free(X8,X9,X9) | ~organization(sK2,sK9) | ~organization(X8,X9)) )),
  inference(subsumption_resolution,[],[f67,f26])).
fof(f26,plain,(
  reorganization_free(sK2,sK9,sK9)),
  inference(cnf_transformation,[],[f19])).
fof(f67,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~greater(X5,X6) | greater(sK5,X7) | ~survival_chance(X8,X7,X9) | ~inertia(sK2,X5,sK9) | ~inertia(X8,X6,X9) | ~reorganization_free(sK2,sK9,sK9) | ~reorganization_free(X8,X9,X9) | ~organization(sK2,sK9) | ~organization(X8,X9)) )),
  inference(resolution,[],[f22,f30])).
fof(f30,plain,(
  survival_chance(sK2,sK5,sK9)),
  inference(cnf_transformation,[],[f19])).
fof(f22,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~survival_chance(X1,X7,X3) | ~greater(X5,X4) | greater(X7,X6) | ~survival_chance(X0,X6,X2) | ~inertia(X1,X5,X3) | ~inertia(X0,X4,X2) | ~reorganization_free(X1,X3,X3) | ~reorganization_free(X0,X2,X2) | ~organization(X1,X3) | ~organization(X0,X2)) )),
  inference(cnf_transformation,[],[f13])).
fof(f13,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (greater(X7,X6) | ~greater(X5,X4) | ~survival_chance(X1,X7,X3) | ~survival_chance(X0,X6,X2) | ~inertia(X1,X5,X3) | ~inertia(X0,X4,X2) | ~reorganization_free(X1,X3,X3) | ~reorganization_free(X0,X2,X2) | ~organization(X1,X3) | ~organization(X0,X2))),
  inference(flattening,[],[f12])).
fof(f12,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : (greater(X7,X6) | (~greater(X5,X4) | ~survival_chance(X1,X7,X3) | ~survival_chance(X0,X6,X2) | ~inertia(X1,X5,X3) | ~inertia(X0,X4,X2) | ~reorganization_free(X1,X3,X3) | ~reorganization_free(X0,X2,X2) | ~organization(X1,X3) | ~organization(X0,X2)))),
  inference(ennf_transformation,[],[f7])).
fof(f7,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : ((greater(X5,X4) & survival_chance(X1,X7,X3) & survival_chance(X0,X6,X2) & inertia(X1,X5,X3) & inertia(X0,X4,X2) & reorganization_free(X1,X3,X3) & reorganization_free(X0,X2,X2) & organization(X1,X3) & organization(X0,X2)) => greater(X7,X6))),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X0,X3,X9,X10,X7,X8,X11,X12] : ((greater(X8,X7) & survival_chance(X3,X12,X10) & survival_chance(X0,X11,X9) & inertia(X3,X8,X10) & inertia(X0,X7,X9) & reorganization_free(X3,X10,X10) & reorganization_free(X0,X9,X9) & organization(X3,X10) & organization(X0,X9)) => greater(X12,X11))),
  file('./p/MGT008+1.tptp',unknown)).
fof(f115,plain,(
  ( ! [X2,X0,X1] : (~inertia(sK2,X0,sK9) | ~inertia(sK1,X1,sK8) | greater(X0,X1) | ~class(sK2,X2,sK9) | ~class(sK1,X2,sK8)) )),
  inference(subsumption_resolution,[],[f114,f23])).
fof(f114,plain,(
  ( ! [X2,X0,X1] : (~inertia(sK2,X0,sK9) | ~inertia(sK1,X1,sK8) | greater(X0,X1) | ~class(sK2,X2,sK9) | ~class(sK1,X2,sK8) | ~organization(sK1,sK8)) )),
  inference(subsumption_resolution,[],[f111,f33])).
fof(f33,plain,(
  greater(sK7,sK6)),
  inference(cnf_transformation,[],[f19])).
fof(f111,plain,(
  ( ! [X2,X0,X1] : (~inertia(sK2,X0,sK9) | ~inertia(sK1,X1,sK8) | greater(X0,X1) | ~greater(sK7,sK6) | ~class(sK2,X2,sK9) | ~class(sK1,X2,sK8) | ~organization(sK1,sK8)) )),
  inference(resolution,[],[f38,f31])).
fof(f31,plain,(
  size(sK1,sK6,sK8)),
  inference(cnf_transformation,[],[f19])).
fof(f38,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~size(X8,X6,X10) | ~inertia(sK2,X7,sK9) | ~inertia(X8,X9,X10) | greater(X7,X9) | ~greater(sK7,X6) | ~class(sK2,X11,sK9) | ~class(X8,X11,X10) | ~organization(X8,X10)) )),
  inference(subsumption_resolution,[],[f36,f24])).
fof(f36,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~greater(sK7,X6) | ~inertia(sK2,X7,sK9) | ~inertia(X8,X9,X10) | greater(X7,X9) | ~size(X8,X6,X10) | ~class(sK2,X11,sK9) | ~class(X8,X11,X10) | ~organization(sK2,sK9) | ~organization(X8,X10)) )),
  inference(resolution,[],[f21,f32])).
fof(f32,plain,(
  size(sK2,sK7,sK9)),
  inference(cnf_transformation,[],[f19])).
fof(f21,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~size(X1,X4,X8) | ~greater(X4,X3) | ~inertia(X1,X6,X8) | ~inertia(X0,X5,X7) | greater(X6,X5) | ~size(X0,X3,X7) | ~class(X1,X2,X8) | ~class(X0,X2,X7) | ~organization(X1,X8) | ~organization(X0,X7)) )),
  inference(cnf_transformation,[],[f11])).
fof(f11,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (greater(X6,X5) | ~greater(X4,X3) | ~inertia(X1,X6,X8) | ~inertia(X0,X5,X7) | ~size(X1,X4,X8) | ~size(X0,X3,X7) | ~class(X1,X2,X8) | ~class(X0,X2,X7) | ~organization(X1,X8) | ~organization(X0,X7))),
  inference(flattening,[],[f10])).
fof(f10,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (greater(X6,X5) | (~greater(X4,X3) | ~inertia(X1,X6,X8) | ~inertia(X0,X5,X7) | ~size(X1,X4,X8) | ~size(X0,X3,X7) | ~class(X1,X2,X8) | ~class(X0,X2,X7) | ~organization(X1,X8) | ~organization(X0,X7)))),
  inference(ennf_transformation,[],[f6])).
fof(f6,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((greater(X4,X3) & inertia(X1,X6,X8) & inertia(X0,X5,X7) & size(X1,X4,X8) & size(X0,X3,X7) & class(X1,X2,X8) & class(X0,X2,X7) & organization(X1,X8) & organization(X0,X7)) => greater(X6,X5))),
  inference(rectify,[],[f2])).
fof(f2,axiom,(
  ! [X0,X3,X4,X5,X6,X7,X8,X9,X10] : ((greater(X6,X5) & inertia(X3,X8,X10) & inertia(X0,X7,X9) & size(X3,X6,X10) & size(X0,X5,X9) & class(X3,X4,X10) & class(X0,X4,X9) & organization(X3,X10) & organization(X0,X9)) => greater(X8,X7))),
  file('./p/MGT008+1.tptp',unknown)).
% SZS output end Proof for MGT008+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.007 s
% ------------------------------
% ------------------------------
