% Refutation found. Thanks to Tanya!
% SZS status Theorem for NUM459+1
% SZS output start Proof for NUM459+1
fof(f250439,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f662,f709,f7436,f219384,f219423,f250436])).
fof(f250436,plain,(
  ~spl1_13 | ~spl1_16 | spl1_44 | ~spl1_640),
  inference(avatar_contradiction_clause,[],[f250435])).
fof(f250435,plain,(
  $false | (~spl1_13 | ~spl1_16 | spl1_44 | ~spl1_640)),
  inference(subsumption_resolution,[],[f250434,f575])).
fof(f575,plain,(
  aNaturalNumber0(sdtmndt0(xn,xm)) | ~spl1_13),
  inference(avatar_component_clause,[],[f574])).
fof(f574,plain,(
  spl1_13 <=> aNaturalNumber0(sdtmndt0(xn,xm))),
  introduced(avatar_definition,[new_symbols(naming,[spl1_13])])).
fof(f250434,plain,(
  ~aNaturalNumber0(sdtmndt0(xn,xm)) | (~spl1_13 | ~spl1_16 | spl1_44 | ~spl1_640)),
  inference(subsumption_resolution,[],[f250433,f606])).
fof(f606,plain,(
  aNaturalNumber0(sdtmndt0(xm,xn)) | ~spl1_16),
  inference(avatar_component_clause,[],[f605])).
fof(f605,plain,(
  spl1_16 <=> aNaturalNumber0(sdtmndt0(xm,xn))),
  introduced(avatar_definition,[new_symbols(naming,[spl1_16])])).
fof(f250433,plain,(
  ~aNaturalNumber0(sdtmndt0(xm,xn)) | ~aNaturalNumber0(sdtmndt0(xn,xm)) | (~spl1_13 | ~spl1_16 | spl1_44 | ~spl1_640)),
  inference(subsumption_resolution,[],[f250423,f7307])).
fof(f7307,plain,(
  sz00 != sdtmndt0(xm,xn) | spl1_44),
  inference(avatar_component_clause,[],[f7306])).
fof(f7306,plain,(
  spl1_44 <=> sz00 = sdtmndt0(xm,xn)),
  introduced(avatar_definition,[new_symbols(naming,[spl1_44])])).
fof(f250423,plain,(
  sz00 = sdtmndt0(xm,xn) | ~aNaturalNumber0(sdtmndt0(xm,xn)) | ~aNaturalNumber0(sdtmndt0(xn,xm)) | (~spl1_13 | ~spl1_16 | ~spl1_640)),
  inference(trivial_inequality_removal,[],[f250421])).
fof(f250421,plain,(
  sz00 != sz00 | sz00 = sdtmndt0(xm,xn) | ~aNaturalNumber0(sdtmndt0(xm,xn)) | ~aNaturalNumber0(sdtmndt0(xn,xm)) | (~spl1_13 | ~spl1_16 | ~spl1_640)),
  inference(superposition,[],[f85,f219426])).
fof(f219426,plain,(
  sz00 = sdtpldt0(sdtmndt0(xn,xm),sdtmndt0(xm,xn)) | (~spl1_13 | ~spl1_16 | ~spl1_640)),
  inference(backward_demodulation,[],[f219383,f76244])).
fof(f76244,plain,(
  sdtpldt0(sdtmndt0(xn,xm),sdtmndt0(xm,xn)) = sdtpldt0(sdtmndt0(xm,xn),sdtmndt0(xn,xm)) | (~spl1_13 | ~spl1_16)),
  inference(resolution,[],[f675,f606])).
fof(f675,plain,(
  ( ! [X0] : (~aNaturalNumber0(X0) | sdtpldt0(X0,sdtmndt0(xn,xm)) = sdtpldt0(sdtmndt0(xn,xm),X0)) ) | ~spl1_13),
  inference(resolution,[],[f575,f68])).
fof(f68,plain,(
  ( ! [X0,X1] : (~aNaturalNumber0(X1) | sdtpldt0(X0,X1) = sdtpldt0(X1,X0) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : (sdtpldt0(X0,X1) = sdtpldt0(X1,X0) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : (sdtpldt0(X0,X1) = sdtpldt0(X1,X0) | (~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0,X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => sdtpldt0(X0,X1) = sdtpldt0(X1,X0))),
  file('./p/NUM459+1.tptp',unknown)).
fof(f219383,plain,(
  sz00 = sdtpldt0(sdtmndt0(xm,xn),sdtmndt0(xn,xm)) | ~spl1_640),
  inference(avatar_component_clause,[],[f219381])).
fof(f219381,plain,(
  spl1_640 <=> sz00 = sdtpldt0(sdtmndt0(xm,xn),sdtmndt0(xn,xm))),
  introduced(avatar_definition,[new_symbols(naming,[spl1_640])])).
fof(f85,plain,(
  ( ! [X0,X1] : (sz00 != sdtpldt0(X0,X1) | sz00 = X1 | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ((sz00 = X1 & sz00 = X0) | sz00 != sdtpldt0(X0,X1) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : (((sz00 = X1 & sz00 = X0) | sz00 != sdtpldt0(X0,X1)) | (~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f16])).
fof(f16,axiom,(
  ! [X0,X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (sz00 = sdtpldt0(X0,X1) => (sz00 = X1 & sz00 = X0)))),
  file('./p/NUM459+1.tptp',unknown)).
fof(f219423,plain,(
  ~spl1_13 | ~spl1_16 | spl1_639),
  inference(avatar_contradiction_clause,[],[f219422])).
fof(f219422,plain,(
  $false | (~spl1_13 | ~spl1_16 | spl1_639)),
  inference(subsumption_resolution,[],[f219421,f606])).
fof(f219421,plain,(
  ~aNaturalNumber0(sdtmndt0(xm,xn)) | (~spl1_13 | spl1_639)),
  inference(subsumption_resolution,[],[f219420,f575])).
fof(f219420,plain,(
  ~aNaturalNumber0(sdtmndt0(xn,xm)) | ~aNaturalNumber0(sdtmndt0(xm,xn)) | spl1_639),
  inference(resolution,[],[f219379,f66])).
fof(f66,plain,(
  ( ! [X0,X1] : (aNaturalNumber0(sdtpldt0(X0,X1)) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : (aNaturalNumber0(sdtpldt0(X0,X1)) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : (aNaturalNumber0(sdtpldt0(X0,X1)) | (~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => aNaturalNumber0(sdtpldt0(X0,X1)))),
  file('./p/NUM459+1.tptp',unknown)).
fof(f219379,plain,(
  ~aNaturalNumber0(sdtpldt0(sdtmndt0(xm,xn),sdtmndt0(xn,xm))) | spl1_639),
  inference(avatar_component_clause,[],[f219377])).
fof(f219377,plain,(
  spl1_639 <=> aNaturalNumber0(sdtpldt0(sdtmndt0(xm,xn),sdtmndt0(xn,xm)))),
  introduced(avatar_definition,[new_symbols(naming,[spl1_639])])).
fof(f219384,plain,(
  ~spl1_639 | spl1_640 | ~spl1_13 | ~spl1_16),
  inference(avatar_split_clause,[],[f219375,f605,f574,f219381,f219377])).
fof(f219375,plain,(
  sz00 = sdtpldt0(sdtmndt0(xm,xn),sdtmndt0(xn,xm)) | ~aNaturalNumber0(sdtpldt0(sdtmndt0(xm,xn),sdtmndt0(xn,xm))) | (~spl1_13 | ~spl1_16)),
  inference(trivial_inequality_removal,[],[f219366])).
fof(f219366,plain,(
  xm != xm | sz00 = sdtpldt0(sdtmndt0(xm,xn),sdtmndt0(xn,xm)) | ~aNaturalNumber0(sdtpldt0(sdtmndt0(xm,xn),sdtmndt0(xn,xm))) | (~spl1_13 | ~spl1_16)),
  inference(superposition,[],[f333,f76723])).
fof(f76723,plain,(
  xm = sdtpldt0(sdtpldt0(sdtmndt0(xm,xn),sdtmndt0(xn,xm)),xm) | (~spl1_13 | ~spl1_16)),
  inference(forward_demodulation,[],[f76659,f745])).
fof(f745,plain,(
  xm = sdtpldt0(sdtmndt0(xm,xn),xn) | ~spl1_16),
  inference(forward_demodulation,[],[f741,f255])).
fof(f255,plain,(
  xm = sdtpldt0(xn,sdtmndt0(xm,xn))),
  inference(subsumption_resolution,[],[f254,f95])).
fof(f95,plain,(
  aNaturalNumber0(xn)),
  inference(cnf_transformation,[],[f21])).
fof(f21,axiom,(
  aNaturalNumber0(xn) & aNaturalNumber0(xm)),
  file('./p/NUM459+1.tptp',unknown)).
fof(f254,plain,(
  xm = sdtpldt0(xn,sdtmndt0(xm,xn)) | ~aNaturalNumber0(xn)),
  inference(subsumption_resolution,[],[f244,f94])).
fof(f94,plain,(
  aNaturalNumber0(xm)),
  inference(cnf_transformation,[],[f21])).
fof(f244,plain,(
  xm = sdtpldt0(xn,sdtmndt0(xm,xn)) | ~aNaturalNumber0(xm) | ~aNaturalNumber0(xn)),
  inference(resolution,[],[f101,f97])).
fof(f97,plain,(
  sdtlseqdt0(xn,xm)),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  xm != xn & sdtlseqdt0(xn,xm) & sdtlseqdt0(xm,xn)),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  xm != xn & (sdtlseqdt0(xn,xm) & sdtlseqdt0(xm,xn))),
  inference(ennf_transformation,[],[f23])).
fof(f23,negated_conjecture,(
  ~((sdtlseqdt0(xn,xm) & sdtlseqdt0(xm,xn)) => xm = xn)),
  inference(negated_conjecture,[],[f22])).
fof(f22,conjecture,(
  (sdtlseqdt0(xn,xm) & sdtlseqdt0(xm,xn)) => xm = xn),
  file('./p/NUM459+1.tptp',unknown)).
fof(f101,plain,(
  ( ! [X0,X1] : (~sdtlseqdt0(X0,X1) | sdtpldt0(X0,sdtmndt0(X1,X0)) = X1 | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(equality_resolution,[],[f91])).
fof(f91,plain,(
  ( ! [X2,X0,X1] : (sdtpldt0(X0,X2) = X1 | sdtmndt0(X1,X0) != X2 | ~sdtlseqdt0(X0,X1) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : (! [X2] : ((sdtmndt0(X1,X0) = X2 | sdtpldt0(X0,X2) != X1 | ~aNaturalNumber0(X2)) & ((sdtpldt0(X0,X2) = X1 & aNaturalNumber0(X2)) | sdtmndt0(X1,X0) != X2)) | ~sdtlseqdt0(X0,X1) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : (! [X2] : ((sdtmndt0(X1,X0) = X2 | (sdtpldt0(X0,X2) != X1 | ~aNaturalNumber0(X2))) & ((sdtpldt0(X0,X2) = X1 & aNaturalNumber0(X2)) | sdtmndt0(X1,X0) != X2)) | ~sdtlseqdt0(X0,X1) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (! [X2] : (sdtmndt0(X1,X0) = X2 <=> (sdtpldt0(X0,X2) = X1 & aNaturalNumber0(X2))) | ~sdtlseqdt0(X0,X1) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : ((! [X2] : (sdtmndt0(X1,X0) = X2 <=> (sdtpldt0(X0,X2) = X1 & aNaturalNumber0(X2))) | ~sdtlseqdt0(X0,X1)) | (~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (sdtlseqdt0(X0,X1) => ! [X2] : (sdtmndt0(X1,X0) = X2 <=> (sdtpldt0(X0,X2) = X1 & aNaturalNumber0(X2)))))),
  file('./p/NUM459+1.tptp',unknown)).
fof(f741,plain,(
  sdtpldt0(xn,sdtmndt0(xm,xn)) = sdtpldt0(sdtmndt0(xm,xn),xn) | ~spl1_16),
  inference(resolution,[],[f606,f163])).
fof(f163,plain,(
  ( ! [X9] : (~aNaturalNumber0(X9) | sdtpldt0(X9,xn) = sdtpldt0(xn,X9)) )),
  inference(resolution,[],[f68,f95])).
fof(f76659,plain,(
  sdtpldt0(sdtmndt0(xm,xn),xn) = sdtpldt0(sdtpldt0(sdtmndt0(xm,xn),sdtmndt0(xn,xm)),xm) | (~spl1_13 | ~spl1_16)),
  inference(resolution,[],[f1452,f606])).
fof(f1452,plain,(
  ( ! [X14] : (~aNaturalNumber0(X14) | sdtpldt0(sdtpldt0(X14,sdtmndt0(xn,xm)),xm) = sdtpldt0(X14,xn)) ) | ~spl1_13),
  inference(forward_demodulation,[],[f1441,f691])).
fof(f691,plain,(
  xn = sdtpldt0(sdtmndt0(xn,xm),xm) | ~spl1_13),
  inference(forward_demodulation,[],[f688,f253])).
fof(f253,plain,(
  xn = sdtpldt0(xm,sdtmndt0(xn,xm))),
  inference(subsumption_resolution,[],[f252,f94])).
fof(f252,plain,(
  xn = sdtpldt0(xm,sdtmndt0(xn,xm)) | ~aNaturalNumber0(xm)),
  inference(subsumption_resolution,[],[f243,f95])).
fof(f243,plain,(
  xn = sdtpldt0(xm,sdtmndt0(xn,xm)) | ~aNaturalNumber0(xn) | ~aNaturalNumber0(xm)),
  inference(resolution,[],[f101,f96])).
fof(f96,plain,(
  sdtlseqdt0(xm,xn)),
  inference(cnf_transformation,[],[f56])).
fof(f688,plain,(
  sdtpldt0(xm,sdtmndt0(xn,xm)) = sdtpldt0(sdtmndt0(xn,xm),xm) | ~spl1_13),
  inference(resolution,[],[f575,f162])).
fof(f162,plain,(
  ( ! [X8] : (~aNaturalNumber0(X8) | sdtpldt0(X8,xm) = sdtpldt0(xm,X8)) )),
  inference(resolution,[],[f68,f94])).
fof(f1441,plain,(
  ( ! [X14] : (sdtpldt0(sdtpldt0(X14,sdtmndt0(xn,xm)),xm) = sdtpldt0(X14,sdtpldt0(sdtmndt0(xn,xm),xm)) | ~aNaturalNumber0(X14)) ) | ~spl1_13),
  inference(resolution,[],[f343,f575])).
fof(f343,plain,(
  ( ! [X17,X16] : (~aNaturalNumber0(X17) | sdtpldt0(sdtpldt0(X16,X17),xm) = sdtpldt0(X16,sdtpldt0(X17,xm)) | ~aNaturalNumber0(X16)) )),
  inference(resolution,[],[f69,f94])).
fof(f69,plain,(
  ( ! [X2,X0,X1] : (~aNaturalNumber0(X2) | sdtpldt0(sdtpldt0(X0,X1),X2) = sdtpldt0(X0,sdtpldt0(X1,X2)) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2] : (sdtpldt0(sdtpldt0(X0,X1),X2) = sdtpldt0(X0,sdtpldt0(X1,X2)) | ~aNaturalNumber0(X2) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2] : (sdtpldt0(sdtpldt0(X0,X1),X2) = sdtpldt0(X0,sdtpldt0(X1,X2)) | (~aNaturalNumber0(X2) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1,X2] : ((aNaturalNumber0(X2) & aNaturalNumber0(X1) & aNaturalNumber0(X0)) => sdtpldt0(sdtpldt0(X0,X1),X2) = sdtpldt0(X0,sdtpldt0(X1,X2)))),
  file('./p/NUM459+1.tptp',unknown)).
fof(f333,plain,(
  ( ! [X5] : (xm != sdtpldt0(X5,xm) | sz00 = X5 | ~aNaturalNumber0(X5)) )),
  inference(subsumption_resolution,[],[f332,f94])).
fof(f332,plain,(
  ( ! [X5] : (xm != sdtpldt0(X5,xm) | sz00 = X5 | ~aNaturalNumber0(X5) | ~aNaturalNumber0(xm)) )),
  inference(subsumption_resolution,[],[f311,f63])).
fof(f63,plain,(
  aNaturalNumber0(sz00)),
  inference(cnf_transformation,[],[f2])).
fof(f2,axiom,(
  aNaturalNumber0(sz00)),
  file('./p/NUM459+1.tptp',unknown)).
fof(f311,plain,(
  ( ! [X5] : (xm != sdtpldt0(X5,xm) | sz00 = X5 | ~aNaturalNumber0(X5) | ~aNaturalNumber0(sz00) | ~aNaturalNumber0(xm)) )),
  inference(superposition,[],[f81,f112])).
fof(f112,plain,(
  xm = sdtpldt0(sz00,xm)),
  inference(resolution,[],[f71,f94])).
fof(f71,plain,(
  ( ! [X0] : (~aNaturalNumber0(X0) | sdtpldt0(sz00,X0) = X0) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0] : ((sdtpldt0(sz00,X0) = X0 & sdtpldt0(X0,sz00) = X0) | ~aNaturalNumber0(X0))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0] : (aNaturalNumber0(X0) => (sdtpldt0(sz00,X0) = X0 & sdtpldt0(X0,sz00) = X0))),
  file('./p/NUM459+1.tptp',unknown)).
fof(f81,plain,(
  ( ! [X2,X0,X1] : (sdtpldt0(X1,X0) != sdtpldt0(X2,X0) | X1 = X2 | ~aNaturalNumber0(X2) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2] : (X1 = X2 | (sdtpldt0(X1,X0) != sdtpldt0(X2,X0) & sdtpldt0(X0,X1) != sdtpldt0(X0,X2)) | ~aNaturalNumber0(X2) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2] : ((X1 = X2 | (sdtpldt0(X1,X0) != sdtpldt0(X2,X0) & sdtpldt0(X0,X1) != sdtpldt0(X0,X2))) | (~aNaturalNumber0(X2) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f14])).
fof(f14,axiom,(
  ! [X0,X1,X2] : ((aNaturalNumber0(X2) & aNaturalNumber0(X1) & aNaturalNumber0(X0)) => ((sdtpldt0(X1,X0) = sdtpldt0(X2,X0) | sdtpldt0(X0,X1) = sdtpldt0(X0,X2)) => X1 = X2))),
  file('./p/NUM459+1.tptp',unknown)).
fof(f7436,plain,(
  ~spl1_44),
  inference(avatar_contradiction_clause,[],[f7435])).
fof(f7435,plain,(
  $false | ~spl1_44),
  inference(subsumption_resolution,[],[f7434,f98])).
fof(f98,plain,(
  xm != xn),
  inference(cnf_transformation,[],[f56])).
fof(f7434,plain,(
  xm = xn | ~spl1_44),
  inference(backward_demodulation,[],[f7316,f109])).
fof(f109,plain,(
  xn = sdtpldt0(xn,sz00)),
  inference(resolution,[],[f70,f95])).
fof(f70,plain,(
  ( ! [X0] : (~aNaturalNumber0(X0) | sdtpldt0(X0,sz00) = X0) )),
  inference(cnf_transformation,[],[f33])).
fof(f7316,plain,(
  xm = sdtpldt0(xn,sz00) | ~spl1_44),
  inference(backward_demodulation,[],[f7308,f255])).
fof(f7308,plain,(
  sz00 = sdtmndt0(xm,xn) | ~spl1_44),
  inference(avatar_component_clause,[],[f7306])).
fof(f709,plain,(
  spl1_16),
  inference(avatar_contradiction_clause,[],[f708])).
fof(f708,plain,(
  $false | spl1_16),
  inference(subsumption_resolution,[],[f707,f95])).
fof(f707,plain,(
  ~aNaturalNumber0(xn) | spl1_16),
  inference(subsumption_resolution,[],[f706,f94])).
fof(f706,plain,(
  ~aNaturalNumber0(xm) | ~aNaturalNumber0(xn) | spl1_16),
  inference(subsumption_resolution,[],[f705,f97])).
fof(f705,plain,(
  ~sdtlseqdt0(xn,xm) | ~aNaturalNumber0(xm) | ~aNaturalNumber0(xn) | spl1_16),
  inference(resolution,[],[f607,f102])).
fof(f102,plain,(
  ( ! [X0,X1] : (aNaturalNumber0(sdtmndt0(X1,X0)) | ~sdtlseqdt0(X0,X1) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(equality_resolution,[],[f90])).
fof(f90,plain,(
  ( ! [X2,X0,X1] : (aNaturalNumber0(X2) | sdtmndt0(X1,X0) != X2 | ~sdtlseqdt0(X0,X1) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f62])).
fof(f607,plain,(
  ~aNaturalNumber0(sdtmndt0(xm,xn)) | spl1_16),
  inference(avatar_component_clause,[],[f605])).
fof(f662,plain,(
  spl1_13),
  inference(avatar_contradiction_clause,[],[f661])).
fof(f661,plain,(
  $false | spl1_13),
  inference(subsumption_resolution,[],[f660,f94])).
fof(f660,plain,(
  ~aNaturalNumber0(xm) | spl1_13),
  inference(subsumption_resolution,[],[f659,f95])).
fof(f659,plain,(
  ~aNaturalNumber0(xn) | ~aNaturalNumber0(xm) | spl1_13),
  inference(subsumption_resolution,[],[f658,f96])).
fof(f658,plain,(
  ~sdtlseqdt0(xm,xn) | ~aNaturalNumber0(xn) | ~aNaturalNumber0(xm) | spl1_13),
  inference(resolution,[],[f576,f102])).
fof(f576,plain,(
  ~aNaturalNumber0(sdtmndt0(xn,xm)) | spl1_13),
  inference(avatar_component_clause,[],[f574])).
% SZS output end Proof for NUM459+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 128313
% Time elapsed: 4.234 s
% ------------------------------
% ------------------------------
