% Refutation found. Thanks to Tanya!
% SZS status Theorem for NUM471+2
% SZS output start Proof for NUM471+2
fof(f6058,plain,(
  $false),
  inference(subsumption_resolution,[],[f6057,f168])).
fof(f168,plain,(
  aNaturalNumber0(xm)),
  inference(cnf_transformation,[],[f34])).
fof(f34,axiom,(
  aNaturalNumber0(xn) & aNaturalNumber0(xm) & aNaturalNumber0(xl)),
  file('./p/NUM471+2.tptp',unknown)).
fof(f6057,plain,(
  ~aNaturalNumber0(xm)),
  inference(subsumption_resolution,[],[f6054,f169])).
fof(f169,plain,(
  aNaturalNumber0(xn)),
  inference(cnf_transformation,[],[f34])).
fof(f6054,plain,(
  ~aNaturalNumber0(xn) | ~aNaturalNumber0(xm)),
  inference(resolution,[],[f5722,f196])).
fof(f196,plain,(
  ( ! [X2,X0] : (sdtlseqdt0(X0,sdtpldt0(X0,X2)) | ~aNaturalNumber0(X2) | ~aNaturalNumber0(X0)) )),
  inference(subsumption_resolution,[],[f185,f116])).
fof(f116,plain,(
  ( ! [X0,X1] : (aNaturalNumber0(sdtpldt0(X0,X1)) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : (aNaturalNumber0(sdtpldt0(X0,X1)) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : (aNaturalNumber0(sdtpldt0(X0,X1)) | (~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => aNaturalNumber0(sdtpldt0(X0,X1)))),
  file('./p/NUM471+2.tptp',unknown)).
fof(f185,plain,(
  ( ! [X2,X0] : (sdtlseqdt0(X0,sdtpldt0(X0,X2)) | ~aNaturalNumber0(X2) | ~aNaturalNumber0(sdtpldt0(X0,X2)) | ~aNaturalNumber0(X0)) )),
  inference(equality_resolution,[],[f139])).
fof(f139,plain,(
  ( ! [X2,X0,X1] : (sdtlseqdt0(X0,X1) | sdtpldt0(X0,X2) != X1 | ~aNaturalNumber0(X2) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : (((sdtlseqdt0(X0,X1) | ! [X2] : (sdtpldt0(X0,X2) != X1 | ~aNaturalNumber0(X2))) & ((sdtpldt0(X0,sK0(X0,X1)) = X1 & aNaturalNumber0(sK0(X0,X1))) | ~sdtlseqdt0(X0,X1))) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f99,f100])).
fof(f100,plain,(
  ! [X1,X0] : (? [X3] : (sdtpldt0(X0,X3) = X1 & aNaturalNumber0(X3)) => (sdtpldt0(X0,sK0(X0,X1)) = X1 & aNaturalNumber0(sK0(X0,X1))))),
  introduced(choice_axiom,[])).
fof(f99,plain,(
  ! [X0,X1] : (((sdtlseqdt0(X0,X1) | ! [X2] : (sdtpldt0(X0,X2) != X1 | ~aNaturalNumber0(X2))) & (? [X3] : (sdtpldt0(X0,X3) = X1 & aNaturalNumber0(X3)) | ~sdtlseqdt0(X0,X1))) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(rectify,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : (((sdtlseqdt0(X0,X1) | ! [X2] : (sdtpldt0(X0,X2) != X1 | ~aNaturalNumber0(X2))) & (? [X2] : (sdtpldt0(X0,X2) = X1 & aNaturalNumber0(X2)) | ~sdtlseqdt0(X0,X1))) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ((sdtlseqdt0(X0,X1) <=> ? [X2] : (sdtpldt0(X0,X2) = X1 & aNaturalNumber0(X2))) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : ((sdtlseqdt0(X0,X1) <=> ? [X2] : (sdtpldt0(X0,X2) = X1 & aNaturalNumber0(X2))) | (~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (sdtlseqdt0(X0,X1) <=> ? [X2] : (sdtpldt0(X0,X2) = X1 & aNaturalNumber0(X2))))),
  file('./p/NUM471+2.tptp',unknown)).
fof(f5722,plain,(
  ~sdtlseqdt0(xm,sdtpldt0(xm,xn))),
  inference(subsumption_resolution,[],[f5721,f168])).
fof(f5721,plain,(
  ~sdtlseqdt0(xm,sdtpldt0(xm,xn)) | ~aNaturalNumber0(xm)),
  inference(subsumption_resolution,[],[f5720,f280])).
fof(f280,plain,(
  aNaturalNumber0(sdtpldt0(xm,xn))),
  inference(subsumption_resolution,[],[f279,f167])).
fof(f167,plain,(
  aNaturalNumber0(xl)),
  inference(cnf_transformation,[],[f34])).
fof(f279,plain,(
  aNaturalNumber0(sdtpldt0(xm,xn)) | ~aNaturalNumber0(xl)),
  inference(subsumption_resolution,[],[f276,f180])).
fof(f180,plain,(
  aNaturalNumber0(xq)),
  inference(cnf_transformation,[],[f38])).
fof(f38,axiom,(
  xq = sdtsldt0(sdtpldt0(xm,xn),xl) & sdtpldt0(xm,xn) = sdtasdt0(xl,xq) & aNaturalNumber0(xq)),
  file('./p/NUM471+2.tptp',unknown)).
fof(f276,plain,(
  aNaturalNumber0(sdtpldt0(xm,xn)) | ~aNaturalNumber0(xq) | ~aNaturalNumber0(xl)),
  inference(superposition,[],[f117,f181])).
fof(f181,plain,(
  sdtpldt0(xm,xn) = sdtasdt0(xl,xq)),
  inference(cnf_transformation,[],[f38])).
fof(f117,plain,(
  ( ! [X0,X1] : (aNaturalNumber0(sdtasdt0(X0,X1)) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : (aNaturalNumber0(sdtasdt0(X0,X1)) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : (aNaturalNumber0(sdtasdt0(X0,X1)) | (~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => aNaturalNumber0(sdtasdt0(X0,X1)))),
  file('./p/NUM471+2.tptp',unknown)).
fof(f5720,plain,(
  ~sdtlseqdt0(xm,sdtpldt0(xm,xn)) | ~aNaturalNumber0(sdtpldt0(xm,xn)) | ~aNaturalNumber0(xm)),
  inference(subsumption_resolution,[],[f5713,f4521])).
fof(f4521,plain,(
  xm != sdtpldt0(xm,xn)),
  inference(subsumption_resolution,[],[f4520,f180])).
fof(f4520,plain,(
  xm != sdtpldt0(xm,xn) | ~aNaturalNumber0(xq)),
  inference(subsumption_resolution,[],[f4512,f248])).
fof(f248,plain,(
  xp != xq),
  inference(subsumption_resolution,[],[f247,f113])).
fof(f113,plain,(
  aNaturalNumber0(sz00)),
  inference(cnf_transformation,[],[f2])).
fof(f2,axiom,(
  aNaturalNumber0(sz00)),
  file('./p/NUM471+2.tptp',unknown)).
fof(f247,plain,(
  xp != xq | ~aNaturalNumber0(sz00)),
  inference(superposition,[],[f183,f207])).
fof(f207,plain,(
  xp = sdtpldt0(xp,sz00)),
  inference(resolution,[],[f120,f177])).
fof(f177,plain,(
  aNaturalNumber0(xp)),
  inference(cnf_transformation,[],[f37])).
fof(f37,axiom,(
  xp = sdtsldt0(xm,xl) & xm = sdtasdt0(xl,xp) & aNaturalNumber0(xp)),
  file('./p/NUM471+2.tptp',unknown)).
fof(f120,plain,(
  ( ! [X0] : (~aNaturalNumber0(X0) | sdtpldt0(X0,sz00) = X0) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : ((sdtpldt0(sz00,X0) = X0 & sdtpldt0(X0,sz00) = X0) | ~aNaturalNumber0(X0))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0] : (aNaturalNumber0(X0) => (sdtpldt0(sz00,X0) = X0 & sdtpldt0(X0,sz00) = X0))),
  file('./p/NUM471+2.tptp',unknown)).
fof(f183,plain,(
  ( ! [X0] : (xq != sdtpldt0(xp,X0) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ~sdtlseqdt0(xp,xq) & ! [X0] : (xq != sdtpldt0(xp,X0) | ~aNaturalNumber0(X0))),
  inference(ennf_transformation,[],[f40])).
fof(f40,negated_conjecture,(
  ~(sdtlseqdt0(xp,xq) | ? [X0] : (xq = sdtpldt0(xp,X0) & aNaturalNumber0(X0)))),
  inference(negated_conjecture,[],[f39])).
fof(f39,conjecture,(
  sdtlseqdt0(xp,xq) | ? [X0] : (xq = sdtpldt0(xp,X0) & aNaturalNumber0(X0))),
  file('./p/NUM471+2.tptp',unknown)).
fof(f4512,plain,(
  xm != sdtpldt0(xm,xn) | xp = xq | ~aNaturalNumber0(xq)),
  inference(superposition,[],[f1020,f181])).
fof(f1020,plain,(
  ( ! [X1] : (xm != sdtasdt0(xl,X1) | xp = X1 | ~aNaturalNumber0(X1)) )),
  inference(subsumption_resolution,[],[f1019,f167])).
fof(f1019,plain,(
  ( ! [X1] : (xm != sdtasdt0(xl,X1) | xp = X1 | ~aNaturalNumber0(X1) | ~aNaturalNumber0(xl)) )),
  inference(subsumption_resolution,[],[f1018,f176])).
fof(f176,plain,(
  sz00 != xl),
  inference(cnf_transformation,[],[f36])).
fof(f36,axiom,(
  sz00 != xl),
  file('./p/NUM471+2.tptp',unknown)).
fof(f1018,plain,(
  ( ! [X1] : (xm != sdtasdt0(xl,X1) | xp = X1 | ~aNaturalNumber0(X1) | sz00 = xl | ~aNaturalNumber0(xl)) )),
  inference(subsumption_resolution,[],[f1002,f177])).
fof(f1002,plain,(
  ( ! [X1] : (xm != sdtasdt0(xl,X1) | xp = X1 | ~aNaturalNumber0(X1) | ~aNaturalNumber0(xp) | sz00 = xl | ~aNaturalNumber0(xl)) )),
  inference(superposition,[],[f132,f178])).
fof(f178,plain,(
  xm = sdtasdt0(xl,xp)),
  inference(cnf_transformation,[],[f37])).
fof(f132,plain,(
  ( ! [X2,X0,X1] : (sdtasdt0(X0,X1) != sdtasdt0(X0,X2) | X1 = X2 | ~aNaturalNumber0(X2) | ~aNaturalNumber0(X1) | sz00 = X0 | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : (! [X1,X2] : (X1 = X2 | (sdtasdt0(X1,X0) != sdtasdt0(X2,X0) & sdtasdt0(X0,X1) != sdtasdt0(X0,X2)) | ~aNaturalNumber0(X2) | ~aNaturalNumber0(X1)) | sz00 = X0 | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ((! [X1,X2] : ((X1 = X2 | (sdtasdt0(X1,X0) != sdtasdt0(X2,X0) & sdtasdt0(X0,X1) != sdtasdt0(X0,X2))) | (~aNaturalNumber0(X2) | ~aNaturalNumber0(X1))) | sz00 = X0) | ~aNaturalNumber0(X0))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0] : (aNaturalNumber0(X0) => (sz00 != X0 => ! [X1,X2] : ((aNaturalNumber0(X2) & aNaturalNumber0(X1)) => ((sdtasdt0(X1,X0) = sdtasdt0(X2,X0) | sdtasdt0(X0,X1) = sdtasdt0(X0,X2)) => X1 = X2))))),
  file('./p/NUM471+2.tptp',unknown)).
fof(f5713,plain,(
  xm = sdtpldt0(xm,xn) | ~sdtlseqdt0(xm,sdtpldt0(xm,xn)) | ~aNaturalNumber0(sdtpldt0(xm,xn)) | ~aNaturalNumber0(xm)),
  inference(resolution,[],[f5709,f144])).
fof(f144,plain,(
  ( ! [X0,X1] : (~sdtlseqdt0(X1,X0) | X0 = X1 | ~sdtlseqdt0(X0,X1) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1] : (X0 = X1 | ~sdtlseqdt0(X1,X0) | ~sdtlseqdt0(X0,X1) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ((X0 = X1 | (~sdtlseqdt0(X1,X0) | ~sdtlseqdt0(X0,X1))) | (~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f21])).
fof(f21,axiom,(
  ! [X0,X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => ((sdtlseqdt0(X1,X0) & sdtlseqdt0(X0,X1)) => X0 = X1))),
  file('./p/NUM471+2.tptp',unknown)).
fof(f5709,plain,(
  sdtlseqdt0(sdtpldt0(xm,xn),xm)),
  inference(subsumption_resolution,[],[f5708,f180])).
fof(f5708,plain,(
  sdtlseqdt0(sdtpldt0(xm,xn),xm) | ~aNaturalNumber0(xq)),
  inference(subsumption_resolution,[],[f5707,f248])).
fof(f5707,plain,(
  sdtlseqdt0(sdtpldt0(xm,xn),xm) | xp = xq | ~aNaturalNumber0(xq)),
  inference(subsumption_resolution,[],[f5700,f305])).
fof(f305,plain,(
  sdtlseqdt0(xq,xp)),
  inference(subsumption_resolution,[],[f304,f180])).
fof(f304,plain,(
  sdtlseqdt0(xq,xp) | ~aNaturalNumber0(xq)),
  inference(subsumption_resolution,[],[f300,f177])).
fof(f300,plain,(
  sdtlseqdt0(xq,xp) | ~aNaturalNumber0(xp) | ~aNaturalNumber0(xq)),
  inference(resolution,[],[f147,f184])).
fof(f184,plain,(
  ~sdtlseqdt0(xp,xq)),
  inference(cnf_transformation,[],[f97])).
fof(f147,plain,(
  ( ! [X0,X1] : (sdtlseqdt0(X1,X0) | sdtlseqdt0(X0,X1) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : ((sdtlseqdt0(X1,X0) & X0 != X1) | sdtlseqdt0(X0,X1) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : (((sdtlseqdt0(X1,X0) & X0 != X1) | sdtlseqdt0(X0,X1)) | (~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f23])).
fof(f23,axiom,(
  ! [X0,X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => ((sdtlseqdt0(X1,X0) & X0 != X1) | sdtlseqdt0(X0,X1)))),
  file('./p/NUM471+2.tptp',unknown)).
fof(f5700,plain,(
  sdtlseqdt0(sdtpldt0(xm,xn),xm) | ~sdtlseqdt0(xq,xp) | xp = xq | ~aNaturalNumber0(xq)),
  inference(superposition,[],[f1619,f181])).
fof(f1619,plain,(
  ( ! [X1] : (sdtlseqdt0(sdtasdt0(xl,X1),xm) | ~sdtlseqdt0(X1,xp) | xp = X1 | ~aNaturalNumber0(X1)) )),
  inference(subsumption_resolution,[],[f1618,f167])).
fof(f1618,plain,(
  ( ! [X1] : (sdtlseqdt0(sdtasdt0(xl,X1),xm) | ~sdtlseqdt0(X1,xp) | xp = X1 | ~aNaturalNumber0(X1) | ~aNaturalNumber0(xl)) )),
  inference(subsumption_resolution,[],[f1617,f177])).
fof(f1617,plain,(
  ( ! [X1] : (sdtlseqdt0(sdtasdt0(xl,X1),xm) | ~sdtlseqdt0(X1,xp) | xp = X1 | ~aNaturalNumber0(xp) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(xl)) )),
  inference(subsumption_resolution,[],[f1586,f176])).
fof(f1586,plain,(
  ( ! [X1] : (sdtlseqdt0(sdtasdt0(xl,X1),xm) | ~sdtlseqdt0(X1,xp) | xp = X1 | sz00 = xl | ~aNaturalNumber0(xp) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(xl)) )),
  inference(superposition,[],[f153,f178])).
fof(f153,plain,(
  ( ! [X2,X0,X1] : (sdtlseqdt0(sdtasdt0(X0,X1),sdtasdt0(X0,X2)) | ~sdtlseqdt0(X1,X2) | X1 = X2 | sz00 = X0 | ~aNaturalNumber0(X2) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2] : ((sdtlseqdt0(sdtasdt0(X1,X0),sdtasdt0(X2,X0)) & sdtasdt0(X1,X0) != sdtasdt0(X2,X0) & sdtlseqdt0(sdtasdt0(X0,X1),sdtasdt0(X0,X2)) & sdtasdt0(X0,X1) != sdtasdt0(X0,X2)) | ~sdtlseqdt0(X1,X2) | X1 = X2 | sz00 = X0 | ~aNaturalNumber0(X2) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2] : (((sdtlseqdt0(sdtasdt0(X1,X0),sdtasdt0(X2,X0)) & sdtasdt0(X1,X0) != sdtasdt0(X2,X0) & sdtlseqdt0(sdtasdt0(X0,X1),sdtasdt0(X0,X2)) & sdtasdt0(X0,X1) != sdtasdt0(X0,X2)) | (~sdtlseqdt0(X1,X2) | X1 = X2 | sz00 = X0)) | (~aNaturalNumber0(X2) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0,X1,X2] : ((aNaturalNumber0(X2) & aNaturalNumber0(X1) & aNaturalNumber0(X0)) => ((sdtlseqdt0(X1,X2) & X1 != X2 & sz00 != X0) => (sdtlseqdt0(sdtasdt0(X1,X0),sdtasdt0(X2,X0)) & sdtasdt0(X1,X0) != sdtasdt0(X2,X0) & sdtlseqdt0(sdtasdt0(X0,X1),sdtasdt0(X0,X2)) & sdtasdt0(X0,X1) != sdtasdt0(X0,X2))))),
  file('./p/NUM471+2.tptp',unknown)).
% SZS output end Proof for NUM471+2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 8059
% Time elapsed: 0.054 s
% ------------------------------
% ------------------------------
