% Refutation found. Thanks to Tanya!
% SZS status Theorem for NUM432+1
% SZS output start Proof for NUM432+1
fof(f4303,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f324,f4301])).
fof(f4301,plain,(
  ~spl1_5),
  inference(avatar_contradiction_clause,[],[f4300])).
fof(f4300,plain,(
  $false | ~spl1_5),
  inference(subsumption_resolution,[],[f4299,f97])).
fof(f97,plain,(
  aInteger0(xa)),
  inference(cnf_transformation,[],[f22])).
fof(f22,axiom,(
  aInteger0(xc) & sz00 != xq & aInteger0(xq) & aInteger0(xb) & aInteger0(xa)),
  file('./p/NUM432+1.tptp',unknown)).
fof(f4299,plain,(
  ~aInteger0(xa) | ~spl1_5),
  inference(subsumption_resolution,[],[f4298,f101])).
fof(f101,plain,(
  aInteger0(xc)),
  inference(cnf_transformation,[],[f22])).
fof(f4298,plain,(
  ~aInteger0(xc) | ~aInteger0(xa) | ~spl1_5),
  inference(subsumption_resolution,[],[f4297,f99])).
fof(f99,plain,(
  aInteger0(xq)),
  inference(cnf_transformation,[],[f22])).
fof(f4297,plain,(
  ~aInteger0(xq) | ~aInteger0(xc) | ~aInteger0(xa) | ~spl1_5),
  inference(subsumption_resolution,[],[f4296,f100])).
fof(f100,plain,(
  sz00 != xq),
  inference(cnf_transformation,[],[f22])).
fof(f4296,plain,(
  sz00 = xq | ~aInteger0(xq) | ~aInteger0(xc) | ~aInteger0(xa) | ~spl1_5),
  inference(subsumption_resolution,[],[f4293,f109])).
fof(f109,plain,(
  ~sdteqdtlpzmzozddtrp0(xa,xc,xq)),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ~sdteqdtlpzmzozddtrp0(xa,xc,xq)),
  inference(flattening,[],[f28])).
fof(f28,negated_conjecture,(
  ~sdteqdtlpzmzozddtrp0(xa,xc,xq)),
  inference(negated_conjecture,[],[f27])).
fof(f27,conjecture,(
  sdteqdtlpzmzozddtrp0(xa,xc,xq)),
  file('./p/NUM432+1.tptp',unknown)).
fof(f4293,plain,(
  sdteqdtlpzmzozddtrp0(xa,xc,xq) | sz00 = xq | ~aInteger0(xq) | ~aInteger0(xc) | ~aInteger0(xa) | ~spl1_5),
  inference(resolution,[],[f379,f94])).
fof(f94,plain,(
  ( ! [X2,X0,X1] : (~aDivisorOf0(X2,sdtpldt0(X0,smndt0(X1))) | sdteqdtlpzmzozddtrp0(X0,X1,X2) | sz00 = X2 | ~aInteger0(X2) | ~aInteger0(X1) | ~aInteger0(X0)) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : (((sdteqdtlpzmzozddtrp0(X0,X1,X2) | ~aDivisorOf0(X2,sdtpldt0(X0,smndt0(X1)))) & (aDivisorOf0(X2,sdtpldt0(X0,smndt0(X1))) | ~sdteqdtlpzmzozddtrp0(X0,X1,X2))) | sz00 = X2 | ~aInteger0(X2) | ~aInteger0(X1) | ~aInteger0(X0))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2] : ((sdteqdtlpzmzozddtrp0(X0,X1,X2) <=> aDivisorOf0(X2,sdtpldt0(X0,smndt0(X1)))) | sz00 = X2 | ~aInteger0(X2) | ~aInteger0(X1) | ~aInteger0(X0))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : ((sdteqdtlpzmzozddtrp0(X0,X1,X2) <=> aDivisorOf0(X2,sdtpldt0(X0,smndt0(X1)))) | (sz00 = X2 | ~aInteger0(X2) | ~aInteger0(X1) | ~aInteger0(X0)))),
  inference(ennf_transformation,[],[f19])).
fof(f19,axiom,(
  ! [X0,X1,X2] : ((sz00 != X2 & aInteger0(X2) & aInteger0(X1) & aInteger0(X0)) => (sdteqdtlpzmzozddtrp0(X0,X1,X2) <=> aDivisorOf0(X2,sdtpldt0(X0,smndt0(X1)))))),
  file('./p/NUM432+1.tptp',unknown)).
fof(f379,plain,(
  aDivisorOf0(xq,sdtpldt0(xa,smndt0(xc))) | ~spl1_5),
  inference(subsumption_resolution,[],[f378,f99])).
fof(f378,plain,(
  aDivisorOf0(xq,sdtpldt0(xa,smndt0(xc))) | ~aInteger0(xq) | ~spl1_5),
  inference(subsumption_resolution,[],[f377,f100])).
fof(f377,plain,(
  aDivisorOf0(xq,sdtpldt0(xa,smndt0(xc))) | sz00 = xq | ~aInteger0(xq) | ~spl1_5),
  inference(subsumption_resolution,[],[f375,f314])).
fof(f314,plain,(
  aInteger0(sdtpldt0(xn,xm)) | ~spl1_5),
  inference(avatar_component_clause,[],[f313])).
fof(f313,plain,(
  spl1_5 <=> aInteger0(sdtpldt0(xn,xm))),
  introduced(avatar_definition,[new_symbols(naming,[spl1_5])])).
fof(f375,plain,(
  aDivisorOf0(xq,sdtpldt0(xa,smndt0(xc))) | ~aInteger0(sdtpldt0(xn,xm)) | sz00 = xq | ~aInteger0(xq)),
  inference(superposition,[],[f112,f108])).
fof(f108,plain,(
  sdtasdt0(xq,sdtpldt0(xn,xm)) = sdtpldt0(xa,smndt0(xc))),
  inference(cnf_transformation,[],[f26])).
fof(f26,axiom,(
  sdtasdt0(xq,sdtpldt0(xn,xm)) = sdtpldt0(xa,smndt0(xc))),
  file('./p/NUM432+1.tptp',unknown)).
fof(f112,plain,(
  ( ! [X2,X1] : (aDivisorOf0(X1,sdtasdt0(X1,X2)) | ~aInteger0(X2) | sz00 = X1 | ~aInteger0(X1)) )),
  inference(subsumption_resolution,[],[f110,f70])).
fof(f70,plain,(
  ( ! [X0,X1] : (aInteger0(sdtasdt0(X0,X1)) | ~aInteger0(X1) | ~aInteger0(X0)) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (aInteger0(sdtasdt0(X0,X1)) | ~aInteger0(X1) | ~aInteger0(X0))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (aInteger0(sdtasdt0(X0,X1)) | (~aInteger0(X1) | ~aInteger0(X0)))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0,X1] : ((aInteger0(X1) & aInteger0(X0)) => aInteger0(sdtasdt0(X0,X1)))),
  file('./p/NUM432+1.tptp',unknown)).
fof(f110,plain,(
  ( ! [X2,X1] : (aDivisorOf0(X1,sdtasdt0(X1,X2)) | ~aInteger0(X2) | sz00 = X1 | ~aInteger0(X1) | ~aInteger0(sdtasdt0(X1,X2))) )),
  inference(equality_resolution,[],[f92])).
fof(f92,plain,(
  ( ! [X2,X0,X1] : (aDivisorOf0(X1,X0) | sdtasdt0(X1,X2) != X0 | ~aInteger0(X2) | sz00 = X1 | ~aInteger0(X1) | ~aInteger0(X0)) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : (! [X1] : ((aDivisorOf0(X1,X0) | ! [X2] : (sdtasdt0(X1,X2) != X0 | ~aInteger0(X2)) | sz00 = X1 | ~aInteger0(X1)) & (((sdtasdt0(X1,sK0(X0,X1)) = X0 & aInteger0(sK0(X0,X1))) & sz00 != X1 & aInteger0(X1)) | ~aDivisorOf0(X1,X0))) | ~aInteger0(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f62,f63])).
fof(f63,plain,(
  ! [X1,X0] : (? [X3] : (sdtasdt0(X1,X3) = X0 & aInteger0(X3)) => (sdtasdt0(X1,sK0(X0,X1)) = X0 & aInteger0(sK0(X0,X1))))),
  introduced(choice_axiom,[])).
fof(f62,plain,(
  ! [X0] : (! [X1] : ((aDivisorOf0(X1,X0) | ! [X2] : (sdtasdt0(X1,X2) != X0 | ~aInteger0(X2)) | sz00 = X1 | ~aInteger0(X1)) & ((? [X3] : (sdtasdt0(X1,X3) = X0 & aInteger0(X3)) & sz00 != X1 & aInteger0(X1)) | ~aDivisorOf0(X1,X0))) | ~aInteger0(X0))),
  inference(rectify,[],[f61])).
fof(f61,plain,(
  ! [X0] : (! [X1] : ((aDivisorOf0(X1,X0) | ! [X2] : (sdtasdt0(X1,X2) != X0 | ~aInteger0(X2)) | sz00 = X1 | ~aInteger0(X1)) & ((? [X2] : (sdtasdt0(X1,X2) = X0 & aInteger0(X2)) & sz00 != X1 & aInteger0(X1)) | ~aDivisorOf0(X1,X0))) | ~aInteger0(X0))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : (! [X1] : ((aDivisorOf0(X1,X0) | (! [X2] : (sdtasdt0(X1,X2) != X0 | ~aInteger0(X2)) | sz00 = X1 | ~aInteger0(X1))) & ((? [X2] : (sdtasdt0(X1,X2) = X0 & aInteger0(X2)) & sz00 != X1 & aInteger0(X1)) | ~aDivisorOf0(X1,X0))) | ~aInteger0(X0))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (! [X1] : (aDivisorOf0(X1,X0) <=> (? [X2] : (sdtasdt0(X1,X2) = X0 & aInteger0(X2)) & sz00 != X1 & aInteger0(X1))) | ~aInteger0(X0))),
  inference(ennf_transformation,[],[f18])).
fof(f18,axiom,(
  ! [X0] : (aInteger0(X0) => ! [X1] : (aDivisorOf0(X1,X0) <=> (? [X2] : (sdtasdt0(X1,X2) = X0 & aInteger0(X2)) & sz00 != X1 & aInteger0(X1))))),
  file('./p/NUM432+1.tptp',unknown)).
fof(f324,plain,(
  spl1_5),
  inference(avatar_contradiction_clause,[],[f323])).
fof(f323,plain,(
  $false | spl1_5),
  inference(subsumption_resolution,[],[f322,f104])).
fof(f104,plain,(
  aInteger0(xn)),
  inference(cnf_transformation,[],[f24])).
fof(f24,axiom,(
  sdtasdt0(xq,xn) = sdtpldt0(xa,smndt0(xb)) & aInteger0(xn)),
  file('./p/NUM432+1.tptp',unknown)).
fof(f322,plain,(
  ~aInteger0(xn) | spl1_5),
  inference(subsumption_resolution,[],[f321,f106])).
fof(f106,plain,(
  aInteger0(xm)),
  inference(cnf_transformation,[],[f25])).
fof(f25,axiom,(
  sdtasdt0(xq,xm) = sdtpldt0(xb,smndt0(xc)) & aInteger0(xm)),
  file('./p/NUM432+1.tptp',unknown)).
fof(f321,plain,(
  ~aInteger0(xm) | ~aInteger0(xn) | spl1_5),
  inference(resolution,[],[f315,f69])).
fof(f69,plain,(
  ( ! [X0,X1] : (aInteger0(sdtpldt0(X0,X1)) | ~aInteger0(X1) | ~aInteger0(X0)) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : (aInteger0(sdtpldt0(X0,X1)) | ~aInteger0(X1) | ~aInteger0(X0))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : (aInteger0(sdtpldt0(X0,X1)) | (~aInteger0(X1) | ~aInteger0(X0)))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1] : ((aInteger0(X1) & aInteger0(X0)) => aInteger0(sdtpldt0(X0,X1)))),
  file('./p/NUM432+1.tptp',unknown)).
fof(f315,plain,(
  ~aInteger0(sdtpldt0(xn,xm)) | spl1_5),
  inference(avatar_component_clause,[],[f313])).
% SZS output end Proof for NUM432+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 6908
% Time elapsed: 0.038 s
% ------------------------------
% ------------------------------
