% Refutation found. Thanks to Tanya!
% SZS status Theorem for NUM578+3
% SZS output start Proof for NUM578+3
fof(f18296,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f666,f667,f4194,f14641,f18295])).
fof(f18295,plain,(
  ~spl44_2 | ~spl44_145),
  inference(avatar_contradiction_clause,[],[f18294])).
fof(f18294,plain,(
  $false | (~spl44_2 | ~spl44_145)),
  inference(subsumption_resolution,[],[f18292,f14660])).
fof(f14660,plain,(
  sP12(xi,xj) | (~spl44_2 | ~spl44_145)),
  inference(subsumption_resolution,[],[f14659,f573])).
fof(f573,plain,(
  aElementOf0(xj,szNzAzT0)),
  inference(cnf_transformation,[],[f84])).
fof(f84,axiom,(
  aElementOf0(xj,szNzAzT0) & aElementOf0(xi,szNzAzT0)),
  file('./p/NUM578+3.tptp',unknown)).
fof(f14659,plain,(
  sP12(xi,xj) | ~aElementOf0(xj,szNzAzT0) | (~spl44_2 | ~spl44_145)),
  inference(subsumption_resolution,[],[f14654,f4191])).
fof(f4191,plain,(
  aElementOf0(xi,szNzAzT0) | ~spl44_145),
  inference(avatar_component_clause,[],[f4190])).
fof(f4190,plain,(
  spl44_145 <=> aElementOf0(xi,szNzAzT0)),
  introduced(avatar_definition,[new_symbols(naming,[spl44_145])])).
fof(f14654,plain,(
  sP12(xi,xj) | ~aElementOf0(xi,szNzAzT0) | ~aElementOf0(xj,szNzAzT0) | ~spl44_2),
  inference(resolution,[],[f661,f591])).
fof(f591,plain,(
  ( ! [X2,X3] : (~sdtlseqdt0(szszuzczcdt0(X2),X3) | sP12(X3,X2) | ~aElementOf0(X3,szNzAzT0) | ~aElementOf0(X2,szNzAzT0)) )),
  inference(cnf_transformation,[],[f347])).
fof(f347,plain,(
  szmzizndt0(sdtlpdtrp0(xN,xi)) = szmzizndt0(sdtlpdtrp0(xN,xj)) & ! [X0] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X0) | ~aElementOf0(X0,sdtlpdtrp0(xN,xj))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xj)) & ! [X1] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X1) | ~aElementOf0(X1,sdtlpdtrp0(xN,xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xi)) & xi != xj & ! [X2,X3] : (sP12(X3,X2) | ~sdtlseqdt0(szszuzczcdt0(X2),X3) | ~aElementOf0(X3,szNzAzT0) | ~aElementOf0(X2,szNzAzT0))),
  inference(rectify,[],[f229])).
fof(f229,plain,(
  szmzizndt0(sdtlpdtrp0(xN,xi)) = szmzizndt0(sdtlpdtrp0(xN,xj)) & ! [X7] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X7) | ~aElementOf0(X7,sdtlpdtrp0(xN,xj))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xj)) & ! [X8] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X8) | ~aElementOf0(X8,sdtlpdtrp0(xN,xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xi)) & xi != xj & ! [X0,X1] : (sP12(X1,X0) | ~sdtlseqdt0(szszuzczcdt0(X0),X1) | ~aElementOf0(X1,szNzAzT0) | ~aElementOf0(X0,szNzAzT0))),
  inference(definition_folding,[],[f211,f228,f227,f226])).
fof(f226,plain,(
  ! [X0] : (! [X5] : (aElementOf0(X5,sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) <=> (szmzizndt0(sdtlpdtrp0(xN,X0)) != X5 & aElementOf0(X5,sdtlpdtrp0(xN,X0)) & aElement0(X5))) | ~sP10(X0))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f227,plain,(
  ! [X1,X0] : (? [X3] : (~sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X3) & aElementOf0(X3,sdtlpdtrp0(xN,X0))) | ~aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X0)) | ~sP11(X1,X0))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f228,plain,(
  ! [X1,X0] : ((szmzizndt0(sdtlpdtrp0(xN,X0)) != szmzizndt0(sdtlpdtrp0(xN,X1)) & sP11(X1,X0) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X2) | ~aElementOf0(X2,sdtlpdtrp0(xN,X1))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X1)) & aSubsetOf0(sdtlpdtrp0(xN,X1),sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X4] : (aElementOf0(X4,sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) | ~aElementOf0(X4,sdtlpdtrp0(xN,X1))) & sP10(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X6] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X0)),X6) | ~aElementOf0(X6,sdtlpdtrp0(xN,X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X0)),sdtlpdtrp0(xN,X0))) | ~sP12(X1,X0))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f211,plain,(
  szmzizndt0(sdtlpdtrp0(xN,xi)) = szmzizndt0(sdtlpdtrp0(xN,xj)) & ! [X7] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X7) | ~aElementOf0(X7,sdtlpdtrp0(xN,xj))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xj)) & ! [X8] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X8) | ~aElementOf0(X8,sdtlpdtrp0(xN,xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xi)) & xi != xj & ! [X0,X1] : ((szmzizndt0(sdtlpdtrp0(xN,X0)) != szmzizndt0(sdtlpdtrp0(xN,X1)) & (? [X3] : (~sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X3) & aElementOf0(X3,sdtlpdtrp0(xN,X0))) | ~aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X0))) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X2) | ~aElementOf0(X2,sdtlpdtrp0(xN,X1))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X1)) & aSubsetOf0(sdtlpdtrp0(xN,X1),sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X4] : (aElementOf0(X4,sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) | ~aElementOf0(X4,sdtlpdtrp0(xN,X1))) & ! [X5] : (aElementOf0(X5,sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) <=> (szmzizndt0(sdtlpdtrp0(xN,X0)) != X5 & aElementOf0(X5,sdtlpdtrp0(xN,X0)) & aElement0(X5))) & aSet0(sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X6] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X0)),X6) | ~aElementOf0(X6,sdtlpdtrp0(xN,X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X0)),sdtlpdtrp0(xN,X0))) | ~sdtlseqdt0(szszuzczcdt0(X0),X1) | ~aElementOf0(X1,szNzAzT0) | ~aElementOf0(X0,szNzAzT0))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ((szmzizndt0(sdtlpdtrp0(xN,xi)) = szmzizndt0(sdtlpdtrp0(xN,xj)) & ! [X7] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X7) | ~aElementOf0(X7,sdtlpdtrp0(xN,xj))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xj)) & ! [X8] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X8) | ~aElementOf0(X8,sdtlpdtrp0(xN,xi))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xi))) & xi != xj) & ! [X0,X1] : (((((szmzizndt0(sdtlpdtrp0(xN,X0)) != szmzizndt0(sdtlpdtrp0(xN,X1)) & (? [X3] : (~sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X3) & aElementOf0(X3,sdtlpdtrp0(xN,X0))) | ~aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X0)))) & (! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X2) | ~aElementOf0(X2,sdtlpdtrp0(xN,X1))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X1)))) & aSubsetOf0(sdtlpdtrp0(xN,X1),sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X4] : (aElementOf0(X4,sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) | ~aElementOf0(X4,sdtlpdtrp0(xN,X1))) & ! [X5] : (aElementOf0(X5,sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) <=> (szmzizndt0(sdtlpdtrp0(xN,X0)) != X5 & aElementOf0(X5,sdtlpdtrp0(xN,X0)) & aElement0(X5))) & aSet0(sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X6] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X0)),X6) | ~aElementOf0(X6,sdtlpdtrp0(xN,X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X0)),sdtlpdtrp0(xN,X0))) | ~sdtlseqdt0(szszuzczcdt0(X0),X1)) | (~aElementOf0(X1,szNzAzT0) | ~aElementOf0(X0,szNzAzT0)))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ~(! [X0,X1] : ((aElementOf0(X1,szNzAzT0) & aElementOf0(X0,szNzAzT0)) => (sdtlseqdt0(szszuzczcdt0(X0),X1) => (~((! [X2] : (aElementOf0(X2,sdtlpdtrp0(xN,X1)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X1))) => (szmzizndt0(sdtlpdtrp0(xN,X0)) = szmzizndt0(sdtlpdtrp0(xN,X1)) | (! [X3] : (aElementOf0(X3,sdtlpdtrp0(xN,X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X3)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X0))))) & aSubsetOf0(sdtlpdtrp0(xN,X1),sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X4] : (aElementOf0(X4,sdtlpdtrp0(xN,X1)) => aElementOf0(X4,sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0))))) & ! [X5] : (aElementOf0(X5,sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) <=> (szmzizndt0(sdtlpdtrp0(xN,X0)) != X5 & aElementOf0(X5,sdtlpdtrp0(xN,X0)) & aElement0(X5))) & aSet0(sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X6] : (aElementOf0(X6,sdtlpdtrp0(xN,X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X0)),X6)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X0)),sdtlpdtrp0(xN,X0))))) => (xi != xj => ~(szmzizndt0(sdtlpdtrp0(xN,xi)) = szmzizndt0(sdtlpdtrp0(xN,xj)) & ! [X7] : (aElementOf0(X7,sdtlpdtrp0(xN,xj)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X7)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xj)) & ! [X8] : (aElementOf0(X8,sdtlpdtrp0(xN,xi)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X8)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xi)))))),
  inference(rectify,[],[f87])).
fof(f87,negated_conjecture,(
  ~(! [X0,X1] : ((aElementOf0(X1,szNzAzT0) & aElementOf0(X0,szNzAzT0)) => (sdtlseqdt0(szszuzczcdt0(X0),X1) => (~((! [X2] : (aElementOf0(X2,sdtlpdtrp0(xN,X1)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X1))) => (szmzizndt0(sdtlpdtrp0(xN,X0)) = szmzizndt0(sdtlpdtrp0(xN,X1)) | (! [X2] : (aElementOf0(X2,sdtlpdtrp0(xN,X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X0))))) & aSubsetOf0(sdtlpdtrp0(xN,X1),sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X2] : (aElementOf0(X2,sdtlpdtrp0(xN,X1)) => aElementOf0(X2,sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0))))) & ! [X2] : (aElementOf0(X2,sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) <=> (szmzizndt0(sdtlpdtrp0(xN,X0)) != X2 & aElementOf0(X2,sdtlpdtrp0(xN,X0)) & aElement0(X2))) & aSet0(sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X2] : (aElementOf0(X2,sdtlpdtrp0(xN,X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X0)),X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X0)),sdtlpdtrp0(xN,X0))))) => (xi != xj => ~(szmzizndt0(sdtlpdtrp0(xN,xi)) = szmzizndt0(sdtlpdtrp0(xN,xj)) & ! [X0] : (aElementOf0(X0,sdtlpdtrp0(xN,xj)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xj)) & ! [X0] : (aElementOf0(X0,sdtlpdtrp0(xN,xi)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xi)))))),
  inference(negated_conjecture,[],[f86])).
fof(f86,conjecture,(
  ! [X0,X1] : ((aElementOf0(X1,szNzAzT0) & aElementOf0(X0,szNzAzT0)) => (sdtlseqdt0(szszuzczcdt0(X0),X1) => (~((! [X2] : (aElementOf0(X2,sdtlpdtrp0(xN,X1)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X1))) => (szmzizndt0(sdtlpdtrp0(xN,X0)) = szmzizndt0(sdtlpdtrp0(xN,X1)) | (! [X2] : (aElementOf0(X2,sdtlpdtrp0(xN,X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X0))))) & aSubsetOf0(sdtlpdtrp0(xN,X1),sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X2] : (aElementOf0(X2,sdtlpdtrp0(xN,X1)) => aElementOf0(X2,sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0))))) & ! [X2] : (aElementOf0(X2,sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) <=> (szmzizndt0(sdtlpdtrp0(xN,X0)) != X2 & aElementOf0(X2,sdtlpdtrp0(xN,X0)) & aElement0(X2))) & aSet0(sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X2] : (aElementOf0(X2,sdtlpdtrp0(xN,X0)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X0)),X2)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X0)),sdtlpdtrp0(xN,X0))))) => (xi != xj => ~(szmzizndt0(sdtlpdtrp0(xN,xi)) = szmzizndt0(sdtlpdtrp0(xN,xj)) & ! [X0] : (aElementOf0(X0,sdtlpdtrp0(xN,xj)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xj)) & ! [X0] : (aElementOf0(X0,sdtlpdtrp0(xN,xi)) => sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,xi)),X0)) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,xi)),sdtlpdtrp0(xN,xi))))),
  file('./p/NUM578+3.tptp',unknown)).
fof(f661,plain,(
  sdtlseqdt0(szszuzczcdt0(xj),xi) | ~spl44_2),
  inference(avatar_component_clause,[],[f659])).
fof(f659,plain,(
  spl44_2 <=> sdtlseqdt0(szszuzczcdt0(xj),xi)),
  introduced(avatar_definition,[new_symbols(naming,[spl44_2])])).
fof(f18292,plain,(
  ~sP12(xi,xj)),
  inference(equality_resolution,[],[f1962])).
fof(f1962,plain,(
  ( ! [X0] : (szmzizndt0(sdtlpdtrp0(xN,X0)) != szmzizndt0(sdtlpdtrp0(xN,xi)) | ~sP12(X0,xj)) )),
  inference(superposition,[],[f584,f597])).
fof(f597,plain,(
  szmzizndt0(sdtlpdtrp0(xN,xi)) = szmzizndt0(sdtlpdtrp0(xN,xj))),
  inference(cnf_transformation,[],[f347])).
fof(f584,plain,(
  ( ! [X0,X1] : (szmzizndt0(sdtlpdtrp0(xN,X0)) != szmzizndt0(sdtlpdtrp0(xN,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X0,X1] : ((szmzizndt0(sdtlpdtrp0(xN,X0)) != szmzizndt0(sdtlpdtrp0(xN,X1)) & sP11(X0,X1) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X0)),X2) | ~aElementOf0(X2,sdtlpdtrp0(xN,X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X0)),sdtlpdtrp0(xN,X0)) & aSubsetOf0(sdtlpdtrp0(xN,X0),sdtmndt0(sdtlpdtrp0(xN,X1),szmzizndt0(sdtlpdtrp0(xN,X1)))) & ! [X3] : (aElementOf0(X3,sdtmndt0(sdtlpdtrp0(xN,X1),szmzizndt0(sdtlpdtrp0(xN,X1)))) | ~aElementOf0(X3,sdtlpdtrp0(xN,X0))) & sP10(X1) & aSet0(sdtmndt0(sdtlpdtrp0(xN,X1),szmzizndt0(sdtlpdtrp0(xN,X1)))) & ! [X4] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X4) | ~aElementOf0(X4,sdtlpdtrp0(xN,X1))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X1))) | ~sP12(X0,X1))),
  inference(rectify,[],[f338])).
fof(f338,plain,(
  ! [X1,X0] : ((szmzizndt0(sdtlpdtrp0(xN,X0)) != szmzizndt0(sdtlpdtrp0(xN,X1)) & sP11(X1,X0) & ! [X2] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X1)),X2) | ~aElementOf0(X2,sdtlpdtrp0(xN,X1))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X1)),sdtlpdtrp0(xN,X1)) & aSubsetOf0(sdtlpdtrp0(xN,X1),sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X4] : (aElementOf0(X4,sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) | ~aElementOf0(X4,sdtlpdtrp0(xN,X1))) & sP10(X0) & aSet0(sdtmndt0(sdtlpdtrp0(xN,X0),szmzizndt0(sdtlpdtrp0(xN,X0)))) & ! [X6] : (sdtlseqdt0(szmzizndt0(sdtlpdtrp0(xN,X0)),X6) | ~aElementOf0(X6,sdtlpdtrp0(xN,X0))) & aElementOf0(szmzizndt0(sdtlpdtrp0(xN,X0)),sdtlpdtrp0(xN,X0))) | ~sP12(X1,X0))),
  inference(nnf_transformation,[],[f228])).
fof(f14641,plain,(
  ~spl44_3 | ~spl44_145),
  inference(avatar_contradiction_clause,[],[f14640])).
fof(f14640,plain,(
  $false | (~spl44_3 | ~spl44_145)),
  inference(subsumption_resolution,[],[f14638,f10510])).
fof(f10510,plain,(
  sP12(xj,xi) | (~spl44_3 | ~spl44_145)),
  inference(subsumption_resolution,[],[f10509,f4191])).
fof(f10509,plain,(
  sP12(xj,xi) | ~aElementOf0(xi,szNzAzT0) | ~spl44_3),
  inference(subsumption_resolution,[],[f10504,f573])).
fof(f10504,plain,(
  sP12(xj,xi) | ~aElementOf0(xj,szNzAzT0) | ~aElementOf0(xi,szNzAzT0) | ~spl44_3),
  inference(resolution,[],[f665,f591])).
fof(f665,plain,(
  sdtlseqdt0(szszuzczcdt0(xi),xj) | ~spl44_3),
  inference(avatar_component_clause,[],[f663])).
fof(f663,plain,(
  spl44_3 <=> sdtlseqdt0(szszuzczcdt0(xi),xj)),
  introduced(avatar_definition,[new_symbols(naming,[spl44_3])])).
fof(f14638,plain,(
  ~sP12(xj,xi)),
  inference(equality_resolution,[],[f1960])).
fof(f1960,plain,(
  ( ! [X0] : (szmzizndt0(sdtlpdtrp0(xN,X0)) != szmzizndt0(sdtlpdtrp0(xN,xi)) | ~sP12(xj,X0)) )),
  inference(superposition,[],[f584,f597])).
fof(f4194,plain,(
  spl44_145),
  inference(avatar_split_clause,[],[f572,f4190])).
fof(f572,plain,(
  aElementOf0(xi,szNzAzT0)),
  inference(cnf_transformation,[],[f84])).
fof(f667,plain,(
  ~spl44_1),
  inference(avatar_split_clause,[],[f592,f655])).
fof(f655,plain,(
  spl44_1 <=> xi = xj),
  introduced(avatar_definition,[new_symbols(naming,[spl44_1])])).
fof(f592,plain,(
  xi != xj),
  inference(cnf_transformation,[],[f347])).
fof(f666,plain,(
  spl44_1 | spl44_2 | spl44_3),
  inference(avatar_split_clause,[],[f574,f663,f659,f655])).
fof(f574,plain,(
  sdtlseqdt0(szszuzczcdt0(xi),xj) | sdtlseqdt0(szszuzczcdt0(xj),xi) | xi = xj),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  sdtlseqdt0(szszuzczcdt0(xi),xj) | sdtlseqdt0(szszuzczcdt0(xj),xi) | xi = xj),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  (sdtlseqdt0(szszuzczcdt0(xi),xj) | sdtlseqdt0(szszuzczcdt0(xj),xi)) | xi = xj),
  inference(ennf_transformation,[],[f85])).
fof(f85,axiom,(
  xi != xj => (sdtlseqdt0(szszuzczcdt0(xi),xj) | sdtlseqdt0(szszuzczcdt0(xj),xi))),
  file('./p/NUM578+3.tptp',unknown)).
% SZS output end Proof for NUM578+3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 14456
% Time elapsed: 0.211 s
% ------------------------------
% ------------------------------
