% Refutation found. Thanks to Tanya!
% SZS status Theorem for NUM556+3
% SZS output start Proof for NUM556+3
fof(f5116,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f471,f476,f509,f3943,f5069,f5086,f5106])).
fof(f5106,plain,(
  spl20_5 | ~spl20_269),
  inference(avatar_contradiction_clause,[],[f5105])).
fof(f5105,plain,(
  $false | (spl20_5 | ~spl20_269)),
  inference(subsumption_resolution,[],[f5091,f380])).
fof(f380,plain,(
  aElementOf0(xx,xS)),
  inference(cnf_transformation,[],[f64])).
fof(f64,axiom,(
  aElementOf0(xx,xS)),
  file('./p/NUM556+3.tptp',unknown)).
fof(f5091,plain,(
  ~aElementOf0(xx,xS) | (spl20_5 | ~spl20_269)),
  inference(backward_demodulation,[],[f5068,f470])).
fof(f470,plain,(
  ~aElementOf0(sK19,xS) | spl20_5),
  inference(avatar_component_clause,[],[f468])).
fof(f468,plain,(
  spl20_5 <=> aElementOf0(sK19,xS)),
  introduced(avatar_definition,[new_symbols(naming,[spl20_5])])).
fof(f5068,plain,(
  xx = sK19 | ~spl20_269),
  inference(avatar_component_clause,[],[f5066])).
fof(f5066,plain,(
  spl20_269 <=> xx = sK19),
  introduced(avatar_definition,[new_symbols(naming,[spl20_269])])).
fof(f5086,plain,(
  spl20_5 | ~spl20_268),
  inference(avatar_contradiction_clause,[],[f5085])).
fof(f5085,plain,(
  $false | (spl20_5 | ~spl20_268)),
  inference(subsumption_resolution,[],[f5081,f470])).
fof(f5081,plain,(
  aElementOf0(sK19,xS) | ~spl20_268),
  inference(resolution,[],[f5064,f382])).
fof(f382,plain,(
  ( ! [X0] : (~aElementOf0(X0,xQ) | aElementOf0(X0,xS)) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  aElementOf0(xQ,slbdtsldtrb0(xS,xk)) & xk = sbrdtbr0(xQ) & aSubsetOf0(xQ,xS) & ! [X0] : (aElementOf0(X0,xS) | ~aElementOf0(X0,xQ)) & aSet0(xQ)),
  inference(ennf_transformation,[],[f65])).
fof(f65,axiom,(
  aElementOf0(xQ,slbdtsldtrb0(xS,xk)) & xk = sbrdtbr0(xQ) & aSubsetOf0(xQ,xS) & ! [X0] : (aElementOf0(X0,xQ) => aElementOf0(X0,xS)) & aSet0(xQ)),
  file('./p/NUM556+3.tptp',unknown)).
fof(f5064,plain,(
  aElementOf0(sK19,xQ) | ~spl20_268),
  inference(avatar_component_clause,[],[f5062])).
fof(f5062,plain,(
  spl20_268 <=> aElementOf0(sK19,xQ)),
  introduced(avatar_definition,[new_symbols(naming,[spl20_268])])).
fof(f5069,plain,(
  spl20_268 | spl20_269 | ~spl20_6),
  inference(avatar_split_clause,[],[f5022,f473,f5066,f5062])).
fof(f473,plain,(
  spl20_6 <=> aElementOf0(sK19,xP)),
  introduced(avatar_definition,[new_symbols(naming,[spl20_6])])).
fof(f5022,plain,(
  xx = sK19 | aElementOf0(sK19,xQ) | ~spl20_6),
  inference(resolution,[],[f4160,f475])).
fof(f475,plain,(
  aElementOf0(sK19,xP) | ~spl20_6),
  inference(avatar_component_clause,[],[f473])).
fof(f4160,plain,(
  ( ! [X1] : (~aElementOf0(X1,xP) | xx = X1 | aElementOf0(X1,xQ)) )),
  inference(resolution,[],[f400,f395])).
fof(f395,plain,(
  ( ! [X1] : (~aElementOf0(X1,sdtmndt0(xQ,xy)) | aElementOf0(X1,xQ)) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  xP = sdtpldt0(sdtmndt0(xQ,xy),xx) & ! [X0] : ((aElementOf0(X0,xP) | (xx != X0 & ~aElementOf0(X0,sdtmndt0(xQ,xy))) | ~aElement0(X0)) & (((xx = X0 | aElementOf0(X0,sdtmndt0(xQ,xy))) & aElement0(X0)) | ~aElementOf0(X0,xP))) & aSet0(xP) & ! [X1] : ((aElementOf0(X1,sdtmndt0(xQ,xy)) | xy = X1 | ~aElementOf0(X1,xQ) | ~aElement0(X1)) & ((xy != X1 & aElementOf0(X1,xQ) & aElement0(X1)) | ~aElementOf0(X1,sdtmndt0(xQ,xy)))) & aSet0(sdtmndt0(xQ,xy))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  xP = sdtpldt0(sdtmndt0(xQ,xy),xx) & ! [X0] : ((aElementOf0(X0,xP) | ((xx != X0 & ~aElementOf0(X0,sdtmndt0(xQ,xy))) | ~aElement0(X0))) & (((xx = X0 | aElementOf0(X0,sdtmndt0(xQ,xy))) & aElement0(X0)) | ~aElementOf0(X0,xP))) & aSet0(xP) & ! [X1] : ((aElementOf0(X1,sdtmndt0(xQ,xy)) | (xy = X1 | ~aElementOf0(X1,xQ) | ~aElement0(X1))) & ((xy != X1 & aElementOf0(X1,xQ) & aElement0(X1)) | ~aElementOf0(X1,sdtmndt0(xQ,xy)))) & aSet0(sdtmndt0(xQ,xy))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  xP = sdtpldt0(sdtmndt0(xQ,xy),xx) & ! [X0] : (aElementOf0(X0,xP) <=> ((xx = X0 | aElementOf0(X0,sdtmndt0(xQ,xy))) & aElement0(X0))) & aSet0(xP) & ! [X1] : (aElementOf0(X1,sdtmndt0(xQ,xy)) <=> (xy != X1 & aElementOf0(X1,xQ) & aElement0(X1))) & aSet0(sdtmndt0(xQ,xy))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  xP = sdtpldt0(sdtmndt0(xQ,xy),xx) & ! [X0] : (aElementOf0(X0,xP) <=> ((xx = X0 | aElementOf0(X0,sdtmndt0(xQ,xy))) & aElement0(X0))) & aSet0(xP) & ! [X0] : (aElementOf0(X0,sdtmndt0(xQ,xy)) <=> (xy != X0 & aElementOf0(X0,xQ) & aElement0(X0))) & aSet0(sdtmndt0(xQ,xy))),
  file('./p/NUM556+3.tptp',unknown)).
fof(f400,plain,(
  ( ! [X0] : (aElementOf0(X0,sdtmndt0(xQ,xy)) | xx = X0 | ~aElementOf0(X0,xP)) )),
  inference(cnf_transformation,[],[f235])).
fof(f3943,plain,(
  ~spl20_1 | spl20_4),
  inference(avatar_contradiction_clause,[],[f3942])).
fof(f3942,plain,(
  $false | (~spl20_1 | spl20_4)),
  inference(subsumption_resolution,[],[f3941,f389])).
fof(f389,plain,(
  aElement0(xy)),
  inference(cnf_transformation,[],[f67])).
fof(f67,axiom,(
  aElementOf0(xy,xQ) & aElement0(xy)),
  file('./p/NUM556+3.tptp',unknown)).
fof(f3941,plain,(
  ~aElement0(xy) | (~spl20_1 | spl20_4)),
  inference(subsumption_resolution,[],[f3940,f381])).
fof(f381,plain,(
  aSet0(xQ)),
  inference(cnf_transformation,[],[f166])).
fof(f3940,plain,(
  ~aSet0(xQ) | ~aElement0(xy) | (~spl20_1 | spl20_4)),
  inference(subsumption_resolution,[],[f3939,f387])).
fof(f387,plain,(
  isFinite0(xQ)),
  inference(cnf_transformation,[],[f66])).
fof(f66,axiom,(
  xk = sbrdtbr0(xQ) & isFinite0(xQ) & aSet0(xQ)),
  file('./p/NUM556+3.tptp',unknown)).
fof(f3939,plain,(
  ~isFinite0(xQ) | ~aSet0(xQ) | ~aElement0(xy) | (~spl20_1 | spl20_4)),
  inference(resolution,[],[f3925,f284])).
fof(f284,plain,(
  ( ! [X0,X1] : (isFinite0(sdtmndt0(X1,X0)) | ~isFinite0(X1) | ~aSet0(X1) | ~aElement0(X0)) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : (! [X1] : (isFinite0(sdtmndt0(X1,X0)) | ~isFinite0(X1) | ~aSet0(X1)) | ~aElement0(X0))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : (! [X1] : (isFinite0(sdtmndt0(X1,X0)) | (~isFinite0(X1) | ~aSet0(X1))) | ~aElement0(X0))),
  inference(ennf_transformation,[],[f22])).
fof(f22,axiom,(
  ! [X0] : (aElement0(X0) => ! [X1] : ((isFinite0(X1) & aSet0(X1)) => isFinite0(sdtmndt0(X1,X0))))),
  file('./p/NUM556+3.tptp',unknown)).
fof(f3925,plain,(
  ~isFinite0(sdtmndt0(xQ,xy)) | (~spl20_1 | spl20_4)),
  inference(subsumption_resolution,[],[f3924,f465])).
fof(f465,plain,(
  xk != sbrdtbr0(xP) | spl20_4),
  inference(avatar_component_clause,[],[f463])).
fof(f463,plain,(
  spl20_4 <=> xk = sbrdtbr0(xP)),
  introduced(avatar_definition,[new_symbols(naming,[spl20_4])])).
fof(f3924,plain,(
  xk = sbrdtbr0(xP) | ~isFinite0(sdtmndt0(xQ,xy)) | ~spl20_1),
  inference(forward_demodulation,[],[f3923,f3901])).
fof(f3901,plain,(
  xk = szszuzczcdt0(sbrdtbr0(sdtmndt0(xQ,xy)))),
  inference(forward_demodulation,[],[f2194,f384])).
fof(f384,plain,(
  xk = sbrdtbr0(xQ)),
  inference(cnf_transformation,[],[f166])).
fof(f2194,plain,(
  sbrdtbr0(xQ) = szszuzczcdt0(sbrdtbr0(sdtmndt0(xQ,xy)))),
  inference(subsumption_resolution,[],[f2193,f381])).
fof(f2193,plain,(
  sbrdtbr0(xQ) = szszuzczcdt0(sbrdtbr0(sdtmndt0(xQ,xy))) | ~aSet0(xQ)),
  inference(subsumption_resolution,[],[f2148,f387])).
fof(f2148,plain,(
  sbrdtbr0(xQ) = szszuzczcdt0(sbrdtbr0(sdtmndt0(xQ,xy))) | ~isFinite0(xQ) | ~aSet0(xQ)),
  inference(resolution,[],[f309,f390])).
fof(f390,plain,(
  aElementOf0(xy,xQ)),
  inference(cnf_transformation,[],[f67])).
fof(f309,plain,(
  ( ! [X0,X1] : (~aElementOf0(X1,X0) | sbrdtbr0(X0) = szszuzczcdt0(sbrdtbr0(sdtmndt0(X0,X1))) | ~isFinite0(X0) | ~aSet0(X0)) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : (! [X1] : (sbrdtbr0(X0) = szszuzczcdt0(sbrdtbr0(sdtmndt0(X0,X1))) | ~aElementOf0(X1,X0) | ~isFinite0(X0)) | ~aSet0(X0))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0] : (! [X1] : (sbrdtbr0(X0) = szszuzczcdt0(sbrdtbr0(sdtmndt0(X0,X1))) | (~aElementOf0(X1,X0) | ~isFinite0(X0))) | ~aSet0(X0))),
  inference(ennf_transformation,[],[f44])).
fof(f44,axiom,(
  ! [X0] : (aSet0(X0) => ! [X1] : ((aElementOf0(X1,X0) & isFinite0(X0)) => sbrdtbr0(X0) = szszuzczcdt0(sbrdtbr0(sdtmndt0(X0,X1)))))),
  file('./p/NUM556+3.tptp',unknown)).
fof(f3923,plain,(
  sbrdtbr0(xP) = szszuzczcdt0(sbrdtbr0(sdtmndt0(xQ,xy))) | ~isFinite0(sdtmndt0(xQ,xy)) | ~spl20_1),
  inference(forward_demodulation,[],[f3922,f403])).
fof(f403,plain,(
  xP = sdtpldt0(sdtmndt0(xQ,xy),xx)),
  inference(cnf_transformation,[],[f235])).
fof(f3922,plain,(
  szszuzczcdt0(sbrdtbr0(sdtmndt0(xQ,xy))) = sbrdtbr0(sdtpldt0(sdtmndt0(xQ,xy),xx)) | ~isFinite0(sdtmndt0(xQ,xy)) | ~spl20_1),
  inference(subsumption_resolution,[],[f3921,f393])).
fof(f393,plain,(
  aSet0(sdtmndt0(xQ,xy))),
  inference(cnf_transformation,[],[f235])).
fof(f3921,plain,(
  szszuzczcdt0(sbrdtbr0(sdtmndt0(xQ,xy))) = sbrdtbr0(sdtpldt0(sdtmndt0(xQ,xy),xx)) | ~isFinite0(sdtmndt0(xQ,xy)) | ~aSet0(sdtmndt0(xQ,xy)) | ~spl20_1),
  inference(subsumption_resolution,[],[f3919,f451])).
fof(f451,plain,(
  aElement0(xx) | ~spl20_1),
  inference(avatar_component_clause,[],[f450])).
fof(f450,plain,(
  spl20_1 <=> aElement0(xx)),
  introduced(avatar_definition,[new_symbols(naming,[spl20_1])])).
fof(f3919,plain,(
  szszuzczcdt0(sbrdtbr0(sdtmndt0(xQ,xy))) = sbrdtbr0(sdtpldt0(sdtmndt0(xQ,xy),xx)) | ~aElement0(xx) | ~isFinite0(sdtmndt0(xQ,xy)) | ~aSet0(sdtmndt0(xQ,xy))),
  inference(resolution,[],[f404,f308])).
fof(f308,plain,(
  ( ! [X0,X1] : (aElementOf0(X1,X0) | sbrdtbr0(sdtpldt0(X0,X1)) = szszuzczcdt0(sbrdtbr0(X0)) | ~aElement0(X1) | ~isFinite0(X0) | ~aSet0(X0)) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : (! [X1] : (sbrdtbr0(sdtpldt0(X0,X1)) = szszuzczcdt0(sbrdtbr0(X0)) | aElementOf0(X1,X0) | ~aElement0(X1)) | ~isFinite0(X0) | ~aSet0(X0))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : (! [X1] : ((sbrdtbr0(sdtpldt0(X0,X1)) = szszuzczcdt0(sbrdtbr0(X0)) | aElementOf0(X1,X0)) | ~aElement0(X1)) | (~isFinite0(X0) | ~aSet0(X0)))),
  inference(ennf_transformation,[],[f43])).
fof(f43,axiom,(
  ! [X0] : ((isFinite0(X0) & aSet0(X0)) => ! [X1] : (aElement0(X1) => (~aElementOf0(X1,X0) => sbrdtbr0(sdtpldt0(X0,X1)) = szszuzczcdt0(sbrdtbr0(X0)))))),
  file('./p/NUM556+3.tptp',unknown)).
fof(f404,plain,(
  ~aElementOf0(xx,sdtmndt0(xQ,xy))),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0] : ((aElementOf0(X0,sdtmndt0(xQ,xy)) | xy = X0 | ~aElementOf0(X0,xQ) | ~aElement0(X0)) & ((xy != X0 & aElementOf0(X0,xQ) & aElement0(X0)) | ~aElementOf0(X0,sdtmndt0(xQ,xy)))) & aSet0(sdtmndt0(xQ,xy)) & ~aElementOf0(xx,sdtmndt0(xQ,xy))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X0] : ((aElementOf0(X0,sdtmndt0(xQ,xy)) | (xy = X0 | ~aElementOf0(X0,xQ) | ~aElement0(X0))) & ((xy != X0 & aElementOf0(X0,xQ) & aElement0(X0)) | ~aElementOf0(X0,sdtmndt0(xQ,xy)))) & aSet0(sdtmndt0(xQ,xy)) & ~aElementOf0(xx,sdtmndt0(xQ,xy))),
  inference(nnf_transformation,[],[f71])).
fof(f71,axiom,(
  ! [X0] : (aElementOf0(X0,sdtmndt0(xQ,xy)) <=> (xy != X0 & aElementOf0(X0,xQ) & aElement0(X0))) & aSet0(sdtmndt0(xQ,xy)) & ~aElementOf0(xx,sdtmndt0(xQ,xy))),
  file('./p/NUM556+3.tptp',unknown)).
fof(f509,plain,(
  spl20_1),
  inference(avatar_split_clause,[],[f508,f450])).
fof(f508,plain,(
  aElement0(xx)),
  inference(subsumption_resolution,[],[f498,f350])).
fof(f350,plain,(
  aSet0(xS)),
  inference(cnf_transformation,[],[f62])).
fof(f62,axiom,(
  sz00 != xk & aSet0(xT) & aSet0(xS)),
  file('./p/NUM556+3.tptp',unknown)).
fof(f498,plain,(
  aElement0(xx) | ~aSet0(xS)),
  inference(resolution,[],[f240,f380])).
fof(f240,plain,(
  ( ! [X0,X1] : (~aElementOf0(X1,X0) | aElement0(X1) | ~aSet0(X0)) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (! [X1] : (aElement0(X1) | ~aElementOf0(X1,X0)) | ~aSet0(X0))),
  inference(ennf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (aSet0(X0) => ! [X1] : (aElementOf0(X1,X0) => aElement0(X1)))),
  file('./p/NUM556+3.tptp',unknown)).
fof(f476,plain,(
  spl20_6 | ~spl20_4),
  inference(avatar_split_clause,[],[f410,f463,f473])).
fof(f410,plain,(
  xk != sbrdtbr0(xP) | aElementOf0(sK19,xP)),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  xk != sbrdtbr0(xP) | (~aSubsetOf0(xP,xS) & (~aElementOf0(sK19,xS) & aElementOf0(sK19,xP)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19])],[f167,f238])).
fof(f238,plain,(
  ? [X0] : (~aElementOf0(X0,xS) & aElementOf0(X0,xP)) => (~aElementOf0(sK19,xS) & aElementOf0(sK19,xP))),
  introduced(choice_axiom,[])).
fof(f167,plain,(
  xk != sbrdtbr0(xP) | (~aSubsetOf0(xP,xS) & ? [X0] : (~aElementOf0(X0,xS) & aElementOf0(X0,xP)))),
  inference(ennf_transformation,[],[f73])).
fof(f73,negated_conjecture,(
  ~(xk = sbrdtbr0(xP) & (aSubsetOf0(xP,xS) | ! [X0] : (aElementOf0(X0,xP) => aElementOf0(X0,xS))))),
  inference(negated_conjecture,[],[f72])).
fof(f72,conjecture,(
  xk = sbrdtbr0(xP) & (aSubsetOf0(xP,xS) | ! [X0] : (aElementOf0(X0,xP) => aElementOf0(X0,xS)))),
  file('./p/NUM556+3.tptp',unknown)).
fof(f471,plain,(
  ~spl20_5 | ~spl20_4),
  inference(avatar_split_clause,[],[f411,f463,f468])).
fof(f411,plain,(
  xk != sbrdtbr0(xP) | ~aElementOf0(sK19,xS)),
  inference(cnf_transformation,[],[f239])).
% SZS output end Proof for NUM556+3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 7675
% Time elapsed: 0.045 s
% ------------------------------
% ------------------------------
