% Refutation found. Thanks to Tanya!
% SZS status Theorem for NUM500+1
% SZS output start Proof for NUM500+1
fof(f1637,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f477,f539,f543,f1207,f1214,f1636])).
fof(f1636,plain,(
  spl4_15),
  inference(avatar_contradiction_clause,[],[f1635])).
fof(f1635,plain,(
  $false | spl4_15),
  inference(subsumption_resolution,[],[f1634,f207])).
fof(f207,plain,(
  aNaturalNumber0(xn)),
  inference(cnf_transformation,[],[f39])).
fof(f39,axiom,(
  aNaturalNumber0(xp) & aNaturalNumber0(xm) & aNaturalNumber0(xn)),
  file('./p/NUM500+1.tptp',unknown)).
fof(f1634,plain,(
  ~aNaturalNumber0(xn) | spl4_15),
  inference(subsumption_resolution,[],[f1633,f208])).
fof(f208,plain,(
  aNaturalNumber0(xm)),
  inference(cnf_transformation,[],[f39])).
fof(f1633,plain,(
  ~aNaturalNumber0(xm) | ~aNaturalNumber0(xn) | spl4_15),
  inference(resolution,[],[f839,f143])).
fof(f143,plain,(
  ( ! [X0,X1] : (aNaturalNumber0(sdtasdt0(X0,X1)) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : (aNaturalNumber0(sdtasdt0(X0,X1)) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : (aNaturalNumber0(sdtasdt0(X0,X1)) | (~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => aNaturalNumber0(sdtasdt0(X0,X1)))),
  file('./p/NUM500+1.tptp',unknown)).
fof(f839,plain,(
  ~aNaturalNumber0(sdtasdt0(xn,xm)) | spl4_15),
  inference(avatar_component_clause,[],[f837])).
fof(f837,plain,(
  spl4_15 <=> aNaturalNumber0(sdtasdt0(xn,xm))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_15])])).
fof(f1214,plain,(
  ~spl4_15 | spl4_9 | spl4_5),
  inference(avatar_split_clause,[],[f1213,f441,f536,f837])).
fof(f536,plain,(
  spl4_9 <=> aNaturalNumber0(xk)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_9])])).
fof(f441,plain,(
  spl4_5 <=> sz00 = xp),
  introduced(avatar_definition,[new_symbols(naming,[spl4_5])])).
fof(f1213,plain,(
  aNaturalNumber0(xk) | ~aNaturalNumber0(sdtasdt0(xn,xm)) | spl4_5),
  inference(subsumption_resolution,[],[f1142,f209])).
fof(f209,plain,(
  aNaturalNumber0(xp)),
  inference(cnf_transformation,[],[f39])).
fof(f1142,plain,(
  aNaturalNumber0(xk) | ~aNaturalNumber0(sdtasdt0(xn,xm)) | ~aNaturalNumber0(xp) | spl4_5),
  inference(subsumption_resolution,[],[f1141,f442])).
fof(f442,plain,(
  sz00 != xp | spl4_5),
  inference(avatar_component_clause,[],[f441])).
fof(f1141,plain,(
  aNaturalNumber0(xk) | sz00 = xp | ~aNaturalNumber0(sdtasdt0(xn,xm)) | ~aNaturalNumber0(xp)),
  inference(subsumption_resolution,[],[f902,f212])).
fof(f212,plain,(
  doDivides0(xp,sdtasdt0(xn,xm))),
  inference(cnf_transformation,[],[f41])).
fof(f41,axiom,(
  doDivides0(xp,sdtasdt0(xn,xm)) & isPrime0(xp)),
  file('./p/NUM500+1.tptp',unknown)).
fof(f902,plain,(
  aNaturalNumber0(xk) | ~doDivides0(xp,sdtasdt0(xn,xm)) | sz00 = xp | ~aNaturalNumber0(sdtasdt0(xn,xm)) | ~aNaturalNumber0(xp)),
  inference(superposition,[],[f234,f219])).
fof(f219,plain,(
  xk = sdtsldt0(sdtasdt0(xn,xm),xp)),
  inference(cnf_transformation,[],[f45])).
fof(f45,axiom,(
  xk = sdtsldt0(sdtasdt0(xn,xm),xp)),
  file('./p/NUM500+1.tptp',unknown)).
fof(f234,plain,(
  ( ! [X0,X1] : (aNaturalNumber0(sdtsldt0(X1,X0)) | ~doDivides0(X0,X1) | sz00 = X0 | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(equality_resolution,[],[f189])).
fof(f189,plain,(
  ( ! [X2,X0,X1] : (aNaturalNumber0(X2) | sdtsldt0(X1,X0) != X2 | ~doDivides0(X0,X1) | sz00 = X0 | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : (! [X2] : ((sdtsldt0(X1,X0) = X2 | sdtasdt0(X0,X2) != X1 | ~aNaturalNumber0(X2)) & ((sdtasdt0(X0,X2) = X1 & aNaturalNumber0(X2)) | sdtsldt0(X1,X0) != X2)) | ~doDivides0(X0,X1) | sz00 = X0 | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : (! [X2] : ((sdtsldt0(X1,X0) = X2 | (sdtasdt0(X0,X2) != X1 | ~aNaturalNumber0(X2))) & ((sdtasdt0(X0,X2) = X1 & aNaturalNumber0(X2)) | sdtsldt0(X1,X0) != X2)) | ~doDivides0(X0,X1) | sz00 = X0 | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(nnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : (! [X2] : (sdtsldt0(X1,X0) = X2 <=> (sdtasdt0(X0,X2) = X1 & aNaturalNumber0(X2))) | ~doDivides0(X0,X1) | sz00 = X0 | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ((! [X2] : (sdtsldt0(X1,X0) = X2 <=> (sdtasdt0(X0,X2) = X1 & aNaturalNumber0(X2))) | (~doDivides0(X0,X1) | sz00 = X0)) | (~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f31])).
fof(f31,axiom,(
  ! [X0,X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => ((doDivides0(X0,X1) & sz00 != X0) => ! [X2] : (sdtsldt0(X1,X0) = X2 <=> (sdtasdt0(X0,X2) = X1 & aNaturalNumber0(X2)))))),
  file('./p/NUM500+1.tptp',unknown)).
fof(f1207,plain,(
  ~spl4_9 | spl4_8),
  inference(avatar_split_clause,[],[f1204,f532,f536])).
fof(f532,plain,(
  spl4_8 <=> isPrime0(sK3(xk))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_8])])).
fof(f1204,plain,(
  ~aNaturalNumber0(xk) | spl4_8),
  inference(subsumption_resolution,[],[f1203,f220])).
fof(f220,plain,(
  sz00 != xk),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  sz10 != xk & sz00 != xk),
  inference(ennf_transformation,[],[f46])).
fof(f46,axiom,(
  ~(sz10 = xk | sz00 = xk)),
  file('./p/NUM500+1.tptp',unknown)).
fof(f1203,plain,(
  sz00 = xk | ~aNaturalNumber0(xk) | spl4_8),
  inference(subsumption_resolution,[],[f1202,f221])).
fof(f221,plain,(
  sz10 != xk),
  inference(cnf_transformation,[],[f118])).
fof(f1202,plain,(
  sz10 = xk | sz00 = xk | ~aNaturalNumber0(xk) | spl4_8),
  inference(resolution,[],[f534,f206])).
fof(f206,plain,(
  ( ! [X0] : (isPrime0(sK3(X0)) | sz10 = X0 | sz00 = X0 | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : ((isPrime0(sK3(X0)) & doDivides0(sK3(X0),X0) & aNaturalNumber0(sK3(X0))) | sz10 = X0 | sz00 = X0 | ~aNaturalNumber0(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f115,f137])).
fof(f137,plain,(
  ! [X0] : (? [X1] : (isPrime0(X1) & doDivides0(X1,X0) & aNaturalNumber0(X1)) => (isPrime0(sK3(X0)) & doDivides0(sK3(X0),X0) & aNaturalNumber0(sK3(X0))))),
  introduced(choice_axiom,[])).
fof(f115,plain,(
  ! [X0] : (? [X1] : (isPrime0(X1) & doDivides0(X1,X0) & aNaturalNumber0(X1)) | sz10 = X0 | sz00 = X0 | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : (? [X1] : (isPrime0(X1) & doDivides0(X1,X0) & aNaturalNumber0(X1)) | (sz10 = X0 | sz00 = X0 | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f38])).
fof(f38,axiom,(
  ! [X0] : ((sz10 != X0 & sz00 != X0 & aNaturalNumber0(X0)) => ? [X1] : (isPrime0(X1) & doDivides0(X1,X0) & aNaturalNumber0(X1)))),
  file('./p/NUM500+1.tptp',unknown)).
fof(f534,plain,(
  ~isPrime0(sK3(xk)) | spl4_8),
  inference(avatar_component_clause,[],[f532])).
fof(f543,plain,(
  ~spl4_9 | spl4_7),
  inference(avatar_split_clause,[],[f542,f528,f536])).
fof(f528,plain,(
  spl4_7 <=> aNaturalNumber0(sK3(xk))),
  introduced(avatar_definition,[new_symbols(naming,[spl4_7])])).
fof(f542,plain,(
  ~aNaturalNumber0(xk) | spl4_7),
  inference(subsumption_resolution,[],[f541,f220])).
fof(f541,plain,(
  sz00 = xk | ~aNaturalNumber0(xk) | spl4_7),
  inference(subsumption_resolution,[],[f540,f221])).
fof(f540,plain,(
  sz10 = xk | sz00 = xk | ~aNaturalNumber0(xk) | spl4_7),
  inference(resolution,[],[f530,f204])).
fof(f204,plain,(
  ( ! [X0] : (aNaturalNumber0(sK3(X0)) | sz10 = X0 | sz00 = X0 | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f138])).
fof(f530,plain,(
  ~aNaturalNumber0(sK3(xk)) | spl4_7),
  inference(avatar_component_clause,[],[f528])).
fof(f539,plain,(
  ~spl4_7 | ~spl4_8 | ~spl4_9),
  inference(avatar_split_clause,[],[f526,f536,f532,f528])).
fof(f526,plain,(
  ~aNaturalNumber0(xk) | ~isPrime0(sK3(xk)) | ~aNaturalNumber0(sK3(xk))),
  inference(subsumption_resolution,[],[f525,f220])).
fof(f525,plain,(
  sz00 = xk | ~aNaturalNumber0(xk) | ~isPrime0(sK3(xk)) | ~aNaturalNumber0(sK3(xk))),
  inference(subsumption_resolution,[],[f524,f221])).
fof(f524,plain,(
  sz10 = xk | sz00 = xk | ~aNaturalNumber0(xk) | ~isPrime0(sK3(xk)) | ~aNaturalNumber0(sK3(xk))),
  inference(resolution,[],[f205,f224])).
fof(f224,plain,(
  ( ! [X0] : (~doDivides0(X0,xk) | ~isPrime0(X0) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0] : (~isPrime0(X0) | ~doDivides0(X0,xk) | ~aNaturalNumber0(X0))),
  inference(ennf_transformation,[],[f49])).
fof(f49,negated_conjecture,(
  ~? [X0] : (isPrime0(X0) & doDivides0(X0,xk) & aNaturalNumber0(X0))),
  inference(negated_conjecture,[],[f48])).
fof(f48,conjecture,(
  ? [X0] : (isPrime0(X0) & doDivides0(X0,xk) & aNaturalNumber0(X0))),
  file('./p/NUM500+1.tptp',unknown)).
fof(f205,plain,(
  ( ! [X0] : (doDivides0(sK3(X0),X0) | sz10 = X0 | sz00 = X0 | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f138])).
fof(f477,plain,(
  ~spl4_5),
  inference(avatar_contradiction_clause,[],[f476])).
fof(f476,plain,(
  $false | ~spl4_5),
  inference(subsumption_resolution,[],[f454,f305])).
fof(f305,plain,(
  sdtlseqdt0(sz00,xm)),
  inference(subsumption_resolution,[],[f304,f139])).
fof(f139,plain,(
  aNaturalNumber0(sz00)),
  inference(cnf_transformation,[],[f2])).
fof(f2,axiom,(
  aNaturalNumber0(sz00)),
  file('./p/NUM500+1.tptp',unknown)).
fof(f304,plain,(
  sdtlseqdt0(sz00,xm) | ~aNaturalNumber0(sz00)),
  inference(subsumption_resolution,[],[f300,f208])).
fof(f300,plain,(
  sdtlseqdt0(sz00,xm) | ~aNaturalNumber0(xm) | ~aNaturalNumber0(sz00)),
  inference(superposition,[],[f238,f254])).
fof(f254,plain,(
  xm = sdtpldt0(sz00,xm)),
  inference(resolution,[],[f147,f208])).
fof(f147,plain,(
  ( ! [X0] : (~aNaturalNumber0(X0) | sdtpldt0(sz00,X0) = X0) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : ((sdtpldt0(sz00,X0) = X0 & sdtpldt0(X0,sz00) = X0) | ~aNaturalNumber0(X0))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0] : (aNaturalNumber0(X0) => (sdtpldt0(sz00,X0) = X0 & sdtpldt0(X0,sz00) = X0))),
  file('./p/NUM500+1.tptp',unknown)).
fof(f238,plain,(
  ( ! [X2,X0] : (sdtlseqdt0(X0,sdtpldt0(X0,X2)) | ~aNaturalNumber0(X2) | ~aNaturalNumber0(X0)) )),
  inference(subsumption_resolution,[],[f225,f142])).
fof(f142,plain,(
  ( ! [X0,X1] : (aNaturalNumber0(sdtpldt0(X0,X1)) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (aNaturalNumber0(sdtpldt0(X0,X1)) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : (aNaturalNumber0(sdtpldt0(X0,X1)) | (~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => aNaturalNumber0(sdtpldt0(X0,X1)))),
  file('./p/NUM500+1.tptp',unknown)).
fof(f225,plain,(
  ( ! [X2,X0] : (sdtlseqdt0(X0,sdtpldt0(X0,X2)) | ~aNaturalNumber0(X2) | ~aNaturalNumber0(sdtpldt0(X0,X2)) | ~aNaturalNumber0(X0)) )),
  inference(equality_resolution,[],[f165])).
fof(f165,plain,(
  ( ! [X2,X0,X1] : (sdtlseqdt0(X0,X1) | sdtpldt0(X0,X2) != X1 | ~aNaturalNumber0(X2) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : (((sdtlseqdt0(X0,X1) | ! [X2] : (sdtpldt0(X0,X2) != X1 | ~aNaturalNumber0(X2))) & ((sdtpldt0(X0,sK0(X0,X1)) = X1 & aNaturalNumber0(sK0(X0,X1))) | ~sdtlseqdt0(X0,X1))) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f121,f122])).
fof(f122,plain,(
  ! [X1,X0] : (? [X3] : (sdtpldt0(X0,X3) = X1 & aNaturalNumber0(X3)) => (sdtpldt0(X0,sK0(X0,X1)) = X1 & aNaturalNumber0(sK0(X0,X1))))),
  introduced(choice_axiom,[])).
fof(f121,plain,(
  ! [X0,X1] : (((sdtlseqdt0(X0,X1) | ! [X2] : (sdtpldt0(X0,X2) != X1 | ~aNaturalNumber0(X2))) & (? [X3] : (sdtpldt0(X0,X3) = X1 & aNaturalNumber0(X3)) | ~sdtlseqdt0(X0,X1))) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(rectify,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : (((sdtlseqdt0(X0,X1) | ! [X2] : (sdtpldt0(X0,X2) != X1 | ~aNaturalNumber0(X2))) & (? [X2] : (sdtpldt0(X0,X2) = X1 & aNaturalNumber0(X2)) | ~sdtlseqdt0(X0,X1))) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ((sdtlseqdt0(X0,X1) <=> ? [X2] : (sdtpldt0(X0,X2) = X1 & aNaturalNumber0(X2))) | ~aNaturalNumber0(X1) | ~aNaturalNumber0(X0))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : ((sdtlseqdt0(X0,X1) <=> ? [X2] : (sdtpldt0(X0,X2) = X1 & aNaturalNumber0(X2))) | (~aNaturalNumber0(X1) | ~aNaturalNumber0(X0)))),
  inference(ennf_transformation,[],[f18])).
fof(f18,axiom,(
  ! [X0,X1] : ((aNaturalNumber0(X1) & aNaturalNumber0(X0)) => (sdtlseqdt0(X0,X1) <=> ? [X2] : (sdtpldt0(X0,X2) = X1 & aNaturalNumber0(X2))))),
  file('./p/NUM500+1.tptp',unknown)).
fof(f454,plain,(
  ~sdtlseqdt0(sz00,xm) | ~spl4_5),
  inference(backward_demodulation,[],[f443,f214])).
fof(f214,plain,(
  ~sdtlseqdt0(xp,xm)),
  inference(cnf_transformation,[],[f43])).
fof(f43,axiom,(
  ~sdtlseqdt0(xp,xm)),
  file('./p/NUM500+1.tptp',unknown)).
fof(f443,plain,(
  sz00 = xp | ~spl4_5),
  inference(avatar_component_clause,[],[f441])).
% SZS output end Proof for NUM500+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5756
% Time elapsed: 0.019 s
% ------------------------------
% ------------------------------
