% Refutation found. Thanks to Tanya!
% SZS status Theorem for NUM851+2
% SZS output start Proof for NUM851+2
fof(f95383,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f95382])).
fof(f95382,plain,(
  vmul(vd469,vd470) != vmul(vd469,vd470)),
  inference(forward_demodulation,[],[f95268,f69])).
fof(f69,plain,(
  vplus(vd471,vd473) = vd470),
  inference(cnf_transformation,[],[f4])).
fof(f4,axiom,(
  vplus(vd471,vd473) = vd470),
  file('./p/NUM851+2.tptp',unknown)).
fof(f95268,plain,(
  vmul(vd469,vd470) != vmul(vd469,vplus(vd471,vd473))),
  inference(superposition,[],[f106,f3411])).
fof(f3411,plain,(
  ( ! [X2,X0,X1] : (vmul(X0,vplus(X2,X1)) = vplus(vmul(X2,X0),vmul(X1,X0))) )),
  inference(superposition,[],[f365,f73])).
fof(f73,plain,(
  ( ! [X0,X1] : (vmul(X0,X1) = vmul(X1,X0)) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : vmul(X0,X1) = vmul(X1,X0)),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7] : vmul(X6,X7) = vmul(X7,X6)),
  file('./p/NUM851+2.tptp',unknown)).
fof(f365,plain,(
  ( ! [X4,X2,X3] : (vmul(X2,vplus(X3,X4)) = vplus(vmul(X3,X2),vmul(X2,X4))) )),
  inference(superposition,[],[f72,f73])).
fof(f72,plain,(
  ( ! [X2,X0,X1] : (vmul(X0,vplus(X1,X2)) = vplus(vmul(X0,X1),vmul(X0,X2))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : vmul(X0,vplus(X1,X2)) = vplus(vmul(X0,X1),vmul(X0,X2))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X3,X4,X5] : vmul(X3,vplus(X4,X5)) = vplus(vmul(X3,X4),vmul(X3,X5))),
  file('./p/NUM851+2.tptp',unknown)).
fof(f106,plain,(
  vplus(vmul(vd471,vd469),vmul(vd473,vd469)) != vmul(vd469,vd470)),
  inference(backward_demodulation,[],[f73,f105])).
fof(f105,plain,(
  vplus(vmul(vd471,vd469),vmul(vd473,vd469)) != vmul(vd470,vd469)),
  inference(backward_demodulation,[],[f68,f67])).
fof(f67,plain,(
  vmul(vplus(vd471,vd473),vd469) != vplus(vmul(vd471,vd469),vmul(vd473,vd469))),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  vmul(vplus(vd471,vd473),vd469) != vplus(vmul(vd471,vd469),vmul(vd473,vd469))),
  inference(flattening,[],[f2])).
fof(f2,negated_conjecture,(
  ~vmul(vplus(vd471,vd473),vd469) = vplus(vmul(vd471,vd469),vmul(vd473,vd469))),
  inference(negated_conjecture,[],[f1])).
fof(f1,conjecture,(
  vmul(vplus(vd471,vd473),vd469) = vplus(vmul(vd471,vd469),vmul(vd473,vd469))),
  file('./p/NUM851+2.tptp',unknown)).
fof(f68,plain,(
  vmul(vplus(vd471,vd473),vd469) = vmul(vd470,vd469)),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  vmul(vplus(vd471,vd473),vd469) = vmul(vd470,vd469)),
  file('./p/NUM851+2.tptp',unknown)).
% SZS output end Proof for NUM851+2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 159144
% Time elapsed: 6.217 s
% ------------------------------
% ------------------------------
