% Refutation found. Thanks to Tanya!
% SZS status Theorem for NUM842+2
% SZS output start Proof for NUM842+2
fof(f23022,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f199,f204,f210,f301,f22351,f22357,f22373,f22374,f22996])).
fof(f22996,plain,(
  ~spl4_2 | ~spl4_5),
  inference(avatar_contradiction_clause,[],[f22995])).
fof(f22995,plain,(
  $false | (~spl4_2 | ~spl4_5)),
  inference(subsumption_resolution,[],[f22987,f198])).
fof(f198,plain,(
  greater(vd353,vd354) | ~spl4_2),
  inference(avatar_component_clause,[],[f196])).
fof(f196,plain,(
  spl4_2 <=> greater(vd353,vd354)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_2])])).
fof(f22987,plain,(
  ~greater(vd353,vd354) | ~spl4_5),
  inference(resolution,[],[f22776,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (greater(vplus(X2,X1),vplus(X1,X0)) | ~greater(X2,X0)) )),
  inference(superposition,[],[f132,f162])).
fof(f162,plain,(
  ( ! [X0,X1] : (vplus(X0,X1) = vplus(X1,X0)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : vplus(X0,X1) = vplus(X1,X0)),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X50,X51] : vplus(X51,X50) = vplus(X50,X51)),
  file('./p/NUM842+2.tptp',unknown)).
fof(f132,plain,(
  ( ! [X2,X0,X1] : (greater(vplus(X0,X2),vplus(X1,X2)) | ~greater(X0,X1)) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : (greater(vplus(X0,X2),vplus(X1,X2)) | ~greater(X0,X1))),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : (greater(X0,X1) => greater(vplus(X0,X2),vplus(X1,X2)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X7,X8,X9] : (greater(X7,X8) => greater(vplus(X7,X9),vplus(X8,X9)))),
  file('./p/NUM842+2.tptp',unknown)).
fof(f22776,plain,(
  ~greater(vplus(vd353,vd355),vplus(vd355,vd354)) | ~spl4_5),
  inference(forward_demodulation,[],[f22377,f162])).
fof(f22377,plain,(
  ~greater(vplus(vd353,vd355),vplus(vd354,vd355)) | ~spl4_5),
  inference(backward_demodulation,[],[f291,f121])).
fof(f121,plain,(
  ~greater(vplus(vd353,vd355),vplus(vd354,vd356))),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ~greater(vplus(vd353,vd355),vplus(vd354,vd356))),
  inference(flattening,[],[f2])).
fof(f2,negated_conjecture,(
  ~greater(vplus(vd353,vd355),vplus(vd354,vd356))),
  inference(negated_conjecture,[],[f1])).
fof(f1,conjecture,(
  greater(vplus(vd353,vd355),vplus(vd354,vd356))),
  file('./p/NUM842+2.tptp',unknown)).
fof(f291,plain,(
  vd355 = vd356 | ~spl4_5),
  inference(avatar_component_clause,[],[f289])).
fof(f289,plain,(
  spl4_5 <=> vd355 = vd356),
  introduced(avatar_definition,[new_symbols(naming,[spl4_5])])).
fof(f22374,plain,(
  spl4_5 | spl4_3 | ~spl4_4),
  inference(avatar_split_clause,[],[f1521,f206,f201,f289])).
fof(f201,plain,(
  spl4_3 <=> greater(vd355,vd356)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_3])])).
fof(f206,plain,(
  spl4_4 <=> geq(vd355,vd356)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_4])])).
fof(f1521,plain,(
  greater(vd355,vd356) | vd355 = vd356 | ~spl4_4),
  inference(resolution,[],[f208,f143])).
fof(f143,plain,(
  ( ! [X0,X1] : (~geq(X1,X0) | greater(X1,X0) | X0 = X1) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((geq(X1,X0) | (X0 != X1 & ~greater(X1,X0))) & (X0 = X1 | greater(X1,X0) | ~geq(X1,X0)))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((geq(X1,X0) | (X0 != X1 & ~greater(X1,X0))) & ((X0 = X1 | greater(X1,X0)) | ~geq(X1,X0)))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : (geq(X1,X0) <=> (X0 = X1 | greater(X1,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X27,X28] : (geq(X28,X27) <=> (X27 = X28 | greater(X28,X27)))),
  file('./p/NUM842+2.tptp',unknown)).
fof(f208,plain,(
  geq(vd355,vd356) | ~spl4_4),
  inference(avatar_component_clause,[],[f206])).
fof(f22373,plain,(
  ~spl4_3 | ~spl4_2),
  inference(avatar_split_clause,[],[f1531,f196,f201])).
fof(f1531,plain,(
  ~greater(vd353,vd354) | ~greater(vd355,vd356)),
  inference(resolution,[],[f121,f126])).
fof(f126,plain,(
  ( ! [X2,X0,X3,X1] : (greater(vplus(X0,X2),vplus(X1,X3)) | ~greater(X0,X1) | ~greater(X2,X3)) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3] : (greater(vplus(X0,X2),vplus(X1,X3)) | ~greater(X0,X1) | ~greater(X2,X3))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3] : (greater(vplus(X0,X2),vplus(X1,X3)) | (~greater(X0,X1) | ~greater(X2,X3)))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1,X2,X3] : ((greater(X0,X1) & greater(X2,X3)) => greater(vplus(X0,X2),vplus(X1,X3)))),
  file('./p/NUM842+2.tptp',unknown)).
fof(f22357,plain,(
  spl4_17 | spl4_2 | ~spl4_1),
  inference(avatar_split_clause,[],[f307,f192,f196,f1914])).
fof(f1914,plain,(
  spl4_17 <=> vd353 = vd354),
  introduced(avatar_definition,[new_symbols(naming,[spl4_17])])).
fof(f192,plain,(
  spl4_1 <=> geq(vd353,vd354)),
  introduced(avatar_definition,[new_symbols(naming,[spl4_1])])).
fof(f307,plain,(
  greater(vd353,vd354) | vd353 = vd354 | ~spl4_1),
  inference(resolution,[],[f143,f194])).
fof(f194,plain,(
  geq(vd353,vd354) | ~spl4_1),
  inference(avatar_component_clause,[],[f192])).
fof(f22351,plain,(
  ~spl4_3 | ~spl4_17),
  inference(avatar_contradiction_clause,[],[f22350])).
fof(f22350,plain,(
  $false | (~spl4_3 | ~spl4_17)),
  inference(subsumption_resolution,[],[f22349,f2274])).
fof(f2274,plain,(
  ( ! [X19] : (greater(vplus(X19,vd355),vplus(X19,vd356))) ) | ~spl4_3),
  inference(superposition,[],[f458,f1525])).
fof(f1525,plain,(
  vd355 = vplus(vd356,sK1(vd356,vd355)) | ~spl4_3),
  inference(resolution,[],[f203,f154])).
fof(f154,plain,(
  ( ! [X0,X1] : (~greater(X1,X0) | vplus(X0,sK1(X0,X1)) = X1) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ((greater(X1,X0) | ! [X2] : vplus(X0,X2) != X1) & (vplus(X0,sK1(X0,X1)) = X1 | ~greater(X1,X0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f115,f116])).
fof(f116,plain,(
  ! [X1,X0] : (? [X3] : vplus(X0,X3) = X1 => vplus(X0,sK1(X0,X1)) = X1)),
  introduced(choice_axiom,[])).
fof(f115,plain,(
  ! [X0,X1] : ((greater(X1,X0) | ! [X2] : vplus(X0,X2) != X1) & (? [X3] : vplus(X0,X3) = X1 | ~greater(X1,X0)))),
  inference(rectify,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : ((greater(X1,X0) | ! [X2] : vplus(X0,X2) != X1) & (? [X2] : vplus(X0,X2) = X1 | ~greater(X1,X0)))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : (greater(X1,X0) <=> ? [X2] : vplus(X0,X2) = X1)),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X38,X39] : (greater(X39,X38) <=> ? [X40] : vplus(X38,X40) = X39)),
  file('./p/NUM842+2.tptp',unknown)).
fof(f203,plain,(
  greater(vd355,vd356) | ~spl4_3),
  inference(avatar_component_clause,[],[f201])).
fof(f458,plain,(
  ( ! [X45,X43,X44] : (greater(vplus(X43,vplus(X44,X45)),vplus(X43,X44))) )),
  inference(superposition,[],[f133,f165])).
fof(f165,plain,(
  ( ! [X2,X0,X1] : (vplus(vplus(X0,X1),X2) = vplus(X0,vplus(X1,X2))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2] : vplus(vplus(X0,X1),X2) = vplus(X0,vplus(X1,X2))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X55,X56,X57] : vplus(vplus(X55,X56),X57) = vplus(X55,vplus(X56,X57))),
  file('./p/NUM842+2.tptp',unknown)).
fof(f133,plain,(
  ( ! [X0,X1] : (greater(vplus(X0,X1),X0)) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : greater(vplus(X0,X1),X0)),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X10,X11] : greater(vplus(X10,X11),X10)),
  file('./p/NUM842+2.tptp',unknown)).
fof(f22349,plain,(
  ~greater(vplus(vd353,vd355),vplus(vd353,vd356)) | ~spl4_17),
  inference(forward_demodulation,[],[f121,f1916])).
fof(f1916,plain,(
  vd353 = vd354 | ~spl4_17),
  inference(avatar_component_clause,[],[f1914])).
fof(f301,plain,(
  ~spl4_3 | spl4_4),
  inference(avatar_contradiction_clause,[],[f300])).
fof(f300,plain,(
  $false | (~spl4_3 | spl4_4)),
  inference(subsumption_resolution,[],[f298,f207])).
fof(f207,plain,(
  ~geq(vd355,vd356) | spl4_4),
  inference(avatar_component_clause,[],[f206])).
fof(f298,plain,(
  geq(vd355,vd356) | ~spl4_3),
  inference(resolution,[],[f203,f144])).
fof(f144,plain,(
  ( ! [X0,X1] : (~greater(X1,X0) | geq(X1,X0)) )),
  inference(cnf_transformation,[],[f109])).
fof(f210,plain,(
  spl4_3 | spl4_4),
  inference(avatar_split_clause,[],[f122,f206,f201])).
fof(f122,plain,(
  geq(vd355,vd356) | greater(vd355,vd356)),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  (greater(vd353,vd354) & geq(vd355,vd356)) | (geq(vd353,vd354) & greater(vd355,vd356))),
  file('./p/NUM842+2.tptp',unknown)).
fof(f204,plain,(
  spl4_3 | spl4_2),
  inference(avatar_split_clause,[],[f124,f196,f201])).
fof(f124,plain,(
  greater(vd353,vd354) | greater(vd355,vd356)),
  inference(cnf_transformation,[],[f3])).
fof(f199,plain,(
  spl4_1 | spl4_2),
  inference(avatar_split_clause,[],[f125,f196,f192])).
fof(f125,plain,(
  greater(vd353,vd354) | geq(vd353,vd354)),
  inference(cnf_transformation,[],[f3])).
% SZS output end Proof for NUM842+2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 42600
% Time elapsed: 1.154 s
% ------------------------------
% ------------------------------
