% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for PLA008-1
% SZS output start Proof for PLA008-1
fof(f2576,plain,(
  $false),
  inference(subsumption_resolution,[],[f2575,f26])).
fof(f26,axiom,(
  holds(clear(a),s0)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f2575,plain,(
  ~holds(clear(a),s0)),
  inference(subsumption_resolution,[],[f2574,f27])).
fof(f27,axiom,(
  holds(clear(b),s0)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f2574,plain,(
  ~holds(clear(b),s0) | ~holds(clear(a),s0)),
  inference(subsumption_resolution,[],[f2573,f29])).
fof(f29,axiom,(
  holds(empty,s0)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f2573,plain,(
  ~holds(empty,s0) | ~holds(clear(b),s0) | ~holds(clear(a),s0)),
  inference(subsumption_resolution,[],[f2570,f28])).
fof(f28,axiom,(
  holds(clear(c),s0)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f2570,plain,(
  ~holds(clear(c),s0) | ~holds(empty,s0) | ~holds(clear(b),s0) | ~holds(clear(a),s0)),
  inference(resolution,[],[f2565,f24])).
fof(f24,axiom,(
  holds(on(c,d),s0)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f2565,plain,(
  ( ! [X0] : (~holds(on(c,d),X0) | ~holds(clear(c),X0) | ~holds(empty,X0) | ~holds(clear(b),X0) | ~holds(clear(a),X0)) )),
  inference(subsumption_resolution,[],[f2564,f20])).
fof(f20,axiom,(
  differ(c,table)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f2564,plain,(
  ( ! [X0] : (~holds(clear(c),X0) | ~holds(on(c,d),X0) | ~holds(empty,X0) | ~holds(clear(b),X0) | ~holds(clear(a),X0) | ~differ(c,table)) )),
  inference(subsumption_resolution,[],[f2563,f13])).
fof(f13,axiom,(
  differ(a,c)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f2563,plain,(
  ( ! [X0] : (~holds(clear(c),X0) | ~holds(on(c,d),X0) | ~holds(empty,X0) | ~holds(clear(b),X0) | ~holds(clear(a),X0) | ~differ(a,c) | ~differ(c,table)) )),
  inference(subsumption_resolution,[],[f2562,f16])).
fof(f16,axiom,(
  differ(b,c)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f2562,plain,(
  ( ! [X0] : (~holds(clear(c),X0) | ~holds(on(c,d),X0) | ~holds(empty,X0) | ~holds(clear(b),X0) | ~differ(b,c) | ~holds(clear(a),X0) | ~differ(a,c) | ~differ(c,table)) )),
  inference(duplicate_literal_removal,[],[f2561])).
fof(f2561,plain,(
  ( ! [X0] : (~holds(clear(c),X0) | ~holds(on(c,d),X0) | ~holds(empty,X0) | ~holds(clear(b),X0) | ~differ(b,c) | ~holds(clear(a),X0) | ~differ(a,c) | ~holds(clear(c),X0) | ~holds(empty,X0) | ~differ(c,table)) )),
  inference(resolution,[],[f2462,f2])).
fof(f2,axiom,(
  ( ! [X2,X0] : (holds(holding(X0),do(pickup(X0),X2)) | ~holds(clear(X0),X2) | ~holds(empty,X2) | ~differ(X0,table)) )),
  file('./p/PLA008-1.tptp',unknown)).
fof(f2462,plain,(
  ( ! [X2,X3] : (~holds(holding(c),do(pickup(X2),X3)) | ~holds(clear(X2),X3) | ~holds(on(X2,d),X3) | ~holds(empty,X3) | ~holds(clear(b),X3) | ~differ(b,X2) | ~holds(clear(a),X3) | ~differ(a,X2)) )),
  inference(resolution,[],[f1646,f5])).
fof(f5,axiom,(
  ( ! [X2,X0,X3] : (holds(clear(X0),do(pickup(X3),X2)) | ~holds(clear(X0),X2) | ~differ(X0,X3)) )),
  file('./p/PLA008-1.tptp',unknown)).
fof(f1646,plain,(
  ( ! [X2,X3] : (~holds(clear(a),do(pickup(X2),X3)) | ~holds(holding(c),do(pickup(X2),X3)) | ~holds(clear(X2),X3) | ~holds(on(X2,d),X3) | ~holds(empty,X3) | ~holds(clear(b),X3) | ~differ(b,X2)) )),
  inference(resolution,[],[f1216,f5])).
fof(f1216,plain,(
  ( ! [X0,X1] : (~holds(clear(b),do(pickup(X0),X1)) | ~holds(holding(c),do(pickup(X0),X1)) | ~holds(clear(a),do(pickup(X0),X1)) | ~holds(clear(X0),X1) | ~holds(on(X0,d),X1) | ~holds(empty,X1)) )),
  inference(resolution,[],[f1212,f3])).
fof(f3,axiom,(
  ( ! [X2,X0,X1] : (holds(clear(X1),do(pickup(X0),X2)) | ~holds(clear(X0),X2) | ~holds(on(X0,X1),X2) | ~holds(empty,X2)) )),
  file('./p/PLA008-1.tptp',unknown)).
fof(f1212,plain,(
  ( ! [X0] : (~holds(clear(d),X0) | ~holds(clear(a),X0) | ~holds(holding(c),X0) | ~holds(clear(b),X0)) )),
  inference(subsumption_resolution,[],[f1211,f33])).
fof(f33,plain,(
  differ(b,a)),
  inference(resolution,[],[f11,f12])).
fof(f12,axiom,(
  differ(a,b)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f11,axiom,(
  ( ! [X0,X1] : (~differ(X1,X0) | differ(X0,X1)) )),
  file('./p/PLA008-1.tptp',unknown)).
fof(f1211,plain,(
  ( ! [X0] : (~holds(holding(c),X0) | ~holds(clear(a),X0) | ~holds(clear(d),X0) | ~holds(clear(b),X0) | ~differ(b,a)) )),
  inference(resolution,[],[f1206,f10])).
fof(f10,axiom,(
  ( ! [X2,X0,X3,X1] : (holds(clear(X3),do(putdown(X0,X1),X2)) | ~holds(clear(X3),X2) | ~differ(X3,X1)) )),
  file('./p/PLA008-1.tptp',unknown)).
fof(f1206,plain,(
  ( ! [X0] : (~holds(clear(b),do(putdown(c,a),X0)) | ~holds(holding(c),X0) | ~holds(clear(a),X0) | ~holds(clear(d),X0)) )),
  inference(subsumption_resolution,[],[f1205,f35])).
fof(f35,plain,(
  differ(d,a)),
  inference(resolution,[],[f11,f14])).
fof(f14,axiom,(
  differ(a,d)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f1205,plain,(
  ( ! [X0] : (~holds(clear(b),do(putdown(c,a),X0)) | ~holds(holding(c),X0) | ~holds(clear(a),X0) | ~holds(clear(d),X0) | ~differ(d,a)) )),
  inference(resolution,[],[f1201,f10])).
fof(f1201,plain,(
  ( ! [X3] : (~holds(clear(d),do(putdown(c,a),X3)) | ~holds(clear(b),do(putdown(c,a),X3)) | ~holds(holding(c),X3) | ~holds(clear(a),X3)) )),
  inference(subsumption_resolution,[],[f1200,f6])).
fof(f6,axiom,(
  ( ! [X2,X0,X1] : (holds(empty,do(putdown(X0,X1),X2)) | ~holds(holding(X0),X2) | ~holds(clear(X1),X2)) )),
  file('./p/PLA008-1.tptp',unknown)).
fof(f1200,plain,(
  ( ! [X3] : (~holds(clear(d),do(putdown(c,a),X3)) | ~holds(clear(b),do(putdown(c,a),X3)) | ~holds(empty,do(putdown(c,a),X3)) | ~holds(holding(c),X3) | ~holds(clear(a),X3)) )),
  inference(subsumption_resolution,[],[f1198,f8])).
fof(f8,axiom,(
  ( ! [X2,X0,X1] : (holds(clear(X0),do(putdown(X0,X1),X2)) | ~holds(holding(X0),X2) | ~holds(clear(X1),X2)) )),
  file('./p/PLA008-1.tptp',unknown)).
fof(f1198,plain,(
  ( ! [X3] : (~holds(clear(d),do(putdown(c,a),X3)) | ~holds(clear(c),do(putdown(c,a),X3)) | ~holds(clear(b),do(putdown(c,a),X3)) | ~holds(empty,do(putdown(c,a),X3)) | ~holds(holding(c),X3) | ~holds(clear(a),X3)) )),
  inference(resolution,[],[f1192,f7])).
fof(f7,axiom,(
  ( ! [X2,X0,X1] : (holds(on(X0,X1),do(putdown(X0,X1),X2)) | ~holds(holding(X0),X2) | ~holds(clear(X1),X2)) )),
  file('./p/PLA008-1.tptp',unknown)).
fof(f1192,plain,(
  ( ! [X0] : (~holds(on(c,a),X0) | ~holds(clear(d),X0) | ~holds(clear(c),X0) | ~holds(clear(b),X0) | ~holds(empty,X0)) )),
  inference(subsumption_resolution,[],[f1191,f18])).
fof(f18,axiom,(
  differ(b,table)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f1191,plain,(
  ( ! [X0] : (~holds(on(c,a),X0) | ~holds(clear(d),X0) | ~holds(clear(c),X0) | ~holds(clear(b),X0) | ~holds(empty,X0) | ~differ(b,table)) )),
  inference(subsumption_resolution,[],[f1190,f37])).
fof(f37,plain,(
  differ(d,b)),
  inference(resolution,[],[f11,f17])).
fof(f17,axiom,(
  differ(b,d)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f1190,plain,(
  ( ! [X0] : (~holds(on(c,a),X0) | ~holds(clear(d),X0) | ~differ(d,b) | ~holds(clear(c),X0) | ~holds(clear(b),X0) | ~holds(empty,X0) | ~differ(b,table)) )),
  inference(subsumption_resolution,[],[f1189,f36])).
fof(f36,plain,(
  differ(c,b)),
  inference(resolution,[],[f11,f16])).
fof(f1189,plain,(
  ( ! [X0] : (~holds(on(c,a),X0) | ~differ(c,b) | ~holds(clear(d),X0) | ~differ(d,b) | ~holds(clear(c),X0) | ~holds(clear(b),X0) | ~holds(empty,X0) | ~differ(b,table)) )),
  inference(resolution,[],[f1184,f2])).
fof(f1184,plain,(
  ( ! [X2,X3] : (~holds(holding(b),do(pickup(X2),X3)) | ~holds(on(c,a),X3) | ~differ(c,X2) | ~holds(clear(d),X3) | ~differ(d,X2) | ~holds(clear(c),X3)) )),
  inference(duplicate_literal_removal,[],[f1183])).
fof(f1183,plain,(
  ( ! [X2,X3] : (~holds(holding(b),do(pickup(X2),X3)) | ~holds(on(c,a),X3) | ~differ(c,X2) | ~holds(clear(d),X3) | ~differ(d,X2) | ~holds(clear(c),X3) | ~differ(c,X2)) )),
  inference(resolution,[],[f978,f5])).
fof(f978,plain,(
  ( ! [X2,X3] : (~holds(clear(c),do(pickup(X2),X3)) | ~holds(holding(b),do(pickup(X2),X3)) | ~holds(on(c,a),X3) | ~differ(c,X2) | ~holds(clear(d),X3) | ~differ(d,X2)) )),
  inference(resolution,[],[f592,f5])).
fof(f592,plain,(
  ( ! [X4,X5] : (~holds(clear(d),do(pickup(X4),X5)) | ~holds(holding(b),do(pickup(X4),X5)) | ~holds(clear(c),do(pickup(X4),X5)) | ~holds(on(c,a),X5) | ~differ(c,X4)) )),
  inference(resolution,[],[f586,f4])).
fof(f4,axiom,(
  ( ! [X2,X0,X3,X1] : (holds(on(X0,X1),do(pickup(X3),X2)) | ~holds(on(X0,X1),X2) | ~differ(X0,X3)) )),
  file('./p/PLA008-1.tptp',unknown)).
fof(f586,plain,(
  ( ! [X0] : (~holds(on(c,a),X0) | ~holds(holding(b),X0) | ~holds(clear(d),X0) | ~holds(clear(c),X0)) )),
  inference(subsumption_resolution,[],[f585,f19])).
fof(f19,axiom,(
  differ(c,d)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f585,plain,(
  ( ! [X0] : (~holds(on(c,a),X0) | ~holds(holding(b),X0) | ~holds(clear(d),X0) | ~holds(clear(c),X0) | ~differ(c,d)) )),
  inference(resolution,[],[f583,f10])).
fof(f583,plain,(
  ( ! [X3] : (~holds(clear(c),do(putdown(b,d),X3)) | ~holds(on(c,a),X3) | ~holds(holding(b),X3) | ~holds(clear(d),X3)) )),
  inference(subsumption_resolution,[],[f582,f6])).
fof(f582,plain,(
  ( ! [X3] : (~holds(clear(c),do(putdown(b,d),X3)) | ~holds(empty,do(putdown(b,d),X3)) | ~holds(on(c,a),X3) | ~holds(holding(b),X3) | ~holds(clear(d),X3)) )),
  inference(resolution,[],[f327,f7])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (~holds(on(b,d),do(putdown(X0,X1),X2)) | ~holds(clear(c),do(putdown(X0,X1),X2)) | ~holds(empty,do(putdown(X0,X1),X2)) | ~holds(on(c,a),X2)) )),
  inference(resolution,[],[f324,f9])).
fof(f9,axiom,(
  ( ! [X4,X2,X0,X3,X1] : (holds(on(X0,X1),do(putdown(X3,X4),X2)) | ~holds(on(X0,X1),X2)) )),
  file('./p/PLA008-1.tptp',unknown)).
fof(f324,plain,(
  ( ! [X0] : (~holds(on(c,a),X0) | ~holds(clear(c),X0) | ~holds(on(b,d),X0) | ~holds(empty,X0)) )),
  inference(subsumption_resolution,[],[f323,f20])).
fof(f323,plain,(
  ( ! [X0] : (~holds(on(b,d),X0) | ~holds(clear(c),X0) | ~holds(on(c,a),X0) | ~holds(empty,X0) | ~differ(c,table)) )),
  inference(subsumption_resolution,[],[f322,f16])).
fof(f322,plain,(
  ( ! [X0] : (~holds(on(b,d),X0) | ~differ(b,c) | ~holds(clear(c),X0) | ~holds(on(c,a),X0) | ~holds(empty,X0) | ~differ(c,table)) )),
  inference(duplicate_literal_removal,[],[f321])).
fof(f321,plain,(
  ( ! [X0] : (~holds(on(b,d),X0) | ~differ(b,c) | ~holds(clear(c),X0) | ~holds(on(c,a),X0) | ~holds(empty,X0) | ~holds(clear(c),X0) | ~holds(empty,X0) | ~differ(c,table)) )),
  inference(resolution,[],[f299,f2])).
fof(f299,plain,(
  ( ! [X6,X7] : (~holds(holding(c),do(pickup(X6),X7)) | ~holds(on(b,d),X7) | ~differ(b,X6) | ~holds(clear(X6),X7) | ~holds(on(X6,a),X7) | ~holds(empty,X7)) )),
  inference(subsumption_resolution,[],[f297,f15])).
fof(f15,axiom,(
  differ(a,table)),
  file('./p/PLA008-1.tptp',unknown)).
fof(f297,plain,(
  ( ! [X6,X7] : (~holds(holding(c),do(pickup(X6),X7)) | ~differ(a,table) | ~holds(on(b,d),X7) | ~differ(b,X6) | ~holds(clear(X6),X7) | ~holds(on(X6,a),X7) | ~holds(empty,X7)) )),
  inference(resolution,[],[f185,f30])).
fof(f30,axiom,(
  ( ! [X2] : (holds(clear(table),X2)) )),
  file('./p/PLA008-1.tptp',unknown)).
fof(f185,plain,(
  ( ! [X2,X0,X1] : (~holds(clear(X0),do(pickup(X1),X2)) | ~holds(holding(c),do(pickup(X1),X2)) | ~differ(a,X0) | ~holds(on(b,d),X2) | ~differ(b,X1) | ~holds(clear(X1),X2) | ~holds(on(X1,a),X2) | ~holds(empty,X2)) )),
  inference(resolution,[],[f127,f3])).
fof(f127,plain,(
  ( ! [X6,X8,X7] : (~holds(clear(a),do(pickup(X6),X7)) | ~holds(clear(X8),do(pickup(X6),X7)) | ~holds(holding(c),do(pickup(X6),X7)) | ~differ(a,X8) | ~holds(on(b,d),X7) | ~differ(b,X6)) )),
  inference(resolution,[],[f122,f4])).
fof(f122,plain,(
  ( ! [X0,X1] : (~holds(on(b,d),X0) | ~holds(holding(c),X0) | ~holds(clear(X1),X0) | ~holds(clear(a),X0) | ~differ(a,X1)) )),
  inference(resolution,[],[f110,f10])).
fof(f110,plain,(
  ( ! [X4,X3] : (~holds(clear(a),do(putdown(c,X3),X4)) | ~holds(on(b,d),X4) | ~holds(holding(c),X4) | ~holds(clear(X3),X4)) )),
  inference(subsumption_resolution,[],[f109,f6])).
fof(f109,plain,(
  ( ! [X4,X3] : (~holds(clear(a),do(putdown(c,X3),X4)) | ~holds(empty,do(putdown(c,X3),X4)) | ~holds(on(b,d),X4) | ~holds(holding(c),X4) | ~holds(clear(X3),X4)) )),
  inference(resolution,[],[f82,f8])).
fof(f82,plain,(
  ( ! [X2,X0,X1] : (~holds(clear(c),do(putdown(X0,X1),X2)) | ~holds(clear(a),do(putdown(X0,X1),X2)) | ~holds(empty,do(putdown(X0,X1),X2)) | ~holds(on(b,d),X2)) )),
  inference(resolution,[],[f81,f9])).
fof(f81,plain,(
  ( ! [X0] : (~holds(on(b,d),X0) | ~holds(clear(c),X0) | ~holds(clear(a),X0) | ~holds(empty,X0)) )),
  inference(subsumption_resolution,[],[f80,f15])).
fof(f80,plain,(
  ( ! [X0] : (~holds(on(b,d),X0) | ~holds(clear(c),X0) | ~holds(clear(a),X0) | ~holds(empty,X0) | ~differ(a,table)) )),
  inference(subsumption_resolution,[],[f79,f34])).
fof(f34,plain,(
  differ(c,a)),
  inference(resolution,[],[f11,f13])).
fof(f79,plain,(
  ( ! [X0] : (~holds(on(b,d),X0) | ~holds(clear(c),X0) | ~differ(c,a) | ~holds(clear(a),X0) | ~holds(empty,X0) | ~differ(a,table)) )),
  inference(subsumption_resolution,[],[f78,f33])).
fof(f78,plain,(
  ( ! [X0] : (~holds(on(b,d),X0) | ~differ(b,a) | ~holds(clear(c),X0) | ~differ(c,a) | ~holds(clear(a),X0) | ~holds(empty,X0) | ~differ(a,table)) )),
  inference(resolution,[],[f71,f2])).
fof(f71,plain,(
  ( ! [X2,X3] : (~holds(holding(a),do(pickup(X2),X3)) | ~holds(on(b,d),X3) | ~differ(b,X2) | ~holds(clear(c),X3) | ~differ(c,X2)) )),
  inference(resolution,[],[f61,f5])).
fof(f61,plain,(
  ( ! [X4,X5] : (~holds(clear(c),do(pickup(X4),X5)) | ~holds(holding(a),do(pickup(X4),X5)) | ~holds(on(b,d),X5) | ~differ(b,X4)) )),
  inference(resolution,[],[f57,f4])).
fof(f57,plain,(
  ( ! [X1] : (~holds(on(b,d),X1) | ~holds(holding(a),X1) | ~holds(clear(c),X1)) )),
  inference(resolution,[],[f53,f7])).
fof(f53,plain,(
  ( ! [X2,X0,X1] : (~holds(on(a,c),do(putdown(X0,X1),X2)) | ~holds(on(b,d),X2)) )),
  inference(resolution,[],[f44,f9])).
fof(f44,plain,(
  ( ! [X0] : (~holds(on(b,d),X0) | ~holds(on(a,c),X0)) )),
  inference(resolution,[],[f1,f31])).
fof(f31,axiom,(
  ( ! [X2] : (~holds(and(on(b,d),on(a,c)),X2)) )),
  file('./p/PLA008-1.tptp',unknown)).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (holds(and(X0,X1),X2) | ~holds(X0,X2) | ~holds(X1,X2)) )),
  file('./p/PLA008-1.tptp',unknown)).
% SZS output end Proof for PLA008-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 6396
% Time elapsed: 0.103 s
% ------------------------------
% ------------------------------
