% Refutation found. Thanks to Tanya!
% SZS status Theorem for PUZ060+1
% SZS output start Proof for PUZ060+1
fof(f23,plain,(
  $false),
  inference(subsumption_resolution,[],[f22,f13])).
fof(f13,plain,(
  alive(sK2)),
  inference(cnf_transformation,[],[f7])).
fof(f7,plain,(
  ~likes(sK1,sK0) & ! [X6] : (! [X7] : not_killed_by(X6,X7) | ~alive(X6)) & ! [X8] : (eats(sK3,X8) | ~eats(sK2,X8)) & alive(sK2) & eats(sK2,sK0) & ! [X9] : (food(X9) | ! [X10] : (~not_killed_by(X10,X9) | ~eats(X10,X9))) & food(sK5) & food(sK4) & ! [X11] : (likes(sK1,X11) | ~food(X11))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f5,f6])).
fof(f6,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~likes(X1,X0) & ! [X6] : (! [X7] : not_killed_by(X6,X7) | ~alive(X6)) & ! [X8] : (eats(X3,X8) | ~eats(X2,X8)) & alive(X2) & eats(X2,X0) & ! [X9] : (food(X9) | ! [X10] : (~not_killed_by(X10,X9) | ~eats(X10,X9))) & food(X5) & food(X4) & ! [X11] : (likes(X1,X11) | ~food(X11))) => (~likes(sK1,sK0) & ! [X6] : (! [X7] : not_killed_by(X6,X7) | ~alive(X6)) & ! [X8] : (eats(sK3,X8) | ~eats(sK2,X8)) & alive(sK2) & eats(sK2,sK0) & ! [X9] : (food(X9) | ! [X10] : (~not_killed_by(X10,X9) | ~eats(X10,X9))) & food(sK5) & food(sK4) & ! [X11] : (likes(sK1,X11) | ~food(X11)))),
  introduced(choice_axiom,[])).
fof(f5,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~likes(X1,X0) & ! [X6] : (! [X7] : not_killed_by(X6,X7) | ~alive(X6)) & ! [X8] : (eats(X3,X8) | ~eats(X2,X8)) & alive(X2) & eats(X2,X0) & ! [X9] : (food(X9) | ! [X10] : (~not_killed_by(X10,X9) | ~eats(X10,X9))) & food(X5) & food(X4) & ! [X11] : (likes(X1,X11) | ~food(X11)))),
  inference(flattening,[],[f4])).
fof(f4,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (~likes(X1,X0) & (! [X6] : (! [X7] : not_killed_by(X6,X7) | ~alive(X6)) & ! [X8] : (eats(X3,X8) | ~eats(X2,X8)) & alive(X2) & eats(X2,X0) & ! [X9] : (food(X9) | ! [X10] : (~not_killed_by(X10,X9) | ~eats(X10,X9))) & food(X5) & food(X4) & ! [X11] : (likes(X1,X11) | ~food(X11))))),
  inference(ennf_transformation,[],[f3])).
fof(f3,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : ((! [X6] : (alive(X6) => ! [X7] : not_killed_by(X6,X7)) & ! [X8] : (eats(X2,X8) => eats(X3,X8)) & alive(X2) & eats(X2,X0) & ! [X9] : (? [X10] : (not_killed_by(X10,X9) & eats(X10,X9)) => food(X9)) & food(X5) & food(X4) & ! [X11] : (food(X11) => likes(X1,X11))) => likes(X1,X0))),
  inference(rectify,[],[f2])).
fof(f2,negated_conjecture,(
  ~! [X0,X1,X2,X3,X4,X5] : ((! [X7] : (alive(X7) => ! [X6] : not_killed_by(X7,X6)) & ! [X6] : (eats(X2,X6) => eats(X3,X6)) & alive(X2) & eats(X2,X0) & ! [X6] : (? [X7] : (not_killed_by(X7,X6) & eats(X7,X6)) => food(X6)) & food(X5) & food(X4) & ! [X6] : (food(X6) => likes(X1,X6))) => likes(X1,X0))),
  inference(negated_conjecture,[],[f1])).
fof(f1,conjecture,(
  ! [X0,X1,X2,X3,X4,X5] : ((! [X7] : (alive(X7) => ! [X6] : not_killed_by(X7,X6)) & ! [X6] : (eats(X2,X6) => eats(X3,X6)) & alive(X2) & eats(X2,X0) & ! [X6] : (? [X7] : (not_killed_by(X7,X6) & eats(X7,X6)) => food(X6)) & food(X5) & food(X4) & ! [X6] : (food(X6) => likes(X1,X6))) => likes(X1,X0))),
  file('./p/PUZ060+1.tptp',unknown)).
fof(f22,plain,(
  ~alive(sK2)),
  inference(subsumption_resolution,[],[f20,f17])).
fof(f17,plain,(
  ~food(sK0)),
  inference(resolution,[],[f8,f16])).
fof(f16,plain,(
  ~likes(sK1,sK0)),
  inference(cnf_transformation,[],[f7])).
fof(f8,plain,(
  ( ! [X11] : (likes(sK1,X11) | ~food(X11)) )),
  inference(cnf_transformation,[],[f7])).
fof(f20,plain,(
  food(sK0) | ~alive(sK2)),
  inference(resolution,[],[f19,f12])).
fof(f12,plain,(
  eats(sK2,sK0)),
  inference(cnf_transformation,[],[f7])).
fof(f19,plain,(
  ( ! [X0,X1] : (~eats(X1,X0) | food(X0) | ~alive(X1)) )),
  inference(resolution,[],[f11,f15])).
fof(f15,plain,(
  ( ! [X6,X7] : (not_killed_by(X6,X7) | ~alive(X6)) )),
  inference(cnf_transformation,[],[f7])).
fof(f11,plain,(
  ( ! [X10,X9] : (~not_killed_by(X10,X9) | food(X9) | ~eats(X10,X9)) )),
  inference(cnf_transformation,[],[f7])).
% SZS output end Proof for PUZ060+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
