% Refutation found. Thanks to Tanya!
% SZS status Theorem for REL014+1
% SZS output start Proof for REL014+1
fof(f445,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f439])).
fof(f439,plain,(
  sK0 != sK0),
  inference(superposition,[],[f35,f386])).
fof(f386,plain,(
  ( ! [X2] : (composition(one,X2) = X2) )),
  inference(backward_demodulation,[],[f374,f365])).
fof(f365,plain,(
  ( ! [X2] : (composition(converse(one),X2) = X2) )),
  inference(forward_demodulation,[],[f344,f26])).
fof(f26,plain,(
  ( ! [X0] : (converse(converse(X0)) = X0) )),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0] : converse(converse(X0)) = X0),
  file('./p/REL014+1.tptp',unknown)).
fof(f344,plain,(
  ( ! [X2] : (composition(converse(one),X2) = converse(converse(X2))) )),
  inference(superposition,[],[f44,f24])).
fof(f24,plain,(
  ( ! [X0] : (composition(X0,one) = X0) )),
  inference(cnf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0] : composition(X0,one) = X0),
  file('./p/REL014+1.tptp',unknown)).
fof(f44,plain,(
  ( ! [X0,X1] : (converse(composition(converse(X0),X1)) = composition(converse(X1),X0)) )),
  inference(superposition,[],[f28,f26])).
fof(f28,plain,(
  ( ! [X0,X1] : (converse(composition(X0,X1)) = composition(converse(X1),converse(X0))) )),
  inference(cnf_transformation,[],[f10])).
fof(f10,axiom,(
  ! [X0,X1] : converse(composition(X0,X1)) = composition(converse(X1),converse(X0))),
  file('./p/REL014+1.tptp',unknown)).
fof(f374,plain,(
  one = converse(one)),
  inference(superposition,[],[f365,f24])).
fof(f35,plain,(
  sK0 != composition(one,sK0)),
  inference(subsumption_resolution,[],[f32,f24])).
fof(f32,plain,(
  sK0 != composition(one,sK0) | sK0 != composition(sK0,one)),
  inference(cnf_transformation,[],[f18])).
fof(f18,plain,(
  sK0 != composition(one,sK0) | sK0 != composition(sK0,one)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f16,f17])).
fof(f17,plain,(
  ? [X0] : (composition(one,X0) != X0 | composition(X0,one) != X0) => (sK0 != composition(one,sK0) | sK0 != composition(sK0,one))),
  introduced(choice_axiom,[])).
fof(f16,plain,(
  ? [X0] : (composition(one,X0) != X0 | composition(X0,one) != X0)),
  inference(ennf_transformation,[],[f15])).
fof(f15,negated_conjecture,(
  ~! [X0] : (composition(one,X0) = X0 & composition(X0,one) = X0)),
  inference(negated_conjecture,[],[f14])).
fof(f14,conjecture,(
  ! [X0] : (composition(one,X0) = X0 & composition(X0,one) = X0)),
  file('./p/REL014+1.tptp',unknown)).
% SZS output end Proof for REL014+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5373
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
