% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SCT021-1
% SZS output start Proof for SCT021-1
fof(f1071134,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f1118,f94494,f94498,f94806,f440526,f623094,f655281,f678041,f693021,f707009,f764887,f765482,f765806,f765808,f772661,f776045,f776911,f778198,f812841,f832346,f869160,f869705,f874461,f889510,f891189,f901871,f974499,f991254,f1034210,f1035745,f1041223,f1055002,f1064870,f1071108])).
fof(f1071108,plain,(
  spl0_10 | ~spl0_614 | ~spl0_1377 | ~spl0_1448 | ~spl0_1825),
  inference(avatar_contradiction_clause,[],[f1071107])).
fof(f1071107,plain,(
  $false | (spl0_10 | ~spl0_614 | ~spl0_1377 | ~spl0_1448 | ~spl0_1825)),
  inference(subsumption_resolution,[],[f1071106,f1065157])).
fof(f1065157,plain,(
  ( ! [X13] : (v_x != c_Wellfounded_Oacc(v_x,X13)) ) | (spl0_10 | ~spl0_614 | ~spl0_1448 | ~spl0_1825)),
  inference(forward_demodulation,[],[f904089,f927481])).
fof(f927481,plain,(
  v_x = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) | ~spl0_1825),
  inference(avatar_component_clause,[],[f927479])).
fof(f927479,plain,(
  spl0_1825 <=> v_x = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1825])])).
fof(f904089,plain,(
  ( ! [X13] : (c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) != c_Wellfounded_Oacc(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X13)) ) | (spl0_10 | ~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f903538,f903334])).
fof(f903334,plain,(
  ( ! [X13] : (c_HOL_Ominus__class_Ominus(v_x,tc_bool) != c_Wellfounded_Oacc(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X13)) ) | (spl0_10 | ~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f902199,f902086])).
fof(f902086,plain,(
  ( ! [X12,X13] : (c_Orderings_Obot__class_Obot(tc_fun(X12,tc_bool)) != c_Wellfounded_Oacc(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X13)) ) | (~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f692919,f119548])).
fof(f119548,plain,(
  ( ! [X12,X13] : (c_Orderings_Obot__class_Obot(tc_fun(X12,tc_bool)) != c_Wellfounded_Oacc(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X13)) ) | ~spl0_614),
  inference(superposition,[],[f459,f96764])).
fof(f96764,plain,(
  ( ! [X6,X7] : (c_Orderings_Otop__class_Otop(tc_fun(X6,tc_bool)) = c_Wellfounded_Oacc(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X7)) ) | ~spl0_614),
  inference(resolution,[],[f96630,f3530])).
fof(f3530,plain,(
  ( ! [X17,X15,X16] : (~c_Wellfounded_Owf(X16,X17) | c_Orderings_Otop__class_Otop(tc_fun(X15,tc_bool)) = c_Wellfounded_Oacc(X16,X17)) )),
  inference(resolution,[],[f1667,f949])).
fof(f949,plain,(
  ( ! [X2,X0,X1] : (hBOOL(hAPP(c_Wellfounded_Oacc(X0,X1),X2)) | ~c_Wellfounded_Owf(X0,X1)) )),
  inference(resolution,[],[f506,f678])).
fof(f678,axiom,(
  ( ! [X0,X1,X13] : (~hBOOL(c_in(X1,X13,X0)) | hBOOL(hAPP(X13,X1))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f506,axiom,(
  ( ! [X2,X0,X1] : (hBOOL(c_in(X1,c_Wellfounded_Oacc(X2,X0),X0)) | ~c_Wellfounded_Owf(X2,X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f1667,plain,(
  ( ! [X2,X1] : (~hBOOL(hAPP(X2,c_ATP__Linkup_Osko__Set__XUNIV__eq__I__1__1(X2,X1))) | c_Orderings_Otop__class_Otop(tc_fun(X1,tc_bool)) = X2) )),
  inference(resolution,[],[f625,f677])).
fof(f677,axiom,(
  ( ! [X0,X1,X13] : (hBOOL(c_in(X1,X13,X0)) | ~hBOOL(hAPP(X13,X1))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f625,axiom,(
  ( ! [X4,X0] : (~hBOOL(c_in(c_ATP__Linkup_Osko__Set__XUNIV__eq__I__1__1(X4,X0),X4,X0)) | c_Orderings_Otop__class_Otop(tc_fun(X0,tc_bool)) = X4) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f96630,plain,(
  ( ! [X206] : (c_Wellfounded_Owf(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X206)) ) | ~spl0_614),
  inference(resolution,[],[f95216,f404])).
fof(f404,axiom,(
  ( ! [X2,X0] : (~hBOOL(c_in(c_ATP__Linkup_Osko__Wellfounded__Xwf__acc__iff__1__1(X2,X0),c_Wellfounded_Oacc(X2,X0),X0)) | c_Wellfounded_Owf(X2,X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f95216,plain,(
  ( ! [X17,X18] : (hBOOL(c_in(X17,c_Wellfounded_Oacc(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X18),X18))) ) | ~spl0_614),
  inference(resolution,[],[f95127,f426])).
fof(f426,axiom,(
  ( ! [X0,X12,X1] : (hBOOL(c_in(c_Pair(c_ATP__Linkup_Osko__Wellfounded__Xnot__acc__down__1__1(X12,X1,X0),X1,X0,X0),X12,tc_prod(X0,X0))) | hBOOL(c_in(X1,c_Wellfounded_Oacc(X12,X0),X0))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f95127,plain,(
  ( ! [X0,X1] : (~hBOOL(c_in(X0,c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X1))) ) | ~spl0_614),
  inference(resolution,[],[f95076,f5734])).
fof(f5734,plain,(
  ( ! [X2,X0,X1] : (~c_Relation_Oirrefl(c_Relation_OId__on(X1,X2),X2) | ~hBOOL(c_in(X0,X1,X2))) )),
  inference(resolution,[],[f653,f672])).
fof(f672,axiom,(
  ( ! [X2,X0,X1] : (~hBOOL(c_in(c_Pair(X1,X1,X0,X0),X2,tc_prod(X0,X0))) | ~c_Relation_Oirrefl(X2,X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f653,axiom,(
  ( ! [X4,X0,X1] : (hBOOL(c_in(c_Pair(X1,X1,X0,X0),c_Relation_OId__on(X4,X0),tc_prod(X0,X0))) | ~hBOOL(c_in(X1,X4,X0))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f95076,plain,(
  ( ! [X10,X11] : (c_Relation_Oirrefl(c_Relation_OId__on(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X10),X11)) ) | ~spl0_614),
  inference(resolution,[],[f94497,f76700])).
fof(f76700,plain,(
  ( ! [X14,X12,X13] : (c_Relation_Oirrefl(X12,X13) | hBOOL(hAPP(X12,c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__2(X12,X14)))) )),
  inference(resolution,[],[f76580,f515])).
fof(f515,axiom,(
  ( ! [X0,X1,X15] : (~hBOOL(hAPP(X15,X1)) | hBOOL(hAPP(X15,c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__2(X15,X0)))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f76580,plain,(
  ( ! [X17,X15,X16] : (hBOOL(hAPP(X15,c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__1(X15,X16))) | c_Relation_Oirrefl(X15,X17)) )),
  inference(resolution,[],[f40557,f5756])).
fof(f5756,plain,(
  ( ! [X10,X9] : (c_FunDef_Oin__rel(X9,c_ATP__Linkup_Osko__Relation__Xirrefl__def__1__1(X9,X10),c_ATP__Linkup_Osko__Relation__Xirrefl__def__1__1(X9,X10),X10,X10) | c_Relation_Oirrefl(X9,X10)) )),
  inference(resolution,[],[f666,f590])).
fof(f590,axiom,(
  ( ! [X6,X0,X12,X10,X1] : (~hBOOL(c_in(c_Pair(X1,X10,X0,X6),X12,tc_prod(X0,X6))) | c_FunDef_Oin__rel(X12,X1,X10,X0,X6)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f666,axiom,(
  ( ! [X2,X0] : (hBOOL(c_in(c_Pair(c_ATP__Linkup_Osko__Relation__Xirrefl__def__1__1(X2,X0),c_ATP__Linkup_Osko__Relation__Xirrefl__def__1__1(X2,X0),X0,X0),X2,tc_prod(X0,X0))) | c_Relation_Oirrefl(X2,X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f40557,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (~c_FunDef_Oin__rel(X33,X34,X35,X36,X37) | hBOOL(hAPP(X33,c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__1(X33,X38)))) )),
  inference(resolution,[],[f4992,f5582])).
fof(f5582,plain,(
  ( ! [X2,X0,X1] : (~hBOOL(hAPP(X0,X1)) | hBOOL(hAPP(X0,c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__1(X0,X2)))) )),
  inference(resolution,[],[f622,f675])).
fof(f675,axiom,(
  ( ! [X0,X1] : (hBOOL(c_in(X1,c_Orderings_Otop__class_Otop(tc_fun(X0,tc_bool)),X0))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f622,axiom,(
  ( ! [X0,X23,X15] : (~hBOOL(c_in(X23,c_Orderings_Otop__class_Otop(tc_fun(X0,tc_bool)),X0)) | ~hBOOL(hAPP(X15,X23)) | hBOOL(hAPP(X15,c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__1(X15,X0)))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f4992,plain,(
  ( ! [X47,X45,X48,X46,X49] : (hBOOL(hAPP(X45,c_Pair(X46,X47,X48,X49))) | ~c_FunDef_Oin__rel(X45,X46,X47,X48,X49)) )),
  inference(resolution,[],[f591,f678])).
fof(f591,axiom,(
  ( ! [X6,X0,X12,X10,X1] : (hBOOL(c_in(c_Pair(X1,X10,X0,X6),X12,tc_prod(X0,X6))) | ~c_FunDef_Oin__rel(X12,X1,X10,X0,X6)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f94497,plain,(
  ( ! [X0,X1] : (~hBOOL(hAPP(c_Relation_OId__on(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X1),X0))) ) | ~spl0_614),
  inference(avatar_component_clause,[],[f94496])).
fof(f94496,plain,(
  spl0_614 <=> ! [X1,X0] : ~hBOOL(hAPP(c_Relation_OId__on(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X1),X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_614])])).
fof(f459,axiom,(
  ( ! [X0] : (c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)) != c_Orderings_Otop__class_Otop(tc_fun(X0,tc_bool))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f692919,plain,(
  v_x = c_Orderings_Obot__class_Obot(tc_bool) | ~spl0_1448),
  inference(avatar_component_clause,[],[f692918])).
fof(f692918,plain,(
  spl0_1448 <=> v_x = c_Orderings_Obot__class_Obot(tc_bool)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1448])])).
fof(f902199,plain,(
  ( ! [X0] : (c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)) = c_HOL_Ominus__class_Ominus(v_x,tc_bool)) ) | (spl0_10 | ~spl0_1448)),
  inference(backward_demodulation,[],[f692919,f486755])).
fof(f486755,plain,(
  ( ! [X0] : (c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)) = c_HOL_Ominus__class_Ominus(c_Orderings_Obot__class_Obot(tc_bool),tc_bool)) ) | spl0_10),
  inference(resolution,[],[f486754,f217320])).
fof(f217320,plain,(
  ( ! [X387,X388] : (c_Wellfounded_Owf(c_Relation_OId__on(c_HOL_Ominus__class_Ominus(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X387),X388)) ) | spl0_10),
  inference(resolution,[],[f147775,f404])).
fof(f147775,plain,(
  ( ! [X10,X11,X9] : (hBOOL(c_in(X9,c_Wellfounded_Oacc(c_Relation_OId__on(c_HOL_Ominus__class_Ominus(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X10),X11),X11))) ) | spl0_10),
  inference(resolution,[],[f146863,f7535])).
fof(f7535,plain,(
  ( ! [X52,X53,X51] : (hBOOL(hAPP(X52,c_Pair(c_ATP__Linkup_Osko__Wellfounded__Xnot__acc__down__1__1(X52,X51,X53),X51,X53,X53))) | hBOOL(c_in(X51,c_Wellfounded_Oacc(X52,X53),X53))) )),
  inference(resolution,[],[f426,f678])).
fof(f146863,plain,(
  ( ! [X208,X207] : (~hBOOL(hAPP(c_Relation_OId__on(c_HOL_Ominus__class_Ominus(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X208),X207))) ) | spl0_10),
  inference(subsumption_resolution,[],[f146835,f1047])).
fof(f1047,plain,(
  ~hBOOL(c_Orderings_Obot__class_Obot(tc_bool)) | spl0_10),
  inference(avatar_component_clause,[],[f1045])).
fof(f1045,plain,(
  spl0_10 <=> hBOOL(c_Orderings_Obot__class_Obot(tc_bool))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f146835,plain,(
  ( ! [X208,X207] : (hBOOL(c_Orderings_Obot__class_Obot(tc_bool)) | ~hBOOL(hAPP(c_Relation_OId__on(c_HOL_Ominus__class_Ominus(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X208),X207))) )),
  inference(superposition,[],[f40426,f146462])).
fof(f146462,plain,(
  ( ! [X89] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_HOL_Ominus__class_Ominus(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X89)) )),
  inference(superposition,[],[f145546,f48])).
fof(f48,axiom,(
  ( ! [X6,X0,X15,X22] : (hAPP(c_COMBK(X15,X0,X6),X22) = X15) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f145546,plain,(
  ( ! [X11] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_HOL_Ominus__class_Ominus(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),hAPP(X11,v_x))) )),
  inference(forward_demodulation,[],[f145492,f1099])).
fof(f1099,plain,(
  hAPP(c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_bool)),v_x) = c_Orderings_Obot__class_Obot(tc_bool)),
  inference(resolution,[],[f540,f709])).
fof(f709,axiom,(
  class_Orderings_Obot(tc_bool)),
  file('./p/SCT021-1.tptp',unknown)).
fof(f540,axiom,(
  ( ! [X6] : (~class_Orderings_Obot(X6) | hAPP(c_Orderings_Obot__class_Obot(tc_fun(t_a,X6)),v_x) = c_Orderings_Obot__class_Obot(X6)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f145492,plain,(
  ( ! [X11] : (hAPP(c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_bool)),v_x) = hAPP(c_HOL_Ominus__class_Ominus(hAPP(c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_bool)),v_x),tc_bool),hAPP(X11,v_x))) )),
  inference(superposition,[],[f7150,f106])).
fof(f106,axiom,(
  ( ! [X4,X0] : (c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)) = hAPP(c_HOL_Ominus__class_Ominus(c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)),tc_fun(X0,tc_bool)),X4)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f7150,plain,(
  ( ! [X0,X1] : (hAPP(hAPP(c_HOL_Ominus__class_Ominus(X0,tc_fun(t_a,tc_bool)),X1),v_x) = hAPP(c_HOL_Ominus__class_Ominus(hAPP(X0,v_x),tc_bool),hAPP(X1,v_x))) )),
  inference(resolution,[],[f206,f710])).
fof(f710,axiom,(
  class_HOL_Ominus(tc_bool)),
  file('./p/SCT021-1.tptp',unknown)).
fof(f206,axiom,(
  ( ! [X6,X4,X7] : (~class_HOL_Ominus(X6) | hAPP(hAPP(c_HOL_Ominus__class_Ominus(X4,tc_fun(t_a,X6)),X7),v_x) = hAPP(c_HOL_Ominus__class_Ominus(hAPP(X4,v_x),X6),hAPP(X7,v_x))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f40426,plain,(
  ( ! [X50,X48,X49] : (hBOOL(hAPP(X48,c_ATP__Linkup_Osko__Relation__XId__onE__1__1(X48,X50,X49))) | ~hBOOL(hAPP(c_Relation_OId__on(X48,X49),X50))) )),
  inference(resolution,[],[f4729,f678])).
fof(f4729,plain,(
  ( ! [X2,X0,X1] : (hBOOL(c_in(c_ATP__Linkup_Osko__Relation__XId__onE__1__1(X0,X1,X2),X0,X2)) | ~hBOOL(hAPP(c_Relation_OId__on(X0,X2),X1))) )),
  inference(resolution,[],[f406,f677])).
fof(f406,axiom,(
  ( ! [X28,X4,X0] : (~hBOOL(c_in(X28,c_Relation_OId__on(X4,X0),tc_prod(X0,X0))) | hBOOL(c_in(c_ATP__Linkup_Osko__Relation__XId__onE__1__1(X4,X28,X0),X4,X0))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f486754,plain,(
  ( ! [X0,X1] : (~c_Wellfounded_Owf(c_Relation_OId__on(X0,X1),X1) | c_Orderings_Obot__class_Obot(tc_fun(X1,tc_bool)) = X0) )),
  inference(subsumption_resolution,[],[f486748,f287])).
fof(f287,axiom,(
  ( ! [X0,X1] : (c_lessequals(X1,X1,tc_fun(X0,tc_bool))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f486748,plain,(
  ( ! [X0,X1] : (~c_lessequals(X0,X0,tc_fun(X1,tc_bool)) | ~c_Wellfounded_Owf(c_Relation_OId__on(X0,X1),X1) | c_Orderings_Obot__class_Obot(tc_fun(X1,tc_bool)) = X0) )),
  inference(superposition,[],[f6387,f175])).
fof(f175,axiom,(
  ( ! [X4,X0] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(X4,tc_fun(X0,tc_bool)),X4) = X4) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f6387,plain,(
  ( ! [X6,X4,X5] : (~c_lessequals(X6,hAPP(c_Lattices_Olower__semilattice__class_Oinf(X4,tc_fun(X5,tc_bool)),X6),tc_fun(X5,tc_bool)) | ~c_Wellfounded_Owf(c_Relation_OId__on(X4,X5),X5) | c_Orderings_Obot__class_Obot(tc_fun(X5,tc_bool)) = X6) )),
  inference(superposition,[],[f403,f374])).
fof(f374,axiom,(
  ( ! [X4,X0,X7] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(X4,tc_fun(X0,tc_bool)),X7) = hAPP(c_Relation_OImage(c_Relation_OId__on(X4,X0),X0,X0),X7)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f403,axiom,(
  ( ! [X4,X0,X12] : (~c_lessequals(X4,hAPP(c_Relation_OImage(X12,X0,X0),X4),tc_fun(X0,tc_bool)) | ~c_Wellfounded_Owf(X12,X0) | c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)) = X4) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f903538,plain,(
  c_HOL_Ominus__class_Ominus(v_x,tc_bool) = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) | (spl0_10 | ~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f902200,f902199])).
fof(f902200,plain,(
  ( ! [X1] : (c_Orderings_Obot__class_Obot(tc_fun(X1,tc_bool)) = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool)) ) | (~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f692919,f486756])).
fof(f486756,plain,(
  ( ! [X1] : (c_Orderings_Obot__class_Obot(tc_fun(X1,tc_bool)) = c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool)) ) | ~spl0_614),
  inference(resolution,[],[f486754,f141493])).
fof(f141493,plain,(
  ( ! [X316,X317] : (c_Wellfounded_Owf(c_Relation_OId__on(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X316),X317)) ) | ~spl0_614),
  inference(resolution,[],[f101050,f404])).
fof(f101050,plain,(
  ( ! [X30,X31,X29] : (hBOOL(c_in(X29,c_Wellfounded_Oacc(c_Relation_OId__on(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X30),X31),X31))) ) | ~spl0_614),
  inference(resolution,[],[f100982,f426])).
fof(f100982,plain,(
  ( ! [X2,X0,X1] : (~hBOOL(c_in(X0,c_Relation_OId__on(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X1),X2))) ) | ~spl0_614),
  inference(resolution,[],[f100775,f5734])).
fof(f100775,plain,(
  ( ! [X14,X15,X13] : (c_Relation_Oirrefl(c_Relation_OId__on(c_Relation_OId__on(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X13),X14),X15)) ) | ~spl0_614),
  inference(resolution,[],[f95075,f76700])).
fof(f95075,plain,(
  ( ! [X8,X7,X9] : (~hBOOL(hAPP(c_Relation_OId__on(c_Relation_OId__on(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X7),X8),X9))) ) | ~spl0_614),
  inference(resolution,[],[f94497,f40426])).
fof(f1071106,plain,(
  ( ! [X0] : (v_x = c_Wellfounded_Oacc(v_x,X0)) ) | (spl0_10 | ~spl0_614 | ~spl0_1377 | ~spl0_1448 | ~spl0_1825)),
  inference(forward_demodulation,[],[f1071030,f1055709])).
fof(f1055709,plain,(
  ( ! [X107] : (v_x = hAPP(v_x,X107)) ) | (~spl0_1448 | ~spl0_1825)),
  inference(backward_demodulation,[],[f927481,f1029798])).
fof(f1029798,plain,(
  ( ! [X107] : (v_x = hAPP(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X107)) ) | ~spl0_1448),
  inference(forward_demodulation,[],[f1029797,f902468])).
fof(f902468,plain,(
  ( ! [X66] : (v_x = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X66,tc_bool),v_x)) ) | ~spl0_1448),
  inference(backward_demodulation,[],[f692919,f685555])).
fof(f685555,plain,(
  ( ! [X66] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X66,tc_bool),c_Orderings_Obot__class_Obot(tc_bool))) )),
  inference(subsumption_resolution,[],[f685554,f709])).
fof(f685554,plain,(
  ( ! [X66] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X66,tc_bool),c_Orderings_Obot__class_Obot(tc_bool)) | ~class_Orderings_Obot(tc_bool)) )),
  inference(subsumption_resolution,[],[f685516,f707])).
fof(f707,axiom,(
  class_Orderings_Oorder(tc_bool)),
  file('./p/SCT021-1.tptp',unknown)).
fof(f685516,plain,(
  ( ! [X66] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X66,tc_bool),c_Orderings_Obot__class_Obot(tc_bool)) | ~class_Orderings_Oorder(tc_bool) | ~class_Orderings_Obot(tc_bool)) )),
  inference(resolution,[],[f1601,f1353])).
fof(f1353,plain,(
  ( ! [X33,X34] : (~c_lessequals(X33,c_Orderings_Obot__class_Obot(X34),X34) | c_Orderings_Obot__class_Obot(X34) = X33 | ~class_Orderings_Oorder(X34) | ~class_Orderings_Obot(X34)) )),
  inference(resolution,[],[f120,f13])).
fof(f13,axiom,(
  ( ! [X0,X1] : (c_lessequals(c_Orderings_Obot__class_Obot(X0),X1,X0) | ~class_Orderings_Obot(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f120,axiom,(
  ( ! [X0,X10,X1] : (~c_lessequals(X10,X1,X0) | ~c_lessequals(X1,X10,X0) | X1 = X10 | ~class_Orderings_Oorder(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f1601,plain,(
  ( ! [X2,X1] : (c_lessequals(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X2,tc_bool),X1),X1,tc_bool)) )),
  inference(superposition,[],[f1535,f1468])).
fof(f1468,plain,(
  ( ! [X0,X1] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(X0,tc_bool),X1) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X1,tc_bool),X0)) )),
  inference(resolution,[],[f179,f702])).
fof(f702,axiom,(
  class_Lattices_Olower__semilattice(tc_bool)),
  file('./p/SCT021-1.tptp',unknown)).
fof(f179,axiom,(
  ( ! [X0,X10,X1] : (~class_Lattices_Olower__semilattice(X0) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(X1,X0),X10) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X10,X0),X1)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f1535,plain,(
  ( ! [X6,X7] : (c_lessequals(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X6,tc_bool),X7),X6,tc_bool)) )),
  inference(subsumption_resolution,[],[f1529,f701])).
fof(f701,axiom,(
  class_Lattices_Oupper__semilattice(tc_bool)),
  file('./p/SCT021-1.tptp',unknown)).
fof(f1529,plain,(
  ( ! [X6,X7] : (c_lessequals(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X6,tc_bool),X7),X6,tc_bool) | ~class_Lattices_Oupper__semilattice(tc_bool)) )),
  inference(superposition,[],[f73,f1336])).
fof(f1336,plain,(
  ( ! [X0,X1] : (hAPP(c_Lattices_Oupper__semilattice__class_Osup(X0,tc_bool),hAPP(c_Lattices_Olower__semilattice__class_Oinf(X0,tc_bool),X1)) = X0) )),
  inference(resolution,[],[f114,f706])).
fof(f706,axiom,(
  class_Lattices_Olattice(tc_bool)),
  file('./p/SCT021-1.tptp',unknown)).
fof(f114,axiom,(
  ( ! [X0,X10,X1] : (~class_Lattices_Olattice(X0) | hAPP(c_Lattices_Oupper__semilattice__class_Osup(X1,X0),hAPP(c_Lattices_Olower__semilattice__class_Oinf(X1,X0),X10)) = X1) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f73,axiom,(
  ( ! [X0,X10,X1] : (c_lessequals(X10,hAPP(c_Lattices_Oupper__semilattice__class_Osup(X1,X0),X10),X0) | ~class_Lattices_Oupper__semilattice(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f1029797,plain,(
  ( ! [X107,X105] : (v_x = hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X105,tc_bool),v_x),tc_bool),X107)) ) | ~spl0_1448),
  inference(forward_demodulation,[],[f1029764,f902468])).
fof(f1029764,plain,(
  ( ! [X107,X105,X106] : (v_x = hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X105,tc_bool),hAPP(c_Lattices_Olower__semilattice__class_Oinf(X106,tc_bool),v_x)),tc_bool),X107)) ) | ~spl0_1448),
  inference(resolution,[],[f23606,f902655])).
fof(f902655,plain,(
  ( ! [X24,X23] : (~c_lessequals(X23,v_x,tc_bool) | v_x = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X23,tc_bool),X24)) ) | ~spl0_1448),
  inference(forward_demodulation,[],[f902024,f692919])).
fof(f902024,plain,(
  ( ! [X24,X23] : (~c_lessequals(X23,v_x,tc_bool) | c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X23,tc_bool),X24)) ) | ~spl0_1448),
  inference(backward_demodulation,[],[f692919,f2233])).
fof(f2233,plain,(
  ( ! [X24,X23] : (~c_lessequals(X23,c_Orderings_Obot__class_Obot(tc_bool),tc_bool) | c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X23,tc_bool),X24)) )),
  inference(subsumption_resolution,[],[f2213,f702])).
fof(f2213,plain,(
  ( ! [X24,X23] : (~c_lessequals(X23,c_Orderings_Obot__class_Obot(tc_bool),tc_bool) | ~class_Lattices_Olower__semilattice(tc_bool) | c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X23,tc_bool),X24)) )),
  inference(resolution,[],[f211,f1365])).
fof(f1365,plain,(
  ( ! [X37] : (~c_lessequals(X37,c_Orderings_Obot__class_Obot(tc_bool),tc_bool) | c_Orderings_Obot__class_Obot(tc_bool) = X37) )),
  inference(subsumption_resolution,[],[f1355,f707])).
fof(f1355,plain,(
  ( ! [X37] : (~c_lessequals(X37,c_Orderings_Obot__class_Obot(tc_bool),tc_bool) | c_Orderings_Obot__class_Obot(tc_bool) = X37 | ~class_Orderings_Oorder(tc_bool)) )),
  inference(resolution,[],[f120,f923])).
fof(f923,plain,(
  ( ! [X1] : (c_lessequals(c_Orderings_Obot__class_Obot(tc_bool),X1,tc_bool)) )),
  inference(subsumption_resolution,[],[f920,f706])).
fof(f920,plain,(
  ( ! [X1] : (c_lessequals(c_Orderings_Obot__class_Obot(tc_bool),X1,tc_bool) | ~class_Lattices_Olattice(tc_bool)) )),
  inference(superposition,[],[f76,f915])).
fof(f915,plain,(
  ( ! [X0] : (hAPP(c_Lattices_Oupper__semilattice__class_Osup(X0,tc_bool),c_Orderings_Obot__class_Obot(tc_bool)) = X0) )),
  inference(resolution,[],[f102,f704])).
fof(f704,axiom,(
  class_Lattices_Obounded__lattice(tc_bool)),
  file('./p/SCT021-1.tptp',unknown)).
fof(f102,axiom,(
  ( ! [X0,X1] : (~class_Lattices_Obounded__lattice(X0) | hAPP(c_Lattices_Oupper__semilattice__class_Osup(X1,X0),c_Orderings_Obot__class_Obot(X0)) = X1) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f76,axiom,(
  ( ! [X0,X10,X1] : (c_lessequals(X10,hAPP(c_Lattices_Oupper__semilattice__class_Osup(X1,X0),X10),X0) | ~class_Lattices_Olattice(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f211,axiom,(
  ( ! [X0,X5,X1,X9] : (c_lessequals(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X5,X0),X9),X1,X0) | ~c_lessequals(X5,X1,X0) | ~class_Lattices_Olower__semilattice(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f23606,plain,(
  ( ! [X14,X15,X16] : (c_lessequals(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X14,tc_bool),hAPP(c_Lattices_Olower__semilattice__class_Oinf(X15,tc_bool),X16)),X16,tc_bool)) )),
  inference(resolution,[],[f2086,f1601])).
fof(f2086,plain,(
  ( ! [X21,X19,X20] : (~c_lessequals(X21,hAPP(c_Lattices_Olower__semilattice__class_Oinf(X20,tc_bool),X19),tc_bool) | c_lessequals(X21,X19,tc_bool)) )),
  inference(subsumption_resolution,[],[f2073,f701])).
fof(f2073,plain,(
  ( ! [X21,X19,X20] : (c_lessequals(X21,X19,tc_bool) | ~c_lessequals(X21,hAPP(c_Lattices_Olower__semilattice__class_Oinf(X20,tc_bool),X19),tc_bool) | ~class_Lattices_Oupper__semilattice(tc_bool)) )),
  inference(superposition,[],[f178,f1602])).
fof(f1602,plain,(
  ( ! [X4,X3] : (hAPP(c_Lattices_Oupper__semilattice__class_Osup(X3,tc_bool),hAPP(c_Lattices_Olower__semilattice__class_Oinf(X4,tc_bool),X3)) = X3) )),
  inference(superposition,[],[f1336,f1468])).
fof(f178,axiom,(
  ( ! [X0,X5,X1,X9] : (c_lessequals(X1,hAPP(c_Lattices_Oupper__semilattice__class_Osup(X5,X0),X9),X0) | ~c_lessequals(X1,X9,X0) | ~class_Lattices_Oupper__semilattice(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f1071030,plain,(
  ( ! [X0,X3] : (c_Wellfounded_Oacc(v_x,X0) = hAPP(v_x,X3)) ) | (spl0_10 | ~spl0_614 | ~spl0_1377 | ~spl0_1448 | ~spl0_1825)),
  inference(superposition,[],[f48,f1069609])).
fof(f1069609,plain,(
  ( ! [X17,X18,X16] : (v_x = c_COMBK(c_Wellfounded_Oacc(v_x,X16),X17,X18)) ) | (spl0_10 | ~spl0_614 | ~spl0_1377 | ~spl0_1448 | ~spl0_1825)),
  inference(resolution,[],[f1069439,f1055551])).
fof(f1055551,plain,(
  ( ! [X2,X0,X3] : (hBOOL(X0) | v_x = c_COMBK(X0,X2,X3)) ) | (spl0_10 | ~spl0_614 | ~spl0_1448 | ~spl0_1825)),
  inference(backward_demodulation,[],[f927481,f904007])).
fof(f904007,plain,(
  ( ! [X2,X0,X3] : (hBOOL(X0) | c_COMBK(X0,X2,X3) = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool)) ) | (spl0_10 | ~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f903538,f903237])).
fof(f903237,plain,(
  ( ! [X2,X0,X3] : (c_COMBK(X0,X2,X3) = c_HOL_Ominus__class_Ominus(v_x,tc_bool) | hBOOL(X0)) ) | (spl0_10 | ~spl0_1448)),
  inference(backward_demodulation,[],[f902199,f545154])).
fof(f545154,plain,(
  ( ! [X2,X0,X3,X1] : (hBOOL(X0) | c_Orderings_Obot__class_Obot(tc_fun(X1,tc_bool)) = c_COMBK(X0,X2,X3)) )),
  inference(resolution,[],[f541983,f486754])).
fof(f541983,plain,(
  ( ! [X243,X241,X244,X242,X240] : (c_Wellfounded_Owf(c_Relation_OId__on(c_COMBK(X240,X241,X242),X243),X244) | hBOOL(X240)) )),
  inference(resolution,[],[f541677,f92710])).
fof(f92710,plain,(
  ( ! [X109,X107,X110,X108,X106] : (~hBOOL(hAPP(c_Relation_OId__on(c_COMBK(X106,X107,X108),X110),X109)) | hBOOL(X106)) )),
  inference(superposition,[],[f40426,f48])).
fof(f541677,plain,(
  ( ! [X290,X289,X291] : (hBOOL(hAPP(X289,c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__2(X289,X290))) | c_Wellfounded_Owf(X289,X291)) )),
  inference(resolution,[],[f119010,f404])).
fof(f119010,plain,(
  ( ! [X23,X21,X22,X20] : (hBOOL(c_in(X20,c_Wellfounded_Oacc(X21,X22),X22)) | hBOOL(hAPP(X21,c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__2(X21,X23)))) )),
  inference(resolution,[],[f7529,f40560])).
fof(f40560,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (~c_FunDef_Oin__rel(X51,X52,X53,X54,X55) | hBOOL(hAPP(X51,c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__2(X51,X56)))) )),
  inference(resolution,[],[f4992,f515])).
fof(f7529,plain,(
  ( ! [X30,X31,X29] : (c_FunDef_Oin__rel(X30,c_ATP__Linkup_Osko__Wellfounded__Xnot__acc__down__1__1(X30,X29,X31),X29,X31,X31) | hBOOL(c_in(X29,c_Wellfounded_Oacc(X30,X31),X31))) )),
  inference(resolution,[],[f426,f590])).
fof(f1069439,plain,(
  ( ! [X70] : (~hBOOL(c_Wellfounded_Oacc(v_x,X70))) ) | (~spl0_614 | ~spl0_1377 | ~spl0_1448 | ~spl0_1825)),
  inference(superposition,[],[f616482,f1065156])).
fof(f1065156,plain,(
  ( ! [X6,X7] : (c_Orderings_Otop__class_Otop(tc_fun(X6,tc_bool)) = c_Wellfounded_Oacc(v_x,X7)) ) | (~spl0_614 | ~spl0_1448 | ~spl0_1825)),
  inference(forward_demodulation,[],[f902080,f927481])).
fof(f902080,plain,(
  ( ! [X6,X7] : (c_Orderings_Otop__class_Otop(tc_fun(X6,tc_bool)) = c_Wellfounded_Oacc(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X7)) ) | (~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f692919,f96764])).
fof(f616482,plain,(
  ( ! [X664] : (~hBOOL(c_Orderings_Otop__class_Otop(tc_fun(X664,tc_bool)))) ) | ~spl0_1377),
  inference(avatar_component_clause,[],[f616481])).
fof(f616481,plain,(
  spl0_1377 <=> ! [X664] : ~hBOOL(c_Orderings_Otop__class_Otop(tc_fun(X664,tc_bool)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1377])])).
fof(f1064870,plain,(
  spl0_10 | spl0_61 | ~spl0_614 | ~spl0_1448 | ~spl0_1825),
  inference(avatar_contradiction_clause,[],[f1064869])).
fof(f1064869,plain,(
  $false | (spl0_10 | spl0_61 | ~spl0_614 | ~spl0_1448 | ~spl0_1825)),
  inference(subsumption_resolution,[],[f1064868,f1055070])).
fof(f1055070,plain,(
  v_x != v_F(v_P) | (~spl0_614 | ~spl0_1448 | ~spl0_1825)),
  inference(backward_demodulation,[],[f927481,f902208])).
fof(f902208,plain,(
  v_F(v_P) != c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) | (~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f692919,f488347])).
fof(f488347,plain,(
  v_F(v_P) != c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool) | ~spl0_614),
  inference(superposition,[],[f13758,f486756])).
fof(f13758,plain,(
  ( ! [X12] : (v_F(v_P) != c_Orderings_Obot__class_Obot(tc_fun(X12,tc_bool))) )),
  inference(superposition,[],[f9,f1493])).
fof(f1493,plain,(
  ( ! [X39,X40] : (v_F(v_P) = hAPP(c_Set_Oinsert(c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__2(v_F(v_P),X39),X40),v_F(v_P))) )),
  inference(resolution,[],[f1242,f1121])).
fof(f1121,plain,(
  ( ! [X0] : (hBOOL(hAPP(v_F(v_P),c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__2(v_F(v_P),X0)))) )),
  inference(resolution,[],[f1119,f515])).
fof(f1119,plain,(
  hBOOL(hAPP(v_F(v_P),c_Pair(v_a,v_b,tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt)))),
  inference(resolution,[],[f682,f678])).
fof(f682,axiom,(
  hBOOL(c_in(c_Pair(v_a,v_b,tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),v_F(v_P),tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt)))),
  file('./p/SCT021-1.tptp',unknown)).
fof(f1242,plain,(
  ( ! [X6,X7,X5] : (~hBOOL(hAPP(X7,X5)) | hAPP(c_Set_Oinsert(X5,X6),X7) = X7) )),
  inference(resolution,[],[f779,f677])).
fof(f779,plain,(
  ( ! [X4,X0,X5] : (~hBOOL(c_in(X5,X4,X0)) | hAPP(c_Set_Oinsert(X5,X0),X4) = X4) )),
  inference(backward_demodulation,[],[f186,f173])).
fof(f173,axiom,(
  ( ! [X4,X0,X5] : (~hBOOL(c_in(X5,X4,X0)) | hAPP(c_Set_Oinsert(X5,X0),hAPP(c_HOL_Ominus__class_Ominus(X4,tc_fun(X0,tc_bool)),hAPP(c_Set_Oinsert(X5,X0),c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool))))) = X4) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f186,axiom,(
  ( ! [X4,X0,X5] : (hAPP(c_Set_Oinsert(X5,X0),X4) = hAPP(c_Set_Oinsert(X5,X0),hAPP(c_HOL_Ominus__class_Ominus(X4,tc_fun(X0,tc_bool)),hAPP(c_Set_Oinsert(X5,X0),c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)))))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f9,axiom,(
  ( ! [X4,X0,X5] : (hAPP(c_Set_Oinsert(X5,X0),X4) != c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f1064868,plain,(
  v_x = v_F(v_P) | (spl0_10 | spl0_61 | ~spl0_614 | ~spl0_1448 | ~spl0_1825)),
  inference(forward_demodulation,[],[f1064774,f1055709])).
fof(f1064774,plain,(
  ( ! [X2] : (v_F(v_P) = hAPP(v_x,X2)) ) | (spl0_10 | spl0_61 | ~spl0_614 | ~spl0_1448 | ~spl0_1825)),
  inference(superposition,[],[f48,f1064765])).
fof(f1064765,plain,(
  ( ! [X2081,X2080] : (v_x = c_COMBK(v_F(v_P),X2080,X2081)) ) | (spl0_10 | spl0_61 | ~spl0_614 | ~spl0_1448 | ~spl0_1825)),
  inference(resolution,[],[f1055551,f13804])).
fof(f13804,plain,(
  ~hBOOL(v_F(v_P)) | spl0_61),
  inference(avatar_component_clause,[],[f13802])).
fof(f13802,plain,(
  spl0_61 <=> hBOOL(v_F(v_P))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_61])])).
fof(f1055002,plain,(
  spl0_1825 | spl0_10 | ~spl0_614 | ~spl0_1389 | ~spl0_1448),
  inference(avatar_split_clause,[],[f1054583,f692918,f624344,f94496,f1045,f927479])).
fof(f624344,plain,(
  spl0_1389 <=> ! [X8] : ~hBOOL(hAPP(v_x,X8))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1389])])).
fof(f1054583,plain,(
  v_x = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) | (spl0_10 | ~spl0_614 | ~spl0_1389 | ~spl0_1448)),
  inference(resolution,[],[f903926,f1041341])).
fof(f1041341,plain,(
  ( ! [X52,X53] : (c_Wellfounded_Owf(c_Relation_OId__on(v_x,X52),X53)) ) | ~spl0_1389),
  inference(resolution,[],[f624345,f541979])).
fof(f541979,plain,(
  ( ! [X227,X225,X226,X224] : (c_Wellfounded_Owf(c_Relation_OId__on(X224,X225),X226) | hBOOL(hAPP(X224,c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__2(X224,X227)))) )),
  inference(resolution,[],[f541677,f92673])).
fof(f92673,plain,(
  ( ! [X17,X15,X18,X16] : (~hBOOL(hAPP(c_Relation_OId__on(X15,X16),X17)) | hBOOL(hAPP(X15,c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__2(X15,X18)))) )),
  inference(resolution,[],[f40426,f515])).
fof(f624345,plain,(
  ( ! [X8] : (~hBOOL(hAPP(v_x,X8))) ) | ~spl0_1389),
  inference(avatar_component_clause,[],[f624344])).
fof(f903926,plain,(
  ( ! [X0,X1] : (~c_Wellfounded_Owf(c_Relation_OId__on(X0,X1),X1) | c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) = X0) ) | (spl0_10 | ~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f903538,f903125])).
fof(f903125,plain,(
  ( ! [X0,X1] : (c_HOL_Ominus__class_Ominus(v_x,tc_bool) = X0 | ~c_Wellfounded_Owf(c_Relation_OId__on(X0,X1),X1)) ) | (spl0_10 | ~spl0_1448)),
  inference(backward_demodulation,[],[f902199,f486754])).
fof(f1041223,plain,(
  spl0_1389 | ~spl0_1421),
  inference(avatar_split_clause,[],[f1041115,f681332,f624344])).
fof(f681332,plain,(
  spl0_1421 <=> ! [X113,X114] : (~c_Equiv__Relations_Oequiv(v_x,X113,X114) | c_Relation_Oirrefl(X113,X114))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1421])])).
fof(f1041115,plain,(
  ( ! [X6] : (~hBOOL(hAPP(v_x,X6))) ) | ~spl0_1421),
  inference(resolution,[],[f1041014,f677])).
fof(f1041014,plain,(
  ( ! [X0,X1] : (~hBOOL(c_in(X0,v_x,X1))) ) | ~spl0_1421),
  inference(resolution,[],[f1041010,f5734])).
fof(f1041010,plain,(
  ( ! [X0] : (c_Relation_Oirrefl(c_Relation_OId__on(v_x,X0),X0)) ) | ~spl0_1421),
  inference(resolution,[],[f681333,f143776])).
fof(f143776,plain,(
  ( ! [X2,X3] : (c_Equiv__Relations_Oequiv(X2,c_Relation_OId__on(X2,X3),X3)) )),
  inference(trivial_inequality_removal,[],[f143767])).
fof(f143767,plain,(
  ( ! [X2,X3] : (c_Relation_OId__on(X2,X3) != c_Relation_OId__on(X2,X3) | c_Equiv__Relations_Oequiv(X2,c_Relation_OId__on(X2,X3),X3)) )),
  inference(backward_demodulation,[],[f143765,f7662])).
fof(f7662,plain,(
  ( ! [X2,X3] : (c_Equiv__Relations_Oequiv(X2,c_Relation_OId__on(X2,X3),X3) | c_Relation_OId__on(X2,X3) != c_Relation_Orel__comp(c_Relation_OId__on(X2,X3),c_Relation_OId__on(X2,X3),X3,X3,X3)) )),
  inference(forward_demodulation,[],[f7661,f836])).
fof(f836,plain,(
  ( ! [X4,X0] : (c_Relation_ORange(c_Relation_OId__on(X4,X0),X0,X0) = X4) )),
  inference(backward_demodulation,[],[f551,f761])).
fof(f761,plain,(
  ( ! [X4,X0] : (c_Relation_ORange(c_Relation_Oconverse(c_Relation_OId__on(X4,X0),X0,X0),X0,X0) = X4) )),
  inference(definition_unfolding,[],[f530,f554])).
fof(f554,axiom,(
  ( ! [X6,X2,X0] : (c_Relation_ODomain(X2,X0,X6) = c_Relation_ORange(c_Relation_Oconverse(X2,X0,X6),X6,X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f530,axiom,(
  ( ! [X4,X0] : (c_Relation_ODomain(c_Relation_OId__on(X4,X0),X0,X0) = X4) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f551,axiom,(
  ( ! [X4,X0] : (c_Relation_OId__on(X4,X0) = c_Relation_Oconverse(c_Relation_OId__on(X4,X0),X0,X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f7661,plain,(
  ( ! [X2,X3] : (c_Equiv__Relations_Oequiv(c_Relation_ORange(c_Relation_OId__on(X2,X3),X3,X3),c_Relation_OId__on(X2,X3),X3) | c_Relation_OId__on(X2,X3) != c_Relation_Orel__comp(c_Relation_OId__on(X2,X3),c_Relation_OId__on(X2,X3),X3,X3,X3)) )),
  inference(forward_demodulation,[],[f7654,f551])).
fof(f7654,plain,(
  ( ! [X2,X3] : (c_Relation_OId__on(X2,X3) != c_Relation_Orel__comp(c_Relation_OId__on(X2,X3),c_Relation_OId__on(X2,X3),X3,X3,X3) | c_Equiv__Relations_Oequiv(c_Relation_ORange(c_Relation_Oconverse(c_Relation_OId__on(X2,X3),X3,X3),X3,X3),c_Relation_OId__on(X2,X3),X3)) )),
  inference(superposition,[],[f766,f551])).
fof(f766,plain,(
  ( ! [X2,X0] : (c_Relation_Orel__comp(c_Relation_Oconverse(X2,X0,X0),X2,X0,X0,X0) != X2 | c_Equiv__Relations_Oequiv(c_Relation_ORange(c_Relation_Oconverse(X2,X0,X0),X0,X0),X2,X0)) )),
  inference(definition_unfolding,[],[f552,f554])).
fof(f552,axiom,(
  ( ! [X2,X0] : (c_Equiv__Relations_Oequiv(c_Relation_ODomain(X2,X0,X0),X2,X0) | c_Relation_Orel__comp(c_Relation_Oconverse(X2,X0,X0),X2,X0,X0,X0) != X2) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f143765,plain,(
  ( ! [X24,X25] : (c_Relation_OId__on(X24,X25) = c_Relation_Orel__comp(c_Relation_OId__on(X24,X25),c_Relation_OId__on(X24,X25),X25,X25,X25)) )),
  inference(forward_demodulation,[],[f143729,f103451])).
fof(f103451,plain,(
  ( ! [X80,X79] : (c_Relation_OId__on(X79,X80) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(c_Relation_OId__on(X79,X80),tc_fun(tc_prod(X80,X80),tc_bool)),c_Relation_Orel__comp(c_Relation_OId__on(X79,X80),c_Relation_OId__on(X79,X80),X80,X80,X80))) )),
  inference(superposition,[],[f834,f1678])).
fof(f1678,plain,(
  ( ! [X4,X3] : (c_Relation_OId__on(X3,X4) = c_Relation_Orel__comp(c_Relation_OId__on(X3,X4),c_Transitive__Closure_Ortrancl(c_Relation_OId__on(X3,X4),X4),X4,X4,X4)) )),
  inference(resolution,[],[f811,f468])).
fof(f468,axiom,(
  ( ! [X4,X0] : (c_Relation_Otrans(c_Relation_OId__on(X4,X0),X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f811,plain,(
  ( ! [X2,X0] : (~c_Relation_Otrans(X2,X0) | c_Relation_Orel__comp(X2,c_Transitive__Closure_Ortrancl(X2,X0),X0,X0,X0) = X2) )),
  inference(backward_demodulation,[],[f539,f747])).
fof(f747,plain,(
  ( ! [X2,X0] : (~c_Relation_Otrans(X2,X0) | c_Relation_Orel__comp(c_Transitive__Closure_Ortrancl(X2,X0),X2,X0,X0,X0) = X2) )),
  inference(definition_unfolding,[],[f469,f573])).
fof(f573,axiom,(
  ( ! [X2,X0] : (c_Transitive__Closure_Otrancl(X2,X0) = c_Relation_Orel__comp(c_Transitive__Closure_Ortrancl(X2,X0),X2,X0,X0,X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f469,axiom,(
  ( ! [X2,X0] : (~c_Relation_Otrans(X2,X0) | c_Transitive__Closure_Otrancl(X2,X0) = X2) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f539,axiom,(
  ( ! [X2,X0] : (c_Relation_Orel__comp(c_Transitive__Closure_Ortrancl(X2,X0),X2,X0,X0,X0) = c_Relation_Orel__comp(X2,c_Transitive__Closure_Ortrancl(X2,X0),X0,X0,X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f834,plain,(
  ( ! [X2,X0] : (c_Relation_Orel__comp(X2,c_Transitive__Closure_Ortrancl(X2,X0),X0,X0,X0) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X2,tc_fun(tc_prod(X0,X0),tc_bool)),c_Relation_Orel__comp(X2,c_Relation_Orel__comp(X2,c_Transitive__Closure_Ortrancl(X2,X0),X0,X0,X0),X0,X0,X0))) )),
  inference(forward_demodulation,[],[f833,f539])).
fof(f833,plain,(
  ( ! [X2,X0] : (c_Relation_Orel__comp(X2,c_Transitive__Closure_Ortrancl(X2,X0),X0,X0,X0) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X2,tc_fun(tc_prod(X0,X0),tc_bool)),c_Relation_Orel__comp(X2,c_Relation_Orel__comp(c_Transitive__Closure_Ortrancl(X2,X0),X2,X0,X0,X0),X0,X0,X0))) )),
  inference(backward_demodulation,[],[f547,f794])).
fof(f794,plain,(
  ( ! [X2,X0] : (c_Relation_Orel__comp(X2,c_Transitive__Closure_Ortrancl(X2,X0),X0,X0,X0) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X2,tc_fun(tc_prod(X0,X0),tc_bool)),c_Relation_Orel__comp(c_Relation_Orel__comp(X2,c_Transitive__Closure_Ortrancl(X2,X0),X0,X0,X0),X2,X0,X0,X0))) )),
  inference(backward_demodulation,[],[f539,f723])).
fof(f723,plain,(
  ( ! [X2,X0] : (c_Relation_Orel__comp(c_Transitive__Closure_Ortrancl(X2,X0),X2,X0,X0,X0) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X2,tc_fun(tc_prod(X0,X0),tc_bool)),c_Relation_Orel__comp(c_Relation_Orel__comp(c_Transitive__Closure_Ortrancl(X2,X0),X2,X0,X0,X0),X2,X0,X0,X0))) )),
  inference(definition_unfolding,[],[f366,f573,f573])).
fof(f366,axiom,(
  ( ! [X2,X0] : (c_Transitive__Closure_Otrancl(X2,X0) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X2,tc_fun(tc_prod(X0,X0),tc_bool)),c_Relation_Orel__comp(c_Transitive__Closure_Otrancl(X2,X0),X2,X0,X0,X0))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f547,axiom,(
  ( ! [X6,X0,X14,X12,X32,X46,X13] : (c_Relation_Orel__comp(c_Relation_Orel__comp(X12,X13,X0,X46,X14),X32,X0,X14,X6) = c_Relation_Orel__comp(X12,c_Relation_Orel__comp(X13,X32,X46,X14,X6),X0,X46,X6)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f143729,plain,(
  ( ! [X24,X25] : (c_Relation_Orel__comp(c_Relation_OId__on(X24,X25),c_Relation_OId__on(X24,X25),X25,X25,X25) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(c_Relation_OId__on(X24,X25),tc_fun(tc_prod(X25,X25),tc_bool)),c_Relation_Orel__comp(c_Relation_OId__on(X24,X25),c_Relation_OId__on(X24,X25),X25,X25,X25))) )),
  inference(resolution,[],[f6408,f128])).
fof(f128,axiom,(
  ( ! [X4,X0,X7] : (~c_lessequals(X4,X7,tc_fun(X0,tc_bool)) | hAPP(c_Lattices_Oupper__semilattice__class_Osup(X4,tc_fun(X0,tc_bool)),X7) = X7) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f6408,plain,(
  ( ! [X0,X1] : (c_lessequals(c_Relation_OId__on(X0,X1),c_Relation_Orel__comp(c_Relation_OId__on(X0,X1),c_Relation_OId__on(X0,X1),X1,X1,X1),tc_fun(tc_prod(X1,X1),tc_bool))) )),
  inference(forward_demodulation,[],[f6392,f551])).
fof(f6392,plain,(
  ( ! [X0,X1] : (c_lessequals(c_Relation_OId__on(X0,X1),c_Relation_Orel__comp(c_Relation_Oconverse(c_Relation_OId__on(X0,X1),X1,X1),c_Relation_OId__on(X0,X1),X1,X1,X1),tc_fun(tc_prod(X1,X1),tc_bool))) )),
  inference(resolution,[],[f414,f536])).
fof(f536,axiom,(
  ( ! [X4,X0] : (c_Relation_Orefl__on(X4,c_Relation_OId__on(X4,X0),X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f414,axiom,(
  ( ! [X4,X2,X0] : (~c_Relation_Orefl__on(X4,X2,X0) | c_lessequals(X2,c_Relation_Orel__comp(c_Relation_Oconverse(X2,X0,X0),X2,X0,X0,X0),tc_fun(tc_prod(X0,X0),tc_bool))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f681333,plain,(
  ( ! [X114,X113] : (~c_Equiv__Relations_Oequiv(v_x,X113,X114) | c_Relation_Oirrefl(X113,X114)) ) | ~spl0_1421),
  inference(avatar_component_clause,[],[f681332])).
fof(f1035745,plain,(
  ~spl0_12 | ~spl0_122 | ~spl0_1376),
  inference(avatar_contradiction_clause,[],[f1035740])).
fof(f1035740,plain,(
  $false | (~spl0_12 | ~spl0_122 | ~spl0_1376)),
  inference(resolution,[],[f1034494,f31933])).
fof(f31933,plain,(
  c_Relation_Oirrefl(c_Orderings_Otop__class_Otop(tc_bool),tc_Arrow__Order__Mirabelle_Oalt) | ~spl0_122),
  inference(avatar_component_clause,[],[f31932])).
fof(f31932,plain,(
  spl0_122 <=> c_Relation_Oirrefl(c_Orderings_Otop__class_Otop(tc_bool),tc_Arrow__Order__Mirabelle_Oalt)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_122])])).
fof(f1034494,plain,(
  ( ! [X30] : (~c_Relation_Oirrefl(c_Orderings_Otop__class_Otop(tc_bool),X30)) ) | (~spl0_12 | ~spl0_1376)),
  inference(backward_demodulation,[],[f1034291,f3237])).
fof(f3237,plain,(
  ( ! [X30,X29] : (~c_Relation_Oirrefl(c_Orderings_Otop__class_Otop(tc_fun(X29,tc_bool)),X30)) )),
  inference(resolution,[],[f2966,f674])).
fof(f674,axiom,(
  ( ! [X0,X1] : (hBOOL(hAPP(c_Orderings_Otop__class_Otop(tc_fun(X0,tc_bool)),X1))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f2966,plain,(
  ( ! [X10,X11,X9] : (~hBOOL(hAPP(X9,c_Pair(X11,X11,X10,X10))) | ~c_Relation_Oirrefl(X9,X10)) )),
  inference(resolution,[],[f672,f677])).
fof(f1034291,plain,(
  ( ! [X35] : (c_Orderings_Otop__class_Otop(tc_bool) = c_Orderings_Otop__class_Otop(tc_fun(X35,tc_bool))) ) | (~spl0_12 | ~spl0_1376)),
  inference(backward_demodulation,[],[f1034208,f3544])).
fof(f3544,plain,(
  ( ! [X35] : (c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool) = c_Orderings_Otop__class_Otop(tc_fun(X35,tc_bool))) ) | ~spl0_12),
  inference(subsumption_resolution,[],[f3539,f1085])).
fof(f1085,plain,(
  hBOOL(c_Orderings_Otop__class_Otop(tc_bool)) | ~spl0_12),
  inference(avatar_component_clause,[],[f1084])).
fof(f1084,plain,(
  spl0_12 <=> hBOOL(c_Orderings_Otop__class_Otop(tc_bool))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_12])])).
fof(f3539,plain,(
  ( ! [X35] : (~hBOOL(c_Orderings_Otop__class_Otop(tc_bool)) | c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool) = c_Orderings_Otop__class_Otop(tc_fun(X35,tc_bool))) )),
  inference(superposition,[],[f1667,f1034])).
fof(f1034,plain,(
  ( ! [X0] : (c_Orderings_Otop__class_Otop(tc_bool) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool),X0)) )),
  inference(resolution,[],[f268,f704])).
fof(f268,axiom,(
  ( ! [X0,X1] : (~class_Lattices_Obounded__lattice(X0) | c_Orderings_Otop__class_Otop(X0) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(X0),X0),X1)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f1034208,plain,(
  c_Orderings_Otop__class_Otop(tc_bool) = c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool) | (~spl0_12 | ~spl0_1376)),
  inference(forward_demodulation,[],[f1034008,f3694])).
fof(f3694,plain,(
  ( ! [X10,X11] : (c_Orderings_Otop__class_Otop(tc_bool) = hAPP(c_Orderings_Otop__class_Otop(tc_fun(X10,tc_bool)),X11)) ) | ~spl0_12),
  inference(superposition,[],[f1034,f3544])).
fof(f1034008,plain,(
  ( ! [X2,X3] : (c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool) = hAPP(c_Orderings_Otop__class_Otop(tc_fun(X2,tc_bool)),X3)) ) | ~spl0_1376),
  inference(superposition,[],[f48,f1001200])).
fof(f1001200,plain,(
  ( ! [X1110,X1108,X1109] : (c_Orderings_Otop__class_Otop(tc_fun(X1108,tc_bool)) = c_COMBK(c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool),X1109,X1110)) ) | ~spl0_1376),
  inference(resolution,[],[f3543,f616475])).
fof(f616475,plain,(
  hBOOL(c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool)) | ~spl0_1376),
  inference(avatar_component_clause,[],[f616474])).
fof(f616474,plain,(
  spl0_1376 <=> hBOOL(c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1376])])).
fof(f3543,plain,(
  ( ! [X45,X43,X44,X42] : (~hBOOL(X42) | c_Orderings_Otop__class_Otop(tc_fun(X45,tc_bool)) = c_COMBK(X42,X43,X44)) )),
  inference(superposition,[],[f1667,f48])).
fof(f1034210,plain,(
  ~spl0_12 | ~spl0_614 | ~spl0_1376 | ~spl0_1448 | ~spl0_1486),
  inference(avatar_contradiction_clause,[],[f1034209])).
fof(f1034209,plain,(
  $false | (~spl0_12 | ~spl0_614 | ~spl0_1376 | ~spl0_1448 | ~spl0_1486)),
  inference(subsumption_resolution,[],[f1034208,f974713])).
fof(f974713,plain,(
  ( ! [X6] : (c_Orderings_Otop__class_Otop(tc_bool) != c_Lattices_Oupper__semilattice__class_Osup(X6,tc_bool)) ) | (~spl0_12 | ~spl0_614 | ~spl0_1448 | ~spl0_1486)),
  inference(backward_demodulation,[],[f974497,f903536])).
fof(f903536,plain,(
  ( ! [X6] : (c_Lattices_Oupper__semilattice__class_Osup(X6,tc_bool) != c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool)) ) | (~spl0_12 | ~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f902200,f9837])).
fof(f9837,plain,(
  ( ! [X6,X5] : (c_Orderings_Obot__class_Obot(tc_fun(X5,tc_bool)) != c_Lattices_Oupper__semilattice__class_Osup(X6,tc_bool)) ) | ~spl0_12),
  inference(superposition,[],[f9,f1516])).
fof(f1516,plain,(
  ( ! [X30,X31] : (c_Lattices_Oupper__semilattice__class_Osup(X30,tc_bool) = hAPP(c_Set_Oinsert(c_Orderings_Otop__class_Otop(tc_bool),X31),c_Lattices_Oupper__semilattice__class_Osup(X30,tc_bool))) ) | ~spl0_12),
  inference(subsumption_resolution,[],[f1507,f1085])).
fof(f1507,plain,(
  ( ! [X30,X31] : (~hBOOL(c_Orderings_Otop__class_Otop(tc_bool)) | c_Lattices_Oupper__semilattice__class_Osup(X30,tc_bool) = hAPP(c_Set_Oinsert(c_Orderings_Otop__class_Otop(tc_bool),X31),c_Lattices_Oupper__semilattice__class_Osup(X30,tc_bool))) )),
  inference(superposition,[],[f1242,f1049])).
fof(f1049,plain,(
  ( ! [X0] : (c_Orderings_Otop__class_Otop(tc_bool) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X0,tc_bool),c_Orderings_Otop__class_Otop(tc_bool))) )),
  inference(resolution,[],[f269,f704])).
fof(f269,axiom,(
  ( ! [X0,X1] : (~class_Lattices_Obounded__lattice(X0) | c_Orderings_Otop__class_Otop(X0) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X1,X0),c_Orderings_Otop__class_Otop(X0))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f974497,plain,(
  c_Orderings_Otop__class_Otop(tc_bool) = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) | (~spl0_12 | ~spl0_1486)),
  inference(forward_demodulation,[],[f974357,f3694])).
fof(f974357,plain,(
  ( ! [X2,X3] : (hAPP(c_Orderings_Otop__class_Otop(tc_fun(X2,tc_bool)),X3) = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool)) ) | ~spl0_1486),
  inference(superposition,[],[f48,f956691])).
fof(f956691,plain,(
  ( ! [X6,X8,X7] : (c_Orderings_Otop__class_Otop(tc_fun(X6,tc_bool)) = c_COMBK(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X7,X8)) ) | ~spl0_1486),
  inference(resolution,[],[f3543,f751279])).
fof(f751279,plain,(
  hBOOL(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool)) | ~spl0_1486),
  inference(avatar_component_clause,[],[f751278])).
fof(f751278,plain,(
  spl0_1486 <=> hBOOL(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1486])])).
fof(f991254,plain,(
  spl0_10 | ~spl0_12 | spl0_56 | spl0_61 | ~spl0_614 | spl0_1376 | spl0_1388 | ~spl0_1448 | ~spl0_1486),
  inference(avatar_contradiction_clause,[],[f991253])).
fof(f991253,plain,(
  $false | (spl0_10 | ~spl0_12 | spl0_56 | spl0_61 | ~spl0_614 | spl0_1376 | spl0_1388 | ~spl0_1448 | ~spl0_1486)),
  inference(subsumption_resolution,[],[f991252,f623093])).
fof(f623093,plain,(
  v_x != c_Orderings_Otop__class_Otop(tc_bool) | spl0_1388),
  inference(avatar_component_clause,[],[f623091])).
fof(f623091,plain,(
  spl0_1388 <=> v_x = c_Orderings_Otop__class_Otop(tc_bool)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1388])])).
fof(f991252,plain,(
  v_x = c_Orderings_Otop__class_Otop(tc_bool) | (spl0_10 | ~spl0_12 | spl0_56 | spl0_61 | ~spl0_614 | spl0_1376 | ~spl0_1448 | ~spl0_1486)),
  inference(backward_demodulation,[],[f991251,f983587])).
fof(f983587,plain,(
  ( ! [X4,X3] : (c_Orderings_Otop__class_Otop(tc_bool) = c_COMBK(c_in(c_Pair(v_b,v_a,tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),v_x,tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt)),X3,X4)) ) | (spl0_10 | ~spl0_12 | spl0_56 | spl0_61 | ~spl0_614 | ~spl0_1448 | ~spl0_1486)),
  inference(resolution,[],[f982917,f975102])).
fof(f975102,plain,(
  ( ! [X2,X0,X3] : (hBOOL(X0) | c_Orderings_Otop__class_Otop(tc_bool) = c_COMBK(X0,X2,X3)) ) | (spl0_10 | ~spl0_12 | ~spl0_614 | ~spl0_1448 | ~spl0_1486)),
  inference(backward_demodulation,[],[f974497,f904007])).
fof(f982917,plain,(
  ~hBOOL(c_in(c_Pair(v_b,v_a,tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),v_x,tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt))) | (spl0_10 | ~spl0_12 | spl0_56 | spl0_61 | ~spl0_614 | ~spl0_1448 | ~spl0_1486)),
  inference(backward_demodulation,[],[f982885,f9925])).
fof(f9925,plain,(
  ~hBOOL(c_in(c_Pair(v_b,v_a,tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),v_F(v_P),tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt))) | spl0_56),
  inference(avatar_component_clause,[],[f9923])).
fof(f9923,plain,(
  spl0_56 <=> hBOOL(c_in(c_Pair(v_b,v_a,tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),v_F(v_P),tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_56])])).
fof(f982885,plain,(
  v_x = v_F(v_P) | (spl0_10 | ~spl0_12 | spl0_61 | ~spl0_614 | ~spl0_1448 | ~spl0_1486)),
  inference(forward_demodulation,[],[f982814,f975145])).
fof(f975145,plain,(
  ( ! [X27] : (v_x = hAPP(c_Orderings_Otop__class_Otop(tc_bool),X27)) ) | (spl0_10 | ~spl0_12 | ~spl0_614 | ~spl0_1448 | ~spl0_1486)),
  inference(backward_demodulation,[],[f974497,f904083])).
fof(f904083,plain,(
  ( ! [X27] : (v_x = hAPP(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X27)) ) | (spl0_10 | ~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f903538,f903328])).
fof(f903328,plain,(
  ( ! [X27] : (v_x = hAPP(c_HOL_Ominus__class_Ominus(v_x,tc_bool),X27)) ) | (spl0_10 | ~spl0_1448)),
  inference(backward_demodulation,[],[f902199,f902030])).
fof(f902030,plain,(
  ( ! [X26,X27] : (v_x = hAPP(c_Orderings_Obot__class_Obot(tc_fun(X26,tc_bool)),X27)) ) | ~spl0_1448),
  inference(backward_demodulation,[],[f692919,f4622])).
fof(f4622,plain,(
  ( ! [X26,X27] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_Orderings_Obot__class_Obot(tc_fun(X26,tc_bool)),X27)) )),
  inference(resolution,[],[f4583,f1365])).
fof(f4583,plain,(
  ( ! [X4,X2,X3] : (c_lessequals(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X2,tc_bool)),X3),X4,tc_bool)) )),
  inference(subsumption_resolution,[],[f4522,f702])).
fof(f4522,plain,(
  ( ! [X4,X2,X3] : (c_lessequals(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X2,tc_bool)),X3),X4,tc_bool) | ~class_Lattices_Olower__semilattice(tc_bool)) )),
  inference(resolution,[],[f3812,f209])).
fof(f209,axiom,(
  ( ! [X0,X5,X1,X9] : (~c_lessequals(X1,hAPP(c_Lattices_Olower__semilattice__class_Oinf(X5,X0),X9),X0) | c_lessequals(X1,X5,X0) | ~class_Lattices_Olower__semilattice(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f3812,plain,(
  ( ! [X72,X71,X73] : (c_lessequals(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X71,tc_bool)),X72),hAPP(X73,X72),tc_bool)) )),
  inference(subsumption_resolution,[],[f3795,f711])).
fof(f711,axiom,(
  class_HOL_Oord(tc_bool)),
  file('./p/SCT021-1.tptp',unknown)).
fof(f3795,plain,(
  ( ! [X72,X71,X73] : (c_lessequals(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X71,tc_bool)),X72),hAPP(X73,X72),tc_bool) | ~class_HOL_Oord(tc_bool)) )),
  inference(resolution,[],[f239,f14])).
fof(f14,axiom,(
  ( ! [X4,X0] : (c_lessequals(c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)),X4,tc_fun(X0,tc_bool))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f239,axiom,(
  ( ! [X6,X0,X17,X1,X16] : (~c_lessequals(X16,X17,tc_fun(X0,X6)) | c_lessequals(hAPP(X16,X1),hAPP(X17,X1),X6) | ~class_HOL_Oord(X6)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f982814,plain,(
  ( ! [X2] : (v_F(v_P) = hAPP(c_Orderings_Otop__class_Otop(tc_bool),X2)) ) | (spl0_10 | ~spl0_12 | spl0_61 | ~spl0_614 | ~spl0_1448 | ~spl0_1486)),
  inference(superposition,[],[f48,f982762])).
fof(f982762,plain,(
  ( ! [X1340,X1339] : (c_Orderings_Otop__class_Otop(tc_bool) = c_COMBK(v_F(v_P),X1339,X1340)) ) | (spl0_10 | ~spl0_12 | spl0_61 | ~spl0_614 | ~spl0_1448 | ~spl0_1486)),
  inference(resolution,[],[f975102,f13804])).
fof(f991251,plain,(
  ( ! [X158,X156,X154,X157] : (v_x = c_COMBK(c_in(X154,v_x,X156),X157,X158)) ) | (spl0_10 | ~spl0_12 | ~spl0_614 | spl0_1376 | ~spl0_1448 | ~spl0_1486)),
  inference(forward_demodulation,[],[f990791,f990530])).
fof(f990530,plain,(
  ( ! [X35] : (v_x = c_Orderings_Otop__class_Otop(tc_fun(X35,tc_bool))) ) | (spl0_10 | ~spl0_12 | ~spl0_614 | spl0_1376 | ~spl0_1448 | ~spl0_1486)),
  inference(backward_demodulation,[],[f990528,f3544])).
fof(f990528,plain,(
  v_x = c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool) | (spl0_10 | ~spl0_12 | ~spl0_614 | spl0_1376 | ~spl0_1448 | ~spl0_1486)),
  inference(forward_demodulation,[],[f990465,f975145])).
fof(f990465,plain,(
  ( ! [X2] : (c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool) = hAPP(c_Orderings_Otop__class_Otop(tc_bool),X2)) ) | (spl0_10 | ~spl0_12 | ~spl0_614 | spl0_1376 | ~spl0_1448 | ~spl0_1486)),
  inference(superposition,[],[f48,f975290])).
fof(f975290,plain,(
  ( ! [X327,X328] : (c_Orderings_Otop__class_Otop(tc_bool) = c_COMBK(c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool),X327,X328)) ) | (spl0_10 | ~spl0_12 | ~spl0_614 | spl0_1376 | ~spl0_1448 | ~spl0_1486)),
  inference(backward_demodulation,[],[f974497,f953072])).
fof(f953072,plain,(
  ( ! [X327,X328] : (c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) = c_COMBK(c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool),X327,X328)) ) | (spl0_10 | ~spl0_614 | spl0_1376 | ~spl0_1448)),
  inference(resolution,[],[f904007,f616476])).
fof(f616476,plain,(
  ~hBOOL(c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool)) | spl0_1376),
  inference(avatar_component_clause,[],[f616474])).
fof(f990791,plain,(
  ( ! [X158,X156,X154,X157,X155] : (v_x = c_COMBK(c_in(X154,c_Orderings_Otop__class_Otop(tc_fun(X155,tc_bool)),X156),X157,X158)) ) | (spl0_10 | ~spl0_12 | ~spl0_614 | spl0_1376 | ~spl0_1448 | ~spl0_1486)),
  inference(backward_demodulation,[],[f990530,f4805])).
fof(f4805,plain,(
  ( ! [X158,X156,X154,X157,X155,X153] : (c_Orderings_Otop__class_Otop(tc_fun(X153,tc_bool)) = c_COMBK(c_in(X154,c_Orderings_Otop__class_Otop(tc_fun(X155,tc_bool)),X156),X157,X158)) )),
  inference(resolution,[],[f3543,f3589])).
fof(f3589,plain,(
  ( ! [X39,X41,X40] : (hBOOL(c_in(X41,c_Orderings_Otop__class_Otop(tc_fun(X40,tc_bool)),X39))) )),
  inference(superposition,[],[f675,f3531])).
fof(f3531,plain,(
  ( ! [X19,X18] : (c_Orderings_Otop__class_Otop(tc_fun(X19,tc_bool)) = c_Orderings_Otop__class_Otop(tc_fun(X18,tc_bool))) )),
  inference(resolution,[],[f1667,f674])).
fof(f974499,plain,(
  ~spl0_12 | ~spl0_61 | ~spl0_614 | ~spl0_1448 | ~spl0_1486),
  inference(avatar_contradiction_clause,[],[f974498])).
fof(f974498,plain,(
  $false | (~spl0_12 | ~spl0_61 | ~spl0_614 | ~spl0_1448 | ~spl0_1486)),
  inference(subsumption_resolution,[],[f974497,f957258])).
fof(f957258,plain,(
  c_Orderings_Otop__class_Otop(tc_bool) != c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) | (~spl0_12 | ~spl0_61 | ~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f957077,f902208])).
fof(f957077,plain,(
  v_F(v_P) = c_Orderings_Otop__class_Otop(tc_bool) | (~spl0_12 | ~spl0_61)),
  inference(forward_demodulation,[],[f956980,f3694])).
fof(f956980,plain,(
  ( ! [X2,X3] : (v_F(v_P) = hAPP(c_Orderings_Otop__class_Otop(tc_fun(X2,tc_bool)),X3)) ) | ~spl0_61),
  inference(superposition,[],[f48,f956846])).
fof(f956846,plain,(
  ( ! [X883,X882,X884] : (c_Orderings_Otop__class_Otop(tc_fun(X882,tc_bool)) = c_COMBK(v_F(v_P),X883,X884)) ) | ~spl0_61),
  inference(resolution,[],[f3543,f13803])).
fof(f13803,plain,(
  hBOOL(v_F(v_P)) | ~spl0_61),
  inference(avatar_component_clause,[],[f13802])).
fof(f901871,plain,(
  ~spl0_12 | ~spl0_795 | ~spl0_819 | ~spl0_1388 | spl0_1448 | ~spl0_1486),
  inference(avatar_contradiction_clause,[],[f901870])).
fof(f901870,plain,(
  $false | (~spl0_12 | ~spl0_795 | ~spl0_819 | ~spl0_1388 | spl0_1448 | ~spl0_1486)),
  inference(subsumption_resolution,[],[f901869,f692920])).
fof(f692920,plain,(
  v_x != c_Orderings_Obot__class_Obot(tc_bool) | spl0_1448),
  inference(avatar_component_clause,[],[f692918])).
fof(f901869,plain,(
  v_x = c_Orderings_Obot__class_Obot(tc_bool) | (~spl0_12 | ~spl0_795 | ~spl0_819 | ~spl0_1388 | ~spl0_1486)),
  inference(forward_demodulation,[],[f901863,f898846])).
fof(f898846,plain,(
  ( ! [X37] : (hAPP(v_x,X37) = X37) ) | (~spl0_12 | ~spl0_795 | ~spl0_1388 | ~spl0_1486)),
  inference(backward_demodulation,[],[f898708,f889346])).
fof(f889346,plain,(
  ( ! [X37] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X37) = X37) ) | ~spl0_1388),
  inference(forward_demodulation,[],[f889345,f832524])).
fof(f832524,plain,(
  ( ! [X32] : (v_x = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X32,tc_bool),v_x)) ) | ~spl0_1388),
  inference(forward_demodulation,[],[f807360,f807477])).
fof(f807477,plain,(
  ( ! [X76] : (v_x = hAPP(c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool),X76)) ) | ~spl0_1388),
  inference(forward_demodulation,[],[f807476,f623092])).
fof(f623092,plain,(
  v_x = c_Orderings_Otop__class_Otop(tc_bool) | ~spl0_1388),
  inference(avatar_component_clause,[],[f623091])).
fof(f807476,plain,(
  ( ! [X76] : (v_x = hAPP(c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool),X76)) ) | ~spl0_1388),
  inference(subsumption_resolution,[],[f807429,f708])).
fof(f708,axiom,(
  class_Orderings_Otop(tc_bool)),
  file('./p/SCT021-1.tptp',unknown)).
fof(f807429,plain,(
  ( ! [X76] : (v_x = hAPP(c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool),X76) | ~class_Orderings_Otop(tc_bool)) ) | ~spl0_1388),
  inference(resolution,[],[f777555,f263])).
fof(f263,axiom,(
  ( ! [X0,X1] : (c_lessequals(X1,c_Orderings_Otop__class_Otop(X0),X0) | ~class_Orderings_Otop(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f777555,plain,(
  ( ! [X24,X23] : (~c_lessequals(v_x,X23,tc_bool) | v_x = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X23,tc_bool),X24)) ) | ~spl0_1388),
  inference(forward_demodulation,[],[f776942,f623092])).
fof(f776942,plain,(
  ( ! [X24,X23] : (~c_lessequals(v_x,X23,tc_bool) | c_Orderings_Otop__class_Otop(tc_bool) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X23,tc_bool),X24)) ) | ~spl0_1388),
  inference(backward_demodulation,[],[f623092,f2458])).
fof(f2458,plain,(
  ( ! [X24,X23] : (~c_lessequals(c_Orderings_Otop__class_Otop(tc_bool),X23,tc_bool) | c_Orderings_Otop__class_Otop(tc_bool) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X23,tc_bool),X24)) )),
  inference(subsumption_resolution,[],[f2437,f701])).
fof(f2437,plain,(
  ( ! [X24,X23] : (~c_lessequals(c_Orderings_Otop__class_Otop(tc_bool),X23,tc_bool) | ~class_Lattices_Oupper__semilattice(tc_bool) | c_Orderings_Otop__class_Otop(tc_bool) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X23,tc_bool),X24)) )),
  inference(resolution,[],[f279,f1364])).
fof(f1364,plain,(
  ( ! [X11] : (~c_lessequals(c_Orderings_Otop__class_Otop(tc_bool),X11,tc_bool) | c_Orderings_Otop__class_Otop(tc_bool) = X11) )),
  inference(subsumption_resolution,[],[f1345,f707])).
fof(f1345,plain,(
  ( ! [X11] : (~c_lessequals(c_Orderings_Otop__class_Otop(tc_bool),X11,tc_bool) | c_Orderings_Otop__class_Otop(tc_bool) = X11 | ~class_Orderings_Oorder(tc_bool)) )),
  inference(resolution,[],[f120,f942])).
fof(f942,plain,(
  ( ! [X0] : (c_lessequals(X0,c_Orderings_Otop__class_Otop(tc_bool),tc_bool)) )),
  inference(subsumption_resolution,[],[f938,f702])).
fof(f938,plain,(
  ( ! [X0] : (c_lessequals(X0,c_Orderings_Otop__class_Otop(tc_bool),tc_bool) | ~class_Lattices_Olower__semilattice(tc_bool)) )),
  inference(superposition,[],[f154,f936])).
fof(f936,plain,(
  ( ! [X0] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Otop__class_Otop(tc_bool),tc_bool),X0) = X0) )),
  inference(resolution,[],[f264,f704])).
fof(f264,axiom,(
  ( ! [X0,X1] : (~class_Lattices_Obounded__lattice(X0) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Otop__class_Otop(X0),X0),X1) = X1) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f154,axiom,(
  ( ! [X0,X10,X1] : (c_lessequals(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X1,X0),X10),X1,X0) | ~class_Lattices_Olower__semilattice(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f279,axiom,(
  ( ! [X0,X5,X1,X9] : (c_lessequals(X1,hAPP(c_Lattices_Oupper__semilattice__class_Osup(X5,X0),X9),X0) | ~c_lessequals(X1,X5,X0) | ~class_Lattices_Oupper__semilattice(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f807360,plain,(
  ( ! [X33,X32] : (v_x = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X32,tc_bool),hAPP(c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool),X33))) ) | ~spl0_1388),
  inference(subsumption_resolution,[],[f807324,f701])).
fof(f807324,plain,(
  ( ! [X33,X32] : (v_x = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X32,tc_bool),hAPP(c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool),X33)) | ~class_Lattices_Oupper__semilattice(tc_bool)) ) | ~spl0_1388),
  inference(resolution,[],[f777554,f72])).
fof(f72,axiom,(
  ( ! [X0,X10,X1] : (c_lessequals(X1,hAPP(c_Lattices_Oupper__semilattice__class_Osup(X1,X0),X10),X0) | ~class_Lattices_Oupper__semilattice(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f777554,plain,(
  ( ! [X24,X23] : (~c_lessequals(v_x,X23,tc_bool) | v_x = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X24,tc_bool),X23)) ) | ~spl0_1388),
  inference(forward_demodulation,[],[f776938,f623092])).
fof(f776938,plain,(
  ( ! [X24,X23] : (~c_lessequals(v_x,X23,tc_bool) | c_Orderings_Otop__class_Otop(tc_bool) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X24,tc_bool),X23)) ) | ~spl0_1388),
  inference(backward_demodulation,[],[f623092,f2083])).
fof(f2083,plain,(
  ( ! [X24,X23] : (~c_lessequals(c_Orderings_Otop__class_Otop(tc_bool),X23,tc_bool) | c_Orderings_Otop__class_Otop(tc_bool) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X24,tc_bool),X23)) )),
  inference(subsumption_resolution,[],[f2063,f701])).
fof(f2063,plain,(
  ( ! [X24,X23] : (~c_lessequals(c_Orderings_Otop__class_Otop(tc_bool),X23,tc_bool) | ~class_Lattices_Oupper__semilattice(tc_bool) | c_Orderings_Otop__class_Otop(tc_bool) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(X24,tc_bool),X23)) )),
  inference(resolution,[],[f178,f1364])).
fof(f889345,plain,(
  ( ! [X37,X36] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(c_Lattices_Oupper__semilattice__class_Osup(X36,tc_bool),v_x),tc_bool),X37) = X37) ) | ~spl0_1388),
  inference(subsumption_resolution,[],[f889344,f702])).
fof(f889344,plain,(
  ( ! [X37,X36] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(c_Lattices_Oupper__semilattice__class_Osup(X36,tc_bool),v_x),tc_bool),X37) = X37 | ~class_Lattices_Olower__semilattice(tc_bool)) ) | ~spl0_1388),
  inference(subsumption_resolution,[],[f889307,f701])).
fof(f889307,plain,(
  ( ! [X37,X36] : (~class_Lattices_Oupper__semilattice(tc_bool) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(c_Lattices_Oupper__semilattice__class_Osup(X36,tc_bool),v_x),tc_bool),X37) = X37 | ~class_Lattices_Olower__semilattice(tc_bool)) ) | ~spl0_1388),
  inference(resolution,[],[f2060,f777530])).
fof(f777530,plain,(
  ( ! [X106] : (c_lessequals(X106,v_x,tc_bool)) ) | ~spl0_1388),
  inference(backward_demodulation,[],[f623092,f765821])).
fof(f765821,plain,(
  ( ! [X106] : (c_lessequals(X106,c_Orderings_Otop__class_Otop(tc_bool),tc_bool)) )),
  inference(forward_demodulation,[],[f752869,f752875])).
fof(f752875,plain,(
  ( ! [X111,X114,X112,X113] : (c_Orderings_Otop__class_Otop(tc_bool) = hAPP(c_Transitive__Closure_Ortrancl(c_COMBK(c_Orderings_Otop__class_Otop(tc_bool),X111,X112),X113),X114)) )),
  inference(subsumption_resolution,[],[f752874,f708])).
fof(f752874,plain,(
  ( ! [X111,X114,X112,X113] : (c_Orderings_Otop__class_Otop(tc_bool) = hAPP(c_Transitive__Closure_Ortrancl(c_COMBK(c_Orderings_Otop__class_Otop(tc_bool),X111,X112),X113),X114) | ~class_Orderings_Otop(tc_bool)) )),
  inference(subsumption_resolution,[],[f752853,f707])).
fof(f752853,plain,(
  ( ! [X111,X114,X112,X113] : (c_Orderings_Otop__class_Otop(tc_bool) = hAPP(c_Transitive__Closure_Ortrancl(c_COMBK(c_Orderings_Otop__class_Otop(tc_bool),X111,X112),X113),X114) | ~class_Orderings_Oorder(tc_bool) | ~class_Orderings_Otop(tc_bool)) )),
  inference(resolution,[],[f244541,f1343])).
fof(f1343,plain,(
  ( ! [X8,X7] : (~c_lessequals(c_Orderings_Otop__class_Otop(X7),X8,X7) | c_Orderings_Otop__class_Otop(X7) = X8 | ~class_Orderings_Oorder(X7) | ~class_Orderings_Otop(X7)) )),
  inference(resolution,[],[f120,f263])).
fof(f244541,plain,(
  ( ! [X681,X683,X679,X680,X682] : (c_lessequals(X679,hAPP(c_Transitive__Closure_Ortrancl(c_COMBK(X679,X680,X681),X683),X682),tc_bool)) )),
  inference(superposition,[],[f244117,f48])).
fof(f244117,plain,(
  ( ! [X111,X109,X110] : (c_lessequals(hAPP(X109,X110),hAPP(c_Transitive__Closure_Ortrancl(X109,X111),X110),tc_bool)) )),
  inference(subsumption_resolution,[],[f244028,f711])).
fof(f244028,plain,(
  ( ! [X111,X109,X110] : (c_lessequals(hAPP(X109,X110),hAPP(c_Transitive__Closure_Ortrancl(X109,X111),X110),tc_bool) | ~class_HOL_Oord(tc_bool)) )),
  inference(resolution,[],[f243913,f239])).
fof(f243913,plain,(
  ( ! [X0,X1] : (c_lessequals(X0,c_Transitive__Closure_Ortrancl(X0,X1),tc_fun(tc_prod(X1,X1),tc_bool))) )),
  inference(resolution,[],[f23734,f7722])).
fof(f7722,plain,(
  ( ! [X2,X3] : (c_lessequals(c_Relation_Orel__comp(X3,c_Transitive__Closure_Ortrancl(X3,X2),X2,X2,X2),c_Transitive__Closure_Ortrancl(X3,X2),tc_fun(tc_prod(X2,X2),tc_bool))) )),
  inference(superposition,[],[f68,f790])).
fof(f790,plain,(
  ( ! [X2,X0] : (c_Transitive__Closure_Ortrancl(X2,X0) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(c_Relation_OId(X0),tc_fun(tc_prod(X0,X0),tc_bool)),c_Relation_Orel__comp(X2,c_Transitive__Closure_Ortrancl(X2,X0),X0,X0,X0))) )),
  inference(backward_demodulation,[],[f539,f412])).
fof(f412,axiom,(
  ( ! [X2,X0] : (c_Transitive__Closure_Ortrancl(X2,X0) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(c_Relation_OId(X0),tc_fun(tc_prod(X0,X0),tc_bool)),c_Relation_Orel__comp(c_Transitive__Closure_Ortrancl(X2,X0),X2,X0,X0,X0))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f68,axiom,(
  ( ! [X4,X0,X7] : (c_lessequals(X7,hAPP(c_Lattices_Oupper__semilattice__class_Osup(X4,tc_fun(X0,tc_bool)),X7),tc_fun(X0,tc_bool))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f23734,plain,(
  ( ! [X80,X78,X79] : (~c_lessequals(c_Relation_Orel__comp(X78,c_Transitive__Closure_Ortrancl(X78,X79),X79,X79,X79),X80,tc_fun(tc_prod(X79,X79),tc_bool)) | c_lessequals(X78,X80,tc_fun(tc_prod(X79,X79),tc_bool))) )),
  inference(superposition,[],[f282,f834])).
fof(f282,axiom,(
  ( ! [X4,X0,X7,X11] : (~c_lessequals(hAPP(c_Lattices_Oupper__semilattice__class_Osup(X4,tc_fun(X0,tc_bool)),X7),X11,tc_fun(X0,tc_bool)) | c_lessequals(X4,X11,tc_fun(X0,tc_bool))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f752869,plain,(
  ( ! [X109,X107,X110,X108,X106] : (c_lessequals(X106,hAPP(c_Transitive__Closure_Ortrancl(c_COMBK(c_Orderings_Otop__class_Otop(tc_bool),X107,X108),X109),X110),tc_bool)) )),
  inference(subsumption_resolution,[],[f752868,f708])).
fof(f752868,plain,(
  ( ! [X109,X107,X110,X108,X106] : (c_lessequals(X106,hAPP(c_Transitive__Closure_Ortrancl(c_COMBK(c_Orderings_Otop__class_Otop(tc_bool),X107,X108),X109),X110),tc_bool) | ~class_Orderings_Otop(tc_bool)) )),
  inference(subsumption_resolution,[],[f752852,f707])).
fof(f752852,plain,(
  ( ! [X109,X107,X110,X108,X106] : (c_lessequals(X106,hAPP(c_Transitive__Closure_Ortrancl(c_COMBK(c_Orderings_Otop__class_Otop(tc_bool),X107,X108),X109),X110),tc_bool) | ~class_Orderings_Oorder(tc_bool) | ~class_Orderings_Otop(tc_bool)) )),
  inference(resolution,[],[f244541,f2489])).
fof(f2489,plain,(
  ( ! [X61,X59,X60] : (~c_lessequals(c_Orderings_Otop__class_Otop(X59),X60,X59) | c_lessequals(X61,X60,X59) | ~class_Orderings_Oorder(X59) | ~class_Orderings_Otop(X59)) )),
  inference(resolution,[],[f285,f263])).
fof(f285,axiom,(
  ( ! [X0,X10,X1,X18] : (~c_lessequals(X18,X10,X0) | ~c_lessequals(X10,X1,X0) | c_lessequals(X18,X1,X0) | ~class_Orderings_Oorder(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f2060,plain,(
  ( ! [X14,X12,X15,X13] : (~c_lessequals(X12,X13,X14) | ~class_Lattices_Oupper__semilattice(X14) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(c_Lattices_Oupper__semilattice__class_Osup(X15,X14),X13),X14),X12) = X12 | ~class_Lattices_Olower__semilattice(X14)) )),
  inference(resolution,[],[f178,f66])).
fof(f66,axiom,(
  ( ! [X0,X10,X1] : (~c_lessequals(X10,X1,X0) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(X1,X0),X10) = X10 | ~class_Lattices_Olower__semilattice(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f898708,plain,(
  v_x = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) | (~spl0_12 | ~spl0_795 | ~spl0_1388 | ~spl0_1486)),
  inference(forward_demodulation,[],[f898663,f891813])).
fof(f891813,plain,(
  ( ! [X60] : (v_x = hAPP(c_Orderings_Obot__class_Obot(tc_bool),X60)) ) | (~spl0_795 | ~spl0_1388)),
  inference(backward_demodulation,[],[f172779,f890254])).
fof(f890254,plain,(
  ( ! [X60] : (v_x = hAPP(c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool),X60)) ) | ~spl0_1388),
  inference(superposition,[],[f1367,f889346])).
fof(f1367,plain,(
  ( ! [X0,X1] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(X0,tc_bool),hAPP(c_Lattices_Oupper__semilattice__class_Osup(X0,tc_bool),X1)) = X0) )),
  inference(resolution,[],[f123,f706])).
fof(f123,axiom,(
  ( ! [X0,X10,X1] : (~class_Lattices_Olattice(X0) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(X1,X0),hAPP(c_Lattices_Oupper__semilattice__class_Osup(X1,X0),X10)) = X1) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f172779,plain,(
  c_Orderings_Obot__class_Obot(tc_bool) = c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool) | ~spl0_795),
  inference(avatar_component_clause,[],[f172778])).
fof(f172778,plain,(
  spl0_795 <=> c_Orderings_Obot__class_Obot(tc_bool) = c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_795])])).
fof(f898663,plain,(
  ( ! [X2] : (hAPP(c_Orderings_Obot__class_Obot(tc_bool),X2) = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool)) ) | (~spl0_12 | ~spl0_795 | ~spl0_1388 | ~spl0_1486)),
  inference(superposition,[],[f48,f896173])).
fof(f896173,plain,(
  ( ! [X0,X1] : (c_Orderings_Obot__class_Obot(tc_bool) = c_COMBK(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X0,X1)) ) | (~spl0_12 | ~spl0_795 | ~spl0_1388 | ~spl0_1486)),
  inference(resolution,[],[f891395,f751279])).
fof(f891395,plain,(
  ( ! [X43,X44,X42] : (~hBOOL(X42) | c_Orderings_Obot__class_Obot(tc_bool) = c_COMBK(X42,X43,X44)) ) | (~spl0_12 | ~spl0_795 | ~spl0_1388)),
  inference(backward_demodulation,[],[f172779,f777579])).
fof(f777579,plain,(
  ( ! [X43,X44,X42] : (~hBOOL(X42) | c_COMBK(X42,X43,X44) = c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool)) ) | (~spl0_12 | ~spl0_1388)),
  inference(backward_demodulation,[],[f776945,f3543])).
fof(f776945,plain,(
  ( ! [X35] : (c_Orderings_Otop__class_Otop(tc_fun(X35,tc_bool)) = c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool)) ) | (~spl0_12 | ~spl0_1388)),
  inference(backward_demodulation,[],[f623092,f3544])).
fof(f901863,plain,(
  c_Orderings_Obot__class_Obot(tc_bool) = hAPP(v_x,v_x) | (~spl0_12 | ~spl0_795 | ~spl0_819 | ~spl0_1388)),
  inference(superposition,[],[f777258,f901780])).
fof(f901780,plain,(
  v_x = c_HOL_Ominus__class_Ominus(v_x,tc_bool) | (~spl0_12 | ~spl0_795 | ~spl0_819 | ~spl0_1388)),
  inference(forward_demodulation,[],[f901735,f891813])).
fof(f901735,plain,(
  ( ! [X2] : (hAPP(c_Orderings_Obot__class_Obot(tc_bool),X2) = c_HOL_Ominus__class_Ominus(v_x,tc_bool)) ) | (~spl0_12 | ~spl0_795 | ~spl0_819 | ~spl0_1388)),
  inference(superposition,[],[f48,f896356])).
fof(f896356,plain,(
  ( ! [X899,X898] : (c_Orderings_Obot__class_Obot(tc_bool) = c_COMBK(c_HOL_Ominus__class_Ominus(v_x,tc_bool),X898,X899)) ) | (~spl0_12 | ~spl0_795 | ~spl0_819 | ~spl0_1388)),
  inference(resolution,[],[f891395,f215376])).
fof(f215376,plain,(
  hBOOL(c_HOL_Ominus__class_Ominus(v_x,tc_bool)) | ~spl0_819),
  inference(avatar_component_clause,[],[f215375])).
fof(f215375,plain,(
  spl0_819 <=> hBOOL(c_HOL_Ominus__class_Ominus(v_x,tc_bool))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_819])])).
fof(f777258,plain,(
  ( ! [X87] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_HOL_Ominus__class_Ominus(X87,tc_bool),v_x)) ) | ~spl0_1388),
  inference(backward_demodulation,[],[f623092,f145624])).
fof(f145624,plain,(
  ( ! [X87] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_HOL_Ominus__class_Ominus(X87,tc_bool),c_Orderings_Otop__class_Otop(tc_bool))) )),
  inference(superposition,[],[f145545,f48])).
fof(f145545,plain,(
  ( ! [X10] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_HOL_Ominus__class_Ominus(hAPP(X10,v_x),tc_bool),c_Orderings_Otop__class_Otop(tc_bool))) )),
  inference(forward_demodulation,[],[f145544,f1099])).
fof(f145544,plain,(
  ( ! [X10] : (hAPP(c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_bool)),v_x) = hAPP(c_HOL_Ominus__class_Ominus(hAPP(X10,v_x),tc_bool),c_Orderings_Otop__class_Otop(tc_bool))) )),
  inference(forward_demodulation,[],[f145491,f1112])).
fof(f1112,plain,(
  c_Orderings_Otop__class_Otop(tc_bool) = hAPP(c_Orderings_Otop__class_Otop(tc_fun(t_a,tc_bool)),v_x)),
  inference(resolution,[],[f673,f708])).
fof(f673,axiom,(
  ( ! [X6] : (~class_Orderings_Otop(X6) | hAPP(c_Orderings_Otop__class_Otop(tc_fun(t_a,X6)),v_x) = c_Orderings_Otop__class_Otop(X6)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f145491,plain,(
  ( ! [X10] : (hAPP(c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_bool)),v_x) = hAPP(c_HOL_Ominus__class_Ominus(hAPP(X10,v_x),tc_bool),hAPP(c_Orderings_Otop__class_Otop(tc_fun(t_a,tc_bool)),v_x))) )),
  inference(superposition,[],[f7150,f337])).
fof(f337,axiom,(
  ( ! [X4,X0] : (c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)) = hAPP(c_HOL_Ominus__class_Ominus(X4,tc_fun(X0,tc_bool)),c_Orderings_Otop__class_Otop(tc_fun(X0,tc_bool)))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f891189,plain,(
  spl0_1411 | ~spl0_12 | ~spl0_802 | ~spl0_1387 | ~spl0_1388 | ~spl0_1420 | ~spl0_1486),
  inference(avatar_split_clause,[],[f891109,f751278,f681314,f623091,f623088,f173159,f1084,f681030])).
fof(f681030,plain,(
  spl0_1411 <=> ! [X24] : v_x = X24),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1411])])).
fof(f173159,plain,(
  spl0_802 <=> v_x = c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_802])])).
fof(f623088,plain,(
  spl0_1387 <=> ! [X67] : c_Orderings_Otop__class_Otop(tc_bool) = hAPP(v_x,X67)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1387])])).
fof(f681314,plain,(
  spl0_1420 <=> ! [X91,X90] : v_x = hAPP(c_Set_Oinsert(X90,X91),v_x)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1420])])).
fof(f891109,plain,(
  ( ! [X0] : (v_x = X0) ) | (~spl0_12 | ~spl0_802 | ~spl0_1387 | ~spl0_1388 | ~spl0_1420 | ~spl0_1486)),
  inference(superposition,[],[f869163,f890854])).
fof(f890854,plain,(
  ( ! [X5] : (hAPP(v_x,X5) = X5) ) | (~spl0_12 | ~spl0_802 | ~spl0_1387 | ~spl0_1388 | ~spl0_1420 | ~spl0_1486)),
  inference(forward_demodulation,[],[f890853,f890687])).
fof(f890687,plain,(
  v_x = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) | (~spl0_12 | ~spl0_802 | ~spl0_1387 | ~spl0_1388 | ~spl0_1486)),
  inference(forward_demodulation,[],[f890640,f869163])).
fof(f890640,plain,(
  ( ! [X2] : (c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) = hAPP(v_x,X2)) ) | (~spl0_12 | ~spl0_802 | ~spl0_1388 | ~spl0_1486)),
  inference(superposition,[],[f48,f889605])).
fof(f889605,plain,(
  ( ! [X6,X5] : (v_x = c_COMBK(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X5,X6)) ) | (~spl0_12 | ~spl0_802 | ~spl0_1388 | ~spl0_1486)),
  inference(resolution,[],[f751279,f869229])).
fof(f869229,plain,(
  ( ! [X43,X44,X42] : (~hBOOL(X42) | v_x = c_COMBK(X42,X43,X44)) ) | (~spl0_12 | ~spl0_802 | ~spl0_1388)),
  inference(backward_demodulation,[],[f173160,f777579])).
fof(f173160,plain,(
  v_x = c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool) | ~spl0_802),
  inference(avatar_component_clause,[],[f173159])).
fof(f890853,plain,(
  ( ! [X5] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X5) = X5) ) | (~spl0_12 | ~spl0_802 | ~spl0_1387 | ~spl0_1388 | ~spl0_1420 | ~spl0_1486)),
  inference(forward_demodulation,[],[f890852,f869163])).
fof(f890852,plain,(
  ( ! [X5] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(v_x,X5),tc_bool),X5) = X5) ) | (~spl0_12 | ~spl0_802 | ~spl0_1387 | ~spl0_1388 | ~spl0_1420 | ~spl0_1486)),
  inference(forward_demodulation,[],[f890718,f681315])).
fof(f681315,plain,(
  ( ! [X90,X91] : (v_x = hAPP(c_Set_Oinsert(X90,X91),v_x)) ) | ~spl0_1420),
  inference(avatar_component_clause,[],[f681314])).
fof(f890718,plain,(
  ( ! [X4,X5,X3] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(hAPP(c_Set_Oinsert(X3,X4),v_x),X5),tc_bool),X5) = X5) ) | (~spl0_12 | ~spl0_802 | ~spl0_1387 | ~spl0_1388 | ~spl0_1486)),
  inference(backward_demodulation,[],[f890687,f777200])).
fof(f777200,plain,(
  ( ! [X4,X5,X3] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(hAPP(c_Set_Oinsert(X3,X4),c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool)),X5),tc_bool),X5) = X5) ) | ~spl0_1388),
  inference(backward_demodulation,[],[f623092,f90299])).
fof(f90299,plain,(
  ( ! [X4,X5,X3] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(hAPP(c_Set_Oinsert(X3,X4),c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Otop__class_Otop(tc_bool),tc_bool)),X5),tc_bool),X5) = X5) )),
  inference(subsumption_resolution,[],[f90284,f702])).
fof(f90284,plain,(
  ( ! [X4,X5,X3] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(hAPP(c_Set_Oinsert(X3,X4),c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Otop__class_Otop(tc_bool),tc_bool)),X5),tc_bool),X5) = X5 | ~class_Lattices_Olower__semilattice(tc_bool)) )),
  inference(resolution,[],[f7372,f66])).
fof(f7372,plain,(
  ( ! [X14,X15,X16] : (c_lessequals(X14,hAPP(hAPP(c_Set_Oinsert(X15,X16),c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Otop__class_Otop(tc_bool),tc_bool)),X14),tc_bool)) )),
  inference(superposition,[],[f3808,f936])).
fof(f3808,plain,(
  ( ! [X52,X50,X53,X51] : (c_lessequals(hAPP(X50,X51),hAPP(hAPP(c_Set_Oinsert(X52,X53),X50),X51),tc_bool)) )),
  inference(subsumption_resolution,[],[f3790,f711])).
fof(f3790,plain,(
  ( ! [X52,X50,X53,X51] : (c_lessequals(hAPP(X50,X51),hAPP(hAPP(c_Set_Oinsert(X52,X53),X50),X51),tc_bool) | ~class_HOL_Oord(tc_bool)) )),
  inference(resolution,[],[f239,f82])).
fof(f82,axiom,(
  ( ! [X0,X7,X5] : (c_lessequals(X7,hAPP(c_Set_Oinsert(X5,X0),X7),tc_fun(X0,tc_bool))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f869163,plain,(
  ( ! [X67] : (v_x = hAPP(v_x,X67)) ) | (~spl0_1387 | ~spl0_1388)),
  inference(forward_demodulation,[],[f623089,f623092])).
fof(f623089,plain,(
  ( ! [X67] : (c_Orderings_Otop__class_Otop(tc_bool) = hAPP(v_x,X67)) ) | ~spl0_1387),
  inference(avatar_component_clause,[],[f623088])).
fof(f889510,plain,(
  spl0_10 | ~spl0_1388 | spl0_1486 | ~spl0_1492),
  inference(avatar_contradiction_clause,[],[f889509])).
fof(f889509,plain,(
  $false | (spl0_10 | ~spl0_1388 | spl0_1486 | ~spl0_1492)),
  inference(subsumption_resolution,[],[f889508,f1047])).
fof(f889508,plain,(
  hBOOL(c_Orderings_Obot__class_Obot(tc_bool)) | (~spl0_1388 | spl0_1486 | ~spl0_1492)),
  inference(forward_demodulation,[],[f889461,f812168])).
fof(f812168,plain,(
  ( ! [X60] : (hAPP(c_Orderings_Obot__class_Obot(tc_bool),X60) = X60) ) | (~spl0_1388 | spl0_1486)),
  inference(backward_demodulation,[],[f812095,f777525])).
fof(f777525,plain,(
  ( ! [X60] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X60) = X60) ) | ~spl0_1388),
  inference(backward_demodulation,[],[f623092,f750863])).
fof(f750863,plain,(
  ( ! [X60] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Otop__class_Otop(tc_bool),tc_bool),X60) = X60) )),
  inference(subsumption_resolution,[],[f750827,f708])).
fof(f750827,plain,(
  ( ! [X60] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Otop__class_Otop(tc_bool),tc_bool),X60) = X60 | ~class_Orderings_Otop(tc_bool)) )),
  inference(resolution,[],[f187070,f263])).
fof(f187070,plain,(
  ( ! [X26,X27] : (~c_lessequals(X26,X27,tc_bool) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(X27,tc_bool),X26) = X26) )),
  inference(forward_demodulation,[],[f187069,f179959])).
fof(f179959,plain,(
  ( ! [X28,X26,X29,X27] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(X26,tc_bool),X27) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(hAPP(c_Set_Oinsert(X28,X29),c_Lattices_Olower__semilattice__class_Oinf(X27,tc_bool)),X27),tc_bool),hAPP(c_Lattices_Olower__semilattice__class_Oinf(X26,tc_bool),X27))) )),
  inference(subsumption_resolution,[],[f179785,f702])).
fof(f179785,plain,(
  ( ! [X28,X26,X29,X27] : (~class_Lattices_Olower__semilattice(tc_bool) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(X26,tc_bool),X27) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(hAPP(c_Set_Oinsert(X28,X29),c_Lattices_Olower__semilattice__class_Oinf(X27,tc_bool)),X27),tc_bool),hAPP(c_Lattices_Olower__semilattice__class_Oinf(X26,tc_bool),X27))) )),
  inference(resolution,[],[f2258,f7369])).
fof(f7369,plain,(
  ( ! [X6,X4,X5] : (c_lessequals(X4,hAPP(hAPP(c_Set_Oinsert(X5,X6),c_Lattices_Olower__semilattice__class_Oinf(X4,tc_bool)),X4),tc_bool)) )),
  inference(superposition,[],[f3808,f867])).
fof(f867,plain,(
  ( ! [X0] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(X0,tc_bool),X0) = X0) )),
  inference(resolution,[],[f174,f702])).
fof(f174,axiom,(
  ( ! [X0,X1] : (~class_Lattices_Olower__semilattice(X0) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(X1,X0),X1) = X1) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f2258,plain,(
  ( ! [X14,X12,X15,X13] : (~c_lessequals(X12,X13,X14) | ~class_Lattices_Olower__semilattice(X14) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(X15,X14),X12) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X13,X14),hAPP(c_Lattices_Olower__semilattice__class_Oinf(X15,X14),X12))) )),
  inference(duplicate_literal_removal,[],[f2237])).
fof(f2237,plain,(
  ( ! [X14,X12,X15,X13] : (~c_lessequals(X12,X13,X14) | ~class_Lattices_Olower__semilattice(X14) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(X15,X14),X12) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X13,X14),hAPP(c_Lattices_Olower__semilattice__class_Oinf(X15,X14),X12)) | ~class_Lattices_Olower__semilattice(X14)) )),
  inference(resolution,[],[f212,f66])).
fof(f212,axiom,(
  ( ! [X0,X5,X1,X9] : (c_lessequals(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X5,X0),X9),X1,X0) | ~c_lessequals(X9,X1,X0) | ~class_Lattices_Olower__semilattice(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f187069,plain,(
  ( ! [X28,X26,X29,X27] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(hAPP(c_Set_Oinsert(X28,X29),c_Lattices_Olower__semilattice__class_Oinf(X26,tc_bool)),X26),tc_bool),hAPP(c_Lattices_Olower__semilattice__class_Oinf(X27,tc_bool),X26)) = X26 | ~c_lessequals(X26,X27,tc_bool)) )),
  inference(forward_demodulation,[],[f187068,f6700])).
fof(f6700,plain,(
  ( ! [X2,X0,X1] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X0,tc_bool),X1),tc_bool),X2) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X0,tc_bool),hAPP(c_Lattices_Olower__semilattice__class_Oinf(X1,tc_bool),X2))) )),
  inference(resolution,[],[f96,f706])).
fof(f96,axiom,(
  ( ! [X0,X10,X1,X18] : (~class_Lattices_Olattice(X0) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X1,X0),X10),X0),X18) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X1,X0),hAPP(c_Lattices_Olower__semilattice__class_Oinf(X10,X0),X18))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f187068,plain,(
  ( ! [X28,X26,X29,X27] : (~c_lessequals(X26,X27,tc_bool) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(hAPP(c_Set_Oinsert(X28,X29),c_Lattices_Olower__semilattice__class_Oinf(X26,tc_bool)),X26),tc_bool),X27),tc_bool),X26) = X26) )),
  inference(subsumption_resolution,[],[f186888,f702])).
fof(f186888,plain,(
  ( ! [X28,X26,X29,X27] : (~c_lessequals(X26,X27,tc_bool) | ~class_Lattices_Olower__semilattice(tc_bool) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(hAPP(c_Set_Oinsert(X28,X29),c_Lattices_Olower__semilattice__class_Oinf(X26,tc_bool)),X26),tc_bool),X27),tc_bool),X26) = X26) )),
  inference(resolution,[],[f5124,f7369])).
fof(f5124,plain,(
  ( ! [X14,X12,X15,X13] : (~c_lessequals(X12,X15,X14) | ~c_lessequals(X12,X13,X14) | ~class_Lattices_Olower__semilattice(X14) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X15,X14),X13),X14),X12) = X12) )),
  inference(duplicate_literal_removal,[],[f5080])).
fof(f5080,plain,(
  ( ! [X14,X12,X15,X13] : (~c_lessequals(X12,X13,X14) | ~c_lessequals(X12,X15,X14) | ~class_Lattices_Olower__semilattice(X14) | hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X15,X14),X13),X14),X12) = X12 | ~class_Lattices_Olower__semilattice(X14)) )),
  inference(resolution,[],[f150,f66])).
fof(f150,axiom,(
  ( ! [X0,X10,X1,X18] : (c_lessequals(X1,hAPP(c_Lattices_Olower__semilattice__class_Oinf(X10,X0),X18),X0) | ~c_lessequals(X1,X18,X0) | ~c_lessequals(X1,X10,X0) | ~class_Lattices_Olower__semilattice(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f812095,plain,(
  c_Orderings_Obot__class_Obot(tc_bool) = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) | spl0_1486),
  inference(forward_demodulation,[],[f811955,f4622])).
fof(f811955,plain,(
  ( ! [X2,X3] : (hAPP(c_Orderings_Obot__class_Obot(tc_fun(X2,tc_bool)),X3) = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool)) ) | spl0_1486),
  inference(superposition,[],[f48,f759494])).
fof(f759494,plain,(
  ( ! [X21,X19,X20] : (c_Orderings_Obot__class_Obot(tc_fun(X19,tc_bool)) = c_COMBK(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X20,X21)) ) | spl0_1486),
  inference(resolution,[],[f545154,f751280])).
fof(f751280,plain,(
  ~hBOOL(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool)) | spl0_1486),
  inference(avatar_component_clause,[],[f751278])).
fof(f889461,plain,(
  hBOOL(hAPP(c_Orderings_Obot__class_Obot(tc_bool),c_Orderings_Obot__class_Obot(tc_bool))) | ~spl0_1492),
  inference(superposition,[],[f165,f765805])).
fof(f765805,plain,(
  ( ! [X22] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_Set_Oinsert(c_Orderings_Obot__class_Obot(tc_bool),X22),c_Orderings_Obot__class_Obot(tc_bool))) ) | ~spl0_1492),
  inference(avatar_component_clause,[],[f765804])).
fof(f765804,plain,(
  spl0_1492 <=> ! [X22] : c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_Set_Oinsert(c_Orderings_Obot__class_Obot(tc_bool),X22),c_Orderings_Obot__class_Obot(tc_bool))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1492])])).
fof(f165,axiom,(
  ( ! [X4,X0,X1] : (hBOOL(hAPP(hAPP(c_Set_Oinsert(X1,X0),X4),X1))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f874461,plain,(
  ~spl0_12 | ~spl0_802 | ~spl0_819 | ~spl0_1387 | ~spl0_1388 | spl0_1448),
  inference(avatar_contradiction_clause,[],[f874460])).
fof(f874460,plain,(
  $false | (~spl0_12 | ~spl0_802 | ~spl0_819 | ~spl0_1387 | ~spl0_1388 | spl0_1448)),
  inference(subsumption_resolution,[],[f874459,f692920])).
fof(f874459,plain,(
  v_x = c_Orderings_Obot__class_Obot(tc_bool) | (~spl0_12 | ~spl0_802 | ~spl0_819 | ~spl0_1387 | ~spl0_1388)),
  inference(forward_demodulation,[],[f874453,f869163])).
fof(f874453,plain,(
  c_Orderings_Obot__class_Obot(tc_bool) = hAPP(v_x,v_x) | (~spl0_12 | ~spl0_802 | ~spl0_819 | ~spl0_1387 | ~spl0_1388)),
  inference(superposition,[],[f777258,f874328])).
fof(f874328,plain,(
  v_x = c_HOL_Ominus__class_Ominus(v_x,tc_bool) | (~spl0_12 | ~spl0_802 | ~spl0_819 | ~spl0_1387 | ~spl0_1388)),
  inference(forward_demodulation,[],[f874305,f869163])).
fof(f874305,plain,(
  ( ! [X2] : (c_HOL_Ominus__class_Ominus(v_x,tc_bool) = hAPP(v_x,X2)) ) | (~spl0_12 | ~spl0_802 | ~spl0_819 | ~spl0_1388)),
  inference(superposition,[],[f48,f869666])).
fof(f869666,plain,(
  ( ! [X438,X439] : (v_x = c_COMBK(c_HOL_Ominus__class_Ominus(v_x,tc_bool),X438,X439)) ) | (~spl0_12 | ~spl0_802 | ~spl0_819 | ~spl0_1388)),
  inference(backward_demodulation,[],[f173160,f859693])).
fof(f859693,plain,(
  ( ! [X438,X439] : (c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool) = c_COMBK(c_HOL_Ominus__class_Ominus(v_x,tc_bool),X438,X439)) ) | (~spl0_12 | ~spl0_819 | ~spl0_1388)),
  inference(resolution,[],[f777579,f215376])).
fof(f869705,plain,(
  spl0_1420 | ~spl0_12 | ~spl0_802 | ~spl0_1388),
  inference(avatar_split_clause,[],[f869204,f623091,f173159,f1084,f681314])).
fof(f869204,plain,(
  ( ! [X37,X38] : (v_x = hAPP(c_Set_Oinsert(X37,X38),v_x)) ) | (~spl0_12 | ~spl0_802 | ~spl0_1388)),
  inference(backward_demodulation,[],[f173160,f776930])).
fof(f776930,plain,(
  ( ! [X37,X38] : (c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool) = hAPP(c_Set_Oinsert(X37,X38),c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool))) ) | (~spl0_12 | ~spl0_1388)),
  inference(backward_demodulation,[],[f623092,f1517])).
fof(f1517,plain,(
  ( ! [X37,X38] : (c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool) = hAPP(c_Set_Oinsert(X37,X38),c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool))) ) | ~spl0_12),
  inference(subsumption_resolution,[],[f1510,f1085])).
fof(f1510,plain,(
  ( ! [X37,X38] : (~hBOOL(c_Orderings_Otop__class_Otop(tc_bool)) | c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool) = hAPP(c_Set_Oinsert(X37,X38),c_Lattices_Oupper__semilattice__class_Osup(c_Orderings_Otop__class_Otop(tc_bool),tc_bool))) )),
  inference(superposition,[],[f1242,f1034])).
fof(f869160,plain,(
  spl0_802 | ~spl0_12 | ~spl0_1388 | ~spl0_1487),
  inference(avatar_split_clause,[],[f869156,f765762,f623091,f1084,f173159])).
fof(f765762,plain,(
  spl0_1487 <=> hBOOL(c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1487])])).
fof(f869156,plain,(
  v_x = c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool) | (~spl0_12 | ~spl0_1388 | ~spl0_1487)),
  inference(forward_demodulation,[],[f869134,f807477])).
fof(f869134,plain,(
  ( ! [X2] : (c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool),X2)) ) | (~spl0_12 | ~spl0_1388 | ~spl0_1487)),
  inference(superposition,[],[f48,f859691])).
fof(f859691,plain,(
  ( ! [X430,X431] : (c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool) = c_COMBK(c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool),X430,X431)) ) | (~spl0_12 | ~spl0_1388 | ~spl0_1487)),
  inference(resolution,[],[f777579,f765763])).
fof(f765763,plain,(
  hBOOL(c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool)) | ~spl0_1487),
  inference(avatar_component_clause,[],[f765762])).
fof(f832346,plain,(
  ~spl0_1411),
  inference(avatar_contradiction_clause,[],[f832345])).
fof(f832345,plain,(
  $false | ~spl0_1411),
  inference(subsumption_resolution,[],[f832344,f823958])).
fof(f823958,plain,(
  ( ! [X72] : (c_Relation_Oirrefl(v_x,X72)) ) | ~spl0_1411),
  inference(forward_demodulation,[],[f814256,f681031])).
fof(f681031,plain,(
  ( ! [X24] : (v_x = X24) ) | ~spl0_1411),
  inference(avatar_component_clause,[],[f681030])).
fof(f814256,plain,(
  ( ! [X72,X73] : (c_Relation_Oirrefl(hAPP(v_x,hAPP(c_HOL_Ominus__class_Ominus(X73,tc_fun(tc_prod(X72,X72),tc_bool)),c_Relation_OId(X72))),X72)) ) | ~spl0_1411),
  inference(backward_demodulation,[],[f681031,f12477])).
fof(f12477,plain,(
  ( ! [X72,X71,X73] : (c_Relation_Oirrefl(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X71,tc_fun(tc_prod(X72,X72),tc_bool)),hAPP(c_HOL_Ominus__class_Ominus(X73,tc_fun(tc_prod(X72,X72),tc_bool)),c_Relation_OId(X72))),X72)) )),
  inference(superposition,[],[f277,f300])).
fof(f300,axiom,(
  ( ! [X4,X0,X7,X11] : (hAPP(c_HOL_Ominus__class_Ominus(hAPP(c_Lattices_Olower__semilattice__class_Oinf(X4,tc_fun(X0,tc_bool)),X7),tc_fun(X0,tc_bool)),X11) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(X4,tc_fun(X0,tc_bool)),hAPP(c_HOL_Ominus__class_Ominus(X7,tc_fun(X0,tc_bool)),X11))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f277,axiom,(
  ( ! [X2,X0] : (c_Relation_Oirrefl(hAPP(c_HOL_Ominus__class_Ominus(X2,tc_fun(tc_prod(X0,X0),tc_bool)),c_Relation_OId(X0)),X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f832344,plain,(
  ( ! [X8] : (~c_Relation_Oirrefl(v_x,X8)) ) | ~spl0_1411),
  inference(forward_demodulation,[],[f822396,f681031])).
fof(f822396,plain,(
  ( ! [X6,X8,X7] : (~c_Relation_Oirrefl(c_Relation_Oinv__image(v_x,X7,X6,X8),X8)) ) | ~spl0_1411),
  inference(backward_demodulation,[],[f681031,f120548])).
fof(f120548,plain,(
  ( ! [X6,X8,X7] : (~c_Relation_Oirrefl(c_Relation_Oinv__image(c_Relation_OId(X6),X7,X6,X8),X8)) )),
  inference(resolution,[],[f23038,f672])).
fof(f23038,plain,(
  ( ! [X47,X45,X48,X46] : (hBOOL(c_in(c_Pair(X45,X45,X46,X46),c_Relation_Oinv__image(c_Relation_OId(X47),X48,X47,X46),tc_prod(X46,X46)))) )),
  inference(resolution,[],[f612,f619])).
fof(f619,axiom,(
  ( ! [X0,X1] : (hBOOL(c_in(c_Pair(X1,X1,X0,X0),c_Relation_OId(X0),tc_prod(X0,X0)))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f612,axiom,(
  ( ! [X6,X2,X0,X10,X1,X16] : (~hBOOL(c_in(c_Pair(hAPP(X16,X1),hAPP(X16,X10),X6,X6),X2,tc_prod(X6,X6))) | hBOOL(c_in(c_Pair(X1,X10,X0,X0),c_Relation_Oinv__image(X2,X16,X6,X0),tc_prod(X0,X0)))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f812841,plain,(
  spl0_1411 | ~spl0_12 | ~spl0_795 | ~spl0_1388 | spl0_1486),
  inference(avatar_split_clause,[],[f812756,f751278,f623091,f172778,f1084,f681030])).
fof(f812756,plain,(
  ( ! [X0] : (v_x = X0) ) | (~spl0_12 | ~spl0_795 | ~spl0_1388 | spl0_1486)),
  inference(superposition,[],[f778604,f812254])).
fof(f812254,plain,(
  ( ! [X5] : (hAPP(c_Orderings_Obot__class_Obot(tc_bool),X5) = X5) ) | (~spl0_12 | ~spl0_795 | ~spl0_1388 | spl0_1486)),
  inference(forward_demodulation,[],[f812253,f812095])).
fof(f812253,plain,(
  ( ! [X5] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X5) = X5) ) | (~spl0_12 | ~spl0_795 | ~spl0_1388 | spl0_1486)),
  inference(forward_demodulation,[],[f812252,f778604])).
fof(f812252,plain,(
  ( ! [X5] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(c_Orderings_Obot__class_Obot(tc_bool),X5),tc_bool),X5) = X5) ) | (~spl0_12 | ~spl0_795 | ~spl0_1388 | spl0_1486)),
  inference(forward_demodulation,[],[f812123,f778218])).
fof(f778218,plain,(
  ( ! [X37,X38] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_Set_Oinsert(X37,X38),c_Orderings_Obot__class_Obot(tc_bool))) ) | (~spl0_12 | ~spl0_795 | ~spl0_1388)),
  inference(backward_demodulation,[],[f172779,f776930])).
fof(f812123,plain,(
  ( ! [X4,X5,X3] : (hAPP(c_Lattices_Olower__semilattice__class_Oinf(hAPP(hAPP(c_Set_Oinsert(X3,X4),c_Orderings_Obot__class_Obot(tc_bool)),X5),tc_bool),X5) = X5) ) | (~spl0_1388 | spl0_1486)),
  inference(backward_demodulation,[],[f812095,f777200])).
fof(f778604,plain,(
  ( ! [X11] : (v_x = hAPP(c_Orderings_Obot__class_Obot(tc_bool),X11)) ) | (~spl0_12 | ~spl0_795 | ~spl0_1388)),
  inference(backward_demodulation,[],[f172779,f778079])).
fof(f778079,plain,(
  ( ! [X11] : (v_x = hAPP(c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool),X11)) ) | (~spl0_12 | ~spl0_1388)),
  inference(forward_demodulation,[],[f776950,f776945])).
fof(f776950,plain,(
  ( ! [X10,X11] : (v_x = hAPP(c_Orderings_Otop__class_Otop(tc_fun(X10,tc_bool)),X11)) ) | (~spl0_12 | ~spl0_1388)),
  inference(backward_demodulation,[],[f623092,f3694])).
fof(f778198,plain,(
  ~spl0_1376 | ~spl0_1388 | spl0_1487),
  inference(avatar_contradiction_clause,[],[f778197])).
fof(f778197,plain,(
  $false | (~spl0_1376 | ~spl0_1388 | spl0_1487)),
  inference(subsumption_resolution,[],[f777504,f765764])).
fof(f765764,plain,(
  ~hBOOL(c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool)) | spl0_1487),
  inference(avatar_component_clause,[],[f765762])).
fof(f777504,plain,(
  hBOOL(c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool)) | (~spl0_1376 | ~spl0_1388)),
  inference(backward_demodulation,[],[f623092,f616475])).
fof(f776911,plain,(
  spl0_10 | ~spl0_614 | ~spl0_819 | ~spl0_1448 | spl0_1486),
  inference(avatar_contradiction_clause,[],[f776910])).
fof(f776910,plain,(
  $false | (spl0_10 | ~spl0_614 | ~spl0_819 | ~spl0_1448 | spl0_1486)),
  inference(subsumption_resolution,[],[f776909,f751280])).
fof(f776909,plain,(
  hBOOL(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool)) | (spl0_10 | ~spl0_614 | ~spl0_819 | ~spl0_1448)),
  inference(forward_demodulation,[],[f215376,f774230])).
fof(f774230,plain,(
  c_HOL_Ominus__class_Ominus(v_x,tc_bool) = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) | (spl0_10 | ~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f772883,f772882])).
fof(f772882,plain,(
  ( ! [X0] : (c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)) = c_HOL_Ominus__class_Ominus(v_x,tc_bool)) ) | (spl0_10 | ~spl0_1448)),
  inference(backward_demodulation,[],[f692919,f486755])).
fof(f772883,plain,(
  ( ! [X1] : (c_Orderings_Obot__class_Obot(tc_fun(X1,tc_bool)) = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool)) ) | (~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f692919,f486756])).
fof(f776045,plain,(
  spl0_1421 | spl0_10 | ~spl0_613 | ~spl0_614 | spl0_819 | ~spl0_1448),
  inference(avatar_split_clause,[],[f775680,f692918,f215375,f94496,f94492,f1045,f681332])).
fof(f94492,plain,(
  spl0_613 <=> ! [X42,X41,X43] : ~hBOOL(hAPP(c_Relation_OId__on(c_Orderings_Obot__class_Obot(tc_fun(X41,tc_bool)),X43),X42))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_613])])).
fof(f775680,plain,(
  ( ! [X116,X115] : (~c_Equiv__Relations_Oequiv(v_x,X115,X116) | c_Relation_Oirrefl(X115,X116)) ) | (spl0_10 | ~spl0_613 | ~spl0_614 | spl0_819 | ~spl0_1448)),
  inference(backward_demodulation,[],[f775124,f774583])).
fof(f774583,plain,(
  ( ! [X116,X115] : (~c_Equiv__Relations_Oequiv(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),X115,X116) | c_Relation_Oirrefl(X115,X116)) ) | (spl0_10 | ~spl0_613 | ~spl0_614 | ~spl0_1448)),
  inference(backward_demodulation,[],[f774230,f773768])).
fof(f773768,plain,(
  ( ! [X116,X115] : (~c_Equiv__Relations_Oequiv(c_HOL_Ominus__class_Ominus(v_x,tc_bool),X115,X116) | c_Relation_Oirrefl(X115,X116)) ) | (spl0_10 | ~spl0_613 | ~spl0_1448)),
  inference(backward_demodulation,[],[f772882,f96172])).
fof(f96172,plain,(
  ( ! [X116,X114,X115] : (~c_Equiv__Relations_Oequiv(c_Orderings_Obot__class_Obot(tc_fun(X114,tc_bool)),X115,X116) | c_Relation_Oirrefl(X115,X116)) ) | ~spl0_613),
  inference(resolution,[],[f95961,f6430])).
fof(f6430,plain,(
  ( ! [X14,X12,X13] : (hBOOL(c_in(c_ATP__Linkup_Osko__Relation__Xirrefl__def__1__1(X13,X14),X12,X14)) | ~c_Equiv__Relations_Oequiv(X12,X13,X14) | c_Relation_Oirrefl(X13,X14)) )),
  inference(resolution,[],[f654,f666])).
fof(f654,axiom,(
  ( ! [X4,X2,X0,X10,X1] : (~hBOOL(c_in(c_Pair(X1,X10,X0,X0),X2,tc_prod(X0,X0))) | ~c_Equiv__Relations_Oequiv(X4,X2,X0) | hBOOL(c_in(X10,X4,X0))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f95961,plain,(
  ( ! [X2,X0,X1] : (~hBOOL(c_in(X0,c_Orderings_Obot__class_Obot(tc_fun(X1,tc_bool)),X2))) ) | ~spl0_613),
  inference(resolution,[],[f95611,f5734])).
fof(f95611,plain,(
  ( ! [X14,X15,X13] : (c_Relation_Oirrefl(c_Relation_OId__on(c_Orderings_Obot__class_Obot(tc_fun(X13,tc_bool)),X14),X15)) ) | ~spl0_613),
  inference(resolution,[],[f94493,f76700])).
fof(f94493,plain,(
  ( ! [X43,X41,X42] : (~hBOOL(hAPP(c_Relation_OId__on(c_Orderings_Obot__class_Obot(tc_fun(X41,tc_bool)),X43),X42))) ) | ~spl0_613),
  inference(avatar_component_clause,[],[f94492])).
fof(f775124,plain,(
  v_x = c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool) | (spl0_10 | ~spl0_614 | spl0_819 | ~spl0_1448)),
  inference(backward_demodulation,[],[f773313,f774230])).
fof(f773313,plain,(
  v_x = c_HOL_Ominus__class_Ominus(v_x,tc_bool) | (spl0_819 | ~spl0_1448)),
  inference(backward_demodulation,[],[f692919,f763871])).
fof(f763871,plain,(
  c_Orderings_Obot__class_Obot(tc_bool) = c_HOL_Ominus__class_Ominus(v_x,tc_bool) | spl0_819),
  inference(forward_demodulation,[],[f763729,f4622])).
fof(f763729,plain,(
  ( ! [X2,X3] : (hAPP(c_Orderings_Obot__class_Obot(tc_fun(X2,tc_bool)),X3) = c_HOL_Ominus__class_Ominus(v_x,tc_bool)) ) | spl0_819),
  inference(superposition,[],[f48,f644953])).
fof(f644953,plain,(
  ( ! [X1037,X1038,X1036] : (c_Orderings_Obot__class_Obot(tc_fun(X1036,tc_bool)) = c_COMBK(c_HOL_Ominus__class_Ominus(v_x,tc_bool),X1037,X1038)) ) | spl0_819),
  inference(resolution,[],[f545154,f215377])).
fof(f215377,plain,(
  ~hBOOL(c_HOL_Ominus__class_Ominus(v_x,tc_bool)) | spl0_819),
  inference(avatar_component_clause,[],[f215375])).
fof(f772661,plain,(
  spl0_1448 | spl0_787),
  inference(avatar_split_clause,[],[f772551,f167778,f692918])).
fof(f167778,plain,(
  spl0_787 <=> hBOOL(v_x)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_787])])).
fof(f772551,plain,(
  v_x = c_Orderings_Obot__class_Obot(tc_bool) | spl0_787),
  inference(forward_demodulation,[],[f772401,f48])).
fof(f772401,plain,(
  ( ! [X43,X41,X42] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_COMBK(v_x,X41,X42),X43)) ) | spl0_787),
  inference(superposition,[],[f4622,f765912])).
fof(f765912,plain,(
  ( ! [X2,X0,X1] : (c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)) = c_COMBK(v_x,X1,X2)) ) | spl0_787),
  inference(resolution,[],[f167780,f545154])).
fof(f167780,plain,(
  ~hBOOL(v_x) | spl0_787),
  inference(avatar_component_clause,[],[f167778])).
fof(f765808,plain,(
  spl0_1377 | ~spl0_12 | spl0_1376),
  inference(avatar_split_clause,[],[f687340,f616474,f1084,f616481])).
fof(f687340,plain,(
  ( ! [X51] : (~hBOOL(c_Orderings_Otop__class_Otop(tc_fun(X51,tc_bool)))) ) | (~spl0_12 | spl0_1376)),
  inference(superposition,[],[f616476,f3544])).
fof(f765806,plain,(
  spl0_1492 | ~spl0_787 | spl0_819),
  inference(avatar_split_clause,[],[f764118,f215375,f167778,f765804])).
fof(f764118,plain,(
  ( ! [X22] : (~hBOOL(v_x) | c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_Set_Oinsert(c_Orderings_Obot__class_Obot(tc_bool),X22),c_Orderings_Obot__class_Obot(tc_bool))) ) | spl0_819),
  inference(superposition,[],[f1242,f764037])).
fof(f764037,plain,(
  v_x = hAPP(c_Orderings_Obot__class_Obot(tc_bool),c_Orderings_Obot__class_Obot(tc_bool)) | spl0_819),
  inference(superposition,[],[f167564,f763871])).
fof(f167564,plain,(
  ( ! [X89] : (hAPP(c_HOL_Ominus__class_Ominus(X89,tc_bool),c_Orderings_Obot__class_Obot(tc_bool)) = X89) )),
  inference(superposition,[],[f145537,f48])).
fof(f145537,plain,(
  ( ! [X1] : (hAPP(X1,v_x) = hAPP(c_HOL_Ominus__class_Ominus(hAPP(X1,v_x),tc_bool),c_Orderings_Obot__class_Obot(tc_bool))) )),
  inference(forward_demodulation,[],[f145486,f1099])).
fof(f145486,plain,(
  ( ! [X1] : (hAPP(X1,v_x) = hAPP(c_HOL_Ominus__class_Ominus(hAPP(X1,v_x),tc_bool),hAPP(c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_bool)),v_x))) )),
  inference(superposition,[],[f7150,f88])).
fof(f88,axiom,(
  ( ! [X4,X0] : (hAPP(c_HOL_Ominus__class_Ominus(X4,tc_fun(X0,tc_bool)),c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool))) = X4) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f765482,plain,(
  ~spl0_795 | spl0_819 | ~spl0_1388 | spl0_1448),
  inference(avatar_contradiction_clause,[],[f765481])).
fof(f765481,plain,(
  $false | (~spl0_795 | spl0_819 | ~spl0_1388 | spl0_1448)),
  inference(subsumption_resolution,[],[f765474,f692920])).
fof(f765474,plain,(
  v_x = c_Orderings_Obot__class_Obot(tc_bool) | (~spl0_795 | spl0_819 | ~spl0_1388)),
  inference(backward_demodulation,[],[f764896,f764036])).
fof(f764036,plain,(
  c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_Orderings_Obot__class_Obot(tc_bool),v_x) | spl0_819),
  inference(superposition,[],[f166789,f763871])).
fof(f166789,plain,(
  ( ! [X89] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_HOL_Ominus__class_Ominus(X89,tc_bool),X89)) )),
  inference(superposition,[],[f145536,f48])).
fof(f145536,plain,(
  ( ! [X0] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_HOL_Ominus__class_Ominus(hAPP(X0,v_x),tc_bool),hAPP(X0,v_x))) )),
  inference(forward_demodulation,[],[f145485,f1099])).
fof(f145485,plain,(
  ( ! [X0] : (hAPP(c_Orderings_Obot__class_Obot(tc_fun(t_a,tc_bool)),v_x) = hAPP(c_HOL_Ominus__class_Ominus(hAPP(X0,v_x),tc_bool),hAPP(X0,v_x))) )),
  inference(superposition,[],[f7150,f89])).
fof(f89,axiom,(
  ( ! [X4,X0] : (c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)) = hAPP(c_HOL_Ominus__class_Ominus(X4,tc_fun(X0,tc_bool)),X4)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f764896,plain,(
  ( ! [X0] : (v_x = hAPP(c_Orderings_Obot__class_Obot(tc_bool),X0)) ) | (~spl0_795 | ~spl0_1388)),
  inference(backward_demodulation,[],[f172779,f691115])).
fof(f691115,plain,(
  ( ! [X0] : (v_x = hAPP(c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool),X0)) ) | ~spl0_1388),
  inference(backward_demodulation,[],[f623092,f1034])).
fof(f764887,plain,(
  spl0_795 | ~spl0_12 | ~spl0_1377 | ~spl0_1388),
  inference(avatar_split_clause,[],[f764849,f623091,f616481,f1084,f172778])).
fof(f764849,plain,(
  c_Orderings_Obot__class_Obot(tc_bool) = c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool) | (~spl0_12 | ~spl0_1377 | ~spl0_1388)),
  inference(forward_demodulation,[],[f764707,f4622])).
fof(f764707,plain,(
  ( ! [X2,X3] : (hAPP(c_Orderings_Obot__class_Obot(tc_fun(X2,tc_bool)),X3) = c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool)) ) | (~spl0_12 | ~spl0_1377 | ~spl0_1388)),
  inference(superposition,[],[f48,f692597])).
fof(f692597,plain,(
  ( ! [X1106,X1108,X1109] : (c_Orderings_Obot__class_Obot(tc_fun(X1106,tc_bool)) = c_COMBK(c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool),X1108,X1109)) ) | (~spl0_12 | ~spl0_1377 | ~spl0_1388)),
  inference(backward_demodulation,[],[f691143,f644967])).
fof(f644967,plain,(
  ( ! [X1106,X1108,X1107,X1109] : (c_Orderings_Obot__class_Obot(tc_fun(X1106,tc_bool)) = c_COMBK(c_Orderings_Otop__class_Otop(tc_fun(X1107,tc_bool)),X1108,X1109)) ) | ~spl0_1377),
  inference(resolution,[],[f545154,f616482])).
fof(f691143,plain,(
  ( ! [X35] : (c_Orderings_Otop__class_Otop(tc_fun(X35,tc_bool)) = c_Lattices_Oupper__semilattice__class_Osup(v_x,tc_bool)) ) | (~spl0_12 | ~spl0_1388)),
  inference(backward_demodulation,[],[f623092,f3544])).
fof(f707009,plain,(
  ~spl0_32),
  inference(avatar_contradiction_clause,[],[f707008])).
fof(f707008,plain,(
  $false | ~spl0_32),
  inference(subsumption_resolution,[],[f696689,f2123])).
fof(f2123,plain,(
  ( ! [X14] : (c_Orderings_Obot__class_Obot(tc_bool) = X14) ) | ~spl0_32),
  inference(avatar_component_clause,[],[f2122])).
fof(f2122,plain,(
  spl0_32 <=> ! [X14] : c_Orderings_Obot__class_Obot(tc_bool) = X14),
  introduced(avatar_definition,[new_symbols(naming,[spl0_32])])).
fof(f696689,plain,(
  ( ! [X14,X15] : (c_Orderings_Obot__class_Obot(tc_bool) != c_Transitive__Closure_Ortrancl(X14,X15)) ) | ~spl0_32),
  inference(backward_demodulation,[],[f2123,f68238])).
fof(f68238,plain,(
  ( ! [X14,X15] : (c_Transitive__Closure_Ortrancl(X14,X15) != c_Orderings_Obot__class_Obot(tc_fun(tc_prod(X15,X15),tc_bool))) )),
  inference(subsumption_resolution,[],[f5136,f21550])).
fof(f21550,plain,(
  ( ! [X17,X15] : (c_Orderings_Obot__class_Obot(tc_fun(X17,tc_bool)) != c_Relation_OId(X15)) )),
  inference(superposition,[],[f9,f1491])).
fof(f1491,plain,(
  ( ! [X37,X35,X36] : (c_Relation_OId(X35) = hAPP(c_Set_Oinsert(c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__2(c_Relation_OId(X35),X36),X37),c_Relation_OId(X35))) )),
  inference(resolution,[],[f1242,f1231])).
fof(f1231,plain,(
  ( ! [X0,X1] : (hBOOL(hAPP(c_Relation_OId(X0),c_ATP__Linkup_Osko__Set__Xbex__UNIV__1__2(c_Relation_OId(X0),X1)))) )),
  inference(resolution,[],[f1230,f515])).
fof(f1230,plain,(
  ( ! [X0,X1] : (hBOOL(hAPP(c_Relation_OId(X0),c_Pair(X1,X1,X0,X0)))) )),
  inference(resolution,[],[f619,f678])).
fof(f5136,plain,(
  ( ! [X14,X15] : (c_Transitive__Closure_Ortrancl(X14,X15) != c_Orderings_Obot__class_Obot(tc_fun(tc_prod(X15,X15),tc_bool)) | c_Relation_OId(X15) = c_Orderings_Obot__class_Obot(tc_fun(tc_prod(X15,X15),tc_bool))) )),
  inference(superposition,[],[f198,f367])).
fof(f367,axiom,(
  ( ! [X0,X12] : (c_Transitive__Closure_Ortrancl(X12,X0) = hAPP(c_Lattices_Oupper__semilattice__class_Osup(c_Transitive__Closure_Ortrancl(X12,X0),tc_fun(tc_prod(X0,X0),tc_bool)),c_Relation_OId(X0))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f198,axiom,(
  ( ! [X4,X0,X7] : (c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)) != hAPP(c_Lattices_Oupper__semilattice__class_Osup(X4,tc_fun(X0,tc_bool)),X7) | c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)) = X7) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f693021,plain,(
  spl0_32 | ~spl0_1448 | ~spl0_1388),
  inference(avatar_split_clause,[],[f693020,f623091,f692918,f2122])).
fof(f693020,plain,(
  ( ! [X18] : (v_x != c_Orderings_Obot__class_Obot(tc_bool) | c_Orderings_Obot__class_Obot(tc_bool) = X18) ) | ~spl0_1388),
  inference(forward_demodulation,[],[f693019,f623092])).
fof(f693019,plain,(
  ( ! [X18] : (c_Orderings_Obot__class_Obot(tc_bool) != c_Orderings_Otop__class_Otop(tc_bool) | c_Orderings_Obot__class_Obot(tc_bool) = X18) )),
  inference(subsumption_resolution,[],[f2140,f704])).
fof(f2140,plain,(
  ( ! [X18] : (c_Orderings_Obot__class_Obot(tc_bool) != c_Orderings_Otop__class_Otop(tc_bool) | c_Orderings_Obot__class_Obot(tc_bool) = X18 | ~class_Lattices_Obounded__lattice(tc_bool)) )),
  inference(superposition,[],[f201,f1034])).
fof(f201,axiom,(
  ( ! [X4,X0,X7] : (c_Orderings_Obot__class_Obot(X0) != hAPP(c_Lattices_Oupper__semilattice__class_Osup(X4,X0),X7) | c_Orderings_Obot__class_Obot(X0) = X7 | ~class_Lattices_Obounded__lattice(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f678041,plain,(
  spl0_122 | ~spl0_12 | ~spl0_61),
  inference(avatar_split_clause,[],[f677626,f13802,f1084,f31932])).
fof(f677626,plain,(
  c_Relation_Oirrefl(c_Orderings_Otop__class_Otop(tc_bool),tc_Arrow__Order__Mirabelle_Oalt) | (~spl0_12 | ~spl0_61)),
  inference(backward_demodulation,[],[f677624,f857])).
fof(f857,plain,(
  c_Relation_Oirrefl(v_F(v_P),tc_Arrow__Order__Mirabelle_Oalt)),
  inference(subsumption_resolution,[],[f856,f681])).
fof(f681,axiom,(
  ( ! [X1] : (c_Relation_Ototal__on(c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_bool)),hAPP(v_P,X1),tc_Arrow__Order__Mirabelle_Oalt)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f856,plain,(
  ~c_Relation_Ototal__on(c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_bool)),hAPP(v_P,v_x),tc_Arrow__Order__Mirabelle_Oalt) | c_Relation_Oirrefl(v_F(v_P),tc_Arrow__Order__Mirabelle_Oalt)),
  inference(subsumption_resolution,[],[f855,f680])).
fof(f680,axiom,(
  ( ! [X1] : (c_Relation_Oirrefl(hAPP(v_P,X1),tc_Arrow__Order__Mirabelle_Oalt)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f855,plain,(
  ~c_Relation_Oirrefl(hAPP(v_P,v_x),tc_Arrow__Order__Mirabelle_Oalt) | ~c_Relation_Ototal__on(c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_bool)),hAPP(v_P,v_x),tc_Arrow__Order__Mirabelle_Oalt) | c_Relation_Oirrefl(v_F(v_P),tc_Arrow__Order__Mirabelle_Oalt)),
  inference(subsumption_resolution,[],[f685,f679])).
fof(f679,axiom,(
  ( ! [X1] : (c_Relation_Otrans(hAPP(v_P,X1),tc_Arrow__Order__Mirabelle_Oalt)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f685,axiom,(
  ~c_Relation_Otrans(hAPP(v_P,v_x),tc_Arrow__Order__Mirabelle_Oalt) | ~c_Relation_Oirrefl(hAPP(v_P,v_x),tc_Arrow__Order__Mirabelle_Oalt) | ~c_Relation_Ototal__on(c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_bool)),hAPP(v_P,v_x),tc_Arrow__Order__Mirabelle_Oalt) | c_Relation_Oirrefl(v_F(v_P),tc_Arrow__Order__Mirabelle_Oalt)),
  file('./p/SCT021-1.tptp',unknown)).
fof(f677624,plain,(
  v_F(v_P) = c_Orderings_Otop__class_Otop(tc_bool) | (~spl0_12 | ~spl0_61)),
  inference(forward_demodulation,[],[f677442,f3694])).
fof(f677442,plain,(
  ( ! [X2,X3] : (v_F(v_P) = hAPP(c_Orderings_Otop__class_Otop(tc_fun(X2,tc_bool)),X3)) ) | ~spl0_61),
  inference(superposition,[],[f48,f670522])).
fof(f670522,plain,(
  ( ! [X4,X2,X3] : (c_Orderings_Otop__class_Otop(tc_fun(X2,tc_bool)) = c_COMBK(v_F(v_P),X3,X4)) ) | ~spl0_61),
  inference(resolution,[],[f13803,f3543])).
fof(f655281,plain,(
  ~spl0_12 | ~spl0_787 | spl0_1388),
  inference(avatar_contradiction_clause,[],[f655280])).
fof(f655280,plain,(
  $false | (~spl0_12 | ~spl0_787 | spl0_1388)),
  inference(subsumption_resolution,[],[f655279,f623093])).
fof(f655279,plain,(
  v_x = c_Orderings_Otop__class_Otop(tc_bool) | (~spl0_12 | ~spl0_787)),
  inference(backward_demodulation,[],[f655088,f3694])).
fof(f655088,plain,(
  ( ! [X2,X3] : (v_x = hAPP(c_Orderings_Otop__class_Otop(tc_fun(X2,tc_bool)),X3)) ) | ~spl0_787),
  inference(superposition,[],[f48,f648352])).
fof(f648352,plain,(
  ( ! [X4,X2,X3] : (c_Orderings_Otop__class_Otop(tc_fun(X2,tc_bool)) = c_COMBK(v_x,X3,X4)) ) | ~spl0_787),
  inference(resolution,[],[f167779,f3543])).
fof(f167779,plain,(
  hBOOL(v_x) | ~spl0_787),
  inference(avatar_component_clause,[],[f167778])).
fof(f623094,plain,(
  spl0_1387 | ~spl0_1388 | ~spl0_802),
  inference(avatar_split_clause,[],[f623086,f173159,f623091,f623088])).
fof(f623086,plain,(
  ( ! [X67] : (v_x != c_Orderings_Otop__class_Otop(tc_bool) | c_Orderings_Otop__class_Otop(tc_bool) = hAPP(v_x,X67)) ) | ~spl0_802),
  inference(subsumption_resolution,[],[f623026,f704])).
fof(f623026,plain,(
  ( ! [X67] : (v_x != c_Orderings_Otop__class_Otop(tc_bool) | c_Orderings_Otop__class_Otop(tc_bool) = hAPP(v_x,X67) | ~class_Lattices_Obounded__lattice(tc_bool)) ) | ~spl0_802),
  inference(superposition,[],[f267,f617008])).
fof(f617008,plain,(
  ( ! [X2] : (v_x = hAPP(c_Lattices_Olower__semilattice__class_Oinf(v_x,tc_bool),hAPP(v_x,X2))) ) | ~spl0_802),
  inference(superposition,[],[f1367,f173160])).
fof(f267,axiom,(
  ( ! [X4,X0,X7] : (c_Orderings_Otop__class_Otop(X0) != hAPP(c_Lattices_Olower__semilattice__class_Oinf(X4,X0),X7) | c_Orderings_Otop__class_Otop(X0) = X7 | ~class_Lattices_Obounded__lattice(X0)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f440526,plain,(
  ~spl0_56),
  inference(avatar_contradiction_clause,[],[f440525])).
fof(f440525,plain,(
  $false | ~spl0_56),
  inference(subsumption_resolution,[],[f440471,f857])).
fof(f440471,plain,(
  ~c_Relation_Oirrefl(v_F(v_P),tc_Arrow__Order__Mirabelle_Oalt) | ~spl0_56),
  inference(resolution,[],[f337770,f672])).
fof(f337770,plain,(
  hBOOL(c_in(c_Pair(v_b,v_b,tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),v_F(v_P),tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt))) | ~spl0_56),
  inference(resolution,[],[f9924,f29319])).
fof(f29319,plain,(
  ( ! [X139] : (~hBOOL(c_in(c_Pair(X139,v_a,tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),v_F(v_P),tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt))) | hBOOL(c_in(c_Pair(X139,v_b,tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),v_F(v_P),tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt)))) )),
  inference(subsumption_resolution,[],[f29295,f854])).
fof(f854,plain,(
  c_Relation_Otrans(v_F(v_P),tc_Arrow__Order__Mirabelle_Oalt)),
  inference(subsumption_resolution,[],[f853,f681])).
fof(f853,plain,(
  ~c_Relation_Ototal__on(c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_bool)),hAPP(v_P,v_x),tc_Arrow__Order__Mirabelle_Oalt) | c_Relation_Otrans(v_F(v_P),tc_Arrow__Order__Mirabelle_Oalt)),
  inference(subsumption_resolution,[],[f852,f680])).
fof(f852,plain,(
  ~c_Relation_Oirrefl(hAPP(v_P,v_x),tc_Arrow__Order__Mirabelle_Oalt) | ~c_Relation_Ototal__on(c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_bool)),hAPP(v_P,v_x),tc_Arrow__Order__Mirabelle_Oalt) | c_Relation_Otrans(v_F(v_P),tc_Arrow__Order__Mirabelle_Oalt)),
  inference(subsumption_resolution,[],[f684,f679])).
fof(f684,axiom,(
  ~c_Relation_Otrans(hAPP(v_P,v_x),tc_Arrow__Order__Mirabelle_Oalt) | ~c_Relation_Oirrefl(hAPP(v_P,v_x),tc_Arrow__Order__Mirabelle_Oalt) | ~c_Relation_Ototal__on(c_Orderings_Otop__class_Otop(tc_fun(tc_Arrow__Order__Mirabelle_Oalt,tc_bool)),hAPP(v_P,v_x),tc_Arrow__Order__Mirabelle_Oalt) | c_Relation_Otrans(v_F(v_P),tc_Arrow__Order__Mirabelle_Oalt)),
  file('./p/SCT021-1.tptp',unknown)).
fof(f29295,plain,(
  ( ! [X139] : (~hBOOL(c_in(c_Pair(X139,v_a,tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),v_F(v_P),tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt))) | ~c_Relation_Otrans(v_F(v_P),tc_Arrow__Order__Mirabelle_Oalt) | hBOOL(c_in(c_Pair(X139,v_b,tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),v_F(v_P),tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt)))) )),
  inference(resolution,[],[f670,f682])).
fof(f670,axiom,(
  ( ! [X2,X0,X10,X1,X18] : (~hBOOL(c_in(c_Pair(X10,X18,X0,X0),X2,tc_prod(X0,X0))) | ~hBOOL(c_in(c_Pair(X1,X10,X0,X0),X2,tc_prod(X0,X0))) | ~c_Relation_Otrans(X2,X0) | hBOOL(c_in(c_Pair(X1,X18,X0,X0),X2,tc_prod(X0,X0)))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f9924,plain,(
  hBOOL(c_in(c_Pair(v_b,v_a,tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt),v_F(v_P),tc_prod(tc_Arrow__Order__Mirabelle_Oalt,tc_Arrow__Order__Mirabelle_Oalt))) | ~spl0_56),
  inference(avatar_component_clause,[],[f9923])).
fof(f94806,plain,(
  ~spl0_10),
  inference(avatar_split_clause,[],[f1102,f1045])).
fof(f1102,plain,(
  ~hBOOL(c_Orderings_Obot__class_Obot(tc_bool))),
  inference(superposition,[],[f563,f1099])).
fof(f563,axiom,(
  ( ! [X0,X1] : (~hBOOL(hAPP(c_Orderings_Obot__class_Obot(tc_fun(X0,tc_bool)),X1))) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f94498,plain,(
  spl0_614 | spl0_10),
  inference(avatar_split_clause,[],[f92689,f1045,f94496])).
fof(f92689,plain,(
  ( ! [X0,X1] : (hBOOL(c_Orderings_Obot__class_Obot(tc_bool)) | ~hBOOL(hAPP(c_Relation_OId__on(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X1),X0))) )),
  inference(superposition,[],[f40426,f1032])).
fof(f1032,plain,(
  ( ! [X0] : (c_Orderings_Obot__class_Obot(tc_bool) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(tc_bool),tc_bool),X0)) )),
  inference(resolution,[],[f79,f704])).
fof(f79,axiom,(
  ( ! [X0,X1] : (~class_Lattices_Obounded__lattice(X0) | c_Orderings_Obot__class_Obot(X0) = hAPP(c_Lattices_Olower__semilattice__class_Oinf(c_Orderings_Obot__class_Obot(X0),X0),X1)) )),
  file('./p/SCT021-1.tptp',unknown)).
fof(f94494,plain,(
  spl0_613 | spl0_10),
  inference(avatar_split_clause,[],[f92697,f1045,f94492])).
fof(f92697,plain,(
  ( ! [X43,X41,X42] : (hBOOL(c_Orderings_Obot__class_Obot(tc_bool)) | ~hBOOL(hAPP(c_Relation_OId__on(c_Orderings_Obot__class_Obot(tc_fun(X41,tc_bool)),X43),X42))) )),
  inference(superposition,[],[f40426,f4622])).
fof(f1118,plain,(
  spl0_12),
  inference(avatar_split_clause,[],[f1114,f1084])).
fof(f1114,plain,(
  hBOOL(c_Orderings_Otop__class_Otop(tc_bool))),
  inference(superposition,[],[f674,f1112])).
% SZS output end Proof for SCT021-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 555982
% Time elapsed: 50.614 s
% ------------------------------
% ------------------------------
