% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SET188-6
% SZS output start Proof for SET188-6
fof(f5490,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f2382,f5489])).
fof(f5489,plain,(
  ~spl0_188),
  inference(avatar_contradiction_clause,[],[f5488])).
fof(f5488,plain,(
  $false | ~spl0_188),
  inference(subsumption_resolution,[],[f5487,f114])).
fof(f114,axiom,(
  ~subclass(x,y)),
  file('./p/SET188-6.tptp',unknown)).
fof(f5487,plain,(
  subclass(x,y) | ~spl0_188),
  inference(duplicate_literal_removal,[],[f5485])).
fof(f5485,plain,(
  subclass(x,y) | subclass(x,y) | ~spl0_188),
  inference(resolution,[],[f5120,f3])).
fof(f3,axiom,(
  ( ! [X0,X1] : (~member(not_subclass_element(X0,X1),X1) | subclass(X0,X1)) )),
  file('./p/SET188-6.tptp',unknown)).
fof(f5120,plain,(
  ( ! [X0] : (member(not_subclass_element(x,X0),y) | subclass(x,X0)) ) | ~spl0_188),
  inference(resolution,[],[f4964,f2])).
fof(f2,axiom,(
  ( ! [X0,X1] : (member(not_subclass_element(X0,X1),X0) | subclass(X0,X1)) )),
  file('./p/SET188-6.tptp',unknown)).
fof(f4964,plain,(
  ( ! [X1] : (~member(X1,x) | member(X1,y)) ) | ~spl0_188),
  inference(subsumption_resolution,[],[f4963,f1585])).
fof(f1585,plain,(
  ( ! [X6,X8,X7] : (~member(X6,X7) | member(X6,universal_class) | ~member(X6,X8)) )),
  inference(resolution,[],[f370,f4])).
fof(f4,axiom,(
  ( ! [X0] : (subclass(X0,universal_class)) )),
  file('./p/SET188-6.tptp',unknown)).
fof(f370,plain,(
  ( ! [X6,X8,X7,X9] : (~subclass(intersection(X8,X7),X9) | ~member(X6,X8) | member(X6,X9) | ~member(X6,X7)) )),
  inference(resolution,[],[f23,f1])).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (~member(X2,X0) | member(X2,X1) | ~subclass(X0,X1)) )),
  file('./p/SET188-6.tptp',unknown)).
fof(f23,axiom,(
  ( ! [X4,X0,X1] : (member(X4,intersection(X0,X1)) | ~member(X4,X1) | ~member(X4,X0)) )),
  file('./p/SET188-6.tptp',unknown)).
fof(f4963,plain,(
  ( ! [X1] : (~member(X1,x) | member(X1,y) | ~member(X1,universal_class)) ) | ~spl0_188),
  inference(subsumption_resolution,[],[f1017,f2404])).
fof(f2404,plain,(
  ( ! [X29] : (~member(X29,null_class)) ) | ~spl0_188),
  inference(backward_demodulation,[],[f2380,f2324])).
fof(f2324,plain,(
  ( ! [X29] : (~member(X29,domain_of(null_class))) )),
  inference(trivial_inequality_removal,[],[f2318])).
fof(f2318,plain,(
  ( ! [X29] : (null_class != null_class | ~member(X29,domain_of(null_class))) )),
  inference(superposition,[],[f135,f1698])).
fof(f1698,plain,(
  ( ! [X0] : (null_class = intersection(null_class,X0)) )),
  inference(resolution,[],[f1430,f1316])).
fof(f1316,plain,(
  ( ! [X0,X1] : (subclass(intersection(X0,X1),X0)) )),
  inference(duplicate_literal_removal,[],[f1296])).
fof(f1296,plain,(
  ( ! [X0,X1] : (subclass(intersection(X0,X1),X0) | subclass(intersection(X0,X1),X0)) )),
  inference(resolution,[],[f188,f3])).
fof(f188,plain,(
  ( ! [X6,X4,X5] : (member(not_subclass_element(intersection(X4,X5),X6),X4) | subclass(intersection(X4,X5),X6)) )),
  inference(resolution,[],[f21,f2])).
fof(f21,axiom,(
  ( ! [X4,X0,X1] : (~member(X4,intersection(X0,X1)) | member(X4,X0)) )),
  file('./p/SET188-6.tptp',unknown)).
fof(f1430,plain,(
  ( ! [X1] : (~subclass(X1,null_class) | null_class = X1) )),
  inference(duplicate_literal_removal,[],[f1426])).
fof(f1426,plain,(
  ( ! [X1] : (null_class = X1 | null_class = X1 | ~subclass(X1,null_class)) )),
  inference(resolution,[],[f1325,f7])).
fof(f7,axiom,(
  ( ! [X0,X1] : (~subclass(X1,X0) | X0 = X1 | ~subclass(X0,X1)) )),
  file('./p/SET188-6.tptp',unknown)).
fof(f1325,plain,(
  ( ! [X0] : (subclass(null_class,X0) | null_class = X0) )),
  inference(superposition,[],[f1316,f67])).
fof(f67,axiom,(
  ( ! [X0] : (null_class = intersection(X0,regular(X0)) | null_class = X0) )),
  file('./p/SET188-6.tptp',unknown)).
fof(f135,plain,(
  ( ! [X4,X0] : (null_class != intersection(X0,cross_product(unordered_pair(X4,X4),universal_class)) | ~member(X4,domain_of(X0))) )),
  inference(definition_unfolding,[],[f30,f28,f12])).
fof(f12,axiom,(
  ( ! [X0] : (unordered_pair(X0,X0) = singleton(X0)) )),
  file('./p/SET188-6.tptp',unknown)).
fof(f28,axiom,(
  ( ! [X0,X5,X1] : (intersection(X5,cross_product(X0,X1)) = restrict(X5,X0,X1)) )),
  file('./p/SET188-6.tptp',unknown)).
fof(f30,axiom,(
  ( ! [X4,X0] : (~member(X4,domain_of(X0)) | restrict(X0,singleton(X4),universal_class) != null_class) )),
  file('./p/SET188-6.tptp',unknown)).
fof(f2380,plain,(
  null_class = domain_of(null_class) | ~spl0_188),
  inference(avatar_component_clause,[],[f2378])).
fof(f2378,plain,(
  spl0_188 <=> null_class = domain_of(null_class)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_188])])).
fof(f1017,plain,(
  ( ! [X1] : (~member(X1,x) | member(X1,null_class) | member(X1,y) | ~member(X1,universal_class)) )),
  inference(resolution,[],[f372,f25])).
fof(f25,axiom,(
  ( ! [X4,X0] : (member(X4,complement(X0)) | member(X4,X0) | ~member(X4,universal_class)) )),
  file('./p/SET188-6.tptp',unknown)).
fof(f372,plain,(
  ( ! [X0] : (~member(X0,complement(y)) | ~member(X0,x) | member(X0,null_class)) )),
  inference(superposition,[],[f23,f113])).
fof(f113,axiom,(
  null_class = intersection(complement(y),x)),
  file('./p/SET188-6.tptp',unknown)).
fof(f2382,plain,(
  spl0_188),
  inference(avatar_split_clause,[],[f2362,f2378])).
fof(f2362,plain,(
  null_class = domain_of(null_class)),
  inference(resolution,[],[f2324,f66])).
fof(f66,axiom,(
  ( ! [X0] : (member(regular(X0),X0) | null_class = X0) )),
  file('./p/SET188-6.tptp',unknown)).
% SZS output end Proof for SET188-6
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 11129
% Time elapsed: 0.138 s
% ------------------------------
% ------------------------------
