% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SET519-6
% SZS output start Proof for SET519-6
fof(f503417,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f1594,f2896,f3242,f4399,f4439,f285437,f285450,f385121,f397525,f420012,f434606,f443583,f495176,f499383,f500718,f501499,f501821,f503283])).
fof(f503283,plain,(
  ~spl0_139 | ~spl0_239 | spl0_3543),
  inference(avatar_contradiction_clause,[],[f503282])).
fof(f503282,plain,(
  $false | (~spl0_139 | ~spl0_239 | spl0_3543)),
  inference(subsumption_resolution,[],[f503281,f327])).
fof(f327,plain,(
  member(x,universal_class)),
  inference(superposition,[],[f11,f176])).
fof(f176,plain,(
  x = unordered_pair(unordered_pair(x,x),unordered_pair(x,unordered_pair(y,y)))),
  inference(definition_unfolding,[],[f113,f118])).
fof(f118,plain,(
  ( ! [X0,X1] : (ordered_pair(X0,X1) = unordered_pair(unordered_pair(X0,X0),unordered_pair(X0,unordered_pair(X1,X1)))) )),
  inference(definition_unfolding,[],[f13,f12,f12])).
fof(f12,axiom,(
  ( ! [X0] : (unordered_pair(X0,X0) = singleton(X0)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f13,axiom,(
  ( ! [X0,X1] : (unordered_pair(singleton(X0),unordered_pair(X0,singleton(X1))) = ordered_pair(X0,X1)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f113,axiom,(
  x = ordered_pair(x,y)),
  file('./p/SET519-6.tptp',unknown)).
fof(f11,axiom,(
  ( ! [X0,X1] : (member(unordered_pair(X0,X1),universal_class)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f503281,plain,(
  ~member(x,universal_class) | (~spl0_139 | ~spl0_239 | spl0_3543)),
  inference(subsumption_resolution,[],[f503210,f501839])).
fof(f501839,plain,(
  ~member(x,regular(x)) | spl0_3543),
  inference(forward_demodulation,[],[f238282,f330707])).
fof(f330707,plain,(
  ( ! [X2] : (complement(complement(X2)) = X2) )),
  inference(subsumption_resolution,[],[f330661,f246742])).
fof(f246742,plain,(
  ( ! [X0] : (subclass(X0,complement(complement(X0)))) )),
  inference(duplicate_literal_removal,[],[f246710])).
fof(f246710,plain,(
  ( ! [X0] : (subclass(X0,complement(complement(X0))) | subclass(X0,complement(complement(X0)))) )),
  inference(resolution,[],[f10892,f2])).
fof(f2,axiom,(
  ( ! [X0,X1] : (member(not_subclass_element(X0,X1),X0) | subclass(X0,X1)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f10892,plain,(
  ( ! [X24,X25] : (~member(not_subclass_element(X24,complement(complement(X25))),X25) | subclass(X24,complement(complement(X25)))) )),
  inference(resolution,[],[f3868,f24])).
fof(f24,axiom,(
  ( ! [X4,X0] : (~member(X4,complement(X0)) | ~member(X4,X0)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f3868,plain,(
  ( ! [X6,X5] : (member(not_subclass_element(X5,complement(X6)),X6) | subclass(X5,complement(X6))) )),
  inference(subsumption_resolution,[],[f3864,f4])).
fof(f4,axiom,(
  ( ! [X0] : (subclass(X0,universal_class)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f3864,plain,(
  ( ! [X6,X5] : (member(not_subclass_element(X5,complement(X6)),X6) | subclass(X5,complement(X6)) | ~subclass(X5,universal_class)) )),
  inference(duplicate_literal_removal,[],[f3861])).
fof(f3861,plain,(
  ( ! [X6,X5] : (member(not_subclass_element(X5,complement(X6)),X6) | subclass(X5,complement(X6)) | ~subclass(X5,universal_class) | subclass(X5,complement(X6))) )),
  inference(resolution,[],[f311,f196])).
fof(f196,plain,(
  ( ! [X10,X8,X9] : (member(not_subclass_element(X8,X9),X10) | ~subclass(X8,X10) | subclass(X8,X9)) )),
  inference(resolution,[],[f1,f2])).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (~member(X2,X0) | member(X2,X1) | ~subclass(X0,X1)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f311,plain,(
  ( ! [X6,X5] : (~member(not_subclass_element(X5,complement(X6)),universal_class) | member(not_subclass_element(X5,complement(X6)),X6) | subclass(X5,complement(X6))) )),
  inference(resolution,[],[f25,f3])).
fof(f3,axiom,(
  ( ! [X0,X1] : (~member(not_subclass_element(X0,X1),X1) | subclass(X0,X1)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f25,axiom,(
  ( ! [X4,X0] : (member(X4,complement(X0)) | member(X4,X0) | ~member(X4,universal_class)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f330661,plain,(
  ( ! [X2] : (complement(complement(X2)) = X2 | ~subclass(X2,complement(complement(X2)))) )),
  inference(resolution,[],[f326839,f7])).
fof(f7,axiom,(
  ( ! [X0,X1] : (~subclass(X1,X0) | X0 = X1 | ~subclass(X0,X1)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f326839,plain,(
  ( ! [X0] : (subclass(complement(complement(X0)),X0)) )),
  inference(duplicate_literal_removal,[],[f326758])).
fof(f326758,plain,(
  ( ! [X0] : (subclass(complement(complement(X0)),X0) | subclass(complement(complement(X0)),X0)) )),
  inference(resolution,[],[f34271,f3])).
fof(f34271,plain,(
  ( ! [X2,X3] : (member(not_subclass_element(complement(complement(X2)),X3),X2) | subclass(complement(complement(X2)),X3)) )),
  inference(subsumption_resolution,[],[f34266,f4])).
fof(f34266,plain,(
  ( ! [X2,X3] : (member(not_subclass_element(complement(complement(X2)),X3),X2) | subclass(complement(complement(X2)),X3) | ~subclass(complement(complement(X2)),universal_class)) )),
  inference(duplicate_literal_removal,[],[f34262])).
fof(f34262,plain,(
  ( ! [X2,X3] : (member(not_subclass_element(complement(complement(X2)),X3),X2) | subclass(complement(complement(X2)),X3) | ~subclass(complement(complement(X2)),universal_class) | subclass(complement(complement(X2)),X3)) )),
  inference(resolution,[],[f1075,f196])).
fof(f1075,plain,(
  ( ! [X4,X5] : (~member(not_subclass_element(complement(complement(X4)),X5),universal_class) | member(not_subclass_element(complement(complement(X4)),X5),X4) | subclass(complement(complement(X4)),X5)) )),
  inference(resolution,[],[f183,f25])).
fof(f183,plain,(
  ( ! [X0,X1] : (~member(not_subclass_element(complement(X0),X1),X0) | subclass(complement(X0),X1)) )),
  inference(resolution,[],[f2,f24])).
fof(f238282,plain,(
  ~member(x,complement(complement(regular(x)))) | spl0_3543),
  inference(avatar_component_clause,[],[f238281])).
fof(f238281,plain,(
  spl0_3543 <=> member(x,complement(complement(regular(x))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3543])])).
fof(f503210,plain,(
  member(x,regular(x)) | ~member(x,universal_class) | (~spl0_139 | ~spl0_239)),
  inference(superposition,[],[f9,f501850])).
fof(f501850,plain,(
  regular(x) = unordered_pair(x,regular(regular(x))) | (~spl0_139 | ~spl0_239)),
  inference(backward_demodulation,[],[f4396,f1589])).
fof(f1589,plain,(
  unordered_pair(x,unordered_pair(y,y)) = regular(x) | ~spl0_139),
  inference(avatar_component_clause,[],[f1587])).
fof(f1587,plain,(
  spl0_139 <=> unordered_pair(x,unordered_pair(y,y)) = regular(x)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_139])])).
fof(f4396,plain,(
  unordered_pair(y,y) = regular(regular(x)) | ~spl0_239),
  inference(avatar_component_clause,[],[f4394])).
fof(f4394,plain,(
  spl0_239 <=> unordered_pair(y,y) = regular(regular(x))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_239])])).
fof(f9,axiom,(
  ( ! [X0,X1] : (member(X0,unordered_pair(X0,X1)) | ~member(X0,universal_class)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f501821,plain,(
  spl0_5709 | ~spl0_139),
  inference(avatar_split_clause,[],[f500684,f1587,f422626])).
fof(f422626,plain,(
  spl0_5709 <=> member(regular(x),x)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5709])])).
fof(f500684,plain,(
  member(regular(x),x) | ~spl0_139),
  inference(forward_demodulation,[],[f328,f1589])).
fof(f328,plain,(
  member(unordered_pair(x,unordered_pair(y,y)),x)),
  inference(subsumption_resolution,[],[f325,f11])).
fof(f325,plain,(
  member(unordered_pair(x,unordered_pair(y,y)),x) | ~member(unordered_pair(x,unordered_pair(y,y)),universal_class)),
  inference(superposition,[],[f10,f176])).
fof(f10,axiom,(
  ( ! [X0,X1] : (member(X1,unordered_pair(X0,X1)) | ~member(X1,universal_class)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f501499,plain,(
  spl0_105 | ~spl0_139 | ~spl0_249),
  inference(avatar_contradiction_clause,[],[f501498])).
fof(f501498,plain,(
  $false | (spl0_105 | ~spl0_139 | ~spl0_249)),
  inference(subsumption_resolution,[],[f500961,f1107])).
fof(f1107,plain,(
  null_class != unordered_pair(x,unordered_pair(y,y)) | spl0_105),
  inference(avatar_component_clause,[],[f1106])).
fof(f1106,plain,(
  spl0_105 <=> null_class = unordered_pair(x,unordered_pair(y,y))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_105])])).
fof(f500961,plain,(
  null_class = unordered_pair(x,unordered_pair(y,y)) | (~spl0_139 | ~spl0_249)),
  inference(backward_demodulation,[],[f5001,f1589])).
fof(f5001,plain,(
  null_class = regular(x) | ~spl0_249),
  inference(avatar_component_clause,[],[f4999])).
fof(f4999,plain,(
  spl0_249 <=> null_class = regular(x)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_249])])).
fof(f500718,plain,(
  spl0_249 | ~spl0_238 | spl0_3543),
  inference(avatar_split_clause,[],[f500551,f238281,f4390,f4999])).
fof(f4390,plain,(
  spl0_238 <=> x = regular(regular(x))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_238])])).
fof(f500551,plain,(
  null_class = regular(x) | (~spl0_238 | spl0_3543)),
  inference(subsumption_resolution,[],[f500532,f499392])).
fof(f499392,plain,(
  ~member(x,regular(x)) | spl0_3543),
  inference(forward_demodulation,[],[f238282,f330707])).
fof(f500532,plain,(
  member(x,regular(x)) | null_class = regular(x) | ~spl0_238),
  inference(superposition,[],[f66,f4392])).
fof(f4392,plain,(
  x = regular(regular(x)) | ~spl0_238),
  inference(avatar_component_clause,[],[f4390])).
fof(f66,axiom,(
  ( ! [X0] : (member(regular(X0),X0) | null_class = X0) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f499383,plain,(
  ~spl0_5709 | spl0_3568),
  inference(avatar_split_clause,[],[f495449,f238496,f422626])).
fof(f238496,plain,(
  spl0_3568 <=> member(regular(x),complement(complement(x)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3568])])).
fof(f495449,plain,(
  ~member(regular(x),x) | spl0_3568),
  inference(forward_demodulation,[],[f238497,f330707])).
fof(f238497,plain,(
  ~member(regular(x),complement(complement(x))) | spl0_3568),
  inference(avatar_component_clause,[],[f238496])).
fof(f495176,plain,(
  ~spl0_178 | spl0_277 | ~spl0_5043 | spl0_5044 | ~spl0_5709),
  inference(avatar_contradiction_clause,[],[f495175])).
fof(f495175,plain,(
  $false | (~spl0_178 | spl0_277 | ~spl0_5043 | spl0_5044 | ~spl0_5709)),
  inference(subsumption_resolution,[],[f495174,f5740])).
fof(f5740,plain,(
  x != regular(x) | spl0_277),
  inference(avatar_component_clause,[],[f5739])).
fof(f5739,plain,(
  spl0_277 <=> x = regular(x)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_277])])).
fof(f495174,plain,(
  x = regular(x) | (~spl0_178 | ~spl0_5043 | spl0_5044 | ~spl0_5709)),
  inference(backward_demodulation,[],[f385116,f445569])).
fof(f445569,plain,(
  regular(x) = regular(unordered_pair(x,regular(x))) | (~spl0_178 | spl0_5044 | ~spl0_5709)),
  inference(subsumption_resolution,[],[f445565,f385119])).
fof(f385119,plain,(
  null_class != unordered_pair(x,regular(x)) | spl0_5044),
  inference(avatar_component_clause,[],[f385118])).
fof(f385118,plain,(
  spl0_5044 <=> null_class = unordered_pair(x,regular(x))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5044])])).
fof(f445565,plain,(
  null_class = unordered_pair(x,regular(x)) | regular(x) = regular(unordered_pair(x,regular(x))) | (~spl0_178 | ~spl0_5709)),
  inference(resolution,[],[f422628,f34491])).
fof(f34491,plain,(
  ( ! [X6,X7] : (~member(X6,X7) | null_class = unordered_pair(X7,X6) | regular(unordered_pair(X7,X6)) = X6) ) | ~spl0_178),
  inference(subsumption_resolution,[],[f34435,f1524])).
fof(f1524,plain,(
  ( ! [X6,X8,X7] : (~member(X6,X8) | member(X6,universal_class) | ~member(X6,X7)) )),
  inference(resolution,[],[f339,f4])).
fof(f339,plain,(
  ( ! [X6,X8,X7,X9] : (~subclass(intersection(X8,X7),X9) | ~member(X6,X8) | member(X6,X9) | ~member(X6,X7)) )),
  inference(resolution,[],[f23,f1])).
fof(f23,axiom,(
  ( ! [X4,X0,X1] : (member(X4,intersection(X0,X1)) | ~member(X4,X1) | ~member(X4,X0)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f34435,plain,(
  ( ! [X6,X7] : (~member(X6,X7) | null_class = unordered_pair(X7,X6) | regular(unordered_pair(X7,X6)) = X6 | ~member(X6,universal_class)) ) | ~spl0_178),
  inference(resolution,[],[f3912,f10])).
fof(f3912,plain,(
  ( ! [X2,X0,X1] : (~member(X2,unordered_pair(X0,X1)) | ~member(X2,X0) | unordered_pair(X0,X1) = null_class | regular(unordered_pair(X0,X1)) = X1) ) | ~spl0_178),
  inference(subsumption_resolution,[],[f3902,f3247])).
fof(f3247,plain,(
  ( ! [X16] : (~member(X16,null_class)) ) | ~spl0_178),
  inference(backward_demodulation,[],[f3240,f3119])).
fof(f3119,plain,(
  ( ! [X16] : (~member(X16,domain_of(null_class))) )),
  inference(trivial_inequality_removal,[],[f3113])).
fof(f3113,plain,(
  ( ! [X16] : (null_class != null_class | ~member(X16,domain_of(null_class))) )),
  inference(superposition,[],[f134,f3089])).
fof(f3089,plain,(
  ( ! [X0] : (null_class = intersection(null_class,X0)) )),
  inference(resolution,[],[f1418,f1337])).
fof(f1337,plain,(
  ( ! [X0,X1] : (subclass(intersection(X0,X1),X0)) )),
  inference(duplicate_literal_removal,[],[f1321])).
fof(f1321,plain,(
  ( ! [X0,X1] : (subclass(intersection(X0,X1),X0) | subclass(intersection(X0,X1),X0)) )),
  inference(resolution,[],[f188,f3])).
fof(f188,plain,(
  ( ! [X6,X4,X5] : (member(not_subclass_element(intersection(X4,X5),X6),X4) | subclass(intersection(X4,X5),X6)) )),
  inference(resolution,[],[f21,f2])).
fof(f21,axiom,(
  ( ! [X4,X0,X1] : (~member(X4,intersection(X0,X1)) | member(X4,X0)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f1418,plain,(
  ( ! [X1] : (~subclass(X1,null_class) | null_class = X1) )),
  inference(duplicate_literal_removal,[],[f1414])).
fof(f1414,plain,(
  ( ! [X1] : (null_class = X1 | null_class = X1 | ~subclass(X1,null_class)) )),
  inference(resolution,[],[f1345,f7])).
fof(f1345,plain,(
  ( ! [X0] : (subclass(null_class,X0) | null_class = X0) )),
  inference(superposition,[],[f1337,f67])).
fof(f67,axiom,(
  ( ! [X0] : (null_class = intersection(X0,regular(X0)) | null_class = X0) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f134,plain,(
  ( ! [X4,X0] : (null_class != intersection(X0,cross_product(unordered_pair(X4,X4),universal_class)) | ~member(X4,domain_of(X0))) )),
  inference(definition_unfolding,[],[f30,f28,f12])).
fof(f28,axiom,(
  ( ! [X0,X5,X1] : (intersection(X5,cross_product(X0,X1)) = restrict(X5,X0,X1)) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f30,axiom,(
  ( ! [X4,X0] : (~member(X4,domain_of(X0)) | restrict(X0,singleton(X4),universal_class) != null_class) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f3240,plain,(
  null_class = domain_of(null_class) | ~spl0_178),
  inference(avatar_component_clause,[],[f3238])).
fof(f3238,plain,(
  spl0_178 <=> null_class = domain_of(null_class)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_178])])).
fof(f3902,plain,(
  ( ! [X2,X0,X1] : (~member(X2,X0) | member(X2,null_class) | ~member(X2,unordered_pair(X0,X1)) | unordered_pair(X0,X1) = null_class | regular(unordered_pair(X0,X1)) = X1) )),
  inference(duplicate_literal_removal,[],[f3889])).
fof(f3889,plain,(
  ( ! [X2,X0,X1] : (~member(X2,X0) | member(X2,null_class) | ~member(X2,unordered_pair(X0,X1)) | unordered_pair(X0,X1) = null_class | regular(unordered_pair(X0,X1)) = X1 | unordered_pair(X0,X1) = null_class) )),
  inference(superposition,[],[f341,f333])).
fof(f333,plain,(
  ( ! [X6,X7] : (regular(unordered_pair(X6,X7)) = X7 | regular(unordered_pair(X6,X7)) = X6 | null_class = unordered_pair(X6,X7)) )),
  inference(resolution,[],[f8,f66])).
fof(f8,axiom,(
  ( ! [X2,X0,X1] : (~member(X2,unordered_pair(X0,X1)) | X0 = X2 | X1 = X2) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f341,plain,(
  ( ! [X0,X1] : (~member(X1,regular(X0)) | member(X1,null_class) | ~member(X1,X0) | null_class = X0) )),
  inference(superposition,[],[f23,f67])).
fof(f422628,plain,(
  member(regular(x),x) | ~spl0_5709),
  inference(avatar_component_clause,[],[f422626])).
fof(f385116,plain,(
  x = regular(unordered_pair(x,regular(x))) | ~spl0_5043),
  inference(avatar_component_clause,[],[f385114])).
fof(f385114,plain,(
  spl0_5043 <=> x = regular(unordered_pair(x,regular(x)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5043])])).
fof(f443583,plain,(
  ~spl0_178 | ~spl0_3543 | spl0_5043 | spl0_5044),
  inference(avatar_contradiction_clause,[],[f443582])).
fof(f443582,plain,(
  $false | (~spl0_178 | ~spl0_3543 | spl0_5043 | spl0_5044)),
  inference(subsumption_resolution,[],[f443581,f385115])).
fof(f385115,plain,(
  x != regular(unordered_pair(x,regular(x))) | spl0_5043),
  inference(avatar_component_clause,[],[f385114])).
fof(f443581,plain,(
  x = regular(unordered_pair(x,regular(x))) | (~spl0_178 | ~spl0_3543 | spl0_5044)),
  inference(subsumption_resolution,[],[f443571,f385119])).
fof(f443571,plain,(
  null_class = unordered_pair(x,regular(x)) | x = regular(unordered_pair(x,regular(x))) | (~spl0_178 | ~spl0_3543)),
  inference(resolution,[],[f434537,f34428])).
fof(f34428,plain,(
  ( ! [X2,X3] : (~member(X2,X3) | null_class = unordered_pair(X2,X3) | regular(unordered_pair(X2,X3)) = X2) ) | ~spl0_178),
  inference(subsumption_resolution,[],[f34370,f1524])).
fof(f34370,plain,(
  ( ! [X2,X3] : (~member(X2,X3) | null_class = unordered_pair(X2,X3) | regular(unordered_pair(X2,X3)) = X2 | ~member(X2,universal_class)) ) | ~spl0_178),
  inference(resolution,[],[f3911,f9])).
fof(f3911,plain,(
  ( ! [X2,X0,X1] : (~member(X2,unordered_pair(X0,X1)) | ~member(X2,X1) | unordered_pair(X0,X1) = null_class | regular(unordered_pair(X0,X1)) = X0) ) | ~spl0_178),
  inference(subsumption_resolution,[],[f3908,f3247])).
fof(f3908,plain,(
  ( ! [X2,X0,X1] : (~member(X2,X1) | member(X2,null_class) | ~member(X2,unordered_pair(X0,X1)) | unordered_pair(X0,X1) = null_class | regular(unordered_pair(X0,X1)) = X0) )),
  inference(duplicate_literal_removal,[],[f3883])).
fof(f3883,plain,(
  ( ! [X2,X0,X1] : (~member(X2,X1) | member(X2,null_class) | ~member(X2,unordered_pair(X0,X1)) | unordered_pair(X0,X1) = null_class | regular(unordered_pair(X0,X1)) = X0 | unordered_pair(X0,X1) = null_class) )),
  inference(superposition,[],[f341,f333])).
fof(f434537,plain,(
  member(x,regular(x)) | ~spl0_3543),
  inference(forward_demodulation,[],[f238283,f330707])).
fof(f238283,plain,(
  member(x,complement(complement(regular(x)))) | ~spl0_3543),
  inference(avatar_component_clause,[],[f238281])).
fof(f434606,plain,(
  ~spl0_278 | spl0_136 | ~spl0_178 | ~spl0_3568),
  inference(avatar_split_clause,[],[f428908,f238496,f3238,f1572,f5744])).
fof(f5744,plain,(
  spl0_278 <=> subclass(x,regular(x))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_278])])).
fof(f1572,plain,(
  spl0_136 <=> null_class = x),
  introduced(avatar_definition,[new_symbols(naming,[spl0_136])])).
fof(f428908,plain,(
  ~subclass(x,regular(x)) | (spl0_136 | ~spl0_178 | ~spl0_3568)),
  inference(forward_demodulation,[],[f428907,f330707])).
fof(f428907,plain,(
  ~subclass(x,regular(complement(complement(x)))) | (spl0_136 | ~spl0_178 | ~spl0_3568)),
  inference(subsumption_resolution,[],[f428906,f1573])).
fof(f1573,plain,(
  null_class != x | spl0_136),
  inference(avatar_component_clause,[],[f1572])).
fof(f428906,plain,(
  null_class = x | ~subclass(x,regular(complement(complement(x)))) | (spl0_136 | ~spl0_178 | ~spl0_3568)),
  inference(forward_demodulation,[],[f318416,f330707])).
fof(f318416,plain,(
  null_class = complement(complement(x)) | ~subclass(x,regular(complement(complement(x)))) | (spl0_136 | ~spl0_178 | ~spl0_3568)),
  inference(subsumption_resolution,[],[f318400,f1573])).
fof(f318400,plain,(
  null_class = complement(complement(x)) | ~subclass(x,regular(complement(complement(x)))) | null_class = x | (~spl0_178 | ~spl0_3568)),
  inference(resolution,[],[f238498,f8840])).
fof(f8840,plain,(
  ( ! [X14,X13] : (~member(regular(X13),X14) | null_class = X14 | ~subclass(X13,regular(X14)) | null_class = X13) ) | ~spl0_178),
  inference(subsumption_resolution,[],[f1427,f3247])).
fof(f1427,plain,(
  ( ! [X14,X13] : (member(regular(X13),null_class) | ~member(regular(X13),X14) | null_class = X14 | ~subclass(X13,regular(X14)) | null_class = X13) )),
  inference(resolution,[],[f341,f195])).
fof(f195,plain,(
  ( ! [X6,X7] : (member(regular(X6),X7) | ~subclass(X6,X7) | null_class = X6) )),
  inference(resolution,[],[f1,f66])).
fof(f238498,plain,(
  member(regular(x),complement(complement(x))) | ~spl0_3568),
  inference(avatar_component_clause,[],[f238496])).
fof(f420012,plain,(
  ~spl0_178 | ~spl0_5044),
  inference(avatar_contradiction_clause,[],[f420011])).
fof(f420011,plain,(
  $false | (~spl0_178 | ~spl0_5044)),
  inference(subsumption_resolution,[],[f420010,f327])).
fof(f420010,plain,(
  ~member(x,universal_class) | (~spl0_178 | ~spl0_5044)),
  inference(subsumption_resolution,[],[f419904,f3247])).
fof(f419904,plain,(
  member(x,null_class) | ~member(x,universal_class) | ~spl0_5044),
  inference(superposition,[],[f9,f385120])).
fof(f385120,plain,(
  null_class = unordered_pair(x,regular(x)) | ~spl0_5044),
  inference(avatar_component_clause,[],[f385118])).
fof(f397525,plain,(
  spl0_5044 | ~spl0_140 | ~spl0_178 | spl0_277 | ~spl0_5043),
  inference(avatar_split_clause,[],[f397524,f385114,f5739,f3238,f1591,f385118])).
fof(f1591,plain,(
  spl0_140 <=> unordered_pair(x,x) = regular(x)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_140])])).
fof(f397524,plain,(
  null_class = unordered_pair(x,regular(x)) | (~spl0_140 | ~spl0_178 | spl0_277 | ~spl0_5043)),
  inference(subsumption_resolution,[],[f397523,f5740])).
fof(f397523,plain,(
  x = regular(x) | null_class = unordered_pair(x,regular(x)) | (~spl0_140 | ~spl0_178 | ~spl0_5043)),
  inference(forward_demodulation,[],[f397402,f385116])).
fof(f397402,plain,(
  null_class = unordered_pair(x,regular(x)) | regular(x) = regular(unordered_pair(x,regular(x))) | (~spl0_140 | ~spl0_178)),
  inference(resolution,[],[f34491,f4471])).
fof(f4471,plain,(
  member(regular(x),x) | ~spl0_140),
  inference(backward_demodulation,[],[f1593,f329])).
fof(f329,plain,(
  member(unordered_pair(x,x),x)),
  inference(subsumption_resolution,[],[f326,f11])).
fof(f326,plain,(
  member(unordered_pair(x,x),x) | ~member(unordered_pair(x,x),universal_class)),
  inference(superposition,[],[f9,f176])).
fof(f1593,plain,(
  unordered_pair(x,x) = regular(x) | ~spl0_140),
  inference(avatar_component_clause,[],[f1591])).
fof(f385121,plain,(
  spl0_5043 | spl0_5044 | ~spl0_140 | ~spl0_178),
  inference(avatar_split_clause,[],[f384919,f3238,f1591,f385118,f385114])).
fof(f384919,plain,(
  null_class = unordered_pair(x,regular(x)) | x = regular(unordered_pair(x,regular(x))) | (~spl0_140 | ~spl0_178)),
  inference(resolution,[],[f34428,f4713])).
fof(f4713,plain,(
  member(x,regular(x)) | ~spl0_140),
  inference(subsumption_resolution,[],[f4639,f327])).
fof(f4639,plain,(
  member(x,regular(x)) | ~member(x,universal_class) | ~spl0_140),
  inference(superposition,[],[f9,f1593])).
fof(f285450,plain,(
  ~spl0_277 | spl0_278),
  inference(avatar_contradiction_clause,[],[f285449])).
fof(f285449,plain,(
  $false | (~spl0_277 | spl0_278)),
  inference(subsumption_resolution,[],[f283491,f177])).
fof(f177,plain,(
  ( ! [X1] : (subclass(X1,X1)) )),
  inference(equality_resolution,[],[f5])).
fof(f5,axiom,(
  ( ! [X0,X1] : (subclass(X0,X1) | X0 != X1) )),
  file('./p/SET519-6.tptp',unknown)).
fof(f283491,plain,(
  ~subclass(x,x) | (~spl0_277 | spl0_278)),
  inference(backward_demodulation,[],[f5741,f5746])).
fof(f5746,plain,(
  ~subclass(x,regular(x)) | spl0_278),
  inference(avatar_component_clause,[],[f5744])).
fof(f5741,plain,(
  x = regular(x) | ~spl0_277),
  inference(avatar_component_clause,[],[f5739])).
fof(f285437,plain,(
  spl0_136 | ~spl0_140 | ~spl0_178 | ~spl0_277),
  inference(avatar_contradiction_clause,[],[f285436])).
fof(f285436,plain,(
  $false | (spl0_136 | ~spl0_140 | ~spl0_178 | ~spl0_277)),
  inference(subsumption_resolution,[],[f283464,f4968])).
fof(f4968,plain,(
  ~member(x,x) | (spl0_136 | ~spl0_140 | ~spl0_178)),
  inference(subsumption_resolution,[],[f4967,f1573])).
fof(f4967,plain,(
  ~member(x,x) | null_class = x | (~spl0_140 | ~spl0_178)),
  inference(subsumption_resolution,[],[f4965,f3247])).
fof(f4965,plain,(
  member(x,null_class) | ~member(x,x) | null_class = x | ~spl0_140),
  inference(resolution,[],[f4713,f341])).
fof(f283464,plain,(
  member(x,x) | (~spl0_140 | ~spl0_277)),
  inference(backward_demodulation,[],[f5741,f4713])).
fof(f4439,plain,(
  ~spl0_105 | ~spl0_139 | ~spl0_178),
  inference(avatar_contradiction_clause,[],[f4438])).
fof(f4438,plain,(
  $false | (~spl0_105 | ~spl0_139 | ~spl0_178)),
  inference(subsumption_resolution,[],[f4427,f3247])).
fof(f4427,plain,(
  member(unordered_pair(y,y),null_class) | (~spl0_105 | ~spl0_139)),
  inference(backward_demodulation,[],[f4402,f4375])).
fof(f4375,plain,(
  member(unordered_pair(y,y),regular(x)) | ~spl0_139),
  inference(subsumption_resolution,[],[f4340,f11])).
fof(f4340,plain,(
  member(unordered_pair(y,y),regular(x)) | ~member(unordered_pair(y,y),universal_class) | ~spl0_139),
  inference(superposition,[],[f10,f1589])).
fof(f4402,plain,(
  null_class = regular(x) | (~spl0_105 | ~spl0_139)),
  inference(backward_demodulation,[],[f1108,f1589])).
fof(f1108,plain,(
  null_class = unordered_pair(x,unordered_pair(y,y)) | ~spl0_105),
  inference(avatar_component_clause,[],[f1106])).
fof(f4399,plain,(
  spl0_105 | spl0_239 | spl0_238 | ~spl0_139),
  inference(avatar_split_clause,[],[f4398,f1587,f4390,f4394,f1106])).
fof(f4398,plain,(
  x = regular(regular(x)) | unordered_pair(y,y) = regular(regular(x)) | null_class = unordered_pair(x,unordered_pair(y,y)) | ~spl0_139),
  inference(forward_demodulation,[],[f4347,f1589])).
fof(f4347,plain,(
  unordered_pair(y,y) = regular(regular(x)) | x = regular(unordered_pair(x,unordered_pair(y,y))) | null_class = unordered_pair(x,unordered_pair(y,y)) | ~spl0_139),
  inference(superposition,[],[f333,f1589])).
fof(f3242,plain,(
  spl0_178),
  inference(avatar_split_clause,[],[f3230,f3238])).
fof(f3230,plain,(
  null_class = domain_of(null_class)),
  inference(resolution,[],[f3119,f66])).
fof(f2896,plain,(
  ~spl0_136),
  inference(avatar_contradiction_clause,[],[f2895])).
fof(f2895,plain,(
  $false | ~spl0_136),
  inference(subsumption_resolution,[],[f2887,f11])).
fof(f2887,plain,(
  ~member(unordered_pair(null_class,null_class),universal_class) | ~spl0_136),
  inference(resolution,[],[f1285,f1610])).
fof(f1610,plain,(
  member(unordered_pair(null_class,null_class),null_class) | ~spl0_136),
  inference(backward_demodulation,[],[f1574,f329])).
fof(f1574,plain,(
  null_class = x | ~spl0_136),
  inference(avatar_component_clause,[],[f1572])).
fof(f1285,plain,(
  ( ! [X0] : (~member(X0,null_class) | ~member(X0,universal_class)) )),
  inference(superposition,[],[f24,f1265])).
fof(f1265,plain,(
  null_class = complement(universal_class)),
  inference(resolution,[],[f1090,f4])).
fof(f1090,plain,(
  ( ! [X0] : (~subclass(complement(X0),X0) | complement(X0) = null_class) )),
  inference(duplicate_literal_removal,[],[f1079])).
fof(f1079,plain,(
  ( ! [X0] : (~subclass(complement(X0),X0) | complement(X0) = null_class | complement(X0) = null_class) )),
  inference(resolution,[],[f195,f182])).
fof(f182,plain,(
  ( ! [X0] : (~member(regular(complement(X0)),X0) | complement(X0) = null_class) )),
  inference(resolution,[],[f66,f24])).
fof(f1594,plain,(
  spl0_136 | spl0_139 | spl0_140),
  inference(avatar_split_clause,[],[f1563,f1591,f1587,f1572])).
fof(f1563,plain,(
  unordered_pair(x,x) = regular(x) | unordered_pair(x,unordered_pair(y,y)) = regular(x) | null_class = x),
  inference(resolution,[],[f335,f66])).
fof(f335,plain,(
  ( ! [X0] : (~member(X0,x) | unordered_pair(x,x) = X0 | unordered_pair(x,unordered_pair(y,y)) = X0) )),
  inference(superposition,[],[f8,f176])).
% SZS output end Proof for SET519-6
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 353981
% Time elapsed: 40.069 s
% ------------------------------
% ------------------------------
