% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SET515-6
% SZS output start Proof for SET515-6
fof(f678594,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f1530,f201354,f270091,f678593])).
fof(f678593,plain,(
  ~spl0_125 | ~spl0_3748 | spl0_3749),
  inference(avatar_contradiction_clause,[],[f678592])).
fof(f678592,plain,(
  $false | (~spl0_125 | ~spl0_3748 | spl0_3749)),
  inference(subsumption_resolution,[],[f678588,f675399])).
fof(f675399,plain,(
  member(x,unordered_pair(x,x)) | (~spl0_3748 | spl0_3749)),
  inference(subsumption_resolution,[],[f675386,f624690])).
fof(f624690,plain,(
  null_class != unordered_pair(x,x) | spl0_3749),
  inference(backward_demodulation,[],[f624410,f201352])).
fof(f201352,plain,(
  null_class != unordered_pair(x,complement(complement(x))) | spl0_3749),
  inference(avatar_component_clause,[],[f201351])).
fof(f201351,plain,(
  spl0_3749 <=> null_class = unordered_pair(x,complement(complement(x)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3749])])).
fof(f624410,plain,(
  ( ! [X2] : (complement(complement(X2)) = X2) )),
  inference(subsumption_resolution,[],[f624372,f117388])).
fof(f117388,plain,(
  ( ! [X0] : (subclass(X0,complement(complement(X0)))) )),
  inference(duplicate_literal_removal,[],[f117353])).
fof(f117353,plain,(
  ( ! [X0] : (subclass(X0,complement(complement(X0))) | subclass(X0,complement(complement(X0)))) )),
  inference(resolution,[],[f4304,f2])).
fof(f2,axiom,(
  ( ! [X0,X1] : (member(not_subclass_element(X0,X1),X0) | subclass(X0,X1)) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f4304,plain,(
  ( ! [X21,X20] : (~member(not_subclass_element(X20,complement(complement(X21))),X21) | subclass(X20,complement(complement(X21)))) )),
  inference(resolution,[],[f1550,f24])).
fof(f24,axiom,(
  ( ! [X4,X0] : (~member(X4,complement(X0)) | ~member(X4,X0)) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f1550,plain,(
  ( ! [X6,X5] : (member(not_subclass_element(X5,complement(X6)),X6) | subclass(X5,complement(X6))) )),
  inference(subsumption_resolution,[],[f1546,f4])).
fof(f4,axiom,(
  ( ! [X0] : (subclass(X0,universal_class)) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f1546,plain,(
  ( ! [X6,X5] : (member(not_subclass_element(X5,complement(X6)),X6) | subclass(X5,complement(X6)) | ~subclass(X5,universal_class)) )),
  inference(duplicate_literal_removal,[],[f1544])).
fof(f1544,plain,(
  ( ! [X6,X5] : (member(not_subclass_element(X5,complement(X6)),X6) | subclass(X5,complement(X6)) | ~subclass(X5,universal_class) | subclass(X5,complement(X6))) )),
  inference(resolution,[],[f323,f193])).
fof(f193,plain,(
  ( ! [X6,X8,X7] : (member(not_subclass_element(X6,X7),X8) | ~subclass(X6,X8) | subclass(X6,X7)) )),
  inference(resolution,[],[f1,f2])).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (~member(X2,X0) | member(X2,X1) | ~subclass(X0,X1)) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f323,plain,(
  ( ! [X6,X5] : (~member(not_subclass_element(X5,complement(X6)),universal_class) | member(not_subclass_element(X5,complement(X6)),X6) | subclass(X5,complement(X6))) )),
  inference(resolution,[],[f25,f3])).
fof(f3,axiom,(
  ( ! [X0,X1] : (~member(not_subclass_element(X0,X1),X1) | subclass(X0,X1)) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f25,axiom,(
  ( ! [X4,X0] : (member(X4,complement(X0)) | member(X4,X0) | ~member(X4,universal_class)) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f624372,plain,(
  ( ! [X2] : (complement(complement(X2)) = X2 | ~subclass(X2,complement(complement(X2)))) )),
  inference(resolution,[],[f622747,f7])).
fof(f7,axiom,(
  ( ! [X0,X1] : (~subclass(X1,X0) | X0 = X1 | ~subclass(X0,X1)) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f622747,plain,(
  ( ! [X5] : (subclass(complement(complement(X5)),X5)) )),
  inference(superposition,[],[f585191,f5464])).
fof(f5464,plain,(
  ( ! [X0] : (intersection(X0,X0) = X0) )),
  inference(resolution,[],[f5461,f1306])).
fof(f1306,plain,(
  ( ! [X2,X3] : (~subclass(X2,intersection(X2,X3)) | intersection(X2,X3) = X2) )),
  inference(resolution,[],[f1276,f7])).
fof(f1276,plain,(
  ( ! [X0,X1] : (subclass(intersection(X0,X1),X0)) )),
  inference(duplicate_literal_removal,[],[f1259])).
fof(f1259,plain,(
  ( ! [X0,X1] : (subclass(intersection(X0,X1),X0) | subclass(intersection(X0,X1),X0)) )),
  inference(resolution,[],[f185,f3])).
fof(f185,plain,(
  ( ! [X2,X0,X1] : (member(not_subclass_element(intersection(X0,X1),X2),X0) | subclass(intersection(X0,X1),X2)) )),
  inference(resolution,[],[f21,f2])).
fof(f21,axiom,(
  ( ! [X4,X0,X1] : (~member(X4,intersection(X0,X1)) | member(X4,X0)) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f5461,plain,(
  ( ! [X0] : (subclass(X0,intersection(X0,X0))) )),
  inference(duplicate_literal_removal,[],[f5444])).
fof(f5444,plain,(
  ( ! [X0] : (subclass(X0,intersection(X0,X0)) | subclass(X0,intersection(X0,X0))) )),
  inference(resolution,[],[f2356,f2])).
fof(f2356,plain,(
  ( ! [X0,X1] : (~member(not_subclass_element(X0,intersection(X1,X0)),X1) | subclass(X0,intersection(X1,X0))) )),
  inference(duplicate_literal_removal,[],[f2338])).
fof(f2338,plain,(
  ( ! [X0,X1] : (~member(not_subclass_element(X0,intersection(X1,X0)),X1) | subclass(X0,intersection(X1,X0)) | subclass(X0,intersection(X1,X0))) )),
  inference(resolution,[],[f334,f2])).
fof(f334,plain,(
  ( ! [X12,X10,X11] : (~member(not_subclass_element(X10,intersection(X11,X12)),X12) | ~member(not_subclass_element(X10,intersection(X11,X12)),X11) | subclass(X10,intersection(X11,X12))) )),
  inference(resolution,[],[f23,f3])).
fof(f23,axiom,(
  ( ! [X4,X0,X1] : (member(X4,intersection(X0,X1)) | ~member(X4,X1) | ~member(X4,X0)) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f585191,plain,(
  ( ! [X0,X1] : (subclass(intersection(complement(complement(X0)),X1),X0)) )),
  inference(duplicate_literal_removal,[],[f585092])).
fof(f585092,plain,(
  ( ! [X0,X1] : (subclass(intersection(complement(complement(X0)),X1),X0) | subclass(intersection(complement(complement(X0)),X1),X0)) )),
  inference(resolution,[],[f17102,f3])).
fof(f17102,plain,(
  ( ! [X21,X19,X20] : (member(not_subclass_element(intersection(complement(complement(X19)),X20),X21),X19) | subclass(intersection(complement(complement(X19)),X20),X21)) )),
  inference(subsumption_resolution,[],[f17085,f11016])).
fof(f11016,plain,(
  ( ! [X6,X7] : (member(not_subclass_element(X6,X7),universal_class) | subclass(X6,X7)) )),
  inference(forward_demodulation,[],[f10985,f10950])).
fof(f10950,plain,(
  ( ! [X1] : (intersection(universal_class,X1) = X1) )),
  inference(resolution,[],[f8312,f4])).
fof(f8312,plain,(
  ( ! [X2,X3] : (~subclass(X2,X3) | intersection(X3,X2) = X2) )),
  inference(subsumption_resolution,[],[f8289,f1297])).
fof(f1297,plain,(
  ( ! [X0,X1] : (subclass(intersection(X0,X1),X1)) )),
  inference(duplicate_literal_removal,[],[f1280])).
fof(f1280,plain,(
  ( ! [X0,X1] : (subclass(intersection(X0,X1),X1) | subclass(intersection(X0,X1),X1)) )),
  inference(resolution,[],[f188,f3])).
fof(f188,plain,(
  ( ! [X2,X0,X1] : (member(not_subclass_element(intersection(X0,X1),X2),X1) | subclass(intersection(X0,X1),X2)) )),
  inference(resolution,[],[f22,f2])).
fof(f22,axiom,(
  ( ! [X4,X0,X1] : (~member(X4,intersection(X0,X1)) | member(X4,X1)) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f8289,plain,(
  ( ! [X2,X3] : (~subclass(X2,X3) | intersection(X3,X2) = X2 | ~subclass(intersection(X3,X2),X2)) )),
  inference(resolution,[],[f5460,f7])).
fof(f5460,plain,(
  ( ! [X2,X1] : (subclass(X1,intersection(X2,X1)) | ~subclass(X1,X2)) )),
  inference(duplicate_literal_removal,[],[f5445])).
fof(f5445,plain,(
  ( ! [X2,X1] : (subclass(X1,intersection(X2,X1)) | ~subclass(X1,X2) | subclass(X1,intersection(X2,X1))) )),
  inference(resolution,[],[f2356,f193])).
fof(f10985,plain,(
  ( ! [X6,X7] : (member(not_subclass_element(X6,X7),universal_class) | subclass(intersection(universal_class,X6),X7)) )),
  inference(superposition,[],[f185,f10950])).
fof(f17085,plain,(
  ( ! [X21,X19,X20] : (subclass(intersection(complement(complement(X19)),X20),X21) | member(not_subclass_element(intersection(complement(complement(X19)),X20),X21),X19) | ~member(not_subclass_element(intersection(complement(complement(X19)),X20),X21),universal_class)) )),
  inference(resolution,[],[f1265,f25])).
fof(f1265,plain,(
  ( ! [X24,X23,X22] : (~member(not_subclass_element(intersection(complement(X22),X23),X24),X22) | subclass(intersection(complement(X22),X23),X24)) )),
  inference(resolution,[],[f185,f24])).
fof(f675386,plain,(
  member(x,unordered_pair(x,x)) | null_class = unordered_pair(x,x) | ~spl0_3748),
  inference(superposition,[],[f66,f624687])).
fof(f624687,plain,(
  x = regular(unordered_pair(x,x)) | ~spl0_3748),
  inference(backward_demodulation,[],[f624410,f201349])).
fof(f201349,plain,(
  x = regular(unordered_pair(x,complement(complement(x)))) | ~spl0_3748),
  inference(avatar_component_clause,[],[f201347])).
fof(f201347,plain,(
  spl0_3748 <=> x = regular(unordered_pair(x,complement(complement(x))))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3748])])).
fof(f66,axiom,(
  ( ! [X0] : (member(regular(X0),X0) | null_class = X0) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f678588,plain,(
  ~member(x,unordered_pair(x,x)) | (~spl0_125 | ~spl0_3748 | spl0_3749)),
  inference(resolution,[],[f677294,f24])).
fof(f677294,plain,(
  member(x,complement(unordered_pair(x,x))) | (~spl0_125 | ~spl0_3748 | spl0_3749)),
  inference(resolution,[],[f675426,f198])).
fof(f198,plain,(
  ( ! [X17] : (~subclass(x,X17) | member(x,X17)) )),
  inference(resolution,[],[f1,f113])).
fof(f113,axiom,(
  member(x,x)),
  file('./p/SET515-6.tptp',unknown)).
fof(f675426,plain,(
  subclass(x,complement(unordered_pair(x,x))) | (~spl0_125 | ~spl0_3748 | spl0_3749)),
  inference(subsumption_resolution,[],[f675396,f624690])).
fof(f675396,plain,(
  subclass(x,complement(unordered_pair(x,x))) | null_class = unordered_pair(x,x) | (~spl0_125 | ~spl0_3748)),
  inference(superposition,[],[f11839,f624687])).
fof(f11839,plain,(
  ( ! [X2] : (subclass(regular(X2),complement(X2)) | null_class = X2) ) | ~spl0_125),
  inference(duplicate_literal_removal,[],[f11827])).
fof(f11827,plain,(
  ( ! [X2] : (null_class = X2 | subclass(regular(X2),complement(X2)) | subclass(regular(X2),complement(X2))) ) | ~spl0_125),
  inference(resolution,[],[f3477,f1550])).
fof(f3477,plain,(
  ( ! [X2,X3] : (~member(not_subclass_element(regular(X2),X3),X2) | null_class = X2 | subclass(regular(X2),X3)) ) | ~spl0_125),
  inference(subsumption_resolution,[],[f1334,f1533])).
fof(f1533,plain,(
  ( ! [X25] : (~member(X25,null_class)) ) | ~spl0_125),
  inference(backward_demodulation,[],[f1528,f1473])).
fof(f1473,plain,(
  ( ! [X25] : (~member(X25,domain_of(null_class))) )),
  inference(trivial_inequality_removal,[],[f1467])).
fof(f1467,plain,(
  ( ! [X25] : (null_class != null_class | ~member(X25,domain_of(null_class))) )),
  inference(superposition,[],[f134,f1366])).
fof(f1366,plain,(
  ( ! [X0] : (null_class = intersection(null_class,X0)) )),
  inference(resolution,[],[f1247,f1276])).
fof(f1247,plain,(
  ( ! [X1] : (~subclass(X1,null_class) | null_class = X1) )),
  inference(resolution,[],[f1245,f7])).
fof(f1245,plain,(
  ( ! [X4] : (subclass(null_class,X4)) )),
  inference(subsumption_resolution,[],[f1242,f4])).
fof(f1242,plain,(
  ( ! [X4] : (~subclass(null_class,universal_class) | subclass(null_class,X4)) )),
  inference(duplicate_literal_removal,[],[f1231])).
fof(f1231,plain,(
  ( ! [X4] : (~subclass(null_class,universal_class) | subclass(null_class,X4) | subclass(null_class,X4)) )),
  inference(resolution,[],[f193,f1037])).
fof(f1037,plain,(
  ( ! [X2] : (~member(not_subclass_element(null_class,X2),universal_class) | subclass(null_class,X2)) )),
  inference(forward_demodulation,[],[f1016,f989])).
fof(f989,plain,(
  null_class = complement(universal_class)),
  inference(resolution,[],[f979,f4])).
fof(f979,plain,(
  ( ! [X0] : (~subclass(complement(X0),X0) | complement(X0) = null_class) )),
  inference(duplicate_literal_removal,[],[f968])).
fof(f968,plain,(
  ( ! [X0] : (~subclass(complement(X0),X0) | complement(X0) = null_class | complement(X0) = null_class) )),
  inference(resolution,[],[f197,f181])).
fof(f181,plain,(
  ( ! [X0] : (~member(regular(complement(X0)),X0) | complement(X0) = null_class) )),
  inference(resolution,[],[f66,f24])).
fof(f197,plain,(
  ( ! [X15,X16] : (member(regular(X15),X16) | ~subclass(X15,X16) | null_class = X15) )),
  inference(resolution,[],[f1,f66])).
fof(f1016,plain,(
  ( ! [X2] : (~member(not_subclass_element(null_class,X2),universal_class) | subclass(complement(universal_class),X2)) )),
  inference(superposition,[],[f182,f989])).
fof(f182,plain,(
  ( ! [X0,X1] : (~member(not_subclass_element(complement(X0),X1),X0) | subclass(complement(X0),X1)) )),
  inference(resolution,[],[f2,f24])).
fof(f134,plain,(
  ( ! [X4,X0] : (null_class != intersection(X0,cross_product(unordered_pair(X4,X4),universal_class)) | ~member(X4,domain_of(X0))) )),
  inference(definition_unfolding,[],[f30,f28,f12])).
fof(f12,axiom,(
  ( ! [X0] : (unordered_pair(X0,X0) = singleton(X0)) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f28,axiom,(
  ( ! [X0,X5,X1] : (intersection(X5,cross_product(X0,X1)) = restrict(X5,X0,X1)) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f30,axiom,(
  ( ! [X4,X0] : (~member(X4,domain_of(X0)) | restrict(X0,singleton(X4),universal_class) != null_class) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f1528,plain,(
  null_class = domain_of(null_class) | ~spl0_125),
  inference(avatar_component_clause,[],[f1526])).
fof(f1526,plain,(
  spl0_125 <=> null_class = domain_of(null_class)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_125])])).
fof(f1334,plain,(
  ( ! [X2,X3] : (member(not_subclass_element(regular(X2),X3),null_class) | ~member(not_subclass_element(regular(X2),X3),X2) | null_class = X2 | subclass(regular(X2),X3)) )),
  inference(resolution,[],[f335,f2])).
fof(f335,plain,(
  ( ! [X0,X1] : (~member(X1,regular(X0)) | member(X1,null_class) | ~member(X1,X0) | null_class = X0) )),
  inference(superposition,[],[f23,f67])).
fof(f67,axiom,(
  ( ! [X0] : (null_class = intersection(X0,regular(X0)) | null_class = X0) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f270091,plain,(
  ~spl0_125 | ~spl0_3749),
  inference(avatar_contradiction_clause,[],[f270090])).
fof(f270090,plain,(
  $false | (~spl0_125 | ~spl0_3749)),
  inference(subsumption_resolution,[],[f270089,f199])).
fof(f199,plain,(
  member(x,universal_class)),
  inference(resolution,[],[f198,f4])).
fof(f270089,plain,(
  ~member(x,universal_class) | (~spl0_125 | ~spl0_3749)),
  inference(subsumption_resolution,[],[f270011,f1533])).
fof(f270011,plain,(
  member(x,null_class) | ~member(x,universal_class) | ~spl0_3749),
  inference(superposition,[],[f9,f201353])).
fof(f201353,plain,(
  null_class = unordered_pair(x,complement(complement(x))) | ~spl0_3749),
  inference(avatar_component_clause,[],[f201351])).
fof(f9,axiom,(
  ( ! [X0,X1] : (member(X0,unordered_pair(X0,X1)) | ~member(X0,universal_class)) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f201354,plain,(
  spl0_3748 | spl0_3749 | ~spl0_125),
  inference(avatar_split_clause,[],[f201017,f1526,f201351,f201347])).
fof(f201017,plain,(
  null_class = unordered_pair(x,complement(complement(x))) | x = regular(unordered_pair(x,complement(complement(x)))) | ~spl0_125),
  inference(resolution,[],[f21547,f117777])).
fof(f117777,plain,(
  member(x,complement(complement(x)))),
  inference(resolution,[],[f117388,f198])).
fof(f21547,plain,(
  ( ! [X2,X3] : (~member(X2,X3) | null_class = unordered_pair(X2,X3) | regular(unordered_pair(X2,X3)) = X2) ) | ~spl0_125),
  inference(subsumption_resolution,[],[f21493,f10982])).
fof(f10982,plain,(
  ( ! [X0,X1] : (~member(X1,X0) | member(X1,universal_class)) )),
  inference(superposition,[],[f21,f10950])).
fof(f21493,plain,(
  ( ! [X2,X3] : (~member(X2,X3) | null_class = unordered_pair(X2,X3) | regular(unordered_pair(X2,X3)) = X2 | ~member(X2,universal_class)) ) | ~spl0_125),
  inference(resolution,[],[f1603,f9])).
fof(f1603,plain,(
  ( ! [X2,X0,X1] : (~member(X2,unordered_pair(X0,X1)) | ~member(X2,X1) | unordered_pair(X0,X1) = null_class | regular(unordered_pair(X0,X1)) = X0) ) | ~spl0_125),
  inference(subsumption_resolution,[],[f1600,f1533])).
fof(f1600,plain,(
  ( ! [X2,X0,X1] : (~member(X2,X1) | member(X2,null_class) | ~member(X2,unordered_pair(X0,X1)) | unordered_pair(X0,X1) = null_class | regular(unordered_pair(X0,X1)) = X0) )),
  inference(duplicate_literal_removal,[],[f1579])).
fof(f1579,plain,(
  ( ! [X2,X0,X1] : (~member(X2,X1) | member(X2,null_class) | ~member(X2,unordered_pair(X0,X1)) | unordered_pair(X0,X1) = null_class | regular(unordered_pair(X0,X1)) = X0 | unordered_pair(X0,X1) = null_class) )),
  inference(superposition,[],[f335,f330])).
fof(f330,plain,(
  ( ! [X10,X9] : (regular(unordered_pair(X9,X10)) = X10 | regular(unordered_pair(X9,X10)) = X9 | null_class = unordered_pair(X9,X10)) )),
  inference(resolution,[],[f8,f66])).
fof(f8,axiom,(
  ( ! [X2,X0,X1] : (~member(X2,unordered_pair(X0,X1)) | X0 = X2 | X1 = X2) )),
  file('./p/SET515-6.tptp',unknown)).
fof(f1530,plain,(
  spl0_125),
  inference(avatar_split_clause,[],[f1517,f1526])).
fof(f1517,plain,(
  null_class = domain_of(null_class)),
  inference(resolution,[],[f1473,f66])).
% SZS output end Proof for SET515-6
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 425750
% Time elapsed: 54.056 s
% ------------------------------
% ------------------------------
