% Refutation found. Thanks to Tanya!
% SZS status Theorem for SET646+3
% SZS output start Proof for SET646+3
fof(f54228,plain,(
  $false),
  inference(subsumption_resolution,[],[f54227,f165])).
fof(f165,plain,(
  member(sK14,sK12)),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  (((~ilf_type(singleton(ordered_pair(sK13,sK14)),relation_type(sK11,sK12)) & member(sK14,sK12) & member(sK13,sK11) & ilf_type(sK14,set_type)) & ilf_type(sK13,set_type)) & ilf_type(sK12,set_type)) & ilf_type(sK11,set_type)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11,sK12,sK13,sK14])],[f60,f104,f103,f102,f101])).
fof(f101,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~ilf_type(singleton(ordered_pair(X2,X3)),relation_type(X0,X1)) & member(X3,X1) & member(X2,X0) & ilf_type(X3,set_type)) & ilf_type(X2,set_type)) & ilf_type(X1,set_type)) & ilf_type(X0,set_type)) => (? [X1] : (? [X2] : (? [X3] : (~ilf_type(singleton(ordered_pair(X2,X3)),relation_type(sK11,X1)) & member(X3,X1) & member(X2,sK11) & ilf_type(X3,set_type)) & ilf_type(X2,set_type)) & ilf_type(X1,set_type)) & ilf_type(sK11,set_type))),
  introduced(choice_axiom,[])).
fof(f102,plain,(
  ? [X1] : (? [X2] : (? [X3] : (~ilf_type(singleton(ordered_pair(X2,X3)),relation_type(sK11,X1)) & member(X3,X1) & member(X2,sK11) & ilf_type(X3,set_type)) & ilf_type(X2,set_type)) & ilf_type(X1,set_type)) => (? [X2] : (? [X3] : (~ilf_type(singleton(ordered_pair(X2,X3)),relation_type(sK11,sK12)) & member(X3,sK12) & member(X2,sK11) & ilf_type(X3,set_type)) & ilf_type(X2,set_type)) & ilf_type(sK12,set_type))),
  introduced(choice_axiom,[])).
fof(f103,plain,(
  ? [X2] : (? [X3] : (~ilf_type(singleton(ordered_pair(X2,X3)),relation_type(sK11,sK12)) & member(X3,sK12) & member(X2,sK11) & ilf_type(X3,set_type)) & ilf_type(X2,set_type)) => (? [X3] : (~ilf_type(singleton(ordered_pair(sK13,X3)),relation_type(sK11,sK12)) & member(X3,sK12) & member(sK13,sK11) & ilf_type(X3,set_type)) & ilf_type(sK13,set_type))),
  introduced(choice_axiom,[])).
fof(f104,plain,(
  ? [X3] : (~ilf_type(singleton(ordered_pair(sK13,X3)),relation_type(sK11,sK12)) & member(X3,sK12) & member(sK13,sK11) & ilf_type(X3,set_type)) => (~ilf_type(singleton(ordered_pair(sK13,sK14)),relation_type(sK11,sK12)) & member(sK14,sK12) & member(sK13,sK11) & ilf_type(sK14,set_type))),
  introduced(choice_axiom,[])).
fof(f60,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (~ilf_type(singleton(ordered_pair(X2,X3)),relation_type(X0,X1)) & member(X3,X1) & member(X2,X0) & ilf_type(X3,set_type)) & ilf_type(X2,set_type)) & ilf_type(X1,set_type)) & ilf_type(X0,set_type))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((~ilf_type(singleton(ordered_pair(X2,X3)),relation_type(X0,X1)) & (member(X3,X1) & member(X2,X0))) & ilf_type(X3,set_type)) & ilf_type(X2,set_type)) & ilf_type(X1,set_type)) & ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f27])).
fof(f27,negated_conjecture,(
  ~! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => ! [X2] : (ilf_type(X2,set_type) => ! [X3] : (ilf_type(X3,set_type) => ((member(X3,X1) & member(X2,X0)) => ilf_type(singleton(ordered_pair(X2,X3)),relation_type(X0,X1)))))))),
  inference(negated_conjecture,[],[f26])).
fof(f26,conjecture,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => ! [X2] : (ilf_type(X2,set_type) => ! [X3] : (ilf_type(X3,set_type) => ((member(X3,X1) & member(X2,X0)) => ilf_type(singleton(ordered_pair(X2,X3)),relation_type(X0,X1)))))))),
  file('./p/SET646+3.tptp',unknown)).
fof(f54227,plain,(
  ~member(sK14,sK12)),
  inference(subsumption_resolution,[],[f54216,f164])).
fof(f164,plain,(
  member(sK13,sK11)),
  inference(cnf_transformation,[],[f105])).
fof(f54216,plain,(
  ~member(sK13,sK11) | ~member(sK14,sK12)),
  inference(resolution,[],[f54164,f1330])).
fof(f1330,plain,(
  ( ! [X2,X0,X3,X1] : (member(ordered_pair(X0,X2),cross_product(X1,X3)) | ~member(X0,X1) | ~member(X2,X3)) )),
  inference(resolution,[],[f744,f220])).
fof(f220,plain,(
  ( ! [X0] : (subset(X0,X0)) )),
  inference(resolution,[],[f138,f159])).
fof(f159,plain,(
  ( ! [X0] : (ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,axiom,(
  ! [X0] : ilf_type(X0,set_type)),
  file('./p/SET646+3.tptp',unknown)).
fof(f138,plain,(
  ( ! [X0] : (~ilf_type(X0,set_type) | subset(X0,X0)) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : (subset(X0,X0) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f16])).
fof(f16,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => subset(X0,X0))),
  file('./p/SET646+3.tptp',unknown)).
fof(f744,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~subset(cross_product(X15,X13),X16) | ~member(X14,X15) | member(ordered_pair(X14,X12),X16) | ~member(X12,X13)) )),
  inference(subsumption_resolution,[],[f743,f159])).
fof(f743,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~member(X12,X13) | ~member(X14,X15) | member(ordered_pair(X14,X12),X16) | ~subset(cross_product(X15,X13),X16) | ~ilf_type(cross_product(X15,X13),set_type)) )),
  inference(subsumption_resolution,[],[f742,f159])).
fof(f742,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~member(X12,X13) | ~member(X14,X15) | member(ordered_pair(X14,X12),X16) | ~subset(cross_product(X15,X13),X16) | ~ilf_type(X16,set_type) | ~ilf_type(cross_product(X15,X13),set_type)) )),
  inference(subsumption_resolution,[],[f741,f159])).
fof(f741,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~member(X12,X13) | ~member(X14,X15) | member(ordered_pair(X14,X12),X16) | ~ilf_type(ordered_pair(X14,X12),set_type) | ~subset(cross_product(X15,X13),X16) | ~ilf_type(X16,set_type) | ~ilf_type(cross_product(X15,X13),set_type)) )),
  inference(subsumption_resolution,[],[f740,f159])).
fof(f740,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~member(X12,X13) | ~member(X14,X15) | ~ilf_type(X14,set_type) | member(ordered_pair(X14,X12),X16) | ~ilf_type(ordered_pair(X14,X12),set_type) | ~subset(cross_product(X15,X13),X16) | ~ilf_type(X16,set_type) | ~ilf_type(cross_product(X15,X13),set_type)) )),
  inference(subsumption_resolution,[],[f739,f159])).
fof(f739,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~member(X12,X13) | ~member(X14,X15) | ~ilf_type(X12,set_type) | ~ilf_type(X14,set_type) | member(ordered_pair(X14,X12),X16) | ~ilf_type(ordered_pair(X14,X12),set_type) | ~subset(cross_product(X15,X13),X16) | ~ilf_type(X16,set_type) | ~ilf_type(cross_product(X15,X13),set_type)) )),
  inference(subsumption_resolution,[],[f738,f159])).
fof(f738,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~member(X12,X13) | ~member(X14,X15) | ~ilf_type(X15,set_type) | ~ilf_type(X12,set_type) | ~ilf_type(X14,set_type) | member(ordered_pair(X14,X12),X16) | ~ilf_type(ordered_pair(X14,X12),set_type) | ~subset(cross_product(X15,X13),X16) | ~ilf_type(X16,set_type) | ~ilf_type(cross_product(X15,X13),set_type)) )),
  inference(subsumption_resolution,[],[f726,f159])).
fof(f726,plain,(
  ( ! [X14,X12,X15,X13,X16] : (~member(X12,X13) | ~member(X14,X15) | ~ilf_type(X13,set_type) | ~ilf_type(X15,set_type) | ~ilf_type(X12,set_type) | ~ilf_type(X14,set_type) | member(ordered_pair(X14,X12),X16) | ~ilf_type(ordered_pair(X14,X12),set_type) | ~subset(cross_product(X15,X13),X16) | ~ilf_type(X16,set_type) | ~ilf_type(cross_product(X15,X13),set_type)) )),
  inference(resolution,[],[f110,f134])).
fof(f134,plain,(
  ( ! [X0,X3,X1] : (~member(X3,X0) | member(X3,X1) | ~ilf_type(X3,set_type) | ~subset(X0,X1) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (! [X1] : (((subset(X0,X1) | (~member(sK4(X0,X1),X1) & member(sK4(X0,X1),X0) & ilf_type(sK4(X0,X1),set_type))) & (! [X3] : (member(X3,X1) | ~member(X3,X0) | ~ilf_type(X3,set_type)) | ~subset(X0,X1))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f81,f82])).
fof(f82,plain,(
  ! [X1,X0] : (? [X2] : (~member(X2,X1) & member(X2,X0) & ilf_type(X2,set_type)) => (~member(sK4(X0,X1),X1) & member(sK4(X0,X1),X0) & ilf_type(sK4(X0,X1),set_type)))),
  introduced(choice_axiom,[])).
fof(f81,plain,(
  ! [X0] : (! [X1] : (((subset(X0,X1) | ? [X2] : (~member(X2,X1) & member(X2,X0) & ilf_type(X2,set_type))) & (! [X3] : (member(X3,X1) | ~member(X3,X0) | ~ilf_type(X3,set_type)) | ~subset(X0,X1))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(rectify,[],[f80])).
fof(f80,plain,(
  ! [X0] : (! [X1] : (((subset(X0,X1) | ? [X2] : (~member(X2,X1) & member(X2,X0) & ilf_type(X2,set_type))) & (! [X2] : (member(X2,X1) | ~member(X2,X0) | ~ilf_type(X2,set_type)) | ~subset(X0,X1))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (! [X1] : ((subset(X0,X1) <=> ! [X2] : (member(X2,X1) | ~member(X2,X0) | ~ilf_type(X2,set_type))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : (! [X1] : ((subset(X0,X1) <=> ! [X2] : ((member(X2,X1) | ~member(X2,X0)) | ~ilf_type(X2,set_type))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f15])).
fof(f15,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => (subset(X0,X1) <=> ! [X2] : (ilf_type(X2,set_type) => (member(X2,X0) => member(X2,X1))))))),
  file('./p/SET646+3.tptp',unknown)).
fof(f110,plain,(
  ( ! [X2,X0,X3,X1] : (member(ordered_pair(X0,X1),cross_product(X2,X3)) | ~member(X1,X3) | ~member(X0,X2) | ~ilf_type(X3,set_type) | ~ilf_type(X2,set_type) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((member(ordered_pair(X0,X1),cross_product(X2,X3)) | ~member(X1,X3) | ~member(X0,X2)) & ((member(X1,X3) & member(X0,X2)) | ~member(ordered_pair(X0,X1),cross_product(X2,X3)))) | ~ilf_type(X3,set_type)) | ~ilf_type(X2,set_type)) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : (((member(ordered_pair(X0,X1),cross_product(X2,X3)) | (~member(X1,X3) | ~member(X0,X2))) & ((member(X1,X3) & member(X0,X2)) | ~member(ordered_pair(X0,X1),cross_product(X2,X3)))) | ~ilf_type(X3,set_type)) | ~ilf_type(X2,set_type)) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : (! [X1] : (! [X2] : (! [X3] : ((member(ordered_pair(X0,X1),cross_product(X2,X3)) <=> (member(X1,X3) & member(X0,X2))) | ~ilf_type(X3,set_type)) | ~ilf_type(X2,set_type)) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => ! [X2] : (ilf_type(X2,set_type) => ! [X3] : (ilf_type(X3,set_type) => (member(ordered_pair(X0,X1),cross_product(X2,X3)) <=> (member(X1,X3) & member(X0,X2)))))))),
  file('./p/SET646+3.tptp',unknown)).
fof(f54164,plain,(
  ~member(ordered_pair(sK13,sK14),cross_product(sK11,sK12))),
  inference(resolution,[],[f43299,f166])).
fof(f166,plain,(
  ~ilf_type(singleton(ordered_pair(sK13,sK14)),relation_type(sK11,sK12))),
  inference(cnf_transformation,[],[f105])).
fof(f43299,plain,(
  ( ! [X12,X13,X11] : (ilf_type(singleton(X11),relation_type(X12,X13)) | ~member(X11,cross_product(X12,X13))) )),
  inference(subsumption_resolution,[],[f43298,f159])).
fof(f43298,plain,(
  ( ! [X12,X13,X11] : (ilf_type(singleton(X11),relation_type(X12,X13)) | ~member(X11,cross_product(X12,X13)) | ~ilf_type(X11,set_type)) )),
  inference(subsumption_resolution,[],[f43292,f159])).
fof(f43292,plain,(
  ( ! [X12,X13,X11] : (ilf_type(singleton(X11),relation_type(X12,X13)) | ~member(X11,cross_product(X12,X13)) | ~ilf_type(cross_product(X12,X13),set_type) | ~ilf_type(X11,set_type)) )),
  inference(resolution,[],[f38888,f107])).
fof(f107,plain,(
  ( ! [X0,X1] : (subset(singleton(X0),X1) | ~member(X0,X1) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : (! [X1] : (((subset(singleton(X0),X1) | ~member(X0,X1)) & (member(X0,X1) | ~subset(singleton(X0),X1))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(nnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0] : (! [X1] : ((subset(singleton(X0),X1) <=> member(X0,X1)) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => (subset(singleton(X0),X1) <=> member(X0,X1))))),
  file('./p/SET646+3.tptp',unknown)).
fof(f38888,plain,(
  ( ! [X4,X2,X3] : (~subset(X2,cross_product(X3,X4)) | ilf_type(X2,relation_type(X3,X4))) )),
  inference(resolution,[],[f26235,f32961])).
fof(f32961,plain,(
  ( ! [X0,X3,X1] : (~ilf_type(X3,subset_type(cross_product(X0,X1))) | ilf_type(X3,relation_type(X0,X1))) )),
  inference(subsumption_resolution,[],[f32960,f159])).
fof(f32960,plain,(
  ( ! [X0,X3,X1] : (~ilf_type(X3,subset_type(cross_product(X0,X1))) | ilf_type(X3,relation_type(X0,X1)) | ~ilf_type(X0,set_type)) )),
  inference(subsumption_resolution,[],[f111,f159])).
fof(f111,plain,(
  ( ! [X0,X3,X1] : (~ilf_type(X3,subset_type(cross_product(X0,X1))) | ilf_type(X3,relation_type(X0,X1)) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (! [X1] : ((! [X2] : (ilf_type(X2,subset_type(cross_product(X0,X1))) | ~ilf_type(X2,relation_type(X0,X1))) & ! [X3] : (ilf_type(X3,relation_type(X0,X1)) | ~ilf_type(X3,subset_type(cross_product(X0,X1))))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => (! [X2] : (ilf_type(X2,relation_type(X0,X1)) => ilf_type(X2,subset_type(cross_product(X0,X1)))) & ! [X3] : (ilf_type(X3,subset_type(cross_product(X0,X1))) => ilf_type(X3,relation_type(X0,X1))))))),
  inference(rectify,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => (! [X3] : (ilf_type(X3,relation_type(X0,X1)) => ilf_type(X3,subset_type(cross_product(X0,X1)))) & ! [X2] : (ilf_type(X2,subset_type(cross_product(X0,X1))) => ilf_type(X2,relation_type(X0,X1))))))),
  file('./p/SET646+3.tptp',unknown)).
fof(f26235,plain,(
  ( ! [X2,X3] : (ilf_type(X3,subset_type(X2)) | ~subset(X3,X2)) )),
  inference(subsumption_resolution,[],[f9115,f8927])).
fof(f8927,plain,(
  ( ! [X14,X15] : (~empty(power_set(X15)) | ~subset(X14,X15)) )),
  inference(subsumption_resolution,[],[f8926,f159])).
fof(f8926,plain,(
  ( ! [X14,X15] : (~subset(X14,X15) | ~empty(power_set(X15)) | ~ilf_type(power_set(X15),set_type)) )),
  inference(subsumption_resolution,[],[f1048,f159])).
fof(f1048,plain,(
  ( ! [X14,X15] : (~subset(X14,X15) | ~ilf_type(X14,set_type) | ~empty(power_set(X15)) | ~ilf_type(power_set(X15),set_type)) )),
  inference(resolution,[],[f1031,f148])).
fof(f148,plain,(
  ( ! [X2,X0] : (~member(X2,X0) | ~ilf_type(X2,set_type) | ~empty(X0) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : (((empty(X0) | (member(sK7(X0),X0) & ilf_type(sK7(X0),set_type))) & (! [X2] : (~member(X2,X0) | ~ilf_type(X2,set_type)) | ~empty(X0))) | ~ilf_type(X0,set_type))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f92,f93])).
fof(f93,plain,(
  ! [X0] : (? [X1] : (member(X1,X0) & ilf_type(X1,set_type)) => (member(sK7(X0),X0) & ilf_type(sK7(X0),set_type)))),
  introduced(choice_axiom,[])).
fof(f92,plain,(
  ! [X0] : (((empty(X0) | ? [X1] : (member(X1,X0) & ilf_type(X1,set_type))) & (! [X2] : (~member(X2,X0) | ~ilf_type(X2,set_type)) | ~empty(X0))) | ~ilf_type(X0,set_type))),
  inference(rectify,[],[f91])).
fof(f91,plain,(
  ! [X0] : (((empty(X0) | ? [X1] : (member(X1,X0) & ilf_type(X1,set_type))) & (! [X1] : (~member(X1,X0) | ~ilf_type(X1,set_type)) | ~empty(X0))) | ~ilf_type(X0,set_type))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : ((empty(X0) <=> ! [X1] : (~member(X1,X0) | ~ilf_type(X1,set_type))) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f21])).
fof(f21,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => (empty(X0) <=> ! [X1] : (ilf_type(X1,set_type) => ~member(X1,X0))))),
  file('./p/SET646+3.tptp',unknown)).
fof(f1031,plain,(
  ( ! [X0,X1] : (member(X0,power_set(X1)) | ~subset(X0,X1)) )),
  inference(subsumption_resolution,[],[f1030,f159])).
fof(f1030,plain,(
  ( ! [X0,X1] : (~subset(X0,X1) | member(X0,power_set(X1)) | ~ilf_type(X0,set_type)) )),
  inference(subsumption_resolution,[],[f1029,f159])).
fof(f1029,plain,(
  ( ! [X0,X1] : (~subset(X0,X1) | member(X0,power_set(X1)) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(duplicate_literal_removal,[],[f1017])).
fof(f1017,plain,(
  ( ! [X0,X1] : (~subset(X0,X1) | member(X0,power_set(X1)) | member(X0,power_set(X1)) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(resolution,[],[f509,f142])).
fof(f142,plain,(
  ( ! [X0,X1] : (~member(sK5(X0,X1),X1) | member(X0,power_set(X1)) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : (! [X1] : (((member(X0,power_set(X1)) | (~member(sK5(X0,X1),X1) & member(sK5(X0,X1),X0) & ilf_type(sK5(X0,X1),set_type))) & (! [X3] : (member(X3,X1) | ~member(X3,X0) | ~ilf_type(X3,set_type)) | ~member(X0,power_set(X1)))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f85,f86])).
fof(f86,plain,(
  ! [X1,X0] : (? [X2] : (~member(X2,X1) & member(X2,X0) & ilf_type(X2,set_type)) => (~member(sK5(X0,X1),X1) & member(sK5(X0,X1),X0) & ilf_type(sK5(X0,X1),set_type)))),
  introduced(choice_axiom,[])).
fof(f85,plain,(
  ! [X0] : (! [X1] : (((member(X0,power_set(X1)) | ? [X2] : (~member(X2,X1) & member(X2,X0) & ilf_type(X2,set_type))) & (! [X3] : (member(X3,X1) | ~member(X3,X0) | ~ilf_type(X3,set_type)) | ~member(X0,power_set(X1)))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(rectify,[],[f84])).
fof(f84,plain,(
  ! [X0] : (! [X1] : (((member(X0,power_set(X1)) | ? [X2] : (~member(X2,X1) & member(X2,X0) & ilf_type(X2,set_type))) & (! [X2] : (member(X2,X1) | ~member(X2,X0) | ~ilf_type(X2,set_type)) | ~member(X0,power_set(X1)))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (! [X1] : ((member(X0,power_set(X1)) <=> ! [X2] : (member(X2,X1) | ~member(X2,X0) | ~ilf_type(X2,set_type))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : (! [X1] : ((member(X0,power_set(X1)) <=> ! [X2] : ((member(X2,X1) | ~member(X2,X0)) | ~ilf_type(X2,set_type))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f17])).
fof(f17,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => (member(X0,power_set(X1)) <=> ! [X2] : (ilf_type(X2,set_type) => (member(X2,X0) => member(X2,X1))))))),
  file('./p/SET646+3.tptp',unknown)).
fof(f509,plain,(
  ( ! [X19,X20,X18] : (member(sK5(X18,X19),X20) | ~subset(X18,X20) | member(X18,power_set(X19))) )),
  inference(subsumption_resolution,[],[f508,f159])).
fof(f508,plain,(
  ( ! [X19,X20,X18] : (member(sK5(X18,X19),X20) | ~subset(X18,X20) | member(X18,power_set(X19)) | ~ilf_type(X19,set_type)) )),
  inference(subsumption_resolution,[],[f507,f159])).
fof(f507,plain,(
  ( ! [X19,X20,X18] : (member(sK5(X18,X19),X20) | ~subset(X18,X20) | ~ilf_type(X18,set_type) | member(X18,power_set(X19)) | ~ilf_type(X19,set_type)) )),
  inference(subsumption_resolution,[],[f506,f159])).
fof(f506,plain,(
  ( ! [X19,X20,X18] : (member(sK5(X18,X19),X20) | ~subset(X18,X20) | ~ilf_type(X20,set_type) | ~ilf_type(X18,set_type) | member(X18,power_set(X19)) | ~ilf_type(X19,set_type)) )),
  inference(subsumption_resolution,[],[f485,f159])).
fof(f485,plain,(
  ( ! [X19,X20,X18] : (member(sK5(X18,X19),X20) | ~ilf_type(sK5(X18,X19),set_type) | ~subset(X18,X20) | ~ilf_type(X20,set_type) | ~ilf_type(X18,set_type) | member(X18,power_set(X19)) | ~ilf_type(X19,set_type)) )),
  inference(duplicate_literal_removal,[],[f477])).
fof(f477,plain,(
  ( ! [X19,X20,X18] : (member(sK5(X18,X19),X20) | ~ilf_type(sK5(X18,X19),set_type) | ~subset(X18,X20) | ~ilf_type(X20,set_type) | ~ilf_type(X18,set_type) | member(X18,power_set(X19)) | ~ilf_type(X19,set_type) | ~ilf_type(X18,set_type)) )),
  inference(resolution,[],[f134,f141])).
fof(f141,plain,(
  ( ! [X0,X1] : (member(sK5(X0,X1),X0) | member(X0,power_set(X1)) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f87])).
fof(f9115,plain,(
  ( ! [X2,X3] : (empty(power_set(X2)) | ilf_type(X3,subset_type(X2)) | ~subset(X3,X2)) )),
  inference(resolution,[],[f8932,f1031])).
fof(f8932,plain,(
  ( ! [X0,X1] : (~member(X0,power_set(X1)) | empty(power_set(X1)) | ilf_type(X0,subset_type(X1))) )),
  inference(subsumption_resolution,[],[f8931,f159])).
fof(f8931,plain,(
  ( ! [X0,X1] : (~member(X0,power_set(X1)) | empty(power_set(X1)) | ilf_type(X0,subset_type(X1)) | ~ilf_type(X1,set_type)) )),
  inference(subsumption_resolution,[],[f400,f159])).
fof(f400,plain,(
  ( ! [X0,X1] : (~member(X0,power_set(X1)) | empty(power_set(X1)) | ~ilf_type(X0,set_type) | ilf_type(X0,subset_type(X1)) | ~ilf_type(X1,set_type)) )),
  inference(subsumption_resolution,[],[f399,f159])).
fof(f399,plain,(
  ( ! [X0,X1] : (~member(X0,power_set(X1)) | ~ilf_type(power_set(X1),set_type) | empty(power_set(X1)) | ~ilf_type(X0,set_type) | ilf_type(X0,subset_type(X1)) | ~ilf_type(X1,set_type)) )),
  inference(duplicate_literal_removal,[],[f396])).
fof(f396,plain,(
  ( ! [X0,X1] : (~member(X0,power_set(X1)) | ~ilf_type(power_set(X1),set_type) | empty(power_set(X1)) | ~ilf_type(X0,set_type) | ilf_type(X0,subset_type(X1)) | ~ilf_type(X0,set_type) | ~ilf_type(X1,set_type)) )),
  inference(resolution,[],[f146,f127])).
fof(f127,plain,(
  ( ! [X0,X1] : (~ilf_type(X1,member_type(power_set(X0))) | ilf_type(X1,subset_type(X0)) | ~ilf_type(X1,set_type) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : (! [X1] : (((ilf_type(X1,subset_type(X0)) | ~ilf_type(X1,member_type(power_set(X0)))) & (ilf_type(X1,member_type(power_set(X0))) | ~ilf_type(X1,subset_type(X0)))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (! [X1] : ((ilf_type(X1,subset_type(X0)) <=> ilf_type(X1,member_type(power_set(X0)))) | ~ilf_type(X1,set_type)) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f12])).
fof(f12,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : (ilf_type(X1,set_type) => (ilf_type(X1,subset_type(X0)) <=> ilf_type(X1,member_type(power_set(X0))))))),
  file('./p/SET646+3.tptp',unknown)).
fof(f146,plain,(
  ( ! [X0,X1] : (ilf_type(X0,member_type(X1)) | ~member(X0,X1) | ~ilf_type(X1,set_type) | empty(X1) | ~ilf_type(X0,set_type)) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : (! [X1] : (((ilf_type(X0,member_type(X1)) | ~member(X0,X1)) & (member(X0,X1) | ~ilf_type(X0,member_type(X1)))) | ~ilf_type(X1,set_type) | empty(X1)) | ~ilf_type(X0,set_type))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (! [X1] : ((ilf_type(X0,member_type(X1)) <=> member(X0,X1)) | ~ilf_type(X1,set_type) | empty(X1)) | ~ilf_type(X0,set_type))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : (! [X1] : ((ilf_type(X0,member_type(X1)) <=> member(X0,X1)) | (~ilf_type(X1,set_type) | empty(X1))) | ~ilf_type(X0,set_type))),
  inference(ennf_transformation,[],[f19])).
fof(f19,axiom,(
  ! [X0] : (ilf_type(X0,set_type) => ! [X1] : ((ilf_type(X1,set_type) & ~empty(X1)) => (ilf_type(X0,member_type(X1)) <=> member(X0,X1))))),
  file('./p/SET646+3.tptp',unknown)).
% SZS output end Proof for SET646+3
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 46822
% Time elapsed: 4.048 s
% ------------------------------
% ------------------------------
