% Refutation found. Thanks to Tanya!
% SZS status Theorem for SET602+4
% SZS output start Proof for SET602+4
fof(f675,plain,(
  $false),
  inference(resolution,[],[f585,f81])).
fof(f81,plain,(
  ~equal_set(difference(sK3,sK3),empty_set)),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ~equal_set(difference(sK3,sK3),empty_set)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f29,f52])).
fof(f52,plain,(
  ? [X0] : ~equal_set(difference(X0,X0),empty_set) => ~equal_set(difference(sK3,sK3),empty_set)),
  introduced(choice_axiom,[])).
fof(f29,plain,(
  ? [X0] : ~equal_set(difference(X0,X0),empty_set)),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0] : equal_set(difference(X0,X0),empty_set)),
  inference(rectify,[],[f13])).
fof(f13,negated_conjecture,(
  ~! [X3] : equal_set(difference(X3,X3),empty_set)),
  inference(negated_conjecture,[],[f12])).
fof(f12,conjecture,(
  ! [X3] : equal_set(difference(X3,X3),empty_set)),
  file('./p/SET602+4.tptp',unknown)).
fof(f585,plain,(
  ( ! [X4] : (equal_set(difference(X4,X4),empty_set)) )),
  inference(resolution,[],[f517,f106])).
fof(f106,plain,(
  ( ! [X0] : (~subset(X0,empty_set) | equal_set(X0,empty_set)) )),
  inference(resolution,[],[f57,f87])).
fof(f87,plain,(
  ( ! [X0] : (subset(empty_set,X0)) )),
  inference(resolution,[],[f55,f66])).
fof(f66,plain,(
  ( ! [X0] : (~member(X0,empty_set)) )),
  inference(cnf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0] : ~member(X0,empty_set)),
  inference(rectify,[],[f6])).
fof(f6,axiom,(
  ! [X2] : ~member(X2,empty_set)),
  file('./p/SET602+4.tptp',unknown)).
fof(f55,plain,(
  ( ! [X0,X1] : (member(sK0(X0,X1),X0) | subset(X0,X1)) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ((subset(X0,X1) | (~member(sK0(X0,X1),X1) & member(sK0(X0,X1),X0))) & (! [X3] : (member(X3,X1) | ~member(X3,X0)) | ~subset(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f31,f32])).
fof(f32,plain,(
  ! [X1,X0] : (? [X2] : (~member(X2,X1) & member(X2,X0)) => (~member(sK0(X0,X1),X1) & member(sK0(X0,X1),X0)))),
  introduced(choice_axiom,[])).
fof(f31,plain,(
  ! [X0,X1] : ((subset(X0,X1) | ? [X2] : (~member(X2,X1) & member(X2,X0))) & (! [X3] : (member(X3,X1) | ~member(X3,X0)) | ~subset(X0,X1)))),
  inference(rectify,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ((subset(X0,X1) | ? [X2] : (~member(X2,X1) & member(X2,X0))) & (! [X2] : (member(X2,X1) | ~member(X2,X0)) | ~subset(X0,X1)))),
  inference(nnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : (subset(X0,X1) <=> ! [X2] : (member(X2,X1) | ~member(X2,X0)))),
  inference(ennf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0,X1] : (subset(X0,X1) <=> ! [X2] : (member(X2,X0) => member(X2,X1)))),
  file('./p/SET602+4.tptp',unknown)).
fof(f57,plain,(
  ( ! [X0,X1] : (~subset(X1,X0) | equal_set(X0,X1) | ~subset(X0,X1)) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : (equal_set(X0,X1) | ~subset(X1,X0) | ~subset(X0,X1))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : (equal_set(X0,X1) | (~subset(X1,X0) | ~subset(X0,X1)))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : ((subset(X1,X0) & subset(X0,X1)) => equal_set(X0,X1))),
  inference(unused_predicate_definition_removal,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1] : (equal_set(X0,X1) <=> (subset(X1,X0) & subset(X0,X1)))),
  file('./p/SET602+4.tptp',unknown)).
fof(f517,plain,(
  ( ! [X0,X1] : (subset(difference(X0,X0),X1)) )),
  inference(duplicate_literal_removal,[],[f498])).
fof(f498,plain,(
  ( ! [X0,X1] : (subset(difference(X0,X0),X1) | subset(difference(X0,X0),X1)) )),
  inference(resolution,[],[f98,f97])).
fof(f97,plain,(
  ( ! [X2,X0,X1] : (member(sK0(difference(X0,X1),X2),X0) | subset(difference(X0,X1),X2)) )),
  inference(resolution,[],[f67,f55])).
fof(f67,plain,(
  ( ! [X2,X0,X1] : (~member(X0,difference(X2,X1)) | member(X0,X2)) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2] : ((member(X0,difference(X2,X1)) | member(X0,X1) | ~member(X0,X2)) & ((~member(X0,X1) & member(X0,X2)) | ~member(X0,difference(X2,X1))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2] : ((member(X0,difference(X2,X1)) | (member(X0,X1) | ~member(X0,X2))) & ((~member(X0,X1) & member(X0,X2)) | ~member(X0,difference(X2,X1))))),
  inference(nnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1,X2] : (member(X0,difference(X2,X1)) <=> (~member(X0,X1) & member(X0,X2)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X1,X0,X3] : (member(X1,difference(X3,X0)) <=> (~member(X1,X0) & member(X1,X3)))),
  file('./p/SET602+4.tptp',unknown)).
fof(f98,plain,(
  ( ! [X2,X0,X1] : (~member(sK0(difference(X0,X1),X2),X1) | subset(difference(X0,X1),X2)) )),
  inference(resolution,[],[f68,f55])).
fof(f68,plain,(
  ( ! [X2,X0,X1] : (~member(X0,difference(X2,X1)) | ~member(X0,X1)) )),
  inference(cnf_transformation,[],[f40])).
% SZS output end Proof for SET602+4
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.008 s
% ------------------------------
% ------------------------------
