% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SET851-2
% SZS output start Proof for SET851-2
fof(f31,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f16,f18,f28])).
fof(f28,plain,(
  ~spl0_2),
  inference(avatar_contradiction_clause,[],[f27])).
fof(f27,plain,(
  $false | ~spl0_2),
  inference(subsumption_resolution,[],[f23,f1])).
fof(f1,axiom,(
  ( ! [X0,X1] : (c_lessequals(X0,X0,tc_set(X1))) )),
  file('./p/SET851-2.tptp',unknown)).
fof(f23,plain,(
  ~c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),c_Zorn_Osucc(v_S,v_x,t_a),tc_set(tc_set(t_a))) | ~spl0_2),
  inference(backward_demodulation,[],[f21,f5])).
fof(f5,axiom,(
  ~c_lessequals(c_Zorn_Osucc(v_S,v_m,t_a),c_Zorn_Osucc(v_S,v_x,t_a),tc_set(tc_set(t_a)))),
  file('./p/SET851-2.tptp',unknown)).
fof(f21,plain,(
  v_x = v_m | ~spl0_2),
  inference(subsumption_resolution,[],[f20,f3])).
fof(f3,axiom,(
  c_in(v_x,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))),
  file('./p/SET851-2.tptp',unknown)).
fof(f20,plain,(
  v_x = v_m | ~c_in(v_x,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~spl0_2),
  inference(subsumption_resolution,[],[f19,f15])).
fof(f15,plain,(
  c_lessequals(v_x,v_m,tc_set(tc_set(t_a))) | ~spl0_2),
  inference(avatar_component_clause,[],[f13])).
fof(f13,plain,(
  spl0_2 <=> c_lessequals(v_x,v_m,tc_set(tc_set(t_a)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f19,plain,(
  ~c_lessequals(v_x,v_m,tc_set(tc_set(t_a))) | v_x = v_m | ~c_in(v_x,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))),
  inference(resolution,[],[f7,f4])).
fof(f4,axiom,(
  ~c_lessequals(c_Zorn_Osucc(v_S,v_x,t_a),v_m,tc_set(tc_set(t_a)))),
  file('./p/SET851-2.tptp',unknown)).
fof(f7,axiom,(
  ( ! [X5] : (c_lessequals(c_Zorn_Osucc(v_S,X5,t_a),v_m,tc_set(tc_set(t_a))) | ~c_lessequals(X5,v_m,tc_set(tc_set(t_a))) | v_m = X5 | ~c_in(X5,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))) )),
  file('./p/SET851-2.tptp',unknown)).
fof(f18,plain,(
  ~spl0_1),
  inference(avatar_split_clause,[],[f17,f9])).
fof(f9,plain,(
  spl0_1 <=> c_lessequals(c_Zorn_Osucc(v_S,v_m,t_a),v_x,tc_set(tc_set(t_a)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f17,plain,(
  ~c_lessequals(c_Zorn_Osucc(v_S,v_m,t_a),v_x,tc_set(tc_set(t_a)))),
  inference(resolution,[],[f2,f5])).
fof(f2,axiom,(
  ( ! [X4,X2,X3,X1] : (c_lessequals(X2,c_Zorn_Osucc(X4,X3,X1),tc_set(tc_set(X1))) | ~c_lessequals(X2,X3,tc_set(tc_set(X1)))) )),
  file('./p/SET851-2.tptp',unknown)).
fof(f16,plain,(
  spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f6,f13,f9])).
fof(f6,axiom,(
  c_lessequals(v_x,v_m,tc_set(tc_set(t_a))) | c_lessequals(c_Zorn_Osucc(v_S,v_m,t_a),v_x,tc_set(tc_set(t_a)))),
  file('./p/SET851-2.tptp',unknown)).
% SZS output end Proof for SET851-2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.001 s
% ------------------------------
% ------------------------------
