% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SET861-2
% SZS output start Proof for SET861-2
fof(f15612,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f202,f1253,f1329,f2581,f2624,f2927,f12350,f15580,f15589,f15611])).
fof(f15611,plain,(
  spl0_50 | ~spl0_488),
  inference(avatar_split_clause,[],[f15608,f15587,f1251])).
fof(f1251,plain,(
  spl0_50 <=> ! [X0] : (c_in(c_Main_OsubsetI__1(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),X0,t_a),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | c_lessequals(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),X0,tc_set(t_a)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_50])])).
fof(f15587,plain,(
  spl0_488 <=> ! [X94] : (c_in(X94,v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | ~c_in(X94,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_488])])).
fof(f15608,plain,(
  ( ! [X41] : (c_in(c_Main_OsubsetI__1(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),X41,t_a),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | c_lessequals(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),X41,tc_set(t_a))) ) | ~spl0_488),
  inference(resolution,[],[f15588,f3])).
fof(f3,axiom,(
  ( ! [X2,X0,X1] : (c_in(c_Main_OsubsetI__1(X1,X0,X2),X1,X2) | c_lessequals(X1,X0,tc_set(X2))) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f15588,plain,(
  ( ! [X94] : (~c_in(X94,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a) | c_in(X94,v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a)) ) | ~spl0_488),
  inference(avatar_component_clause,[],[f15587])).
fof(f15589,plain,(
  spl0_488 | spl0_18 | ~spl0_6 | ~spl0_49),
  inference(avatar_split_clause,[],[f13053,f1248,f185,f431,f15587])).
fof(f431,plain,(
  spl0_18 <=> c_in(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(v_S,t_a),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_18])])).
fof(f185,plain,(
  spl0_6 <=> c_in(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),v_S,tc_set(t_a))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).
fof(f1248,plain,(
  spl0_49 <=> ! [X1,X2] : (c_in(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(v_S,t_a),X1,t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)) | ~c_in(X1,v_S,tc_set(t_a)) | c_in(X2,v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | ~c_in(X2,X1,t_a))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_49])])).
fof(f13053,plain,(
  ( ! [X94] : (c_in(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(v_S,t_a),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)) | c_in(X94,v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | ~c_in(X94,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a)) ) | (~spl0_6 | ~spl0_49)),
  inference(resolution,[],[f1249,f186])).
fof(f186,plain,(
  c_in(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),v_S,tc_set(t_a)) | ~spl0_6),
  inference(avatar_component_clause,[],[f185])).
fof(f1249,plain,(
  ( ! [X2,X1] : (~c_in(X1,v_S,tc_set(t_a)) | c_in(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(v_S,t_a),X1,t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)) | c_in(X2,v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | ~c_in(X2,X1,t_a)) ) | ~spl0_49),
  inference(avatar_component_clause,[],[f1248])).
fof(f15580,plain,(
  spl0_50 | ~spl0_6 | ~spl0_17 | ~spl0_18),
  inference(avatar_split_clause,[],[f15577,f431,f428,f185,f1251])).
fof(f428,plain,(
  spl0_17 <=> ! [X79,X78] : (c_in(X78,X79,t_a) | ~c_in(X78,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a) | c_in(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(v_S,t_a),X79,t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_17])])).
fof(f15577,plain,(
  ( ! [X41] : (c_in(c_Main_OsubsetI__1(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),X41,t_a),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | c_lessequals(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),X41,tc_set(t_a))) ) | (~spl0_6 | ~spl0_17 | ~spl0_18)),
  inference(resolution,[],[f14250,f3])).
fof(f14250,plain,(
  ( ! [X7] : (~c_in(X7,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a) | c_in(X7,v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a)) ) | (~spl0_6 | ~spl0_17 | ~spl0_18)),
  inference(subsumption_resolution,[],[f12453,f429])).
fof(f429,plain,(
  ( ! [X78,X79] : (~c_in(X78,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a) | c_in(X78,X79,t_a) | c_in(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(v_S,t_a),X79,t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a))) ) | ~spl0_17),
  inference(avatar_component_clause,[],[f428])).
fof(f12453,plain,(
  ( ! [X7] : (c_in(X7,v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | ~c_in(X7,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a) | ~c_in(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(v_S,t_a),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a))) ) | (~spl0_6 | ~spl0_18)),
  inference(resolution,[],[f12348,f32])).
fof(f32,plain,(
  ( ! [X0] : (c_lessequals(X0,v_x(c_Zorn_OHausdorff__1(v_S,t_a)),tc_set(t_a)) | ~c_in(X0,c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a))) )),
  inference(resolution,[],[f29,f15])).
fof(f15,axiom,(
  ( ! [X8,X9] : (~c_in(X8,c_Zorn_Ochain(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X9,X8,tc_set(t_a)) | c_lessequals(X9,v_x(X8),tc_set(t_a))) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f29,plain,(
  ( ! [X0,X1] : (c_in(c_Zorn_OHausdorff__1(X0,X1),c_Zorn_Ochain(X0,X1),tc_set(tc_set(X1)))) )),
  inference(resolution,[],[f24,f7])).
fof(f7,axiom,(
  ( ! [X2,X5] : (c_in(c_Zorn_OHausdorff__1(X5,X2),c_Zorn_Omaxchain(X5,X2),tc_set(tc_set(X2)))) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f24,plain,(
  ( ! [X6,X7,X5] : (~c_in(X5,c_Zorn_Omaxchain(X6,X7),tc_set(tc_set(X7))) | c_in(X5,c_Zorn_Ochain(X6,X7),tc_set(tc_set(X7)))) )),
  inference(resolution,[],[f2,f8])).
fof(f8,axiom,(
  ( ! [X2,X5] : (c_lessequals(c_Zorn_Omaxchain(X5,X2),c_Zorn_Ochain(X5,X2),tc_set(tc_set(tc_set(X2))))) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f2,axiom,(
  ( ! [X2,X0,X3,X1] : (~c_lessequals(X1,X0,tc_set(X2)) | c_in(X3,X0,X2) | ~c_in(X3,X1,X2)) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f12348,plain,(
  ( ! [X0,X1] : (~c_lessequals(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(v_S,t_a),X0,t_a),X0,tc_set(t_a)) | c_in(X1,X0,t_a) | ~c_in(X1,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a)) ) | (~spl0_6 | ~spl0_18)),
  inference(subsumption_resolution,[],[f12340,f7])).
fof(f12340,plain,(
  ( ! [X0,X1] : (~c_lessequals(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(v_S,t_a),X0,t_a),X0,tc_set(t_a)) | c_in(X1,X0,t_a) | ~c_in(c_Zorn_OHausdorff__1(v_S,t_a),c_Zorn_Omaxchain(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X1,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a)) ) | (~spl0_6 | ~spl0_18)),
  inference(resolution,[],[f12177,f10])).
fof(f10,axiom,(
  ( ! [X6,X4,X2,X7,X5,X3] : (~c_in(c_union(c_insert(X6,c_emptyset,tc_set(X2)),X3,tc_set(X2)),c_Zorn_Ochain(X5,X2),tc_set(tc_set(X2))) | ~c_lessequals(c_Zorn_Omaxchain__super__lemma__1(X3,X7,X2),X7,tc_set(X2)) | c_in(X4,X7,X2) | ~c_in(X3,c_Zorn_Omaxchain(X5,X2),tc_set(tc_set(X2))) | ~c_in(X4,X6,X2)) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f12177,plain,(
  c_in(c_union(c_insert(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),c_emptyset,tc_set(t_a)),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)),c_Zorn_Ochain(v_S,t_a),tc_set(tc_set(t_a))) | (~spl0_6 | ~spl0_18)),
  inference(subsumption_resolution,[],[f12174,f186])).
fof(f12174,plain,(
  c_in(c_union(c_insert(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),c_emptyset,tc_set(t_a)),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)),c_Zorn_Ochain(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),v_S,tc_set(t_a)) | ~spl0_18),
  inference(resolution,[],[f12131,f35])).
fof(f35,plain,(
  ( ! [X2,X0,X1] : (~c_lessequals(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(X0,X1),X2,X1),X2,tc_set(X1)) | c_in(c_union(c_insert(X2,c_emptyset,tc_set(X1)),c_Zorn_OHausdorff__1(X0,X1),tc_set(X1)),c_Zorn_Ochain(X0,X1),tc_set(tc_set(X1))) | ~c_in(X2,X0,tc_set(X1))) )),
  inference(resolution,[],[f6,f29])).
fof(f6,axiom,(
  ( ! [X4,X2,X5,X3] : (~c_in(X3,c_Zorn_Ochain(X5,X2),tc_set(tc_set(X2))) | ~c_lessequals(c_Zorn_Ochain__extend__1(X3,X4,X2),X4,tc_set(X2)) | c_in(c_union(c_insert(X4,c_emptyset,tc_set(X2)),X3,tc_set(X2)),c_Zorn_Ochain(X5,X2),tc_set(tc_set(X2))) | ~c_in(X4,X5,tc_set(X2))) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f12131,plain,(
  c_lessequals(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(v_S,t_a),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),tc_set(t_a)) | ~spl0_18),
  inference(duplicate_literal_removal,[],[f12115])).
fof(f12115,plain,(
  c_lessequals(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(v_S,t_a),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),tc_set(t_a)) | c_lessequals(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(v_S,t_a),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),tc_set(t_a)) | ~spl0_18),
  inference(resolution,[],[f4391,f42])).
fof(f42,plain,(
  ( ! [X0] : (~c_in(c_Main_OsubsetI__1(X0,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | c_lessequals(X0,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),tc_set(t_a))) )),
  inference(resolution,[],[f37,f4])).
fof(f4,axiom,(
  ( ! [X2,X0,X1] : (~c_in(c_Main_OsubsetI__1(X1,X0,X2),X0,X2) | c_lessequals(X1,X0,tc_set(X2))) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f37,plain,(
  ( ! [X0] : (c_in(X0,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a) | ~c_in(X0,v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a)) )),
  inference(resolution,[],[f33,f22])).
fof(f22,plain,(
  ( ! [X0,X1] : (~c_in(X1,v_S,tc_set(t_a)) | ~c_in(X0,X1,t_a) | c_in(X0,v_xa(X1),t_a)) )),
  inference(resolution,[],[f2,f13])).
fof(f13,axiom,(
  ( ! [X8] : (c_lessequals(X8,v_xa(X8),tc_set(t_a)) | ~c_in(X8,v_S,tc_set(t_a))) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f33,plain,(
  c_in(v_x(c_Zorn_OHausdorff__1(v_S,t_a)),v_S,tc_set(t_a))),
  inference(resolution,[],[f29,f11])).
fof(f11,axiom,(
  ( ! [X8] : (~c_in(X8,c_Zorn_Ochain(v_S,t_a),tc_set(tc_set(t_a))) | c_in(v_x(X8),v_S,tc_set(t_a))) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f4391,plain,(
  ( ! [X0] : (c_in(c_Main_OsubsetI__1(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(v_S,t_a),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a),X0,t_a),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | c_lessequals(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(v_S,t_a),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a),X0,tc_set(t_a))) ) | ~spl0_18),
  inference(resolution,[],[f4226,f3])).
fof(f4226,plain,(
  ( ! [X0] : (~c_in(X0,c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(v_S,t_a),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a),t_a) | c_in(X0,v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a)) ) | ~spl0_18),
  inference(resolution,[],[f433,f38])).
fof(f38,plain,(
  ( ! [X0,X1] : (~c_in(X0,c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)) | c_in(X1,v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | ~c_in(X1,X0,t_a)) )),
  inference(resolution,[],[f32,f2])).
fof(f433,plain,(
  c_in(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(v_S,t_a),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)) | ~spl0_18),
  inference(avatar_component_clause,[],[f431])).
fof(f12350,plain,(
  spl0_17 | ~spl0_6 | ~spl0_18),
  inference(avatar_split_clause,[],[f12349,f431,f185,f428])).
fof(f12349,plain,(
  ( ! [X2,X3] : (c_in(X2,X3,t_a) | c_in(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(v_S,t_a),X3,t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)) | ~c_in(X2,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a)) ) | (~spl0_6 | ~spl0_18)),
  inference(subsumption_resolution,[],[f12341,f7])).
fof(f12341,plain,(
  ( ! [X2,X3] : (c_in(X2,X3,t_a) | c_in(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(v_S,t_a),X3,t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)) | ~c_in(c_Zorn_OHausdorff__1(v_S,t_a),c_Zorn_Omaxchain(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X2,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a)) ) | (~spl0_6 | ~spl0_18)),
  inference(resolution,[],[f12177,f9])).
fof(f9,axiom,(
  ( ! [X6,X4,X2,X7,X5,X3] : (~c_in(c_union(c_insert(X6,c_emptyset,tc_set(X2)),X3,tc_set(X2)),c_Zorn_Ochain(X5,X2),tc_set(tc_set(X2))) | c_in(X4,X7,X2) | c_in(c_Zorn_Omaxchain__super__lemma__1(X3,X7,X2),X3,tc_set(X2)) | ~c_in(X3,c_Zorn_Omaxchain(X5,X2),tc_set(tc_set(X2))) | ~c_in(X4,X6,X2)) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f2927,plain,(
  ~spl0_2),
  inference(avatar_contradiction_clause,[],[f2926])).
fof(f2926,plain,(
  $false | ~spl0_2),
  inference(subsumption_resolution,[],[f2925,f33])).
fof(f2925,plain,(
  ~c_in(v_x(c_Zorn_OHausdorff__1(v_S,t_a)),v_S,tc_set(t_a)) | ~spl0_2),
  inference(subsumption_resolution,[],[f2895,f17])).
fof(f17,plain,(
  ( ! [X0,X1] : (c_lessequals(X0,X0,tc_set(X1))) )),
  inference(duplicate_literal_removal,[],[f16])).
fof(f16,plain,(
  ( ! [X0,X1] : (c_lessequals(X0,X0,tc_set(X1)) | c_lessequals(X0,X0,tc_set(X1))) )),
  inference(resolution,[],[f4,f3])).
fof(f2895,plain,(
  ~c_lessequals(v_x(c_Zorn_OHausdorff__1(v_S,t_a)),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),tc_set(t_a)) | ~c_in(v_x(c_Zorn_OHausdorff__1(v_S,t_a)),v_S,tc_set(t_a)) | ~spl0_2),
  inference(superposition,[],[f20,f99])).
fof(f99,plain,(
  v_x(c_Zorn_OHausdorff__1(v_S,t_a)) = v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))) | ~spl0_2),
  inference(avatar_component_clause,[],[f97])).
fof(f97,plain,(
  spl0_2 <=> v_x(c_Zorn_OHausdorff__1(v_S,t_a)) = v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f20,plain,(
  ( ! [X0] : (~c_lessequals(v_xa(X0),X0,tc_set(t_a)) | ~c_in(X0,v_S,tc_set(t_a))) )),
  inference(subsumption_resolution,[],[f18,f14])).
fof(f14,axiom,(
  ( ! [X8] : (v_xa(X8) != X8 | ~c_in(X8,v_S,tc_set(t_a))) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f18,plain,(
  ( ! [X0] : (v_xa(X0) = X0 | ~c_lessequals(v_xa(X0),X0,tc_set(t_a)) | ~c_in(X0,v_S,tc_set(t_a))) )),
  inference(resolution,[],[f1,f13])).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (~c_lessequals(X1,X0,tc_set(X2)) | X0 = X1 | ~c_lessequals(X0,X1,tc_set(X2))) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f2624,plain,(
  ~spl0_1 | spl0_2),
  inference(avatar_split_clause,[],[f2174,f97,f93])).
fof(f93,plain,(
  spl0_1 <=> c_lessequals(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),tc_set(t_a))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f2174,plain,(
  v_x(c_Zorn_OHausdorff__1(v_S,t_a)) = v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))) | ~c_lessequals(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),tc_set(t_a))),
  inference(resolution,[],[f85,f1])).
fof(f85,plain,(
  c_lessequals(v_x(c_Zorn_OHausdorff__1(v_S,t_a)),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),tc_set(t_a))),
  inference(duplicate_literal_removal,[],[f84])).
fof(f84,plain,(
  c_lessequals(v_x(c_Zorn_OHausdorff__1(v_S,t_a)),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),tc_set(t_a)) | c_lessequals(v_x(c_Zorn_OHausdorff__1(v_S,t_a)),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),tc_set(t_a))),
  inference(resolution,[],[f42,f3])).
fof(f2581,plain,(
  spl0_17 | spl0_18 | ~spl0_6),
  inference(avatar_split_clause,[],[f2132,f185,f431,f428])).
fof(f2132,plain,(
  ( ! [X4,X5] : (c_in(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(v_S,t_a),v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)) | c_in(X4,X5,t_a) | c_in(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(v_S,t_a),X5,t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)) | ~c_in(X4,v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),t_a)) ) | ~spl0_6),
  inference(resolution,[],[f186,f199])).
fof(f199,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~c_in(X7,X5,tc_set(X6)) | c_in(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(X5,X6),X7,X6),c_Zorn_OHausdorff__1(X5,X6),tc_set(X6)) | c_in(X8,X9,X6) | c_in(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(X5,X6),X9,X6),c_Zorn_OHausdorff__1(X5,X6),tc_set(X6)) | ~c_in(X8,X7,X6)) )),
  inference(subsumption_resolution,[],[f193,f7])).
fof(f193,plain,(
  ( ! [X6,X8,X7,X5,X9] : (c_in(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(X5,X6),X7,X6),c_Zorn_OHausdorff__1(X5,X6),tc_set(X6)) | ~c_in(X7,X5,tc_set(X6)) | c_in(X8,X9,X6) | c_in(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(X5,X6),X9,X6),c_Zorn_OHausdorff__1(X5,X6),tc_set(X6)) | ~c_in(c_Zorn_OHausdorff__1(X5,X6),c_Zorn_Omaxchain(X5,X6),tc_set(tc_set(X6))) | ~c_in(X8,X7,X6)) )),
  inference(resolution,[],[f34,f9])).
fof(f34,plain,(
  ( ! [X2,X3,X1] : (c_in(c_union(c_insert(X3,c_emptyset,tc_set(X2)),c_Zorn_OHausdorff__1(X1,X2),tc_set(X2)),c_Zorn_Ochain(X1,X2),tc_set(tc_set(X2))) | c_in(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(X1,X2),X3,X2),c_Zorn_OHausdorff__1(X1,X2),tc_set(X2)) | ~c_in(X3,X1,tc_set(X2))) )),
  inference(resolution,[],[f29,f5])).
fof(f5,axiom,(
  ( ! [X4,X2,X5,X3] : (~c_in(X3,c_Zorn_Ochain(X5,X2),tc_set(tc_set(X2))) | c_in(c_Zorn_Ochain__extend__1(X3,X4,X2),X3,tc_set(X2)) | c_in(c_union(c_insert(X4,c_emptyset,tc_set(X2)),X3,tc_set(X2)),c_Zorn_Ochain(X5,X2),tc_set(tc_set(X2))) | ~c_in(X4,X5,tc_set(X2))) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f1329,plain,(
  spl0_1 | ~spl0_50),
  inference(avatar_split_clause,[],[f1321,f1251,f93])).
fof(f1321,plain,(
  c_lessequals(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),tc_set(t_a)) | ~spl0_50),
  inference(duplicate_literal_removal,[],[f1319])).
fof(f1319,plain,(
  c_lessequals(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),tc_set(t_a)) | c_lessequals(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),tc_set(t_a)) | ~spl0_50),
  inference(resolution,[],[f1252,f4])).
fof(f1252,plain,(
  ( ! [X0] : (c_in(c_Main_OsubsetI__1(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),X0,t_a),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | c_lessequals(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),X0,tc_set(t_a))) ) | ~spl0_50),
  inference(avatar_component_clause,[],[f1251])).
fof(f1253,plain,(
  spl0_49 | spl0_50 | ~spl0_17),
  inference(avatar_split_clause,[],[f1178,f428,f1251,f1248])).
fof(f1178,plain,(
  ( ! [X2,X0,X1] : (c_in(c_Main_OsubsetI__1(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),X0,t_a),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | c_lessequals(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),X0,tc_set(t_a)) | c_in(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(v_S,t_a),X1,t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)) | c_in(X2,v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | ~c_in(X2,X1,t_a) | ~c_in(X1,v_S,tc_set(t_a))) ) | ~spl0_17),
  inference(resolution,[],[f504,f387])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (~c_in(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(X1,t_a),v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)) | c_in(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(X1,t_a),X0,t_a),c_Zorn_OHausdorff__1(X1,t_a),tc_set(t_a)) | c_in(X2,v_x(c_Zorn_OHausdorff__1(v_S,t_a)),t_a) | ~c_in(X2,X0,t_a) | ~c_in(X0,X1,tc_set(t_a))) )),
  inference(resolution,[],[f198,f32])).
fof(f198,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~c_lessequals(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(X0,X1),X3,X1),X3,tc_set(X1)) | ~c_in(X2,X0,tc_set(X1)) | c_in(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(X0,X1),X2,X1),c_Zorn_OHausdorff__1(X0,X1),tc_set(X1)) | c_in(X4,X3,X1) | ~c_in(X4,X2,X1)) )),
  inference(subsumption_resolution,[],[f192,f7])).
fof(f192,plain,(
  ( ! [X4,X2,X0,X3,X1] : (c_in(c_Zorn_Ochain__extend__1(c_Zorn_OHausdorff__1(X0,X1),X2,X1),c_Zorn_OHausdorff__1(X0,X1),tc_set(X1)) | ~c_in(X2,X0,tc_set(X1)) | ~c_lessequals(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(X0,X1),X3,X1),X3,tc_set(X1)) | c_in(X4,X3,X1) | ~c_in(c_Zorn_OHausdorff__1(X0,X1),c_Zorn_Omaxchain(X0,X1),tc_set(tc_set(X1))) | ~c_in(X4,X2,X1)) )),
  inference(resolution,[],[f34,f10])).
fof(f504,plain,(
  ( ! [X2,X3] : (c_in(c_Zorn_Omaxchain__super__lemma__1(c_Zorn_OHausdorff__1(v_S,t_a),X3,t_a),c_Zorn_OHausdorff__1(v_S,t_a),tc_set(t_a)) | c_in(c_Main_OsubsetI__1(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),X2,t_a),X3,t_a) | c_lessequals(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),X2,tc_set(t_a))) ) | ~spl0_17),
  inference(resolution,[],[f429,f3])).
fof(f202,plain,(
  spl0_6),
  inference(avatar_contradiction_clause,[],[f201])).
fof(f201,plain,(
  $false | spl0_6),
  inference(subsumption_resolution,[],[f200,f33])).
fof(f200,plain,(
  ~c_in(v_x(c_Zorn_OHausdorff__1(v_S,t_a)),v_S,tc_set(t_a)) | spl0_6),
  inference(resolution,[],[f187,f12])).
fof(f12,axiom,(
  ( ! [X8] : (c_in(v_xa(X8),v_S,tc_set(t_a)) | ~c_in(X8,v_S,tc_set(t_a))) )),
  file('./p/SET861-2.tptp',unknown)).
fof(f187,plain,(
  ~c_in(v_xa(v_x(c_Zorn_OHausdorff__1(v_S,t_a))),v_S,tc_set(t_a)) | spl0_6),
  inference(avatar_component_clause,[],[f185])).
% SZS output end Proof for SET861-2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 32877
% Time elapsed: 0.923 s
% ------------------------------
% ------------------------------
