% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SET857-2
% SZS output start Proof for SET857-2
fof(f4431,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f33,f851,f4430])).
fof(f4430,plain,(
  spl0_2),
  inference(avatar_contradiction_clause,[],[f4429])).
fof(f4429,plain,(
  $false | spl0_2),
  inference(subsumption_resolution,[],[f4425,f12])).
fof(f12,axiom,(
  ~c_lessequals(v_x,v_n,tc_set(tc_set(t_a)))),
  file('./p/SET857-2.tptp',unknown)).
fof(f4425,plain,(
  c_lessequals(v_x,v_n,tc_set(tc_set(t_a))) | spl0_2),
  inference(duplicate_literal_removal,[],[f4420])).
fof(f4420,plain,(
  c_lessequals(v_x,v_n,tc_set(tc_set(t_a))) | c_lessequals(v_x,v_n,tc_set(tc_set(t_a))) | spl0_2),
  inference(resolution,[],[f4398,f3])).
fof(f3,axiom,(
  ( ! [X2,X3,X1] : (~c_in(c_Main_OsubsetI__1(X1,X3,X2),X3,X2) | c_lessequals(X1,X3,tc_set(X2))) )),
  file('./p/SET857-2.tptp',unknown)).
fof(f4398,plain,(
  ( ! [X33] : (c_in(c_Main_OsubsetI__1(v_x,X33,tc_set(t_a)),v_n,tc_set(t_a)) | c_lessequals(v_x,X33,tc_set(tc_set(t_a)))) ) | spl0_2),
  inference(resolution,[],[f4250,f2])).
fof(f2,axiom,(
  ( ! [X2,X3,X1] : (c_in(c_Main_OsubsetI__1(X1,X3,X2),X1,X2) | c_lessequals(X1,X3,tc_set(X2))) )),
  file('./p/SET857-2.tptp',unknown)).
fof(f4250,plain,(
  ( ! [X0] : (~c_in(X0,v_x,tc_set(t_a)) | c_in(X0,v_n,tc_set(t_a))) ) | spl0_2),
  inference(subsumption_resolution,[],[f4242,f32])).
fof(f32,plain,(
  ~c_lessequals(v_n,v_x,tc_set(tc_set(t_a))) | spl0_2),
  inference(avatar_component_clause,[],[f30])).
fof(f30,plain,(
  spl0_2 <=> c_lessequals(v_n,v_x,tc_set(tc_set(t_a)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f4242,plain,(
  ( ! [X0] : (c_in(X0,v_n,tc_set(t_a)) | c_lessequals(v_n,v_x,tc_set(tc_set(t_a))) | ~c_in(X0,v_x,tc_set(t_a))) )),
  inference(resolution,[],[f4141,f11])).
fof(f11,axiom,(
  c_in(v_x,v_Y,tc_set(tc_set(t_a)))),
  file('./p/SET857-2.tptp',unknown)).
fof(f4141,plain,(
  ( ! [X17,X18] : (~c_in(X17,v_Y,tc_set(tc_set(t_a))) | c_in(X18,v_n,tc_set(t_a)) | c_lessequals(v_n,X17,tc_set(tc_set(t_a))) | ~c_in(X18,X17,tc_set(t_a))) )),
  inference(resolution,[],[f533,f10])).
fof(f10,axiom,(
  c_in(v_n,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a)))),
  file('./p/SET857-2.tptp',unknown)).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (~c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X1,v_Y,tc_set(tc_set(t_a))) | c_in(X2,X0,tc_set(t_a)) | c_lessequals(X0,X1,tc_set(tc_set(t_a))) | ~c_in(X2,X1,tc_set(t_a))) )),
  inference(duplicate_literal_removal,[],[f520])).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (~c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X1,v_Y,tc_set(tc_set(t_a))) | c_in(X2,X0,tc_set(t_a)) | c_lessequals(X0,X1,tc_set(tc_set(t_a))) | ~c_in(X2,X1,tc_set(t_a)) | ~c_in(X1,v_Y,tc_set(tc_set(t_a)))) )),
  inference(resolution,[],[f128,f17])).
fof(f17,plain,(
  ( ! [X0] : (c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X0,v_Y,tc_set(tc_set(t_a)))) )),
  inference(resolution,[],[f1,f9])).
fof(f9,axiom,(
  c_lessequals(v_Y,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(tc_set(t_a))))),
  file('./p/SET857-2.tptp',unknown)).
fof(f1,axiom,(
  ( ! [X2,X0,X3,X1] : (~c_lessequals(X1,X3,tc_set(X2)) | c_in(X0,X3,X2) | ~c_in(X0,X1,X2)) )),
  file('./p/SET857-2.tptp',unknown)).
fof(f128,plain,(
  ( ! [X2,X0,X1] : (~c_in(X1,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X1,v_Y,tc_set(tc_set(t_a))) | c_in(X2,X0,tc_set(t_a)) | c_lessequals(X0,X1,tc_set(tc_set(t_a))) | ~c_in(X2,X1,tc_set(t_a))) )),
  inference(resolution,[],[f78,f18])).
fof(f18,plain,(
  ( ! [X4,X2,X3,X1] : (c_in(X1,c_Zorn_Osucc(X2,X3,X4),tc_set(X4)) | ~c_in(X1,X3,tc_set(X4))) )),
  inference(resolution,[],[f1,f4])).
fof(f4,axiom,(
  ( ! [X4,X2,X5] : (c_lessequals(X4,c_Zorn_Osucc(X5,X4,X2),tc_set(tc_set(X2)))) )),
  file('./p/SET857-2.tptp',unknown)).
fof(f78,plain,(
  ( ! [X4,X2,X3] : (~c_in(X4,c_Zorn_Osucc(v_S,X3,t_a),tc_set(t_a)) | ~c_in(X2,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X3,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X3,v_Y,tc_set(tc_set(t_a))) | c_in(X4,X2,tc_set(t_a)) | c_lessequals(X2,X3,tc_set(tc_set(t_a)))) )),
  inference(resolution,[],[f65,f1])).
fof(f65,plain,(
  ( ! [X0,X1] : (c_lessequals(c_Zorn_Osucc(v_S,X1,t_a),X0,tc_set(tc_set(t_a))) | c_lessequals(X0,X1,tc_set(tc_set(t_a))) | ~c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X1,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X1,v_Y,tc_set(tc_set(t_a)))) )),
  inference(subsumption_resolution,[],[f64,f6])).
fof(f6,axiom,(
  ( ! [X6,X2,X7,X5] : (c_lessequals(c_Zorn_OTFin__linear__lemma1__1(X5,X6,X2),X6,tc_set(tc_set(X2))) | c_lessequals(c_Zorn_Osucc(X5,X6,X2),X7,tc_set(tc_set(X2))) | c_lessequals(X7,X6,tc_set(tc_set(X2))) | ~c_in(X7,c_Zorn_OTFin(X5,X2),tc_set(tc_set(X2))) | ~c_in(X6,c_Zorn_OTFin(X5,X2),tc_set(tc_set(X2)))) )),
  file('./p/SET857-2.tptp',unknown)).
fof(f64,plain,(
  ( ! [X0,X1] : (c_lessequals(X0,X1,tc_set(tc_set(t_a))) | c_lessequals(c_Zorn_Osucc(v_S,X1,t_a),X0,tc_set(tc_set(t_a))) | ~c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X1,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_lessequals(c_Zorn_OTFin__linear__lemma1__1(v_S,X1,t_a),X1,tc_set(tc_set(t_a))) | ~c_in(X1,v_Y,tc_set(tc_set(t_a)))) )),
  inference(subsumption_resolution,[],[f63,f5])).
fof(f5,axiom,(
  ( ! [X6,X2,X7,X5] : (c_in(c_Zorn_OTFin__linear__lemma1__1(X5,X6,X2),c_Zorn_OTFin(X5,X2),tc_set(tc_set(X2))) | c_lessequals(c_Zorn_Osucc(X5,X6,X2),X7,tc_set(tc_set(X2))) | c_lessequals(X7,X6,tc_set(tc_set(X2))) | ~c_in(X7,c_Zorn_OTFin(X5,X2),tc_set(tc_set(X2))) | ~c_in(X6,c_Zorn_OTFin(X5,X2),tc_set(tc_set(X2)))) )),
  file('./p/SET857-2.tptp',unknown)).
fof(f63,plain,(
  ( ! [X0,X1] : (c_lessequals(X0,X1,tc_set(tc_set(t_a))) | c_lessequals(c_Zorn_Osucc(v_S,X1,t_a),X0,tc_set(tc_set(t_a))) | ~c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X1,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(c_Zorn_OTFin__linear__lemma1__1(v_S,X1,t_a),c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_lessequals(c_Zorn_OTFin__linear__lemma1__1(v_S,X1,t_a),X1,tc_set(tc_set(t_a))) | ~c_in(X1,v_Y,tc_set(tc_set(t_a)))) )),
  inference(subsumption_resolution,[],[f58,f7])).
fof(f7,axiom,(
  ( ! [X6,X2,X7,X5] : (c_Zorn_OTFin__linear__lemma1__1(X5,X6,X2) != X6 | c_lessequals(X7,X6,tc_set(tc_set(X2))) | c_lessequals(c_Zorn_Osucc(X5,X6,X2),X7,tc_set(tc_set(X2))) | ~c_in(X7,c_Zorn_OTFin(X5,X2),tc_set(tc_set(X2))) | ~c_in(X6,c_Zorn_OTFin(X5,X2),tc_set(tc_set(X2)))) )),
  file('./p/SET857-2.tptp',unknown)).
fof(f58,plain,(
  ( ! [X0,X1] : (c_lessequals(X0,X1,tc_set(tc_set(t_a))) | c_lessequals(c_Zorn_Osucc(v_S,X1,t_a),X0,tc_set(tc_set(t_a))) | ~c_in(X0,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(X1,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_in(c_Zorn_OTFin__linear__lemma1__1(v_S,X1,t_a),c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_lessequals(c_Zorn_OTFin__linear__lemma1__1(v_S,X1,t_a),X1,tc_set(tc_set(t_a))) | c_Zorn_OTFin__linear__lemma1__1(v_S,X1,t_a) = X1 | ~c_in(X1,v_Y,tc_set(tc_set(t_a)))) )),
  inference(resolution,[],[f8,f14])).
fof(f14,axiom,(
  ( ! [X8,X9] : (c_lessequals(c_Zorn_Osucc(v_S,X8,t_a),X9,tc_set(tc_set(t_a))) | ~c_in(X8,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_lessequals(X8,X9,tc_set(tc_set(t_a))) | X8 = X9 | ~c_in(X9,v_Y,tc_set(tc_set(t_a)))) )),
  file('./p/SET857-2.tptp',unknown)).
fof(f8,axiom,(
  ( ! [X6,X2,X7,X5] : (~c_lessequals(c_Zorn_Osucc(X5,c_Zorn_OTFin__linear__lemma1__1(X5,X6,X2),X2),X6,tc_set(tc_set(X2))) | c_lessequals(X7,X6,tc_set(tc_set(X2))) | c_lessequals(c_Zorn_Osucc(X5,X6,X2),X7,tc_set(tc_set(X2))) | ~c_in(X7,c_Zorn_OTFin(X5,X2),tc_set(tc_set(X2))) | ~c_in(X6,c_Zorn_OTFin(X5,X2),tc_set(tc_set(X2)))) )),
  file('./p/SET857-2.tptp',unknown)).
fof(f851,plain,(
  ~spl0_1),
  inference(avatar_contradiction_clause,[],[f850])).
fof(f850,plain,(
  $false | ~spl0_1),
  inference(subsumption_resolution,[],[f812,f16])).
fof(f16,plain,(
  ( ! [X0,X1] : (c_lessequals(X0,X0,tc_set(X1))) )),
  inference(duplicate_literal_removal,[],[f15])).
fof(f15,plain,(
  ( ! [X0,X1] : (c_lessequals(X0,X0,tc_set(X1)) | c_lessequals(X0,X0,tc_set(X1))) )),
  inference(resolution,[],[f3,f2])).
fof(f812,plain,(
  ~c_lessequals(v_n,v_n,tc_set(tc_set(t_a))) | ~spl0_1),
  inference(backward_demodulation,[],[f28,f12])).
fof(f28,plain,(
  v_n = v_x | ~spl0_1),
  inference(avatar_component_clause,[],[f26])).
fof(f26,plain,(
  spl0_1 <=> v_n = v_x),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f33,plain,(
  spl0_1 | ~spl0_2),
  inference(avatar_split_clause,[],[f24,f30,f26])).
fof(f24,plain,(
  ~c_lessequals(v_n,v_x,tc_set(tc_set(t_a))) | v_n = v_x),
  inference(subsumption_resolution,[],[f23,f11])).
fof(f23,plain,(
  ~c_lessequals(v_n,v_x,tc_set(tc_set(t_a))) | v_n = v_x | ~c_in(v_x,v_Y,tc_set(tc_set(t_a)))),
  inference(subsumption_resolution,[],[f21,f10])).
fof(f21,plain,(
  ~c_in(v_n,c_Zorn_OTFin(v_S,t_a),tc_set(tc_set(t_a))) | ~c_lessequals(v_n,v_x,tc_set(tc_set(t_a))) | v_n = v_x | ~c_in(v_x,v_Y,tc_set(tc_set(t_a)))),
  inference(resolution,[],[f14,f13])).
fof(f13,axiom,(
  ~c_lessequals(c_Zorn_Osucc(v_S,v_n,t_a),v_x,tc_set(tc_set(t_a)))),
  file('./p/SET857-2.tptp',unknown)).
% SZS output end Proof for SET857-2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 13304
% Time elapsed: 0.512 s
% ------------------------------
% ------------------------------
