% Refutation found. Thanks to Tanya!
% SZS status Theorem for SEU158+2
% SZS output start Proof for SEU158+2
fof(f623,plain,(
  $false),
  inference(subsumption_resolution,[],[f615,f430])).
fof(f430,plain,(
  in(sK19,sK20)),
  inference(subsumption_resolution,[],[f389,f429])).
fof(f429,plain,(
  ~subset(unordered_pair(sK19,sK19),sK20)),
  inference(subsumption_resolution,[],[f388,f372])).
fof(f372,plain,(
  ( ! [X0,X1] : (in(X0,X1) | ~subset(unordered_pair(X0,X0),X1)) )),
  inference(definition_unfolding,[],[f285,f336])).
fof(f336,plain,(
  ( ! [X0] : (singleton(X0) = unordered_pair(X0,X0)) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,axiom,(
  ! [X0] : singleton(X0) = unordered_pair(X0,X0)),
  file('./p/SEU158+2.tptp',unknown)).
fof(f285,plain,(
  ( ! [X0,X1] : (in(X0,X1) | ~subset(singleton(X0),X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((subset(singleton(X0),X1) | ~in(X0,X1)) & (in(X0,X1) | ~subset(singleton(X0),X1)))),
  inference(nnf_transformation,[],[f41])).
fof(f41,axiom,(
  ! [X0,X1] : (subset(singleton(X0),X1) <=> in(X0,X1))),
  file('./p/SEU158+2.tptp',unknown)).
fof(f388,plain,(
  ~in(sK19,sK20) | ~subset(unordered_pair(sK19,sK19),sK20)),
  inference(definition_unfolding,[],[f321,f336])).
fof(f321,plain,(
  ~in(sK19,sK20) | ~subset(singleton(sK19),sK20)),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  (~in(sK19,sK20) | ~subset(singleton(sK19),sK20)) & (in(sK19,sK20) | subset(singleton(sK19),sK20))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20])],[f203,f204])).
fof(f204,plain,(
  ? [X0,X1] : ((~in(X0,X1) | ~subset(singleton(X0),X1)) & (in(X0,X1) | subset(singleton(X0),X1))) => ((~in(sK19,sK20) | ~subset(singleton(sK19),sK20)) & (in(sK19,sK20) | subset(singleton(sK19),sK20)))),
  introduced(choice_axiom,[])).
fof(f203,plain,(
  ? [X0,X1] : ((~in(X0,X1) | ~subset(singleton(X0),X1)) & (in(X0,X1) | subset(singleton(X0),X1)))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ? [X0,X1] : (subset(singleton(X0),X1) <~> in(X0,X1))),
  inference(ennf_transformation,[],[f68])).
fof(f68,negated_conjecture,(
  ~! [X0,X1] : (subset(singleton(X0),X1) <=> in(X0,X1))),
  inference(negated_conjecture,[],[f67])).
fof(f67,conjecture,(
  ! [X0,X1] : (subset(singleton(X0),X1) <=> in(X0,X1))),
  file('./p/SEU158+2.tptp',unknown)).
fof(f389,plain,(
  in(sK19,sK20) | subset(unordered_pair(sK19,sK19),sK20)),
  inference(definition_unfolding,[],[f320,f336])).
fof(f320,plain,(
  in(sK19,sK20) | subset(singleton(sK19),sK20)),
  inference(cnf_transformation,[],[f205])).
fof(f615,plain,(
  ~in(sK19,sK20)),
  inference(resolution,[],[f371,f429])).
fof(f371,plain,(
  ( ! [X0,X1] : (subset(unordered_pair(X0,X0),X1) | ~in(X0,X1)) )),
  inference(definition_unfolding,[],[f286,f336])).
fof(f286,plain,(
  ( ! [X0,X1] : (subset(singleton(X0),X1) | ~in(X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
% SZS output end Proof for SEU158+2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.005 s
% ------------------------------
% ------------------------------
