% Refutation found. Thanks to Tanya!
% SZS status Theorem for SEU191+1
% SZS output start Proof for SEU191+1
fof(f730,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f179,f182,f184,f307,f617,f680,f729])).
fof(f729,plain,(
  ~spl15_1 | ~spl15_2 | spl15_3 | ~spl15_4),
  inference(avatar_contradiction_clause,[],[f728])).
fof(f728,plain,(
  $false | (~spl15_1 | ~spl15_2 | spl15_3 | ~spl15_4)),
  inference(subsumption_resolution,[],[f727,f92])).
fof(f92,plain,(
  ( ! [X0] : (relation(identity_relation(X0))) )),
  inference(cnf_transformation,[],[f12])).
fof(f12,axiom,(
  ! [X0] : relation(identity_relation(X0))),
  file('./p/SEU191+1.tptp',unknown)).
fof(f727,plain,(
  ~relation(identity_relation(sK13)) | (~spl15_1 | ~spl15_2 | spl15_3 | ~spl15_4)),
  inference(subsumption_resolution,[],[f726,f169])).
fof(f169,plain,(
  in(sK11,sK13) | ~spl15_1),
  inference(avatar_component_clause,[],[f168])).
fof(f168,plain,(
  spl15_1 <=> in(sK11,sK13)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_1])])).
fof(f726,plain,(
  ~in(sK11,sK13) | ~relation(identity_relation(sK13)) | (~spl15_2 | spl15_3 | ~spl15_4)),
  inference(resolution,[],[f724,f145])).
fof(f145,plain,(
  ( ! [X0,X5] : (in(unordered_pair(singleton(X5),unordered_pair(X5,X5)),identity_relation(X0)) | ~in(X5,X0) | ~relation(identity_relation(X0))) )),
  inference(forward_demodulation,[],[f136,f77])).
fof(f77,plain,(
  ( ! [X0,X1] : (unordered_pair(X0,X1) = unordered_pair(X1,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1] : unordered_pair(X0,X1) = unordered_pair(X1,X0)),
  file('./p/SEU191+1.tptp',unknown)).
fof(f136,plain,(
  ( ! [X0,X5] : (in(unordered_pair(unordered_pair(X5,X5),singleton(X5)),identity_relation(X0)) | ~in(X5,X0) | ~relation(identity_relation(X0))) )),
  inference(equality_resolution,[],[f135])).
fof(f135,plain,(
  ( ! [X0,X5,X1] : (in(unordered_pair(unordered_pair(X5,X5),singleton(X5)),X1) | ~in(X5,X0) | identity_relation(X0) != X1 | ~relation(X1)) )),
  inference(equality_resolution,[],[f122])).
fof(f122,plain,(
  ( ! [X4,X0,X5,X1] : (in(unordered_pair(unordered_pair(X4,X5),singleton(X4)),X1) | X4 != X5 | ~in(X4,X0) | identity_relation(X0) != X1 | ~relation(X1)) )),
  inference(definition_unfolding,[],[f80,f84])).
fof(f84,plain,(
  ( ! [X0,X1] : (ordered_pair(X0,X1) = unordered_pair(unordered_pair(X0,X1),singleton(X0))) )),
  inference(cnf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1] : ordered_pair(X0,X1) = unordered_pair(unordered_pair(X0,X1),singleton(X0))),
  file('./p/SEU191+1.tptp',unknown)).
fof(f80,plain,(
  ( ! [X4,X0,X5,X1] : (in(ordered_pair(X4,X5),X1) | X4 != X5 | ~in(X4,X0) | identity_relation(X0) != X1 | ~relation(X1)) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : (((identity_relation(X0) = X1 | ((sK0(X0,X1) != sK1(X0,X1) | ~in(sK0(X0,X1),X0) | ~in(ordered_pair(sK0(X0,X1),sK1(X0,X1)),X1)) & ((sK0(X0,X1) = sK1(X0,X1) & in(sK0(X0,X1),X0)) | in(ordered_pair(sK0(X0,X1),sK1(X0,X1)),X1)))) & (! [X4,X5] : ((in(ordered_pair(X4,X5),X1) | X4 != X5 | ~in(X4,X0)) & ((X4 = X5 & in(X4,X0)) | ~in(ordered_pair(X4,X5),X1))) | identity_relation(X0) != X1)) | ~relation(X1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f52,f53])).
fof(f53,plain,(
  ! [X1,X0] : (? [X2,X3] : ((X2 != X3 | ~in(X2,X0) | ~in(ordered_pair(X2,X3),X1)) & ((X2 = X3 & in(X2,X0)) | in(ordered_pair(X2,X3),X1))) => ((sK0(X0,X1) != sK1(X0,X1) | ~in(sK0(X0,X1),X0) | ~in(ordered_pair(sK0(X0,X1),sK1(X0,X1)),X1)) & ((sK0(X0,X1) = sK1(X0,X1) & in(sK0(X0,X1),X0)) | in(ordered_pair(sK0(X0,X1),sK1(X0,X1)),X1))))),
  introduced(choice_axiom,[])).
fof(f52,plain,(
  ! [X0,X1] : (((identity_relation(X0) = X1 | ? [X2,X3] : ((X2 != X3 | ~in(X2,X0) | ~in(ordered_pair(X2,X3),X1)) & ((X2 = X3 & in(X2,X0)) | in(ordered_pair(X2,X3),X1)))) & (! [X4,X5] : ((in(ordered_pair(X4,X5),X1) | X4 != X5 | ~in(X4,X0)) & ((X4 = X5 & in(X4,X0)) | ~in(ordered_pair(X4,X5),X1))) | identity_relation(X0) != X1)) | ~relation(X1))),
  inference(rectify,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : (((identity_relation(X0) = X1 | ? [X2,X3] : ((X2 != X3 | ~in(X2,X0) | ~in(ordered_pair(X2,X3),X1)) & ((X2 = X3 & in(X2,X0)) | in(ordered_pair(X2,X3),X1)))) & (! [X2,X3] : ((in(ordered_pair(X2,X3),X1) | X2 != X3 | ~in(X2,X0)) & ((X2 = X3 & in(X2,X0)) | ~in(ordered_pair(X2,X3),X1))) | identity_relation(X0) != X1)) | ~relation(X1))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : (((identity_relation(X0) = X1 | ? [X2,X3] : (((X2 != X3 | ~in(X2,X0)) | ~in(ordered_pair(X2,X3),X1)) & ((X2 = X3 & in(X2,X0)) | in(ordered_pair(X2,X3),X1)))) & (! [X2,X3] : ((in(ordered_pair(X2,X3),X1) | (X2 != X3 | ~in(X2,X0))) & ((X2 = X3 & in(X2,X0)) | ~in(ordered_pair(X2,X3),X1))) | identity_relation(X0) != X1)) | ~relation(X1))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : ((identity_relation(X0) = X1 <=> ! [X2,X3] : (in(ordered_pair(X2,X3),X1) <=> (X2 = X3 & in(X2,X0)))) | ~relation(X1))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1] : (relation(X1) => (identity_relation(X0) = X1 <=> ! [X2,X3] : (in(ordered_pair(X2,X3),X1) <=> (X2 = X3 & in(X2,X0)))))),
  file('./p/SEU191+1.tptp',unknown)).
fof(f724,plain,(
  ~in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK11)),identity_relation(sK13)) | (~spl15_2 | spl15_3 | ~spl15_4)),
  inference(subsumption_resolution,[],[f723,f92])).
fof(f723,plain,(
  ~in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK11)),identity_relation(sK13)) | ~relation(identity_relation(sK13)) | (~spl15_2 | spl15_3 | ~spl15_4)),
  inference(subsumption_resolution,[],[f710,f290])).
fof(f290,plain,(
  relation(relation_composition(identity_relation(sK13),sK14)) | ~spl15_4),
  inference(avatar_component_clause,[],[f289])).
fof(f289,plain,(
  spl15_4 <=> relation(relation_composition(identity_relation(sK13),sK14))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_4])])).
fof(f710,plain,(
  ~in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK11)),identity_relation(sK13)) | ~relation(relation_composition(identity_relation(sK13),sK14)) | ~relation(identity_relation(sK13)) | (~spl15_2 | spl15_3)),
  inference(resolution,[],[f697,f178])).
fof(f178,plain,(
  ~in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),relation_composition(identity_relation(sK13),sK14)) | spl15_3),
  inference(avatar_component_clause,[],[f176])).
fof(f176,plain,(
  spl15_3 <=> in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),relation_composition(identity_relation(sK13),sK14))),
  introduced(avatar_definition,[new_symbols(naming,[spl15_3])])).
fof(f697,plain,(
  ( ! [X0,X1] : (in(unordered_pair(singleton(X0),unordered_pair(X0,sK12)),relation_composition(X1,sK14)) | ~in(unordered_pair(singleton(X0),unordered_pair(X0,sK11)),X1) | ~relation(relation_composition(X1,sK14)) | ~relation(X1)) ) | ~spl15_2),
  inference(subsumption_resolution,[],[f693,f113])).
fof(f113,plain,(
  relation(sK14)),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  (~in(ordered_pair(sK11,sK12),sK14) | ~in(sK11,sK13) | ~in(ordered_pair(sK11,sK12),relation_composition(identity_relation(sK13),sK14))) & ((in(ordered_pair(sK11,sK12),sK14) & in(sK11,sK13)) | in(ordered_pair(sK11,sK12),relation_composition(identity_relation(sK13),sK14))) & relation(sK14)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11,sK12,sK13,sK14])],[f72,f73])).
fof(f73,plain,(
  ? [X0,X1,X2,X3] : ((~in(ordered_pair(X0,X1),X3) | ~in(X0,X2) | ~in(ordered_pair(X0,X1),relation_composition(identity_relation(X2),X3))) & ((in(ordered_pair(X0,X1),X3) & in(X0,X2)) | in(ordered_pair(X0,X1),relation_composition(identity_relation(X2),X3))) & relation(X3)) => ((~in(ordered_pair(sK11,sK12),sK14) | ~in(sK11,sK13) | ~in(ordered_pair(sK11,sK12),relation_composition(identity_relation(sK13),sK14))) & ((in(ordered_pair(sK11,sK12),sK14) & in(sK11,sK13)) | in(ordered_pair(sK11,sK12),relation_composition(identity_relation(sK13),sK14))) & relation(sK14))),
  introduced(choice_axiom,[])).
fof(f72,plain,(
  ? [X0,X1,X2,X3] : ((~in(ordered_pair(X0,X1),X3) | ~in(X0,X2) | ~in(ordered_pair(X0,X1),relation_composition(identity_relation(X2),X3))) & ((in(ordered_pair(X0,X1),X3) & in(X0,X2)) | in(ordered_pair(X0,X1),relation_composition(identity_relation(X2),X3))) & relation(X3))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ? [X0,X1,X2,X3] : ((((~in(ordered_pair(X0,X1),X3) | ~in(X0,X2)) | ~in(ordered_pair(X0,X1),relation_composition(identity_relation(X2),X3))) & ((in(ordered_pair(X0,X1),X3) & in(X0,X2)) | in(ordered_pair(X0,X1),relation_composition(identity_relation(X2),X3)))) & relation(X3))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ? [X0,X1,X2,X3] : ((in(ordered_pair(X0,X1),relation_composition(identity_relation(X2),X3)) <~> (in(ordered_pair(X0,X1),X3) & in(X0,X2))) & relation(X3))),
  inference(ennf_transformation,[],[f30])).
fof(f30,negated_conjecture,(
  ~! [X0,X1,X2,X3] : (relation(X3) => (in(ordered_pair(X0,X1),relation_composition(identity_relation(X2),X3)) <=> (in(ordered_pair(X0,X1),X3) & in(X0,X2))))),
  inference(negated_conjecture,[],[f29])).
fof(f29,conjecture,(
  ! [X0,X1,X2,X3] : (relation(X3) => (in(ordered_pair(X0,X1),relation_composition(identity_relation(X2),X3)) <=> (in(ordered_pair(X0,X1),X3) & in(X0,X2))))),
  file('./p/SEU191+1.tptp',unknown)).
fof(f693,plain,(
  ( ! [X0,X1] : (~in(unordered_pair(singleton(X0),unordered_pair(X0,sK11)),X1) | in(unordered_pair(singleton(X0),unordered_pair(X0,sK12)),relation_composition(X1,sK14)) | ~relation(relation_composition(X1,sK14)) | ~relation(sK14) | ~relation(X1)) ) | ~spl15_2),
  inference(resolution,[],[f173,f158])).
fof(f158,plain,(
  ( ! [X0,X8,X7,X1,X9] : (~in(unordered_pair(singleton(X9),unordered_pair(X9,X8)),X1) | ~in(unordered_pair(singleton(X7),unordered_pair(X7,X9)),X0) | in(unordered_pair(singleton(X7),unordered_pair(X7,X8)),relation_composition(X0,X1)) | ~relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0)) )),
  inference(forward_demodulation,[],[f157,f77])).
fof(f157,plain,(
  ( ! [X0,X8,X7,X1,X9] : (~in(unordered_pair(singleton(X9),unordered_pair(X9,X8)),X1) | in(unordered_pair(singleton(X7),unordered_pair(X7,X8)),relation_composition(X0,X1)) | ~in(unordered_pair(unordered_pair(X7,X9),singleton(X7)),X0) | ~relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0)) )),
  inference(forward_demodulation,[],[f156,f77])).
fof(f156,plain,(
  ( ! [X0,X8,X7,X1,X9] : (in(unordered_pair(singleton(X7),unordered_pair(X7,X8)),relation_composition(X0,X1)) | ~in(unordered_pair(unordered_pair(X9,X8),singleton(X9)),X1) | ~in(unordered_pair(unordered_pair(X7,X9),singleton(X7)),X0) | ~relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0)) )),
  inference(forward_demodulation,[],[f139,f77])).
fof(f139,plain,(
  ( ! [X0,X8,X7,X1,X9] : (in(unordered_pair(unordered_pair(X7,X8),singleton(X7)),relation_composition(X0,X1)) | ~in(unordered_pair(unordered_pair(X9,X8),singleton(X9)),X1) | ~in(unordered_pair(unordered_pair(X7,X9),singleton(X7)),X0) | ~relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0)) )),
  inference(equality_resolution,[],[f128])).
fof(f128,plain,(
  ( ! [X2,X0,X8,X7,X1,X9] : (in(unordered_pair(unordered_pair(X7,X8),singleton(X7)),X2) | ~in(unordered_pair(unordered_pair(X9,X8),singleton(X9)),X1) | ~in(unordered_pair(unordered_pair(X7,X9),singleton(X7)),X0) | relation_composition(X0,X1) != X2 | ~relation(X2) | ~relation(X1) | ~relation(X0)) )),
  inference(definition_unfolding,[],[f87,f84,f84,f84])).
fof(f87,plain,(
  ( ! [X2,X0,X8,X7,X1,X9] : (in(ordered_pair(X7,X8),X2) | ~in(ordered_pair(X9,X8),X1) | ~in(ordered_pair(X7,X9),X0) | relation_composition(X0,X1) != X2 | ~relation(X2) | ~relation(X1) | ~relation(X0)) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : (! [X1] : (! [X2] : (((relation_composition(X0,X1) = X2 | ((! [X5] : (~in(ordered_pair(X5,sK3(X0,X1,X2)),X1) | ~in(ordered_pair(sK2(X0,X1,X2),X5),X0)) | ~in(ordered_pair(sK2(X0,X1,X2),sK3(X0,X1,X2)),X2)) & ((in(ordered_pair(sK4(X0,X1,X2),sK3(X0,X1,X2)),X1) & in(ordered_pair(sK2(X0,X1,X2),sK4(X0,X1,X2)),X0)) | in(ordered_pair(sK2(X0,X1,X2),sK3(X0,X1,X2)),X2)))) & (! [X7,X8] : ((in(ordered_pair(X7,X8),X2) | ! [X9] : (~in(ordered_pair(X9,X8),X1) | ~in(ordered_pair(X7,X9),X0))) & ((in(ordered_pair(sK5(X0,X1,X7,X8),X8),X1) & in(ordered_pair(X7,sK5(X0,X1,X7,X8)),X0)) | ~in(ordered_pair(X7,X8),X2))) | relation_composition(X0,X1) != X2)) | ~relation(X2)) | ~relation(X1)) | ~relation(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3,sK4,sK5])],[f56,f59,f58,f57])).
fof(f57,plain,(
  ! [X2,X1,X0] : (? [X3,X4] : ((! [X5] : (~in(ordered_pair(X5,X4),X1) | ~in(ordered_pair(X3,X5),X0)) | ~in(ordered_pair(X3,X4),X2)) & (? [X6] : (in(ordered_pair(X6,X4),X1) & in(ordered_pair(X3,X6),X0)) | in(ordered_pair(X3,X4),X2))) => ((! [X5] : (~in(ordered_pair(X5,sK3(X0,X1,X2)),X1) | ~in(ordered_pair(sK2(X0,X1,X2),X5),X0)) | ~in(ordered_pair(sK2(X0,X1,X2),sK3(X0,X1,X2)),X2)) & (? [X6] : (in(ordered_pair(X6,sK3(X0,X1,X2)),X1) & in(ordered_pair(sK2(X0,X1,X2),X6),X0)) | in(ordered_pair(sK2(X0,X1,X2),sK3(X0,X1,X2)),X2))))),
  introduced(choice_axiom,[])).
fof(f58,plain,(
  ! [X2,X1,X0] : (? [X6] : (in(ordered_pair(X6,sK3(X0,X1,X2)),X1) & in(ordered_pair(sK2(X0,X1,X2),X6),X0)) => (in(ordered_pair(sK4(X0,X1,X2),sK3(X0,X1,X2)),X1) & in(ordered_pair(sK2(X0,X1,X2),sK4(X0,X1,X2)),X0)))),
  introduced(choice_axiom,[])).
fof(f59,plain,(
  ! [X8,X7,X1,X0] : (? [X10] : (in(ordered_pair(X10,X8),X1) & in(ordered_pair(X7,X10),X0)) => (in(ordered_pair(sK5(X0,X1,X7,X8),X8),X1) & in(ordered_pair(X7,sK5(X0,X1,X7,X8)),X0)))),
  introduced(choice_axiom,[])).
fof(f56,plain,(
  ! [X0] : (! [X1] : (! [X2] : (((relation_composition(X0,X1) = X2 | ? [X3,X4] : ((! [X5] : (~in(ordered_pair(X5,X4),X1) | ~in(ordered_pair(X3,X5),X0)) | ~in(ordered_pair(X3,X4),X2)) & (? [X6] : (in(ordered_pair(X6,X4),X1) & in(ordered_pair(X3,X6),X0)) | in(ordered_pair(X3,X4),X2)))) & (! [X7,X8] : ((in(ordered_pair(X7,X8),X2) | ! [X9] : (~in(ordered_pair(X9,X8),X1) | ~in(ordered_pair(X7,X9),X0))) & (? [X10] : (in(ordered_pair(X10,X8),X1) & in(ordered_pair(X7,X10),X0)) | ~in(ordered_pair(X7,X8),X2))) | relation_composition(X0,X1) != X2)) | ~relation(X2)) | ~relation(X1)) | ~relation(X0))),
  inference(rectify,[],[f55])).
fof(f55,plain,(
  ! [X0] : (! [X1] : (! [X2] : (((relation_composition(X0,X1) = X2 | ? [X3,X4] : ((! [X5] : (~in(ordered_pair(X5,X4),X1) | ~in(ordered_pair(X3,X5),X0)) | ~in(ordered_pair(X3,X4),X2)) & (? [X5] : (in(ordered_pair(X5,X4),X1) & in(ordered_pair(X3,X5),X0)) | in(ordered_pair(X3,X4),X2)))) & (! [X3,X4] : ((in(ordered_pair(X3,X4),X2) | ! [X5] : (~in(ordered_pair(X5,X4),X1) | ~in(ordered_pair(X3,X5),X0))) & (? [X5] : (in(ordered_pair(X5,X4),X1) & in(ordered_pair(X3,X5),X0)) | ~in(ordered_pair(X3,X4),X2))) | relation_composition(X0,X1) != X2)) | ~relation(X2)) | ~relation(X1)) | ~relation(X0))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (! [X1] : (! [X2] : ((relation_composition(X0,X1) = X2 <=> ! [X3,X4] : (in(ordered_pair(X3,X4),X2) <=> ? [X5] : (in(ordered_pair(X5,X4),X1) & in(ordered_pair(X3,X5),X0)))) | ~relation(X2)) | ~relation(X1)) | ~relation(X0))),
  inference(ennf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0] : (relation(X0) => ! [X1] : (relation(X1) => ! [X2] : (relation(X2) => (relation_composition(X0,X1) = X2 <=> ! [X3,X4] : (in(ordered_pair(X3,X4),X2) <=> ? [X5] : (in(ordered_pair(X5,X4),X1) & in(ordered_pair(X3,X5),X0)))))))),
  file('./p/SEU191+1.tptp',unknown)).
fof(f173,plain,(
  in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),sK14) | ~spl15_2),
  inference(avatar_component_clause,[],[f172])).
fof(f172,plain,(
  spl15_2 <=> in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),sK14)),
  introduced(avatar_definition,[new_symbols(naming,[spl15_2])])).
fof(f680,plain,(
  spl15_2 | ~spl15_3 | ~spl15_4),
  inference(avatar_contradiction_clause,[],[f679])).
fof(f679,plain,(
  $false | (spl15_2 | ~spl15_3 | ~spl15_4)),
  inference(subsumption_resolution,[],[f678,f174])).
fof(f174,plain,(
  ~in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),sK14) | spl15_2),
  inference(avatar_component_clause,[],[f172])).
fof(f678,plain,(
  in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),sK14) | (~spl15_3 | ~spl15_4)),
  inference(forward_demodulation,[],[f677,f77])).
fof(f677,plain,(
  in(unordered_pair(singleton(sK11),unordered_pair(sK12,sK11)),sK14) | (~spl15_3 | ~spl15_4)),
  inference(subsumption_resolution,[],[f676,f92])).
fof(f676,plain,(
  in(unordered_pair(singleton(sK11),unordered_pair(sK12,sK11)),sK14) | ~relation(identity_relation(sK13)) | (~spl15_3 | ~spl15_4)),
  inference(subsumption_resolution,[],[f675,f113])).
fof(f675,plain,(
  in(unordered_pair(singleton(sK11),unordered_pair(sK12,sK11)),sK14) | ~relation(sK14) | ~relation(identity_relation(sK13)) | (~spl15_3 | ~spl15_4)),
  inference(subsumption_resolution,[],[f674,f290])).
fof(f674,plain,(
  in(unordered_pair(singleton(sK11),unordered_pair(sK12,sK11)),sK14) | ~relation(relation_composition(identity_relation(sK13),sK14)) | ~relation(sK14) | ~relation(identity_relation(sK13)) | (~spl15_3 | ~spl15_4)),
  inference(subsumption_resolution,[],[f672,f177])).
fof(f177,plain,(
  in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),relation_composition(identity_relation(sK13),sK14)) | ~spl15_3),
  inference(avatar_component_clause,[],[f176])).
fof(f672,plain,(
  in(unordered_pair(singleton(sK11),unordered_pair(sK12,sK11)),sK14) | ~in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),relation_composition(identity_relation(sK13),sK14)) | ~relation(relation_composition(identity_relation(sK13),sK14)) | ~relation(sK14) | ~relation(identity_relation(sK13)) | (~spl15_3 | ~spl15_4)),
  inference(superposition,[],[f161,f668])).
fof(f668,plain,(
  sK11 = sK5(identity_relation(sK13),sK14,sK11,sK12) | (~spl15_3 | ~spl15_4)),
  inference(subsumption_resolution,[],[f667,f113])).
fof(f667,plain,(
  ~relation(sK14) | sK11 = sK5(identity_relation(sK13),sK14,sK11,sK12) | (~spl15_3 | ~spl15_4)),
  inference(subsumption_resolution,[],[f649,f290])).
fof(f649,plain,(
  ~relation(relation_composition(identity_relation(sK13),sK14)) | ~relation(sK14) | sK11 = sK5(identity_relation(sK13),sK14,sK11,sK12) | ~spl15_3),
  inference(resolution,[],[f261,f177])).
fof(f261,plain,(
  ( ! [X6,X4,X7,X5] : (~in(unordered_pair(singleton(X4),unordered_pair(X4,X5)),relation_composition(identity_relation(X6),X7)) | ~relation(relation_composition(identity_relation(X6),X7)) | ~relation(X7) | sK5(identity_relation(X6),X7,X4,X5) = X4) )),
  inference(subsumption_resolution,[],[f258,f92])).
fof(f258,plain,(
  ( ! [X6,X4,X7,X5] : (~in(unordered_pair(singleton(X4),unordered_pair(X4,X5)),relation_composition(identity_relation(X6),X7)) | ~relation(relation_composition(identity_relation(X6),X7)) | ~relation(X7) | ~relation(identity_relation(X6)) | sK5(identity_relation(X6),X7,X4,X5) = X4) )),
  inference(duplicate_literal_removal,[],[f254])).
fof(f254,plain,(
  ( ! [X6,X4,X7,X5] : (~in(unordered_pair(singleton(X4),unordered_pair(X4,X5)),relation_composition(identity_relation(X6),X7)) | ~relation(relation_composition(identity_relation(X6),X7)) | ~relation(X7) | ~relation(identity_relation(X6)) | sK5(identity_relation(X6),X7,X4,X5) = X4 | ~relation(identity_relation(X6))) )),
  inference(resolution,[],[f163,f146])).
fof(f146,plain,(
  ( ! [X4,X0,X5] : (~in(unordered_pair(singleton(X4),unordered_pair(X4,X5)),identity_relation(X0)) | X4 = X5 | ~relation(identity_relation(X0))) )),
  inference(forward_demodulation,[],[f137,f77])).
fof(f137,plain,(
  ( ! [X4,X0,X5] : (X4 = X5 | ~in(unordered_pair(unordered_pair(X4,X5),singleton(X4)),identity_relation(X0)) | ~relation(identity_relation(X0))) )),
  inference(equality_resolution,[],[f123])).
fof(f123,plain,(
  ( ! [X4,X0,X5,X1] : (X4 = X5 | ~in(unordered_pair(unordered_pair(X4,X5),singleton(X4)),X1) | identity_relation(X0) != X1 | ~relation(X1)) )),
  inference(definition_unfolding,[],[f79,f84])).
fof(f79,plain,(
  ( ! [X4,X0,X5,X1] : (X4 = X5 | ~in(ordered_pair(X4,X5),X1) | identity_relation(X0) != X1 | ~relation(X1)) )),
  inference(cnf_transformation,[],[f54])).
fof(f163,plain,(
  ( ! [X0,X8,X7,X1] : (in(unordered_pair(singleton(X7),unordered_pair(X7,sK5(X0,X1,X7,X8))),X0) | ~in(unordered_pair(singleton(X7),unordered_pair(X7,X8)),relation_composition(X0,X1)) | ~relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0)) )),
  inference(forward_demodulation,[],[f162,f77])).
fof(f162,plain,(
  ( ! [X0,X8,X7,X1] : (in(unordered_pair(singleton(X7),unordered_pair(X7,sK5(X0,X1,X7,X8))),X0) | ~in(unordered_pair(unordered_pair(X7,X8),singleton(X7)),relation_composition(X0,X1)) | ~relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0)) )),
  inference(forward_demodulation,[],[f141,f77])).
fof(f141,plain,(
  ( ! [X0,X8,X7,X1] : (in(unordered_pair(unordered_pair(X7,sK5(X0,X1,X7,X8)),singleton(X7)),X0) | ~in(unordered_pair(unordered_pair(X7,X8),singleton(X7)),relation_composition(X0,X1)) | ~relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0)) )),
  inference(equality_resolution,[],[f130])).
fof(f130,plain,(
  ( ! [X2,X0,X8,X7,X1] : (in(unordered_pair(unordered_pair(X7,sK5(X0,X1,X7,X8)),singleton(X7)),X0) | ~in(unordered_pair(unordered_pair(X7,X8),singleton(X7)),X2) | relation_composition(X0,X1) != X2 | ~relation(X2) | ~relation(X1) | ~relation(X0)) )),
  inference(definition_unfolding,[],[f85,f84,f84])).
fof(f85,plain,(
  ( ! [X2,X0,X8,X7,X1] : (in(ordered_pair(X7,sK5(X0,X1,X7,X8)),X0) | ~in(ordered_pair(X7,X8),X2) | relation_composition(X0,X1) != X2 | ~relation(X2) | ~relation(X1) | ~relation(X0)) )),
  inference(cnf_transformation,[],[f60])).
fof(f161,plain,(
  ( ! [X0,X8,X7,X1] : (in(unordered_pair(singleton(sK5(X0,X1,X7,X8)),unordered_pair(X8,sK5(X0,X1,X7,X8))),X1) | ~in(unordered_pair(singleton(X7),unordered_pair(X7,X8)),relation_composition(X0,X1)) | ~relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0)) )),
  inference(forward_demodulation,[],[f160,f77])).
fof(f160,plain,(
  ( ! [X0,X8,X7,X1] : (in(unordered_pair(singleton(sK5(X0,X1,X7,X8)),unordered_pair(X8,sK5(X0,X1,X7,X8))),X1) | ~in(unordered_pair(unordered_pair(X7,X8),singleton(X7)),relation_composition(X0,X1)) | ~relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0)) )),
  inference(forward_demodulation,[],[f159,f77])).
fof(f159,plain,(
  ( ! [X0,X8,X7,X1] : (in(unordered_pair(singleton(sK5(X0,X1,X7,X8)),unordered_pair(sK5(X0,X1,X7,X8),X8)),X1) | ~in(unordered_pair(unordered_pair(X7,X8),singleton(X7)),relation_composition(X0,X1)) | ~relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0)) )),
  inference(forward_demodulation,[],[f140,f77])).
fof(f140,plain,(
  ( ! [X0,X8,X7,X1] : (in(unordered_pair(unordered_pair(sK5(X0,X1,X7,X8),X8),singleton(sK5(X0,X1,X7,X8))),X1) | ~in(unordered_pair(unordered_pair(X7,X8),singleton(X7)),relation_composition(X0,X1)) | ~relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0)) )),
  inference(equality_resolution,[],[f129])).
fof(f129,plain,(
  ( ! [X2,X0,X8,X7,X1] : (in(unordered_pair(unordered_pair(sK5(X0,X1,X7,X8),X8),singleton(sK5(X0,X1,X7,X8))),X1) | ~in(unordered_pair(unordered_pair(X7,X8),singleton(X7)),X2) | relation_composition(X0,X1) != X2 | ~relation(X2) | ~relation(X1) | ~relation(X0)) )),
  inference(definition_unfolding,[],[f86,f84,f84])).
fof(f86,plain,(
  ( ! [X2,X0,X8,X7,X1] : (in(ordered_pair(sK5(X0,X1,X7,X8),X8),X1) | ~in(ordered_pair(X7,X8),X2) | relation_composition(X0,X1) != X2 | ~relation(X2) | ~relation(X1) | ~relation(X0)) )),
  inference(cnf_transformation,[],[f60])).
fof(f617,plain,(
  spl15_1 | ~spl15_3 | ~spl15_4),
  inference(avatar_split_clause,[],[f613,f289,f176,f168])).
fof(f613,plain,(
  in(sK11,sK13) | (~spl15_3 | ~spl15_4)),
  inference(subsumption_resolution,[],[f612,f113])).
fof(f612,plain,(
  ~relation(sK14) | in(sK11,sK13) | (~spl15_3 | ~spl15_4)),
  inference(subsumption_resolution,[],[f595,f290])).
fof(f595,plain,(
  ~relation(relation_composition(identity_relation(sK13),sK14)) | ~relation(sK14) | in(sK11,sK13) | ~spl15_3),
  inference(resolution,[],[f260,f177])).
fof(f260,plain,(
  ( ! [X2,X0,X3,X1] : (~in(unordered_pair(singleton(X0),unordered_pair(X0,X1)),relation_composition(identity_relation(X2),X3)) | ~relation(relation_composition(identity_relation(X2),X3)) | ~relation(X3) | in(X0,X2)) )),
  inference(subsumption_resolution,[],[f259,f92])).
fof(f259,plain,(
  ( ! [X2,X0,X3,X1] : (~in(unordered_pair(singleton(X0),unordered_pair(X0,X1)),relation_composition(identity_relation(X2),X3)) | ~relation(relation_composition(identity_relation(X2),X3)) | ~relation(X3) | ~relation(identity_relation(X2)) | in(X0,X2)) )),
  inference(duplicate_literal_removal,[],[f253])).
fof(f253,plain,(
  ( ! [X2,X0,X3,X1] : (~in(unordered_pair(singleton(X0),unordered_pair(X0,X1)),relation_composition(identity_relation(X2),X3)) | ~relation(relation_composition(identity_relation(X2),X3)) | ~relation(X3) | ~relation(identity_relation(X2)) | in(X0,X2) | ~relation(identity_relation(X2))) )),
  inference(resolution,[],[f163,f147])).
fof(f147,plain,(
  ( ! [X4,X0,X5] : (~in(unordered_pair(singleton(X4),unordered_pair(X4,X5)),identity_relation(X0)) | in(X4,X0) | ~relation(identity_relation(X0))) )),
  inference(forward_demodulation,[],[f138,f77])).
fof(f138,plain,(
  ( ! [X4,X0,X5] : (in(X4,X0) | ~in(unordered_pair(unordered_pair(X4,X5),singleton(X4)),identity_relation(X0)) | ~relation(identity_relation(X0))) )),
  inference(equality_resolution,[],[f124])).
fof(f124,plain,(
  ( ! [X4,X0,X5,X1] : (in(X4,X0) | ~in(unordered_pair(unordered_pair(X4,X5),singleton(X4)),X1) | identity_relation(X0) != X1 | ~relation(X1)) )),
  inference(definition_unfolding,[],[f78,f84])).
fof(f78,plain,(
  ( ! [X4,X0,X5,X1] : (in(X4,X0) | ~in(ordered_pair(X4,X5),X1) | identity_relation(X0) != X1 | ~relation(X1)) )),
  inference(cnf_transformation,[],[f54])).
fof(f307,plain,(
  spl15_4),
  inference(avatar_contradiction_clause,[],[f306])).
fof(f306,plain,(
  $false | spl15_4),
  inference(subsumption_resolution,[],[f305,f92])).
fof(f305,plain,(
  ~relation(identity_relation(sK13)) | spl15_4),
  inference(subsumption_resolution,[],[f304,f113])).
fof(f304,plain,(
  ~relation(sK14) | ~relation(identity_relation(sK13)) | spl15_4),
  inference(resolution,[],[f291,f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0)) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : (relation(relation_composition(X0,X1)) | ~relation(X1) | ~relation(X0))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : (relation(relation_composition(X0,X1)) | (~relation(X1) | ~relation(X0)))),
  inference(ennf_transformation,[],[f11])).
fof(f11,axiom,(
  ! [X0,X1] : ((relation(X1) & relation(X0)) => relation(relation_composition(X0,X1)))),
  file('./p/SEU191+1.tptp',unknown)).
fof(f291,plain,(
  ~relation(relation_composition(identity_relation(sK13),sK14)) | spl15_4),
  inference(avatar_component_clause,[],[f289])).
fof(f184,plain,(
  spl15_1 | spl15_3),
  inference(avatar_split_clause,[],[f183,f176,f168])).
fof(f183,plain,(
  in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),relation_composition(identity_relation(sK13),sK14)) | in(sK11,sK13)),
  inference(forward_demodulation,[],[f134,f77])).
fof(f134,plain,(
  in(sK11,sK13) | in(unordered_pair(unordered_pair(sK11,sK12),singleton(sK11)),relation_composition(identity_relation(sK13),sK14))),
  inference(definition_unfolding,[],[f114,f84])).
fof(f114,plain,(
  in(sK11,sK13) | in(ordered_pair(sK11,sK12),relation_composition(identity_relation(sK13),sK14))),
  inference(cnf_transformation,[],[f74])).
fof(f182,plain,(
  spl15_2 | spl15_3),
  inference(avatar_split_clause,[],[f181,f176,f172])).
fof(f181,plain,(
  in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),relation_composition(identity_relation(sK13),sK14)) | in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),sK14)),
  inference(forward_demodulation,[],[f180,f77])).
fof(f180,plain,(
  in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),sK14) | in(unordered_pair(unordered_pair(sK11,sK12),singleton(sK11)),relation_composition(identity_relation(sK13),sK14))),
  inference(forward_demodulation,[],[f133,f77])).
fof(f133,plain,(
  in(unordered_pair(unordered_pair(sK11,sK12),singleton(sK11)),sK14) | in(unordered_pair(unordered_pair(sK11,sK12),singleton(sK11)),relation_composition(identity_relation(sK13),sK14))),
  inference(definition_unfolding,[],[f115,f84,f84])).
fof(f115,plain,(
  in(ordered_pair(sK11,sK12),sK14) | in(ordered_pair(sK11,sK12),relation_composition(identity_relation(sK13),sK14))),
  inference(cnf_transformation,[],[f74])).
fof(f179,plain,(
  ~spl15_1 | ~spl15_2 | ~spl15_3),
  inference(avatar_split_clause,[],[f166,f176,f172,f168])).
fof(f166,plain,(
  ~in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),relation_composition(identity_relation(sK13),sK14)) | ~in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),sK14) | ~in(sK11,sK13)),
  inference(forward_demodulation,[],[f165,f77])).
fof(f165,plain,(
  ~in(unordered_pair(singleton(sK11),unordered_pair(sK11,sK12)),sK14) | ~in(sK11,sK13) | ~in(unordered_pair(unordered_pair(sK11,sK12),singleton(sK11)),relation_composition(identity_relation(sK13),sK14))),
  inference(forward_demodulation,[],[f132,f77])).
fof(f132,plain,(
  ~in(unordered_pair(unordered_pair(sK11,sK12),singleton(sK11)),sK14) | ~in(sK11,sK13) | ~in(unordered_pair(unordered_pair(sK11,sK12),singleton(sK11)),relation_composition(identity_relation(sK13),sK14))),
  inference(definition_unfolding,[],[f116,f84,f84])).
fof(f116,plain,(
  ~in(ordered_pair(sK11,sK12),sK14) | ~in(sK11,sK13) | ~in(ordered_pair(sK11,sK12),relation_composition(identity_relation(sK13),sK14))),
  inference(cnf_transformation,[],[f74])).
% SZS output end Proof for SEU191+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
