% Refutation found. Thanks to Tanya!
% SZS status Theorem for SEU170+1
% SZS output start Proof for SEU170+1
fof(f30100,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f118,f119,f29513,f30099])).
fof(f30099,plain,(
  spl11_1 | ~spl11_2),
  inference(avatar_contradiction_clause,[],[f30098])).
fof(f30098,plain,(
  $false | (spl11_1 | ~spl11_2)),
  inference(subsumption_resolution,[],[f30097,f113])).
fof(f113,plain,(
  ~disjoint(sK9,sK10) | spl11_1),
  inference(avatar_component_clause,[],[f111])).
fof(f111,plain,(
  spl11_1 <=> disjoint(sK9,sK10)),
  introduced(avatar_definition,[new_symbols(naming,[spl11_1])])).
fof(f30097,plain,(
  disjoint(sK9,sK10) | ~spl11_2),
  inference(resolution,[],[f30056,f94])).
fof(f94,plain,(
  ( ! [X0,X1] : (~disjoint(X0,X1) | disjoint(X1,X0)) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : (disjoint(X1,X0) | ~disjoint(X0,X1))),
  inference(ennf_transformation,[],[f20])).
fof(f20,axiom,(
  ! [X0,X1] : (disjoint(X0,X1) => disjoint(X1,X0))),
  file('./p/SEU170+1.tptp',unknown)).
fof(f30056,plain,(
  disjoint(sK10,sK9) | ~spl11_2),
  inference(superposition,[],[f29801,f275])).
fof(f275,plain,(
  sK10 = set_difference(sK8,set_difference(sK8,sK10))),
  inference(forward_demodulation,[],[f272,f202])).
fof(f202,plain,(
  sK10 = subset_complement(sK8,set_difference(sK8,sK10))),
  inference(forward_demodulation,[],[f196,f179])).
fof(f179,plain,(
  subset_complement(sK8,sK10) = set_difference(sK8,sK10)),
  inference(resolution,[],[f80,f100])).
fof(f100,plain,(
  element(sK10,powerset(sK8))),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ((~subset(sK9,subset_complement(sK8,sK10)) | ~disjoint(sK9,sK10)) & (subset(sK9,subset_complement(sK8,sK10)) | disjoint(sK9,sK10)) & element(sK10,powerset(sK8))) & element(sK9,powerset(sK8))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9,sK10])],[f66,f68,f67])).
fof(f67,plain,(
  ? [X0,X1] : (? [X2] : ((~subset(X1,subset_complement(X0,X2)) | ~disjoint(X1,X2)) & (subset(X1,subset_complement(X0,X2)) | disjoint(X1,X2)) & element(X2,powerset(X0))) & element(X1,powerset(X0))) => (? [X2] : ((~subset(sK9,subset_complement(sK8,X2)) | ~disjoint(sK9,X2)) & (subset(sK9,subset_complement(sK8,X2)) | disjoint(sK9,X2)) & element(X2,powerset(sK8))) & element(sK9,powerset(sK8)))),
  introduced(choice_axiom,[])).
fof(f68,plain,(
  ? [X2] : ((~subset(sK9,subset_complement(sK8,X2)) | ~disjoint(sK9,X2)) & (subset(sK9,subset_complement(sK8,X2)) | disjoint(sK9,X2)) & element(X2,powerset(sK8))) => ((~subset(sK9,subset_complement(sK8,sK10)) | ~disjoint(sK9,sK10)) & (subset(sK9,subset_complement(sK8,sK10)) | disjoint(sK9,sK10)) & element(sK10,powerset(sK8)))),
  introduced(choice_axiom,[])).
fof(f66,plain,(
  ? [X0,X1] : (? [X2] : ((~subset(X1,subset_complement(X0,X2)) | ~disjoint(X1,X2)) & (subset(X1,subset_complement(X0,X2)) | disjoint(X1,X2)) & element(X2,powerset(X0))) & element(X1,powerset(X0)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ? [X0,X1] : (? [X2] : (((~subset(X1,subset_complement(X0,X2)) | ~disjoint(X1,X2)) & (subset(X1,subset_complement(X0,X2)) | disjoint(X1,X2))) & element(X2,powerset(X0))) & element(X1,powerset(X0)))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ? [X0,X1] : (? [X2] : ((disjoint(X1,X2) <~> subset(X1,subset_complement(X0,X2))) & element(X2,powerset(X0))) & element(X1,powerset(X0)))),
  inference(ennf_transformation,[],[f24])).
fof(f24,negated_conjecture,(
  ~! [X0,X1] : (element(X1,powerset(X0)) => ! [X2] : (element(X2,powerset(X0)) => (disjoint(X1,X2) <=> subset(X1,subset_complement(X0,X2)))))),
  inference(negated_conjecture,[],[f23])).
fof(f23,conjecture,(
  ! [X0,X1] : (element(X1,powerset(X0)) => ! [X2] : (element(X2,powerset(X0)) => (disjoint(X1,X2) <=> subset(X1,subset_complement(X0,X2)))))),
  file('./p/SEU170+1.tptp',unknown)).
fof(f80,plain,(
  ( ! [X0,X1] : (~element(X1,powerset(X0)) | set_difference(X0,X1) = subset_complement(X0,X1)) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : (set_difference(X0,X1) = subset_complement(X0,X1) | ~element(X1,powerset(X0)))),
  inference(ennf_transformation,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1] : (element(X1,powerset(X0)) => set_difference(X0,X1) = subset_complement(X0,X1))),
  file('./p/SEU170+1.tptp',unknown)).
fof(f196,plain,(
  sK10 = subset_complement(sK8,subset_complement(sK8,sK10))),
  inference(resolution,[],[f85,f100])).
fof(f85,plain,(
  ( ! [X0,X1] : (~element(X1,powerset(X0)) | subset_complement(X0,subset_complement(X0,X1)) = X1) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : (subset_complement(X0,subset_complement(X0,X1)) = X1 | ~element(X1,powerset(X0)))),
  inference(ennf_transformation,[],[f13])).
fof(f13,axiom,(
  ! [X0,X1] : (element(X1,powerset(X0)) => subset_complement(X0,subset_complement(X0,X1)) = X1)),
  file('./p/SEU170+1.tptp',unknown)).
fof(f272,plain,(
  subset_complement(sK8,set_difference(sK8,sK10)) = set_difference(sK8,set_difference(sK8,sK10))),
  inference(resolution,[],[f264,f80])).
fof(f264,plain,(
  element(set_difference(sK8,sK10),powerset(sK8))),
  inference(subsumption_resolution,[],[f263,f100])).
fof(f263,plain,(
  element(set_difference(sK8,sK10),powerset(sK8)) | ~element(sK10,powerset(sK8))),
  inference(superposition,[],[f81,f179])).
fof(f81,plain,(
  ( ! [X0,X1] : (element(subset_complement(X0,X1),powerset(X0)) | ~element(X1,powerset(X0))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (element(subset_complement(X0,X1),powerset(X0)) | ~element(X1,powerset(X0)))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0,X1] : (element(X1,powerset(X0)) => element(subset_complement(X0,X1),powerset(X0)))),
  file('./p/SEU170+1.tptp',unknown)).
fof(f29801,plain,(
  ( ! [X7] : (disjoint(set_difference(X7,set_difference(sK8,sK10)),sK9)) ) | ~spl11_2),
  inference(duplicate_literal_removal,[],[f29777])).
fof(f29777,plain,(
  ( ! [X7] : (disjoint(set_difference(X7,set_difference(sK8,sK10)),sK9) | disjoint(set_difference(X7,set_difference(sK8,sK10)),sK9)) ) | ~spl11_2),
  inference(resolution,[],[f29562,f97])).
fof(f97,plain,(
  ( ! [X0,X1] : (in(sK7(X0,X1),X1) | disjoint(X0,X1)) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : ((~disjoint(X0,X1) | ! [X2] : (~in(X2,X1) | ~in(X2,X0))) & ((in(sK7(X0,X1),X1) & in(sK7(X0,X1),X0)) | disjoint(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7])],[f39,f63])).
fof(f63,plain,(
  ! [X1,X0] : (? [X3] : (in(X3,X1) & in(X3,X0)) => (in(sK7(X0,X1),X1) & in(sK7(X0,X1),X0)))),
  introduced(choice_axiom,[])).
fof(f39,plain,(
  ! [X0,X1] : ((~disjoint(X0,X1) | ! [X2] : (~in(X2,X1) | ~in(X2,X0))) & (? [X3] : (in(X3,X1) & in(X3,X0)) | disjoint(X0,X1)))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : (~(disjoint(X0,X1) & ? [X2] : (in(X2,X1) & in(X2,X0))) & ~(! [X3] : ~(in(X3,X1) & in(X3,X0)) & ~disjoint(X0,X1)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X0,X1] : (~(disjoint(X0,X1) & ? [X2] : (in(X2,X1) & in(X2,X0))) & ~(! [X2] : ~(in(X2,X1) & in(X2,X0)) & ~disjoint(X0,X1)))),
  file('./p/SEU170+1.tptp',unknown)).
fof(f29562,plain,(
  ( ! [X59,X58] : (~in(sK7(set_difference(X58,set_difference(sK8,sK10)),X59),sK9) | disjoint(set_difference(X58,set_difference(sK8,sK10)),X59)) ) | ~spl11_2),
  inference(resolution,[],[f29535,f138])).
fof(f138,plain,(
  ( ! [X4,X5,X3] : (~in(sK7(set_difference(X3,X4),X5),X4) | disjoint(set_difference(X3,X4),X5)) )),
  inference(resolution,[],[f108,f96])).
fof(f96,plain,(
  ( ! [X0,X1] : (in(sK7(X0,X1),X0) | disjoint(X0,X1)) )),
  inference(cnf_transformation,[],[f64])).
fof(f108,plain,(
  ( ! [X4,X0,X1] : (~in(X4,set_difference(X0,X1)) | ~in(X4,X1)) )),
  inference(equality_resolution,[],[f75])).
fof(f75,plain,(
  ( ! [X4,X2,X0,X1] : (~in(X4,X1) | ~in(X4,X2) | set_difference(X0,X1) != X2) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2] : ((set_difference(X0,X1) = X2 | ((in(sK1(X0,X1,X2),X1) | ~in(sK1(X0,X1,X2),X0) | ~in(sK1(X0,X1,X2),X2)) & ((~in(sK1(X0,X1,X2),X1) & in(sK1(X0,X1,X2),X0)) | in(sK1(X0,X1,X2),X2)))) & (! [X4] : ((in(X4,X2) | in(X4,X1) | ~in(X4,X0)) & ((~in(X4,X1) & in(X4,X0)) | ~in(X4,X2))) | set_difference(X0,X1) != X2))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f50,f51])).
fof(f51,plain,(
  ! [X2,X1,X0] : (? [X3] : ((in(X3,X1) | ~in(X3,X0) | ~in(X3,X2)) & ((~in(X3,X1) & in(X3,X0)) | in(X3,X2))) => ((in(sK1(X0,X1,X2),X1) | ~in(sK1(X0,X1,X2),X0) | ~in(sK1(X0,X1,X2),X2)) & ((~in(sK1(X0,X1,X2),X1) & in(sK1(X0,X1,X2),X0)) | in(sK1(X0,X1,X2),X2))))),
  introduced(choice_axiom,[])).
fof(f50,plain,(
  ! [X0,X1,X2] : ((set_difference(X0,X1) = X2 | ? [X3] : ((in(X3,X1) | ~in(X3,X0) | ~in(X3,X2)) & ((~in(X3,X1) & in(X3,X0)) | in(X3,X2)))) & (! [X4] : ((in(X4,X2) | in(X4,X1) | ~in(X4,X0)) & ((~in(X4,X1) & in(X4,X0)) | ~in(X4,X2))) | set_difference(X0,X1) != X2))),
  inference(rectify,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : ((set_difference(X0,X1) = X2 | ? [X3] : ((in(X3,X1) | ~in(X3,X0) | ~in(X3,X2)) & ((~in(X3,X1) & in(X3,X0)) | in(X3,X2)))) & (! [X3] : ((in(X3,X2) | in(X3,X1) | ~in(X3,X0)) & ((~in(X3,X1) & in(X3,X0)) | ~in(X3,X2))) | set_difference(X0,X1) != X2))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2] : ((set_difference(X0,X1) = X2 | ? [X3] : (((in(X3,X1) | ~in(X3,X0)) | ~in(X3,X2)) & ((~in(X3,X1) & in(X3,X0)) | in(X3,X2)))) & (! [X3] : ((in(X3,X2) | (in(X3,X1) | ~in(X3,X0))) & ((~in(X3,X1) & in(X3,X0)) | ~in(X3,X2))) | set_difference(X0,X1) != X2))),
  inference(nnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0,X1,X2] : (set_difference(X0,X1) = X2 <=> ! [X3] : (in(X3,X2) <=> (~in(X3,X1) & in(X3,X0))))),
  file('./p/SEU170+1.tptp',unknown)).
fof(f29535,plain,(
  ( ! [X0] : (in(X0,set_difference(sK8,sK10)) | ~in(X0,sK9)) ) | ~spl11_2),
  inference(resolution,[],[f29532,f71])).
fof(f71,plain,(
  ( ! [X0,X3,X1] : (~subset(X0,X1) | ~in(X3,X0) | in(X3,X1)) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ((subset(X0,X1) | (~in(sK0(X0,X1),X1) & in(sK0(X0,X1),X0))) & (! [X3] : (in(X3,X1) | ~in(X3,X0)) | ~subset(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f45,f46])).
fof(f46,plain,(
  ! [X1,X0] : (? [X2] : (~in(X2,X1) & in(X2,X0)) => (~in(sK0(X0,X1),X1) & in(sK0(X0,X1),X0)))),
  introduced(choice_axiom,[])).
fof(f45,plain,(
  ! [X0,X1] : ((subset(X0,X1) | ? [X2] : (~in(X2,X1) & in(X2,X0))) & (! [X3] : (in(X3,X1) | ~in(X3,X0)) | ~subset(X0,X1)))),
  inference(rectify,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : ((subset(X0,X1) | ? [X2] : (~in(X2,X1) & in(X2,X0))) & (! [X2] : (in(X2,X1) | ~in(X2,X0)) | ~subset(X0,X1)))),
  inference(nnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : (subset(X0,X1) <=> ! [X2] : (in(X2,X1) | ~in(X2,X0)))),
  inference(ennf_transformation,[],[f2])).
fof(f2,axiom,(
  ! [X0,X1] : (subset(X0,X1) <=> ! [X2] : (in(X2,X0) => in(X2,X1)))),
  file('./p/SEU170+1.tptp',unknown)).
fof(f29532,plain,(
  subset(sK9,set_difference(sK8,sK10)) | ~spl11_2),
  inference(forward_demodulation,[],[f116,f179])).
fof(f116,plain,(
  subset(sK9,subset_complement(sK8,sK10)) | ~spl11_2),
  inference(avatar_component_clause,[],[f115])).
fof(f115,plain,(
  spl11_2 <=> subset(sK9,subset_complement(sK8,sK10))),
  introduced(avatar_definition,[new_symbols(naming,[spl11_2])])).
fof(f29513,plain,(
  ~spl11_1 | spl11_2),
  inference(avatar_contradiction_clause,[],[f29512])).
fof(f29512,plain,(
  $false | (~spl11_1 | spl11_2)),
  inference(subsumption_resolution,[],[f29511,f470])).
fof(f470,plain,(
  ~subset(sK9,set_difference(sK8,sK10)) | spl11_2),
  inference(backward_demodulation,[],[f179,f117])).
fof(f117,plain,(
  ~subset(sK9,subset_complement(sK8,sK10)) | spl11_2),
  inference(avatar_component_clause,[],[f115])).
fof(f29511,plain,(
  subset(sK9,set_difference(sK8,sK10)) | (~spl11_1 | spl11_2)),
  inference(resolution,[],[f28943,f72])).
fof(f72,plain,(
  ( ! [X0,X1] : (in(sK0(X0,X1),X0) | subset(X0,X1)) )),
  inference(cnf_transformation,[],[f47])).
fof(f28943,plain,(
  ~in(sK0(sK9,set_difference(sK8,sK10)),sK9) | (~spl11_1 | spl11_2)),
  inference(subsumption_resolution,[],[f28922,f470])).
fof(f28922,plain,(
  subset(sK9,set_difference(sK8,sK10)) | ~in(sK0(sK9,set_difference(sK8,sK10)),sK9) | ~spl11_1),
  inference(resolution,[],[f835,f471])).
fof(f471,plain,(
  ( ! [X0] : (~in(X0,sK10) | ~in(X0,sK9)) ) | ~spl11_1),
  inference(resolution,[],[f112,f98])).
fof(f98,plain,(
  ( ! [X2,X0,X1] : (~disjoint(X0,X1) | ~in(X2,X1) | ~in(X2,X0)) )),
  inference(cnf_transformation,[],[f64])).
fof(f112,plain,(
  disjoint(sK9,sK10) | ~spl11_1),
  inference(avatar_component_clause,[],[f111])).
fof(f835,plain,(
  ( ! [X9] : (in(sK0(sK9,set_difference(sK8,X9)),X9) | subset(sK9,set_difference(sK8,X9))) )),
  inference(duplicate_literal_removal,[],[f830])).
fof(f830,plain,(
  ( ! [X9] : (in(sK0(sK9,set_difference(sK8,X9)),X9) | subset(sK9,set_difference(sK8,X9)) | subset(sK9,set_difference(sK8,X9))) )),
  inference(resolution,[],[f211,f248])).
fof(f248,plain,(
  ( ! [X0] : (in(sK0(sK9,X0),sK8) | subset(sK9,X0)) )),
  inference(resolution,[],[f168,f72])).
fof(f168,plain,(
  ( ! [X0] : (~in(X0,sK9) | in(X0,sK8)) )),
  inference(resolution,[],[f86,f99])).
fof(f99,plain,(
  element(sK9,powerset(sK8))),
  inference(cnf_transformation,[],[f69])).
fof(f86,plain,(
  ( ! [X2,X0,X1] : (~element(X1,powerset(X0)) | ~in(X2,X1) | in(X2,X0)) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : (! [X2] : (in(X2,X0) | ~in(X2,X1)) | ~element(X1,powerset(X0)))),
  inference(ennf_transformation,[],[f14])).
fof(f14,axiom,(
  ! [X0,X1] : (element(X1,powerset(X0)) => ! [X2] : (in(X2,X1) => in(X2,X0)))),
  file('./p/SEU170+1.tptp',unknown)).
fof(f211,plain,(
  ( ! [X14,X12,X13] : (~in(sK0(X12,set_difference(X13,X14)),X13) | in(sK0(X12,set_difference(X13,X14)),X14) | subset(X12,set_difference(X13,X14))) )),
  inference(resolution,[],[f107,f73])).
fof(f73,plain,(
  ( ! [X0,X1] : (~in(sK0(X0,X1),X1) | subset(X0,X1)) )),
  inference(cnf_transformation,[],[f47])).
fof(f107,plain,(
  ( ! [X4,X0,X1] : (in(X4,set_difference(X0,X1)) | in(X4,X1) | ~in(X4,X0)) )),
  inference(equality_resolution,[],[f76])).
fof(f76,plain,(
  ( ! [X4,X2,X0,X1] : (in(X4,X2) | in(X4,X1) | ~in(X4,X0) | set_difference(X0,X1) != X2) )),
  inference(cnf_transformation,[],[f52])).
fof(f119,plain,(
  spl11_1 | spl11_2),
  inference(avatar_split_clause,[],[f101,f115,f111])).
fof(f101,plain,(
  subset(sK9,subset_complement(sK8,sK10)) | disjoint(sK9,sK10)),
  inference(cnf_transformation,[],[f69])).
fof(f118,plain,(
  ~spl11_1 | ~spl11_2),
  inference(avatar_split_clause,[],[f102,f115,f111])).
fof(f102,plain,(
  ~subset(sK9,subset_complement(sK8,sK10)) | ~disjoint(sK9,sK10)),
  inference(cnf_transformation,[],[f69])).
% SZS output end Proof for SEU170+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 18805
% Time elapsed: 0.417 s
% ------------------------------
% ------------------------------
