% Refutation found. Thanks to Tanya!
% SZS status Theorem for SEU220+1
% SZS output start Proof for SEU220+1
fof(f12201,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f202,f459,f474,f789,f12200])).
fof(f12200,plain,(
  ~spl14_1 | spl14_2 | ~spl14_5 | ~spl14_6),
  inference(avatar_contradiction_clause,[],[f12199])).
fof(f12199,plain,(
  $false | (~spl14_1 | spl14_2 | ~spl14_5 | ~spl14_6)),
  inference(subsumption_resolution,[],[f12198,f201])).
fof(f201,plain,(
  sK12 != apply(relation_composition(function_inverse(sK13),sK13),sK12) | spl14_2),
  inference(avatar_component_clause,[],[f199])).
fof(f199,plain,(
  spl14_2 <=> sK12 = apply(relation_composition(function_inverse(sK13),sK13),sK12)),
  introduced(avatar_definition,[new_symbols(naming,[spl14_2])])).
fof(f12198,plain,(
  sK12 = apply(relation_composition(function_inverse(sK13),sK13),sK12) | (~spl14_1 | ~spl14_5 | ~spl14_6)),
  inference(forward_demodulation,[],[f12197,f196])).
fof(f196,plain,(
  sK12 = apply(sK13,apply(function_inverse(sK13),sK12)) | ~spl14_1),
  inference(avatar_component_clause,[],[f195])).
fof(f195,plain,(
  spl14_1 <=> sK12 = apply(sK13,apply(function_inverse(sK13),sK12))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_1])])).
fof(f12197,plain,(
  apply(relation_composition(function_inverse(sK13),sK13),sK12) = apply(sK13,apply(function_inverse(sK13),sK12)) | (~spl14_5 | ~spl14_6)),
  inference(subsumption_resolution,[],[f12185,f168])).
fof(f168,plain,(
  function(sK13)),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  (sK12 != apply(relation_composition(function_inverse(sK13),sK13),sK12) | sK12 != apply(sK13,apply(function_inverse(sK13),sK12))) & in(sK12,relation_rng(sK13)) & one_to_one(sK13) & function(sK13) & relation(sK13)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f74,f106])).
fof(f106,plain,(
  ? [X0,X1] : ((apply(relation_composition(function_inverse(X1),X1),X0) != X0 | apply(X1,apply(function_inverse(X1),X0)) != X0) & in(X0,relation_rng(X1)) & one_to_one(X1) & function(X1) & relation(X1)) => ((sK12 != apply(relation_composition(function_inverse(sK13),sK13),sK12) | sK12 != apply(sK13,apply(function_inverse(sK13),sK12))) & in(sK12,relation_rng(sK13)) & one_to_one(sK13) & function(sK13) & relation(sK13))),
  introduced(choice_axiom,[])).
fof(f74,plain,(
  ? [X0,X1] : ((apply(relation_composition(function_inverse(X1),X1),X0) != X0 | apply(X1,apply(function_inverse(X1),X0)) != X0) & in(X0,relation_rng(X1)) & one_to_one(X1) & function(X1) & relation(X1))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ? [X0,X1] : (((apply(relation_composition(function_inverse(X1),X1),X0) != X0 | apply(X1,apply(function_inverse(X1),X0)) != X0) & (in(X0,relation_rng(X1)) & one_to_one(X1))) & (function(X1) & relation(X1)))),
  inference(ennf_transformation,[],[f37])).
fof(f37,negated_conjecture,(
  ~! [X0,X1] : ((function(X1) & relation(X1)) => ((in(X0,relation_rng(X1)) & one_to_one(X1)) => (apply(relation_composition(function_inverse(X1),X1),X0) = X0 & apply(X1,apply(function_inverse(X1),X0)) = X0)))),
  inference(negated_conjecture,[],[f36])).
fof(f36,conjecture,(
  ! [X0,X1] : ((function(X1) & relation(X1)) => ((in(X0,relation_rng(X1)) & one_to_one(X1)) => (apply(relation_composition(function_inverse(X1),X1),X0) = X0 & apply(X1,apply(function_inverse(X1),X0)) = X0)))),
  file('./p/SEU220+1.tptp',unknown)).
fof(f12185,plain,(
  ~function(sK13) | apply(relation_composition(function_inverse(sK13),sK13),sK12) = apply(sK13,apply(function_inverse(sK13),sK12)) | (~spl14_5 | ~spl14_6)),
  inference(resolution,[],[f2012,f167])).
fof(f167,plain,(
  relation(sK13)),
  inference(cnf_transformation,[],[f107])).
fof(f2012,plain,(
  ( ! [X0] : (~relation(X0) | ~function(X0) | apply(relation_composition(function_inverse(sK13),X0),sK12) = apply(X0,apply(function_inverse(sK13),sK12))) ) | (~spl14_5 | ~spl14_6)),
  inference(resolution,[],[f497,f170])).
fof(f170,plain,(
  in(sK12,relation_rng(sK13))),
  inference(cnf_transformation,[],[f107])).
fof(f497,plain,(
  ( ! [X0,X1] : (~in(X0,relation_rng(sK13)) | apply(relation_composition(function_inverse(sK13),X1),X0) = apply(X1,apply(function_inverse(sK13),X0)) | ~function(X1) | ~relation(X1)) ) | (~spl14_5 | ~spl14_6)),
  inference(subsumption_resolution,[],[f496,f411])).
fof(f411,plain,(
  relation(function_inverse(sK13)) | ~spl14_5),
  inference(avatar_component_clause,[],[f410])).
fof(f410,plain,(
  spl14_5 <=> relation(function_inverse(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_5])])).
fof(f496,plain,(
  ( ! [X0,X1] : (~in(X0,relation_rng(sK13)) | apply(relation_composition(function_inverse(sK13),X1),X0) = apply(X1,apply(function_inverse(sK13),X0)) | ~function(X1) | ~relation(X1) | ~relation(function_inverse(sK13))) ) | ~spl14_6),
  inference(subsumption_resolution,[],[f490,f415])).
fof(f415,plain,(
  function(function_inverse(sK13)) | ~spl14_6),
  inference(avatar_component_clause,[],[f414])).
fof(f414,plain,(
  spl14_6 <=> function(function_inverse(sK13))),
  introduced(avatar_definition,[new_symbols(naming,[spl14_6])])).
fof(f490,plain,(
  ( ! [X0,X1] : (~in(X0,relation_rng(sK13)) | apply(relation_composition(function_inverse(sK13),X1),X0) = apply(X1,apply(function_inverse(sK13),X0)) | ~function(X1) | ~relation(X1) | ~function(function_inverse(sK13)) | ~relation(function_inverse(sK13))) )),
  inference(superposition,[],[f151,f351])).
fof(f351,plain,(
  relation_rng(sK13) = relation_dom(function_inverse(sK13))),
  inference(subsumption_resolution,[],[f350,f167])).
fof(f350,plain,(
  relation_rng(sK13) = relation_dom(function_inverse(sK13)) | ~relation(sK13)),
  inference(subsumption_resolution,[],[f340,f168])).
fof(f340,plain,(
  relation_rng(sK13) = relation_dom(function_inverse(sK13)) | ~function(sK13) | ~relation(sK13)),
  inference(resolution,[],[f193,f169])).
fof(f169,plain,(
  one_to_one(sK13)),
  inference(cnf_transformation,[],[f107])).
fof(f193,plain,(
  ( ! [X0] : (~one_to_one(X0) | relation_rng(X0) = relation_dom(function_inverse(X0)) | ~function(X0) | ~relation(X0)) )),
  inference(subsumption_resolution,[],[f192,f114])).
fof(f114,plain,(
  ( ! [X0] : (relation(function_inverse(X0)) | ~function(X0) | ~relation(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : ((function(function_inverse(X0)) & relation(function_inverse(X0))) | ~function(X0) | ~relation(X0))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ((function(function_inverse(X0)) & relation(function_inverse(X0))) | (~function(X0) | ~relation(X0)))),
  inference(ennf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0] : ((function(X0) & relation(X0)) => (function(function_inverse(X0)) & relation(function_inverse(X0))))),
  file('./p/SEU220+1.tptp',unknown)).
fof(f192,plain,(
  ( ! [X0] : (relation_rng(X0) = relation_dom(function_inverse(X0)) | ~relation(function_inverse(X0)) | ~one_to_one(X0) | ~function(X0) | ~relation(X0)) )),
  inference(subsumption_resolution,[],[f183,f115])).
fof(f115,plain,(
  ( ! [X0] : (function(function_inverse(X0)) | ~function(X0) | ~relation(X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f183,plain,(
  ( ! [X0] : (relation_rng(X0) = relation_dom(function_inverse(X0)) | ~function(function_inverse(X0)) | ~relation(function_inverse(X0)) | ~one_to_one(X0) | ~function(X0) | ~relation(X0)) )),
  inference(equality_resolution,[],[f158])).
fof(f158,plain,(
  ( ! [X0,X1] : (relation_rng(X0) = relation_dom(X1) | function_inverse(X0) != X1 | ~function(X1) | ~relation(X1) | ~one_to_one(X0) | ~function(X0) | ~relation(X0)) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : (! [X1] : (((function_inverse(X0) = X1 | (((sK11(X0,X1) != apply(X1,sK10(X0,X1)) | ~in(sK10(X0,X1),relation_rng(X0))) & sK10(X0,X1) = apply(X0,sK11(X0,X1)) & in(sK11(X0,X1),relation_dom(X0))) | ~sP0(sK10(X0,X1),sK11(X0,X1),X0,X1)) | relation_rng(X0) != relation_dom(X1)) & ((! [X4,X5] : (((apply(X1,X4) = X5 & in(X4,relation_rng(X0))) | apply(X0,X5) != X4 | ~in(X5,relation_dom(X0))) & sP0(X4,X5,X0,X1)) & relation_rng(X0) = relation_dom(X1)) | function_inverse(X0) != X1)) | ~function(X1) | ~relation(X1)) | ~one_to_one(X0) | ~function(X0) | ~relation(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f103,f104])).
fof(f104,plain,(
  ! [X1,X0] : (? [X2,X3] : (((apply(X1,X2) != X3 | ~in(X2,relation_rng(X0))) & apply(X0,X3) = X2 & in(X3,relation_dom(X0))) | ~sP0(X2,X3,X0,X1)) => (((sK11(X0,X1) != apply(X1,sK10(X0,X1)) | ~in(sK10(X0,X1),relation_rng(X0))) & sK10(X0,X1) = apply(X0,sK11(X0,X1)) & in(sK11(X0,X1),relation_dom(X0))) | ~sP0(sK10(X0,X1),sK11(X0,X1),X0,X1)))),
  introduced(choice_axiom,[])).
fof(f103,plain,(
  ! [X0] : (! [X1] : (((function_inverse(X0) = X1 | ? [X2,X3] : (((apply(X1,X2) != X3 | ~in(X2,relation_rng(X0))) & apply(X0,X3) = X2 & in(X3,relation_dom(X0))) | ~sP0(X2,X3,X0,X1)) | relation_rng(X0) != relation_dom(X1)) & ((! [X4,X5] : (((apply(X1,X4) = X5 & in(X4,relation_rng(X0))) | apply(X0,X5) != X4 | ~in(X5,relation_dom(X0))) & sP0(X4,X5,X0,X1)) & relation_rng(X0) = relation_dom(X1)) | function_inverse(X0) != X1)) | ~function(X1) | ~relation(X1)) | ~one_to_one(X0) | ~function(X0) | ~relation(X0))),
  inference(rectify,[],[f102])).
fof(f102,plain,(
  ! [X0] : (! [X1] : (((function_inverse(X0) = X1 | ? [X2,X3] : (((apply(X1,X2) != X3 | ~in(X2,relation_rng(X0))) & apply(X0,X3) = X2 & in(X3,relation_dom(X0))) | ~sP0(X2,X3,X0,X1)) | relation_rng(X0) != relation_dom(X1)) & ((! [X2,X3] : (((apply(X1,X2) = X3 & in(X2,relation_rng(X0))) | apply(X0,X3) != X2 | ~in(X3,relation_dom(X0))) & sP0(X2,X3,X0,X1)) & relation_rng(X0) = relation_dom(X1)) | function_inverse(X0) != X1)) | ~function(X1) | ~relation(X1)) | ~one_to_one(X0) | ~function(X0) | ~relation(X0))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : (! [X1] : (((function_inverse(X0) = X1 | (? [X2,X3] : (((apply(X1,X2) != X3 | ~in(X2,relation_rng(X0))) & apply(X0,X3) = X2 & in(X3,relation_dom(X0))) | ~sP0(X2,X3,X0,X1)) | relation_rng(X0) != relation_dom(X1))) & ((! [X2,X3] : (((apply(X1,X2) = X3 & in(X2,relation_rng(X0))) | apply(X0,X3) != X2 | ~in(X3,relation_dom(X0))) & sP0(X2,X3,X0,X1)) & relation_rng(X0) = relation_dom(X1)) | function_inverse(X0) != X1)) | ~function(X1) | ~relation(X1)) | ~one_to_one(X0) | ~function(X0) | ~relation(X0))),
  inference(nnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : (! [X1] : ((function_inverse(X0) = X1 <=> (! [X2,X3] : (((apply(X1,X2) = X3 & in(X2,relation_rng(X0))) | apply(X0,X3) != X2 | ~in(X3,relation_dom(X0))) & sP0(X2,X3,X0,X1)) & relation_rng(X0) = relation_dom(X1))) | ~function(X1) | ~relation(X1)) | ~one_to_one(X0) | ~function(X0) | ~relation(X0))),
  inference(definition_folding,[],[f70,f78])).
fof(f78,plain,(
  ! [X2,X3,X0,X1] : (sP0(X2,X3,X0,X1) <=> ((apply(X0,X3) = X2 & in(X3,relation_dom(X0))) | apply(X1,X2) != X3 | ~in(X2,relation_rng(X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f70,plain,(
  ! [X0] : (! [X1] : ((function_inverse(X0) = X1 <=> (! [X2,X3] : (((apply(X1,X2) = X3 & in(X2,relation_rng(X0))) | apply(X0,X3) != X2 | ~in(X3,relation_dom(X0))) & ((apply(X0,X3) = X2 & in(X3,relation_dom(X0))) | apply(X1,X2) != X3 | ~in(X2,relation_rng(X0)))) & relation_rng(X0) = relation_dom(X1))) | ~function(X1) | ~relation(X1)) | ~one_to_one(X0) | ~function(X0) | ~relation(X0))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ((! [X1] : ((function_inverse(X0) = X1 <=> (! [X2,X3] : (((apply(X1,X2) = X3 & in(X2,relation_rng(X0))) | (apply(X0,X3) != X2 | ~in(X3,relation_dom(X0)))) & ((apply(X0,X3) = X2 & in(X3,relation_dom(X0))) | (apply(X1,X2) != X3 | ~in(X2,relation_rng(X0))))) & relation_rng(X0) = relation_dom(X1))) | (~function(X1) | ~relation(X1))) | ~one_to_one(X0)) | (~function(X0) | ~relation(X0)))),
  inference(ennf_transformation,[],[f34])).
fof(f34,axiom,(
  ! [X0] : ((function(X0) & relation(X0)) => (one_to_one(X0) => ! [X1] : ((function(X1) & relation(X1)) => (function_inverse(X0) = X1 <=> (! [X2,X3] : (((apply(X0,X3) = X2 & in(X3,relation_dom(X0))) => (apply(X1,X2) = X3 & in(X2,relation_rng(X0)))) & ((apply(X1,X2) = X3 & in(X2,relation_rng(X0))) => (apply(X0,X3) = X2 & in(X3,relation_dom(X0))))) & relation_rng(X0) = relation_dom(X1))))))),
  file('./p/SEU220+1.tptp',unknown)).
fof(f151,plain,(
  ( ! [X2,X0,X1] : (~in(X0,relation_dom(X1)) | apply(relation_composition(X1,X2),X0) = apply(X2,apply(X1,X0)) | ~function(X2) | ~relation(X2) | ~function(X1) | ~relation(X1)) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : (! [X2] : (apply(relation_composition(X1,X2),X0) = apply(X2,apply(X1,X0)) | ~in(X0,relation_dom(X1)) | ~function(X2) | ~relation(X2)) | ~function(X1) | ~relation(X1))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : (! [X2] : ((apply(relation_composition(X1,X2),X0) = apply(X2,apply(X1,X0)) | ~in(X0,relation_dom(X1))) | (~function(X2) | ~relation(X2))) | (~function(X1) | ~relation(X1)))),
  inference(ennf_transformation,[],[f32])).
fof(f32,axiom,(
  ! [X0,X1] : ((function(X1) & relation(X1)) => ! [X2] : ((function(X2) & relation(X2)) => (in(X0,relation_dom(X1)) => apply(relation_composition(X1,X2),X0) = apply(X2,apply(X1,X0)))))),
  file('./p/SEU220+1.tptp',unknown)).
fof(f789,plain,(
  spl14_1),
  inference(avatar_split_clause,[],[f788,f195])).
fof(f788,plain,(
  sK12 = apply(sK13,apply(function_inverse(sK13),sK12))),
  inference(subsumption_resolution,[],[f787,f167])).
fof(f787,plain,(
  sK12 = apply(sK13,apply(function_inverse(sK13),sK12)) | ~relation(sK13)),
  inference(subsumption_resolution,[],[f786,f168])).
fof(f786,plain,(
  sK12 = apply(sK13,apply(function_inverse(sK13),sK12)) | ~function(sK13) | ~relation(sK13)),
  inference(subsumption_resolution,[],[f770,f169])).
fof(f770,plain,(
  sK12 = apply(sK13,apply(function_inverse(sK13),sK12)) | ~one_to_one(sK13) | ~function(sK13) | ~relation(sK13)),
  inference(resolution,[],[f399,f170])).
fof(f399,plain,(
  ( ! [X6,X7] : (~in(X6,relation_rng(X7)) | apply(X7,apply(function_inverse(X7),X6)) = X6 | ~one_to_one(X7) | ~function(X7) | ~relation(X7)) )),
  inference(resolution,[],[f176,f191])).
fof(f191,plain,(
  ( ! [X4,X0,X5] : (sP0(X4,X5,X0,function_inverse(X0)) | ~one_to_one(X0) | ~function(X0) | ~relation(X0)) )),
  inference(subsumption_resolution,[],[f190,f114])).
fof(f190,plain,(
  ( ! [X4,X0,X5] : (sP0(X4,X5,X0,function_inverse(X0)) | ~relation(function_inverse(X0)) | ~one_to_one(X0) | ~function(X0) | ~relation(X0)) )),
  inference(subsumption_resolution,[],[f182,f115])).
fof(f182,plain,(
  ( ! [X4,X0,X5] : (sP0(X4,X5,X0,function_inverse(X0)) | ~function(function_inverse(X0)) | ~relation(function_inverse(X0)) | ~one_to_one(X0) | ~function(X0) | ~relation(X0)) )),
  inference(equality_resolution,[],[f159])).
fof(f159,plain,(
  ( ! [X4,X0,X5,X1] : (sP0(X4,X5,X0,X1) | function_inverse(X0) != X1 | ~function(X1) | ~relation(X1) | ~one_to_one(X0) | ~function(X0) | ~relation(X0)) )),
  inference(cnf_transformation,[],[f105])).
fof(f176,plain,(
  ( ! [X2,X0,X3] : (~sP0(X0,apply(X3,X0),X2,X3) | ~in(X0,relation_rng(X2)) | apply(X2,apply(X3,X0)) = X0) )),
  inference(equality_resolution,[],[f154])).
fof(f154,plain,(
  ( ! [X2,X0,X3,X1] : (apply(X2,X1) = X0 | apply(X3,X0) != X1 | ~in(X0,relation_rng(X2)) | ~sP0(X0,X1,X2,X3)) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3] : ((sP0(X0,X1,X2,X3) | ((apply(X2,X1) != X0 | ~in(X1,relation_dom(X2))) & apply(X3,X0) = X1 & in(X0,relation_rng(X2)))) & ((apply(X2,X1) = X0 & in(X1,relation_dom(X2))) | apply(X3,X0) != X1 | ~in(X0,relation_rng(X2)) | ~sP0(X0,X1,X2,X3)))),
  inference(rectify,[],[f99])).
fof(f99,plain,(
  ! [X2,X3,X0,X1] : ((sP0(X2,X3,X0,X1) | ((apply(X0,X3) != X2 | ~in(X3,relation_dom(X0))) & apply(X1,X2) = X3 & in(X2,relation_rng(X0)))) & ((apply(X0,X3) = X2 & in(X3,relation_dom(X0))) | apply(X1,X2) != X3 | ~in(X2,relation_rng(X0)) | ~sP0(X2,X3,X0,X1)))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X2,X3,X0,X1] : ((sP0(X2,X3,X0,X1) | ((apply(X0,X3) != X2 | ~in(X3,relation_dom(X0))) & apply(X1,X2) = X3 & in(X2,relation_rng(X0)))) & (((apply(X0,X3) = X2 & in(X3,relation_dom(X0))) | apply(X1,X2) != X3 | ~in(X2,relation_rng(X0))) | ~sP0(X2,X3,X0,X1)))),
  inference(nnf_transformation,[],[f78])).
fof(f474,plain,(
  spl14_6),
  inference(avatar_contradiction_clause,[],[f473])).
fof(f473,plain,(
  $false | spl14_6),
  inference(subsumption_resolution,[],[f472,f167])).
fof(f472,plain,(
  ~relation(sK13) | spl14_6),
  inference(subsumption_resolution,[],[f471,f168])).
fof(f471,plain,(
  ~function(sK13) | ~relation(sK13) | spl14_6),
  inference(resolution,[],[f416,f115])).
fof(f416,plain,(
  ~function(function_inverse(sK13)) | spl14_6),
  inference(avatar_component_clause,[],[f414])).
fof(f459,plain,(
  spl14_5),
  inference(avatar_contradiction_clause,[],[f458])).
fof(f458,plain,(
  $false | spl14_5),
  inference(subsumption_resolution,[],[f457,f167])).
fof(f457,plain,(
  ~relation(sK13) | spl14_5),
  inference(subsumption_resolution,[],[f456,f168])).
fof(f456,plain,(
  ~function(sK13) | ~relation(sK13) | spl14_5),
  inference(resolution,[],[f412,f114])).
fof(f412,plain,(
  ~relation(function_inverse(sK13)) | spl14_5),
  inference(avatar_component_clause,[],[f410])).
fof(f202,plain,(
  ~spl14_1 | ~spl14_2),
  inference(avatar_split_clause,[],[f171,f199,f195])).
fof(f171,plain,(
  sK12 != apply(relation_composition(function_inverse(sK13),sK13),sK12) | sK12 != apply(sK13,apply(function_inverse(sK13),sK12))),
  inference(cnf_transformation,[],[f107])).
% SZS output end Proof for SEU220+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 14967
% Time elapsed: 0.477 s
% ------------------------------
% ------------------------------
