% Refutation found. Thanks to Tanya!
% SZS status Theorem for SWC027+1
% SZS output start Proof for SWC027+1
fof(f1024,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f608,f609,f622,f632,f639,f640,f686,f992,f994,f996])).
fof(f996,plain,(
  spl58_2 | ~spl58_8),
  inference(avatar_split_clause,[],[f995,f634,f605])).
fof(f605,plain,(
  spl58_2 <=> nil = sK54),
  introduced(avatar_definition,[new_symbols(naming,[spl58_2])])).
fof(f634,plain,(
  spl58_8 <=> nil = sK56),
  introduced(avatar_definition,[new_symbols(naming,[spl58_8])])).
fof(f995,plain,(
  nil = sK54 | ~spl58_8),
  inference(backward_demodulation,[],[f636,f552])).
fof(f552,plain,(
  sK54 = sK56),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  (((((nil != sK54 & nil = sK53) | (nil != sK53 & nil = sK54)) & ((nil = sK55 & nil = sK56) | (! [X5] : (~leq(sK57,X5) | ~memberP(sK56,X5) | sK57 = X5 | ~ssItem(X5)) & memberP(sK56,sK57) & sK55 = cons(sK57,nil) & ssItem(sK57))) & sK53 = sK55 & sK54 = sK56 & ssList(sK56)) & ssList(sK55)) & ssList(sK54)) & ssList(sK53)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK53,sK54,sK55,sK56,sK57])],[f222,f347,f346,f345,f344,f343])).
fof(f343,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((nil != X1 & nil = X0) | (nil != X0 & nil = X1)) & ((nil = X2 & nil = X3) | ? [X4] : (! [X5] : (~leq(X4,X5) | ~memberP(X3,X5) | X4 = X5 | ~ssItem(X5)) & memberP(X3,X4) & cons(X4,nil) = X2 & ssItem(X4))) & X0 = X2 & X1 = X3 & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)) => (? [X1] : (? [X2] : (? [X3] : (((nil != X1 & nil = sK53) | (nil != sK53 & nil = X1)) & ((nil = X2 & nil = X3) | ? [X4] : (! [X5] : (~leq(X4,X5) | ~memberP(X3,X5) | X4 = X5 | ~ssItem(X5)) & memberP(X3,X4) & cons(X4,nil) = X2 & ssItem(X4))) & sK53 = X2 & X1 = X3 & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(sK53))),
  introduced(choice_axiom,[])).
fof(f344,plain,(
  ? [X1] : (? [X2] : (? [X3] : (((nil != X1 & nil = sK53) | (nil != sK53 & nil = X1)) & ((nil = X2 & nil = X3) | ? [X4] : (! [X5] : (~leq(X4,X5) | ~memberP(X3,X5) | X4 = X5 | ~ssItem(X5)) & memberP(X3,X4) & cons(X4,nil) = X2 & ssItem(X4))) & sK53 = X2 & X1 = X3 & ssList(X3)) & ssList(X2)) & ssList(X1)) => (? [X2] : (? [X3] : (((nil != sK54 & nil = sK53) | (nil != sK53 & nil = sK54)) & ((nil = X2 & nil = X3) | ? [X4] : (! [X5] : (~leq(X4,X5) | ~memberP(X3,X5) | X4 = X5 | ~ssItem(X5)) & memberP(X3,X4) & cons(X4,nil) = X2 & ssItem(X4))) & sK53 = X2 & sK54 = X3 & ssList(X3)) & ssList(X2)) & ssList(sK54))),
  introduced(choice_axiom,[])).
fof(f345,plain,(
  ? [X2] : (? [X3] : (((nil != sK54 & nil = sK53) | (nil != sK53 & nil = sK54)) & ((nil = X2 & nil = X3) | ? [X4] : (! [X5] : (~leq(X4,X5) | ~memberP(X3,X5) | X4 = X5 | ~ssItem(X5)) & memberP(X3,X4) & cons(X4,nil) = X2 & ssItem(X4))) & sK53 = X2 & sK54 = X3 & ssList(X3)) & ssList(X2)) => (? [X3] : (((nil != sK54 & nil = sK53) | (nil != sK53 & nil = sK54)) & ((nil = sK55 & nil = X3) | ? [X4] : (! [X5] : (~leq(X4,X5) | ~memberP(X3,X5) | X4 = X5 | ~ssItem(X5)) & memberP(X3,X4) & cons(X4,nil) = sK55 & ssItem(X4))) & sK53 = sK55 & sK54 = X3 & ssList(X3)) & ssList(sK55))),
  introduced(choice_axiom,[])).
fof(f346,plain,(
  ? [X3] : (((nil != sK54 & nil = sK53) | (nil != sK53 & nil = sK54)) & ((nil = sK55 & nil = X3) | ? [X4] : (! [X5] : (~leq(X4,X5) | ~memberP(X3,X5) | X4 = X5 | ~ssItem(X5)) & memberP(X3,X4) & cons(X4,nil) = sK55 & ssItem(X4))) & sK53 = sK55 & sK54 = X3 & ssList(X3)) => (((nil != sK54 & nil = sK53) | (nil != sK53 & nil = sK54)) & ((nil = sK55 & nil = sK56) | ? [X4] : (! [X5] : (~leq(X4,X5) | ~memberP(sK56,X5) | X4 = X5 | ~ssItem(X5)) & memberP(sK56,X4) & cons(X4,nil) = sK55 & ssItem(X4))) & sK53 = sK55 & sK54 = sK56 & ssList(sK56))),
  introduced(choice_axiom,[])).
fof(f347,plain,(
  ? [X4] : (! [X5] : (~leq(X4,X5) | ~memberP(sK56,X5) | X4 = X5 | ~ssItem(X5)) & memberP(sK56,X4) & cons(X4,nil) = sK55 & ssItem(X4)) => (! [X5] : (~leq(sK57,X5) | ~memberP(sK56,X5) | sK57 = X5 | ~ssItem(X5)) & memberP(sK56,sK57) & sK55 = cons(sK57,nil) & ssItem(sK57))),
  introduced(choice_axiom,[])).
fof(f222,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((nil != X1 & nil = X0) | (nil != X0 & nil = X1)) & ((nil = X2 & nil = X3) | ? [X4] : (! [X5] : (~leq(X4,X5) | ~memberP(X3,X5) | X4 = X5 | ~ssItem(X5)) & memberP(X3,X4) & cons(X4,nil) = X2 & ssItem(X4))) & X0 = X2 & X1 = X3 & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((nil != X1 & nil = X0) | (nil != X0 & nil = X1)) & ((nil = X2 & nil = X3) | ? [X4] : ((! [X5] : (~leq(X4,X5) | ~memberP(X3,X5) | X4 = X5 | ~ssItem(X5)) & memberP(X3,X4) & cons(X4,nil) = X2) & ssItem(X4))) & X0 = X2 & X1 = X3) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,negated_conjecture,(
  ~! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((nil = X1 | nil != X0) & (nil = X0 | nil != X1)) | ((nil != X2 | nil != X3) & ! [X4] : (ssItem(X4) => (? [X5] : (leq(X4,X5) & memberP(X3,X5) & X4 != X5 & ssItem(X5)) | ~memberP(X3,X4) | cons(X4,nil) != X2))) | X0 != X2 | X1 != X3)))))),
  inference(negated_conjecture,[],[f96])).
fof(f96,conjecture,(
  ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((nil = X1 | nil != X0) & (nil = X0 | nil != X1)) | ((nil != X2 | nil != X3) & ! [X4] : (ssItem(X4) => (? [X5] : (leq(X4,X5) & memberP(X3,X5) & X4 != X5 & ssItem(X5)) | ~memberP(X3,X4) | cons(X4,nil) != X2))) | X0 != X2 | X1 != X3)))))),
  file('./p/SWC027+1.tptp',unknown)).
fof(f636,plain,(
  nil = sK56 | ~spl58_8),
  inference(avatar_component_clause,[],[f634])).
fof(f994,plain,(
  spl58_1 | ~spl58_4),
  inference(avatar_split_clause,[],[f993,f614,f601])).
fof(f601,plain,(
  spl58_1 <=> nil = sK53),
  introduced(avatar_definition,[new_symbols(naming,[spl58_1])])).
fof(f614,plain,(
  spl58_4 <=> nil = sK55),
  introduced(avatar_definition,[new_symbols(naming,[spl58_4])])).
fof(f993,plain,(
  nil = sK53 | ~spl58_4),
  inference(forward_demodulation,[],[f553,f616])).
fof(f616,plain,(
  nil = sK55 | ~spl58_4),
  inference(avatar_component_clause,[],[f614])).
fof(f553,plain,(
  sK53 = sK55),
  inference(cnf_transformation,[],[f348])).
fof(f992,plain,(
  ~spl58_1 | ~spl58_6 | ~spl58_7),
  inference(avatar_contradiction_clause,[],[f991])).
fof(f991,plain,(
  $false | (~spl58_1 | ~spl58_6 | ~spl58_7)),
  inference(subsumption_resolution,[],[f990,f441])).
fof(f441,plain,(
  ssList(nil)),
  inference(cnf_transformation,[],[f17])).
fof(f17,axiom,(
  ssList(nil)),
  file('./p/SWC027+1.tptp',unknown)).
fof(f990,plain,(
  ~ssList(nil) | (~spl58_1 | ~spl58_6 | ~spl58_7)),
  inference(subsumption_resolution,[],[f989,f631])).
fof(f631,plain,(
  ssItem(sK57) | ~spl58_7),
  inference(avatar_component_clause,[],[f629])).
fof(f629,plain,(
  spl58_7 <=> ssItem(sK57)),
  introduced(avatar_definition,[new_symbols(naming,[spl58_7])])).
fof(f989,plain,(
  ~ssItem(sK57) | ~ssList(nil) | (~spl58_1 | ~spl58_6)),
  inference(trivial_inequality_removal,[],[f988])).
fof(f988,plain,(
  nil != nil | ~ssItem(sK57) | ~ssList(nil) | (~spl58_1 | ~spl58_6)),
  inference(superposition,[],[f442,f691])).
fof(f691,plain,(
  nil = cons(sK57,nil) | (~spl58_1 | ~spl58_6)),
  inference(backward_demodulation,[],[f602,f643])).
fof(f643,plain,(
  sK53 = cons(sK57,nil) | ~spl58_6),
  inference(forward_demodulation,[],[f626,f553])).
fof(f626,plain,(
  sK55 = cons(sK57,nil) | ~spl58_6),
  inference(avatar_component_clause,[],[f624])).
fof(f624,plain,(
  spl58_6 <=> sK55 = cons(sK57,nil)),
  introduced(avatar_definition,[new_symbols(naming,[spl58_6])])).
fof(f602,plain,(
  nil = sK53 | ~spl58_1),
  inference(avatar_component_clause,[],[f601])).
fof(f442,plain,(
  ( ! [X0,X1] : (cons(X1,X0) != X0 | ~ssItem(X1) | ~ssList(X0)) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : (! [X1] : (cons(X1,X0) != X0 | ~ssItem(X1)) | ~ssList(X0))),
  inference(ennf_transformation,[],[f18])).
fof(f18,axiom,(
  ! [X0] : (ssList(X0) => ! [X1] : (ssItem(X1) => cons(X1,X0) != X0))),
  file('./p/SWC027+1.tptp',unknown)).
fof(f686,plain,(
  ~spl58_2 | ~spl58_5 | ~spl58_7),
  inference(avatar_contradiction_clause,[],[f685])).
fof(f685,plain,(
  $false | (~spl58_2 | ~spl58_5 | ~spl58_7)),
  inference(subsumption_resolution,[],[f684,f631])).
fof(f684,plain,(
  ~ssItem(sK57) | (~spl58_2 | ~spl58_5)),
  inference(resolution,[],[f471,f647])).
fof(f647,plain,(
  memberP(nil,sK57) | (~spl58_2 | ~spl58_5)),
  inference(backward_demodulation,[],[f606,f644])).
fof(f644,plain,(
  memberP(sK54,sK57) | ~spl58_5),
  inference(forward_demodulation,[],[f621,f552])).
fof(f621,plain,(
  memberP(sK56,sK57) | ~spl58_5),
  inference(avatar_component_clause,[],[f619])).
fof(f619,plain,(
  spl58_5 <=> memberP(sK56,sK57)),
  introduced(avatar_definition,[new_symbols(naming,[spl58_5])])).
fof(f606,plain,(
  nil = sK54 | ~spl58_2),
  inference(avatar_component_clause,[],[f605])).
fof(f471,plain,(
  ( ! [X0] : (~memberP(nil,X0) | ~ssItem(X0)) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0] : (~memberP(nil,X0) | ~ssItem(X0))),
  inference(ennf_transformation,[],[f38])).
fof(f38,axiom,(
  ! [X0] : (ssItem(X0) => ~memberP(nil,X0))),
  file('./p/SWC027+1.tptp',unknown)).
fof(f640,plain,(
  spl58_7 | spl58_8),
  inference(avatar_split_clause,[],[f554,f634,f629])).
fof(f554,plain,(
  nil = sK56 | ssItem(sK57)),
  inference(cnf_transformation,[],[f348])).
fof(f639,plain,(
  spl58_6 | spl58_8),
  inference(avatar_split_clause,[],[f555,f634,f624])).
fof(f555,plain,(
  nil = sK56 | sK55 = cons(sK57,nil)),
  inference(cnf_transformation,[],[f348])).
fof(f632,plain,(
  spl58_7 | spl58_4),
  inference(avatar_split_clause,[],[f558,f614,f629])).
fof(f558,plain,(
  nil = sK55 | ssItem(sK57)),
  inference(cnf_transformation,[],[f348])).
fof(f622,plain,(
  spl58_5 | spl58_4),
  inference(avatar_split_clause,[],[f560,f614,f619])).
fof(f560,plain,(
  nil = sK55 | memberP(sK56,sK57)),
  inference(cnf_transformation,[],[f348])).
fof(f609,plain,(
  spl58_2 | spl58_1),
  inference(avatar_split_clause,[],[f562,f601,f605])).
fof(f562,plain,(
  nil = sK53 | nil = sK54),
  inference(cnf_transformation,[],[f348])).
fof(f608,plain,(
  ~spl58_1 | ~spl58_2),
  inference(avatar_split_clause,[],[f565,f605,f601])).
fof(f565,plain,(
  nil != sK54 | nil != sK53),
  inference(cnf_transformation,[],[f348])).
% SZS output end Proof for SWC027+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5500
% Time elapsed: 0.006 s
% ------------------------------
% ------------------------------
