% Refutation found. Thanks to Tanya!
% SZS status Theorem for SWC039+1
% SZS output start Proof for SWC039+1
fof(f4081,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f615,f622,f628,f3911,f3984,f4026,f4072])).
fof(f4072,plain,(
  ~spl61_105),
  inference(avatar_contradiction_clause,[],[f4071])).
fof(f4071,plain,(
  $false | ~spl61_105),
  inference(subsumption_resolution,[],[f4070,f623])).
fof(f623,plain,(
  ssList(sK57)),
  inference(forward_demodulation,[],[f565,f569])).
fof(f569,plain,(
  sK57 = sK59),
  inference(cnf_transformation,[],[f356])).
fof(f356,plain,(
  (((((nil = sK59 & nil = sK60) | sP6(sK60,sK59)) & nil != sK57 & sK57 = sK59 & sK58 = sK60 & nil = sK58 & ssList(sK60)) & ssList(sK59)) & ssList(sK58)) & ssList(sK57)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK57,sK58,sK59,sK60])],[f234,f355,f354,f353,f352])).
fof(f352,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((nil = X2 & nil = X3) | sP6(X3,X2)) & nil != X0 & X0 = X2 & X1 = X3 & nil = X1 & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)) => (? [X1] : (? [X2] : (? [X3] : (((nil = X2 & nil = X3) | sP6(X3,X2)) & nil != sK57 & sK57 = X2 & X1 = X3 & nil = X1 & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(sK57))),
  introduced(choice_axiom,[])).
fof(f353,plain,(
  ? [X1] : (? [X2] : (? [X3] : (((nil = X2 & nil = X3) | sP6(X3,X2)) & nil != sK57 & sK57 = X2 & X1 = X3 & nil = X1 & ssList(X3)) & ssList(X2)) & ssList(X1)) => (? [X2] : (? [X3] : (((nil = X2 & nil = X3) | sP6(X3,X2)) & nil != sK57 & sK57 = X2 & sK58 = X3 & nil = sK58 & ssList(X3)) & ssList(X2)) & ssList(sK58))),
  introduced(choice_axiom,[])).
fof(f354,plain,(
  ? [X2] : (? [X3] : (((nil = X2 & nil = X3) | sP6(X3,X2)) & nil != sK57 & sK57 = X2 & sK58 = X3 & nil = sK58 & ssList(X3)) & ssList(X2)) => (? [X3] : (((nil = sK59 & nil = X3) | sP6(X3,sK59)) & nil != sK57 & sK57 = sK59 & sK58 = X3 & nil = sK58 & ssList(X3)) & ssList(sK59))),
  introduced(choice_axiom,[])).
fof(f355,plain,(
  ? [X3] : (((nil = sK59 & nil = X3) | sP6(X3,sK59)) & nil != sK57 & sK57 = sK59 & sK58 = X3 & nil = sK58 & ssList(X3)) => (((nil = sK59 & nil = sK60) | sP6(sK60,sK59)) & nil != sK57 & sK57 = sK59 & sK58 = sK60 & nil = sK58 & ssList(sK60))),
  introduced(choice_axiom,[])).
fof(f234,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((nil = X2 & nil = X3) | sP6(X3,X2)) & nil != X0 & X0 = X2 & X1 = X3 & nil = X1 & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0))),
  inference(definition_folding,[],[f223,f233])).
fof(f233,plain,(
  ! [X3,X2] : (? [X4] : (? [X5] : (? [X6] : (! [X7] : (~lt(X7,X4) | ~memberP(X6,X7) | ~ssItem(X7)) & ! [X8] : (~lt(X4,X8) | ~memberP(X5,X8) | ~ssItem(X8)) & app(app(X5,X2),X6) = X3 & cons(X4,nil) = X2 & ssList(X6)) & ssList(X5)) & ssItem(X4)) | ~sP6(X3,X2))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f223,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((nil = X2 & nil = X3) | ? [X4] : (? [X5] : (? [X6] : (! [X7] : (~lt(X7,X4) | ~memberP(X6,X7) | ~ssItem(X7)) & ! [X8] : (~lt(X4,X8) | ~memberP(X5,X8) | ~ssItem(X8)) & app(app(X5,X2),X6) = X3 & cons(X4,nil) = X2 & ssList(X6)) & ssList(X5)) & ssItem(X4))) & nil != X0 & X0 = X2 & X1 = X3 & nil = X1 & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((((nil = X2 & nil = X3) | ? [X4] : (? [X5] : (? [X6] : ((! [X7] : (~lt(X7,X4) | ~memberP(X6,X7) | ~ssItem(X7)) & ! [X8] : (~lt(X4,X8) | ~memberP(X5,X8) | ~ssItem(X8)) & app(app(X5,X2),X6) = X3 & cons(X4,nil) = X2) & ssList(X6)) & ssList(X5)) & ssItem(X4))) & nil != X0 & X0 = X2 & X1 = X3 & nil = X1) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ~! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((nil != X2 | nil != X3) & ! [X4] : (ssItem(X4) => ! [X5] : (ssList(X5) => ! [X6] : (ssList(X6) => (? [X7] : (lt(X7,X4) & memberP(X6,X7) & ssItem(X7)) | ? [X8] : (lt(X4,X8) & memberP(X5,X8) & ssItem(X8)) | app(app(X5,X2),X6) != X3 | cons(X4,nil) != X2))))) | nil = X0 | X0 != X2 | X1 != X3 | nil != X1)))))),
  inference(rectify,[],[f97])).
fof(f97,negated_conjecture,(
  ~! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((nil != X2 | nil != X3) & ! [X4] : (ssItem(X4) => ! [X5] : (ssList(X5) => ! [X6] : (ssList(X6) => (? [X8] : (lt(X8,X4) & memberP(X6,X8) & ssItem(X8)) | ? [X7] : (lt(X4,X7) & memberP(X5,X7) & ssItem(X7)) | app(app(X5,X2),X6) != X3 | cons(X4,nil) != X2))))) | nil = X0 | X0 != X2 | X1 != X3 | nil != X1)))))),
  inference(negated_conjecture,[],[f96])).
fof(f96,conjecture,(
  ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => (((nil != X2 | nil != X3) & ! [X4] : (ssItem(X4) => ! [X5] : (ssList(X5) => ! [X6] : (ssList(X6) => (? [X8] : (lt(X8,X4) & memberP(X6,X8) & ssItem(X8)) | ? [X7] : (lt(X4,X7) & memberP(X5,X7) & ssItem(X7)) | app(app(X5,X2),X6) != X3 | cons(X4,nil) != X2))))) | nil = X0 | X0 != X2 | X1 != X3 | nil != X1)))))),
  file('./p/SWC039+1.tptp',unknown)).
fof(f565,plain,(
  ssList(sK59)),
  inference(cnf_transformation,[],[f356])).
fof(f4070,plain,(
  ~ssList(sK57) | ~spl61_105),
  inference(subsumption_resolution,[],[f4066,f570])).
fof(f570,plain,(
  nil != sK57),
  inference(cnf_transformation,[],[f356])).
fof(f4066,plain,(
  nil = sK57 | ~ssList(sK57) | ~spl61_105),
  inference(resolution,[],[f3910,f503])).
fof(f503,plain,(
  ( ! [X0] : (~segmentP(nil,X0) | nil = X0 | ~ssList(X0)) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0] : (((segmentP(nil,X0) | nil != X0) & (nil = X0 | ~segmentP(nil,X0))) | ~ssList(X0))),
  inference(nnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0] : ((segmentP(nil,X0) <=> nil = X0) | ~ssList(X0))),
  inference(ennf_transformation,[],[f58])).
fof(f58,axiom,(
  ! [X0] : (ssList(X0) => (segmentP(nil,X0) <=> nil = X0))),
  file('./p/SWC039+1.tptp',unknown)).
fof(f3910,plain,(
  segmentP(nil,sK57) | ~spl61_105),
  inference(avatar_component_clause,[],[f3908])).
fof(f3908,plain,(
  spl61_105 <=> segmentP(nil,sK57)),
  introduced(avatar_definition,[new_symbols(naming,[spl61_105])])).
fof(f4026,plain,(
  ~spl61_1 | ~spl61_3 | spl61_103),
  inference(avatar_contradiction_clause,[],[f4025])).
fof(f4025,plain,(
  $false | (~spl61_1 | ~spl61_3 | spl61_103)),
  inference(subsumption_resolution,[],[f4024,f630])).
fof(f630,plain,(
  sP6(nil,sK57) | (~spl61_1 | ~spl61_3)),
  inference(forward_demodulation,[],[f629,f619])).
fof(f619,plain,(
  nil = sK60 | ~spl61_3),
  inference(avatar_component_clause,[],[f617])).
fof(f617,plain,(
  spl61_3 <=> nil = sK60),
  introduced(avatar_definition,[new_symbols(naming,[spl61_3])])).
fof(f629,plain,(
  sP6(sK60,sK57) | ~spl61_1),
  inference(forward_demodulation,[],[f610,f569])).
fof(f610,plain,(
  sP6(sK60,sK59) | ~spl61_1),
  inference(avatar_component_clause,[],[f608])).
fof(f608,plain,(
  spl61_1 <=> sP6(sK60,sK59)),
  introduced(avatar_definition,[new_symbols(naming,[spl61_1])])).
fof(f4024,plain,(
  ~sP6(nil,sK57) | spl61_103),
  inference(resolution,[],[f3902,f557])).
fof(f557,plain,(
  ( ! [X0,X1] : (ssList(sK55(X0,X1)) | ~sP6(X0,X1)) )),
  inference(cnf_transformation,[],[f351])).
fof(f351,plain,(
  ! [X0,X1] : ((((! [X5] : (~lt(X5,sK54(X0,X1)) | ~memberP(sK56(X0,X1),X5) | ~ssItem(X5)) & ! [X6] : (~lt(sK54(X0,X1),X6) | ~memberP(sK55(X0,X1),X6) | ~ssItem(X6)) & app(app(sK55(X0,X1),X1),sK56(X0,X1)) = X0 & cons(sK54(X0,X1),nil) = X1 & ssList(sK56(X0,X1))) & ssList(sK55(X0,X1))) & ssItem(sK54(X0,X1))) | ~sP6(X0,X1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK54,sK55,sK56])],[f347,f350,f349,f348])).
fof(f348,plain,(
  ! [X1,X0] : (? [X2] : (? [X3] : (? [X4] : (! [X5] : (~lt(X5,X2) | ~memberP(X4,X5) | ~ssItem(X5)) & ! [X6] : (~lt(X2,X6) | ~memberP(X3,X6) | ~ssItem(X6)) & app(app(X3,X1),X4) = X0 & cons(X2,nil) = X1 & ssList(X4)) & ssList(X3)) & ssItem(X2)) => (? [X3] : (? [X4] : (! [X5] : (~lt(X5,sK54(X0,X1)) | ~memberP(X4,X5) | ~ssItem(X5)) & ! [X6] : (~lt(sK54(X0,X1),X6) | ~memberP(X3,X6) | ~ssItem(X6)) & app(app(X3,X1),X4) = X0 & cons(sK54(X0,X1),nil) = X1 & ssList(X4)) & ssList(X3)) & ssItem(sK54(X0,X1))))),
  introduced(choice_axiom,[])).
fof(f349,plain,(
  ! [X1,X0] : (? [X3] : (? [X4] : (! [X5] : (~lt(X5,sK54(X0,X1)) | ~memberP(X4,X5) | ~ssItem(X5)) & ! [X6] : (~lt(sK54(X0,X1),X6) | ~memberP(X3,X6) | ~ssItem(X6)) & app(app(X3,X1),X4) = X0 & cons(sK54(X0,X1),nil) = X1 & ssList(X4)) & ssList(X3)) => (? [X4] : (! [X5] : (~lt(X5,sK54(X0,X1)) | ~memberP(X4,X5) | ~ssItem(X5)) & ! [X6] : (~lt(sK54(X0,X1),X6) | ~memberP(sK55(X0,X1),X6) | ~ssItem(X6)) & app(app(sK55(X0,X1),X1),X4) = X0 & cons(sK54(X0,X1),nil) = X1 & ssList(X4)) & ssList(sK55(X0,X1))))),
  introduced(choice_axiom,[])).
fof(f350,plain,(
  ! [X1,X0] : (? [X4] : (! [X5] : (~lt(X5,sK54(X0,X1)) | ~memberP(X4,X5) | ~ssItem(X5)) & ! [X6] : (~lt(sK54(X0,X1),X6) | ~memberP(sK55(X0,X1),X6) | ~ssItem(X6)) & app(app(sK55(X0,X1),X1),X4) = X0 & cons(sK54(X0,X1),nil) = X1 & ssList(X4)) => (! [X5] : (~lt(X5,sK54(X0,X1)) | ~memberP(sK56(X0,X1),X5) | ~ssItem(X5)) & ! [X6] : (~lt(sK54(X0,X1),X6) | ~memberP(sK55(X0,X1),X6) | ~ssItem(X6)) & app(app(sK55(X0,X1),X1),sK56(X0,X1)) = X0 & cons(sK54(X0,X1),nil) = X1 & ssList(sK56(X0,X1))))),
  introduced(choice_axiom,[])).
fof(f347,plain,(
  ! [X0,X1] : (? [X2] : (? [X3] : (? [X4] : (! [X5] : (~lt(X5,X2) | ~memberP(X4,X5) | ~ssItem(X5)) & ! [X6] : (~lt(X2,X6) | ~memberP(X3,X6) | ~ssItem(X6)) & app(app(X3,X1),X4) = X0 & cons(X2,nil) = X1 & ssList(X4)) & ssList(X3)) & ssItem(X2)) | ~sP6(X0,X1))),
  inference(rectify,[],[f346])).
fof(f346,plain,(
  ! [X3,X2] : (? [X4] : (? [X5] : (? [X6] : (! [X7] : (~lt(X7,X4) | ~memberP(X6,X7) | ~ssItem(X7)) & ! [X8] : (~lt(X4,X8) | ~memberP(X5,X8) | ~ssItem(X8)) & app(app(X5,X2),X6) = X3 & cons(X4,nil) = X2 & ssList(X6)) & ssList(X5)) & ssItem(X4)) | ~sP6(X3,X2))),
  inference(nnf_transformation,[],[f233])).
fof(f3902,plain,(
  ~ssList(sK55(nil,sK57)) | spl61_103),
  inference(avatar_component_clause,[],[f3900])).
fof(f3900,plain,(
  spl61_103 <=> ssList(sK55(nil,sK57))),
  introduced(avatar_definition,[new_symbols(naming,[spl61_103])])).
fof(f3984,plain,(
  ~spl61_1 | ~spl61_3 | spl61_104),
  inference(avatar_contradiction_clause,[],[f3983])).
fof(f3983,plain,(
  $false | (~spl61_1 | ~spl61_3 | spl61_104)),
  inference(subsumption_resolution,[],[f3982,f630])).
fof(f3982,plain,(
  ~sP6(nil,sK57) | spl61_104),
  inference(resolution,[],[f3906,f558])).
fof(f558,plain,(
  ( ! [X0,X1] : (ssList(sK56(X0,X1)) | ~sP6(X0,X1)) )),
  inference(cnf_transformation,[],[f351])).
fof(f3906,plain,(
  ~ssList(sK56(nil,sK57)) | spl61_104),
  inference(avatar_component_clause,[],[f3904])).
fof(f3904,plain,(
  spl61_104 <=> ssList(sK56(nil,sK57))),
  introduced(avatar_definition,[new_symbols(naming,[spl61_104])])).
fof(f3911,plain,(
  ~spl61_103 | ~spl61_104 | spl61_105 | ~spl61_1 | ~spl61_3),
  inference(avatar_split_clause,[],[f3898,f617,f608,f3908,f3904,f3900])).
fof(f3898,plain,(
  segmentP(nil,sK57) | ~ssList(sK56(nil,sK57)) | ~ssList(sK55(nil,sK57)) | (~spl61_1 | ~spl61_3)),
  inference(subsumption_resolution,[],[f3897,f449])).
fof(f449,plain,(
  ssList(nil)),
  inference(cnf_transformation,[],[f17])).
fof(f17,axiom,(
  ssList(nil)),
  file('./p/SWC039+1.tptp',unknown)).
fof(f3897,plain,(
  ~ssList(nil) | segmentP(nil,sK57) | ~ssList(sK56(nil,sK57)) | ~ssList(sK55(nil,sK57)) | (~spl61_1 | ~spl61_3)),
  inference(forward_demodulation,[],[f3896,f1707])).
fof(f1707,plain,(
  nil = app(app(sK55(nil,sK57),sK57),sK56(nil,sK57)) | (~spl61_1 | ~spl61_3)),
  inference(resolution,[],[f560,f630])).
fof(f560,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1) | app(app(sK55(X0,X1),X1),sK56(X0,X1)) = X0) )),
  inference(cnf_transformation,[],[f351])).
fof(f3896,plain,(
  segmentP(nil,sK57) | ~ssList(sK56(nil,sK57)) | ~ssList(sK55(nil,sK57)) | ~ssList(app(app(sK55(nil,sK57),sK57),sK56(nil,sK57))) | (~spl61_1 | ~spl61_3)),
  inference(subsumption_resolution,[],[f3873,f623])).
fof(f3873,plain,(
  segmentP(nil,sK57) | ~ssList(sK56(nil,sK57)) | ~ssList(sK55(nil,sK57)) | ~ssList(sK57) | ~ssList(app(app(sK55(nil,sK57),sK57),sK56(nil,sK57))) | (~spl61_1 | ~spl61_3)),
  inference(superposition,[],[f578,f1707])).
fof(f578,plain,(
  ( ! [X2,X3,X1] : (segmentP(app(app(X2,X1),X3),X1) | ~ssList(X3) | ~ssList(X2) | ~ssList(X1) | ~ssList(app(app(X2,X1),X3))) )),
  inference(equality_resolution,[],[f378])).
fof(f378,plain,(
  ( ! [X2,X0,X3,X1] : (segmentP(X0,X1) | app(app(X2,X1),X3) != X0 | ~ssList(X3) | ~ssList(X2) | ~ssList(X1) | ~ssList(X0)) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0] : (! [X1] : (((segmentP(X0,X1) | ! [X2] : (! [X3] : (app(app(X2,X1),X3) != X0 | ~ssList(X3)) | ~ssList(X2))) & (((app(app(sK14(X0,X1),X1),sK15(X0,X1)) = X0 & ssList(sK15(X0,X1))) & ssList(sK14(X0,X1))) | ~segmentP(X0,X1))) | ~ssList(X1)) | ~ssList(X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f257,f259,f258])).
fof(f258,plain,(
  ! [X1,X0] : (? [X4] : (? [X5] : (app(app(X4,X1),X5) = X0 & ssList(X5)) & ssList(X4)) => (? [X5] : (app(app(sK14(X0,X1),X1),X5) = X0 & ssList(X5)) & ssList(sK14(X0,X1))))),
  introduced(choice_axiom,[])).
fof(f259,plain,(
  ! [X1,X0] : (? [X5] : (app(app(sK14(X0,X1),X1),X5) = X0 & ssList(X5)) => (app(app(sK14(X0,X1),X1),sK15(X0,X1)) = X0 & ssList(sK15(X0,X1))))),
  introduced(choice_axiom,[])).
fof(f257,plain,(
  ! [X0] : (! [X1] : (((segmentP(X0,X1) | ! [X2] : (! [X3] : (app(app(X2,X1),X3) != X0 | ~ssList(X3)) | ~ssList(X2))) & (? [X4] : (? [X5] : (app(app(X4,X1),X5) = X0 & ssList(X5)) & ssList(X4)) | ~segmentP(X0,X1))) | ~ssList(X1)) | ~ssList(X0))),
  inference(rectify,[],[f256])).
fof(f256,plain,(
  ! [X0] : (! [X1] : (((segmentP(X0,X1) | ! [X2] : (! [X3] : (app(app(X2,X1),X3) != X0 | ~ssList(X3)) | ~ssList(X2))) & (? [X2] : (? [X3] : (app(app(X2,X1),X3) = X0 & ssList(X3)) & ssList(X2)) | ~segmentP(X0,X1))) | ~ssList(X1)) | ~ssList(X0))),
  inference(nnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : (! [X1] : ((segmentP(X0,X1) <=> ? [X2] : (? [X3] : (app(app(X2,X1),X3) = X0 & ssList(X3)) & ssList(X2))) | ~ssList(X1)) | ~ssList(X0))),
  inference(ennf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => (segmentP(X0,X1) <=> ? [X2] : (? [X3] : (app(app(X2,X1),X3) = X0 & ssList(X3)) & ssList(X2)))))),
  file('./p/SWC039+1.tptp',unknown)).
fof(f628,plain,(
  ~spl61_2),
  inference(avatar_contradiction_clause,[],[f627])).
fof(f627,plain,(
  $false | ~spl61_2),
  inference(subsumption_resolution,[],[f626,f570])).
fof(f626,plain,(
  nil = sK57 | ~spl61_2),
  inference(backward_demodulation,[],[f614,f569])).
fof(f614,plain,(
  nil = sK59 | ~spl61_2),
  inference(avatar_component_clause,[],[f612])).
fof(f612,plain,(
  spl61_2 <=> nil = sK59),
  introduced(avatar_definition,[new_symbols(naming,[spl61_2])])).
fof(f622,plain,(
  spl61_3),
  inference(avatar_split_clause,[],[f621,f617])).
fof(f621,plain,(
  nil = sK60),
  inference(backward_demodulation,[],[f567,f568])).
fof(f568,plain,(
  sK58 = sK60),
  inference(cnf_transformation,[],[f356])).
fof(f567,plain,(
  nil = sK58),
  inference(cnf_transformation,[],[f356])).
fof(f615,plain,(
  spl61_1 | spl61_2),
  inference(avatar_split_clause,[],[f572,f612,f608])).
fof(f572,plain,(
  nil = sK59 | sP6(sK60,sK59)),
  inference(cnf_transformation,[],[f356])).
% SZS output end Proof for SWC039+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 8571
% Time elapsed: 0.039 s
% ------------------------------
% ------------------------------
