% Refutation found. Thanks to Tanya!
% SZS status Theorem for SWC043+1
% SZS output start Proof for SWC043+1
fof(f1346,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f1169,f1345])).
fof(f1345,plain,(
  spl58_20),
  inference(avatar_contradiction_clause,[],[f1344])).
fof(f1344,plain,(
  $false | spl58_20),
  inference(subsumption_resolution,[],[f1343,f611])).
fof(f611,plain,(
  ssList(sK53)),
  inference(forward_demodulation,[],[f550,f554])).
fof(f554,plain,(
  sK53 = sK55),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  ((((nil != sK55 | nil = sK56) & (! [X5] : (! [X6] : (! [X7] : (! [X8] : (~lt(X7,X5) | app(X8,cons(X7,nil)) != sK55 | ~ssList(X8)) | ~ssItem(X7)) | app(cons(X5,nil),X6) != sK57 | ~ssList(X6)) | ~ssItem(X5)) & strictorderedP(sK55) & sK56 = app(sK55,sK57) & ssList(sK57)) & nil != sK53 & sK53 = sK55 & sK54 = sK56 & nil = sK54 & ssList(sK56)) & ssList(sK55)) & ssList(sK54)) & ssList(sK53)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK53,sK54,sK55,sK56,sK57])],[f222,f347,f346,f345,f344,f343])).
fof(f343,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((nil != X2 | nil = X3) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (~lt(X7,X5) | app(X8,cons(X7,nil)) != X2 | ~ssList(X8)) | ~ssItem(X7)) | app(cons(X5,nil),X6) != X4 | ~ssList(X6)) | ~ssItem(X5)) & strictorderedP(X2) & app(X2,X4) = X3 & ssList(X4)) & nil != X0 & X0 = X2 & X1 = X3 & nil = X1 & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0)) => (? [X1] : (? [X2] : (? [X3] : ((nil != X2 | nil = X3) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (~lt(X7,X5) | app(X8,cons(X7,nil)) != X2 | ~ssList(X8)) | ~ssItem(X7)) | app(cons(X5,nil),X6) != X4 | ~ssList(X6)) | ~ssItem(X5)) & strictorderedP(X2) & app(X2,X4) = X3 & ssList(X4)) & nil != sK53 & sK53 = X2 & X1 = X3 & nil = X1 & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(sK53))),
  introduced(choice_axiom,[])).
fof(f344,plain,(
  ? [X1] : (? [X2] : (? [X3] : ((nil != X2 | nil = X3) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (~lt(X7,X5) | app(X8,cons(X7,nil)) != X2 | ~ssList(X8)) | ~ssItem(X7)) | app(cons(X5,nil),X6) != X4 | ~ssList(X6)) | ~ssItem(X5)) & strictorderedP(X2) & app(X2,X4) = X3 & ssList(X4)) & nil != sK53 & sK53 = X2 & X1 = X3 & nil = X1 & ssList(X3)) & ssList(X2)) & ssList(X1)) => (? [X2] : (? [X3] : ((nil != X2 | nil = X3) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (~lt(X7,X5) | app(X8,cons(X7,nil)) != X2 | ~ssList(X8)) | ~ssItem(X7)) | app(cons(X5,nil),X6) != X4 | ~ssList(X6)) | ~ssItem(X5)) & strictorderedP(X2) & app(X2,X4) = X3 & ssList(X4)) & nil != sK53 & sK53 = X2 & sK54 = X3 & nil = sK54 & ssList(X3)) & ssList(X2)) & ssList(sK54))),
  introduced(choice_axiom,[])).
fof(f345,plain,(
  ? [X2] : (? [X3] : ((nil != X2 | nil = X3) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (~lt(X7,X5) | app(X8,cons(X7,nil)) != X2 | ~ssList(X8)) | ~ssItem(X7)) | app(cons(X5,nil),X6) != X4 | ~ssList(X6)) | ~ssItem(X5)) & strictorderedP(X2) & app(X2,X4) = X3 & ssList(X4)) & nil != sK53 & sK53 = X2 & sK54 = X3 & nil = sK54 & ssList(X3)) & ssList(X2)) => (? [X3] : ((nil != sK55 | nil = X3) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (~lt(X7,X5) | app(X8,cons(X7,nil)) != sK55 | ~ssList(X8)) | ~ssItem(X7)) | app(cons(X5,nil),X6) != X4 | ~ssList(X6)) | ~ssItem(X5)) & strictorderedP(sK55) & app(sK55,X4) = X3 & ssList(X4)) & nil != sK53 & sK53 = sK55 & sK54 = X3 & nil = sK54 & ssList(X3)) & ssList(sK55))),
  introduced(choice_axiom,[])).
fof(f346,plain,(
  ? [X3] : ((nil != sK55 | nil = X3) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (~lt(X7,X5) | app(X8,cons(X7,nil)) != sK55 | ~ssList(X8)) | ~ssItem(X7)) | app(cons(X5,nil),X6) != X4 | ~ssList(X6)) | ~ssItem(X5)) & strictorderedP(sK55) & app(sK55,X4) = X3 & ssList(X4)) & nil != sK53 & sK53 = sK55 & sK54 = X3 & nil = sK54 & ssList(X3)) => ((nil != sK55 | nil = sK56) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (~lt(X7,X5) | app(X8,cons(X7,nil)) != sK55 | ~ssList(X8)) | ~ssItem(X7)) | app(cons(X5,nil),X6) != X4 | ~ssList(X6)) | ~ssItem(X5)) & strictorderedP(sK55) & app(sK55,X4) = sK56 & ssList(X4)) & nil != sK53 & sK53 = sK55 & sK54 = sK56 & nil = sK54 & ssList(sK56))),
  introduced(choice_axiom,[])).
fof(f347,plain,(
  ? [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (~lt(X7,X5) | app(X8,cons(X7,nil)) != sK55 | ~ssList(X8)) | ~ssItem(X7)) | app(cons(X5,nil),X6) != X4 | ~ssList(X6)) | ~ssItem(X5)) & strictorderedP(sK55) & app(sK55,X4) = sK56 & ssList(X4)) => (! [X5] : (! [X6] : (! [X7] : (! [X8] : (~lt(X7,X5) | app(X8,cons(X7,nil)) != sK55 | ~ssList(X8)) | ~ssItem(X7)) | app(cons(X5,nil),X6) != sK57 | ~ssList(X6)) | ~ssItem(X5)) & strictorderedP(sK55) & sK56 = app(sK55,sK57) & ssList(sK57))),
  introduced(choice_axiom,[])).
fof(f222,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : ((nil != X2 | nil = X3) & ? [X4] : (! [X5] : (! [X6] : (! [X7] : (! [X8] : (~lt(X7,X5) | app(X8,cons(X7,nil)) != X2 | ~ssList(X8)) | ~ssItem(X7)) | app(cons(X5,nil),X6) != X4 | ~ssList(X6)) | ~ssItem(X5)) & strictorderedP(X2) & app(X2,X4) = X3 & ssList(X4)) & nil != X0 & X0 = X2 & X1 = X3 & nil = X1 & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0))),
  inference(flattening,[],[f221])).
fof(f221,plain,(
  ? [X0] : (? [X1] : (? [X2] : (? [X3] : (((nil != X2 | nil = X3) & ? [X4] : ((! [X5] : (! [X6] : (! [X7] : (! [X8] : (~lt(X7,X5) | app(X8,cons(X7,nil)) != X2 | ~ssList(X8)) | ~ssItem(X7)) | app(cons(X5,nil),X6) != X4 | ~ssList(X6)) | ~ssItem(X5)) & strictorderedP(X2) & app(X2,X4) = X3) & ssList(X4)) & nil != X0 & X0 = X2 & X1 = X3 & nil = X1) & ssList(X3)) & ssList(X2)) & ssList(X1)) & ssList(X0))),
  inference(ennf_transformation,[],[f97])).
fof(f97,negated_conjecture,(
  ~! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => ((nil = X2 & nil != X3) | ! [X4] : (ssList(X4) => (? [X5] : (? [X6] : (? [X7] : (? [X8] : (lt(X7,X5) & app(X8,cons(X7,nil)) = X2 & ssList(X8)) & ssItem(X7)) & app(cons(X5,nil),X6) = X4 & ssList(X6)) & ssItem(X5)) | ~strictorderedP(X2) | app(X2,X4) != X3)) | nil = X0 | X0 != X2 | X1 != X3 | nil != X1)))))),
  inference(negated_conjecture,[],[f96])).
fof(f96,conjecture,(
  ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => ! [X2] : (ssList(X2) => ! [X3] : (ssList(X3) => ((nil = X2 & nil != X3) | ! [X4] : (ssList(X4) => (? [X5] : (? [X6] : (? [X7] : (? [X8] : (lt(X7,X5) & app(X8,cons(X7,nil)) = X2 & ssList(X8)) & ssItem(X7)) & app(cons(X5,nil),X6) = X4 & ssList(X6)) & ssItem(X5)) | ~strictorderedP(X2) | app(X2,X4) != X3)) | nil = X0 | X0 != X2 | X1 != X3 | nil != X1)))))),
  file('./p/SWC043+1.tptp',unknown)).
fof(f550,plain,(
  ssList(sK55)),
  inference(cnf_transformation,[],[f348])).
fof(f1343,plain,(
  ~ssList(sK53) | spl58_20),
  inference(subsumption_resolution,[],[f1342,f556])).
fof(f556,plain,(
  ssList(sK57)),
  inference(cnf_transformation,[],[f348])).
fof(f1342,plain,(
  ~ssList(sK57) | ~ssList(sK53) | spl58_20),
  inference(subsumption_resolution,[],[f1341,f1110])).
fof(f1110,plain,(
  nil != sK57 | spl58_20),
  inference(avatar_component_clause,[],[f1109])).
fof(f1109,plain,(
  spl58_20 <=> nil = sK57),
  introduced(avatar_definition,[new_symbols(naming,[spl58_20])])).
fof(f1341,plain,(
  nil = sK57 | ~ssList(sK57) | ~ssList(sK53)),
  inference(trivial_inequality_removal,[],[f1340])).
fof(f1340,plain,(
  nil != nil | nil = sK57 | ~ssList(sK57) | ~ssList(sK53)),
  inference(superposition,[],[f531,f608])).
fof(f608,plain,(
  nil = app(sK53,sK57)),
  inference(backward_demodulation,[],[f552,f607])).
fof(f607,plain,(
  sK54 = app(sK53,sK57)),
  inference(backward_demodulation,[],[f553,f604])).
fof(f604,plain,(
  sK56 = app(sK53,sK57)),
  inference(backward_demodulation,[],[f554,f557])).
fof(f557,plain,(
  sK56 = app(sK55,sK57)),
  inference(cnf_transformation,[],[f348])).
fof(f553,plain,(
  sK54 = sK56),
  inference(cnf_transformation,[],[f348])).
fof(f552,plain,(
  nil = sK54),
  inference(cnf_transformation,[],[f348])).
fof(f531,plain,(
  ( ! [X0,X1] : (nil != app(X0,X1) | nil = X1 | ~ssList(X1) | ~ssList(X0)) )),
  inference(cnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0] : (! [X1] : (((nil = app(X0,X1) | nil != X0 | nil != X1) & ((nil = X0 & nil = X1) | nil != app(X0,X1))) | ~ssList(X1)) | ~ssList(X0))),
  inference(flattening,[],[f339])).
fof(f339,plain,(
  ! [X0] : (! [X1] : (((nil = app(X0,X1) | (nil != X0 | nil != X1)) & ((nil = X0 & nil = X1) | nil != app(X0,X1))) | ~ssList(X1)) | ~ssList(X0))),
  inference(nnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0] : (! [X1] : ((nil = app(X0,X1) <=> (nil = X0 & nil = X1)) | ~ssList(X1)) | ~ssList(X0))),
  inference(ennf_transformation,[],[f83])).
fof(f83,axiom,(
  ! [X0] : (ssList(X0) => ! [X1] : (ssList(X1) => (nil = app(X0,X1) <=> (nil = X0 & nil = X1))))),
  file('./p/SWC043+1.tptp',unknown)).
fof(f1169,plain,(
  ~spl58_20),
  inference(avatar_contradiction_clause,[],[f1168])).
fof(f1168,plain,(
  $false | ~spl58_20),
  inference(subsumption_resolution,[],[f1167,f555])).
fof(f555,plain,(
  nil != sK53),
  inference(cnf_transformation,[],[f348])).
fof(f1167,plain,(
  nil = sK53 | ~spl58_20),
  inference(backward_demodulation,[],[f1151,f862])).
fof(f862,plain,(
  sK53 = app(sK53,nil)),
  inference(resolution,[],[f534,f611])).
fof(f534,plain,(
  ( ! [X0] : (~ssList(X0) | app(X0,nil) = X0) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : (app(X0,nil) = X0 | ~ssList(X0))),
  inference(ennf_transformation,[],[f84])).
fof(f84,axiom,(
  ! [X0] : (ssList(X0) => app(X0,nil) = X0)),
  file('./p/SWC043+1.tptp',unknown)).
fof(f1151,plain,(
  nil = app(sK53,nil) | ~spl58_20),
  inference(backward_demodulation,[],[f1111,f608])).
fof(f1111,plain,(
  nil = sK57 | ~spl58_20),
  inference(avatar_component_clause,[],[f1109])).
% SZS output end Proof for SWC043+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 6012
% Time elapsed: 0.012 s
% ------------------------------
% ------------------------------
