% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SWC115-1
% SZS output start Proof for SWC115-1
fof(f449668,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f253,f258,f259,f264,f274,f275,f285,f398,f507,f17685,f17726,f449667])).
fof(f449667,plain,(
  ~spl0_5 | ~spl0_6 | spl0_10 | ~spl0_153),
  inference(avatar_contradiction_clause,[],[f449666])).
fof(f449666,plain,(
  $false | (~spl0_5 | ~spl0_6 | spl0_10 | ~spl0_153)),
  inference(subsumption_resolution,[],[f449665,f8])).
fof(f8,axiom,(
  ssList(nil)),
  file('./p/SWC115-1.tptp',unknown)).
fof(f449665,plain,(
  ~ssList(nil) | (~spl0_5 | ~spl0_6 | spl0_10 | ~spl0_153)),
  inference(subsumption_resolution,[],[f449664,f284])).
fof(f284,plain,(
  ~segmentP(sk4,sk3) | spl0_10),
  inference(avatar_component_clause,[],[f282])).
fof(f282,plain,(
  spl0_10 <=> segmentP(sk4,sk3)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f449664,plain,(
  segmentP(sk4,sk3) | ~ssList(nil) | (~spl0_5 | ~spl0_6 | ~spl0_153)),
  inference(superposition,[],[f449654,f324])).
fof(f324,plain,(
  sk4 = app(sk4,nil)),
  inference(resolution,[],[f73,f203])).
fof(f203,plain,(
  ssList(sk4)),
  inference(definition_unfolding,[],[f187,f190])).
fof(f190,axiom,(
  sk2 = sk4),
  file('./p/SWC115-1.tptp',unknown)).
fof(f187,axiom,(
  ssList(sk2)),
  file('./p/SWC115-1.tptp',unknown)).
fof(f73,axiom,(
  ( ! [X0] : (~ssList(X0) | app(X0,nil) = X0) )),
  file('./p/SWC115-1.tptp',unknown)).
fof(f449654,plain,(
  ( ! [X0] : (segmentP(app(sk4,X0),sk3) | ~ssList(X0)) ) | (~spl0_5 | ~spl0_6 | ~spl0_153)),
  inference(subsumption_resolution,[],[f449650,f202])).
fof(f202,plain,(
  ssList(sk3)),
  inference(definition_unfolding,[],[f186,f191])).
fof(f191,axiom,(
  sk1 = sk3),
  file('./p/SWC115-1.tptp',unknown)).
fof(f186,axiom,(
  ssList(sk1)),
  file('./p/SWC115-1.tptp',unknown)).
fof(f449650,plain,(
  ( ! [X0] : (segmentP(app(sk4,X0),sk3) | ~ssList(X0) | ~ssList(sk3)) ) | (~spl0_5 | ~spl0_6 | ~spl0_153)),
  inference(resolution,[],[f7638,f252668])).
fof(f252668,plain,(
  ( ! [X0,X1] : (segmentP(cons(sk5,skaf43(X0,X1)),sk3)) ) | (~spl0_5 | ~spl0_6)),
  inference(subsumption_resolution,[],[f109952,f109955])).
fof(f109955,plain,(
  ( ! [X4,X5] : (ssList(cons(sk5,skaf43(X4,X5)))) ) | (~spl0_5 | ~spl0_6)),
  inference(subsumption_resolution,[],[f109954,f52])).
fof(f52,axiom,(
  ( ! [X0,X1] : (ssList(skaf43(X0,X1))) )),
  file('./p/SWC115-1.tptp',unknown)).
fof(f109954,plain,(
  ( ! [X4,X5] : (ssList(cons(sk5,skaf43(X4,X5))) | ~ssList(skaf43(X4,X5))) ) | (~spl0_5 | ~spl0_6)),
  inference(subsumption_resolution,[],[f109914,f202])).
fof(f109914,plain,(
  ( ! [X4,X5] : (ssList(cons(sk5,skaf43(X4,X5))) | ~ssList(sk3) | ~ssList(skaf43(X4,X5))) ) | (~spl0_5 | ~spl0_6)),
  inference(superposition,[],[f85,f4717])).
fof(f4717,plain,(
  ( ! [X30,X29] : (cons(sk5,skaf43(X29,X30)) = app(sk3,skaf43(X29,X30))) ) | (~spl0_5 | ~spl0_6)),
  inference(resolution,[],[f874,f52])).
fof(f874,plain,(
  ( ! [X36] : (~ssList(X36) | cons(sk5,X36) = app(sk3,X36)) ) | (~spl0_5 | ~spl0_6)),
  inference(forward_demodulation,[],[f743,f263])).
fof(f263,plain,(
  sk3 = cons(sk5,nil) | ~spl0_6),
  inference(avatar_component_clause,[],[f261])).
fof(f261,plain,(
  spl0_6 <=> sk3 = cons(sk5,nil)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).
fof(f743,plain,(
  ( ! [X36] : (~ssList(X36) | cons(sk5,X36) = app(cons(sk5,nil),X36)) ) | ~spl0_5),
  inference(resolution,[],[f120,f257])).
fof(f257,plain,(
  ssItem(sk5) | ~spl0_5),
  inference(avatar_component_clause,[],[f255])).
fof(f255,plain,(
  spl0_5 <=> ssItem(sk5)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f120,axiom,(
  ( ! [X0,X1] : (~ssItem(X0) | ~ssList(X1) | cons(X0,X1) = app(cons(X0,nil),X1)) )),
  file('./p/SWC115-1.tptp',unknown)).
fof(f85,axiom,(
  ( ! [X0,X1] : (ssList(app(X1,X0)) | ~ssList(X1) | ~ssList(X0)) )),
  file('./p/SWC115-1.tptp',unknown)).
fof(f109952,plain,(
  ( ! [X0,X1] : (~ssList(cons(sk5,skaf43(X0,X1))) | segmentP(cons(sk5,skaf43(X0,X1)),sk3)) ) | (~spl0_5 | ~spl0_6)),
  inference(forward_demodulation,[],[f109951,f4717])).
fof(f109951,plain,(
  ( ! [X0,X1] : (segmentP(cons(sk5,skaf43(X0,X1)),sk3) | ~ssList(app(sk3,skaf43(X0,X1)))) ) | (~spl0_5 | ~spl0_6)),
  inference(subsumption_resolution,[],[f109912,f52])).
fof(f109912,plain,(
  ( ! [X0,X1] : (segmentP(cons(sk5,skaf43(X0,X1)),sk3) | ~ssList(app(sk3,skaf43(X0,X1))) | ~ssList(skaf43(X0,X1))) ) | (~spl0_5 | ~spl0_6)),
  inference(superposition,[],[f1963,f4717])).
fof(f1963,plain,(
  ( ! [X3] : (segmentP(app(sk3,X3),sk3) | ~ssList(app(sk3,X3)) | ~ssList(X3)) )),
  inference(forward_demodulation,[],[f1962,f399])).
fof(f399,plain,(
  sk3 = app(nil,sk3)),
  inference(resolution,[],[f202,f74])).
fof(f74,axiom,(
  ( ! [X0] : (~ssList(X0) | app(nil,X0) = X0) )),
  file('./p/SWC115-1.tptp',unknown)).
fof(f1962,plain,(
  ( ! [X3] : (segmentP(app(sk3,X3),sk3) | ~ssList(app(app(nil,sk3),X3)) | ~ssList(X3)) )),
  inference(subsumption_resolution,[],[f1961,f8])).
fof(f1961,plain,(
  ( ! [X3] : (segmentP(app(sk3,X3),sk3) | ~ssList(app(app(nil,sk3),X3)) | ~ssList(nil) | ~ssList(X3)) )),
  inference(subsumption_resolution,[],[f1927,f202])).
fof(f1927,plain,(
  ( ! [X3] : (segmentP(app(sk3,X3),sk3) | ~ssList(app(app(nil,sk3),X3)) | ~ssList(sk3) | ~ssList(nil) | ~ssList(X3)) )),
  inference(superposition,[],[f220,f399])).
fof(f220,plain,(
  ( ! [X2,X0,X1] : (segmentP(app(app(X0,X1),X2),X1) | ~ssList(app(app(X0,X1),X2)) | ~ssList(X1) | ~ssList(X0) | ~ssList(X2)) )),
  inference(equality_resolution,[],[f173])).
fof(f173,axiom,(
  ( ! [X2,X0,X3,X1] : (segmentP(X3,X1) | ~ssList(X3) | ~ssList(X1) | ~ssList(X0) | ~ssList(X2) | app(app(X0,X1),X2) != X3) )),
  file('./p/SWC115-1.tptp',unknown)).
fof(f7638,plain,(
  ( ! [X21,X20] : (~segmentP(cons(sk5,skaf43(sk5,sk4)),X21) | segmentP(app(sk4,X20),X21) | ~ssList(X20) | ~ssList(X21)) ) | ~spl0_153),
  inference(avatar_component_clause,[],[f7637])).
fof(f7637,plain,(
  spl0_153 <=> ! [X20,X21] : (segmentP(app(sk4,X20),X21) | ~segmentP(cons(sk5,skaf43(sk5,sk4)),X21) | ~ssList(X20) | ~ssList(X21))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_153])])).
fof(f17726,plain,(
  ~spl0_145 | spl0_153 | ~spl0_5 | ~spl0_7),
  inference(avatar_split_clause,[],[f17725,f266,f255,f7637,f7591])).
fof(f7591,plain,(
  spl0_145 <=> ssList(cons(sk5,skaf43(sk5,sk4)))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_145])])).
fof(f266,plain,(
  spl0_7 <=> memberP(sk4,sk5)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).
fof(f17725,plain,(
  ( ! [X21,X20] : (segmentP(app(sk4,X20),X21) | ~ssList(cons(sk5,skaf43(sk5,sk4))) | ~ssList(X21) | ~ssList(X20) | ~segmentP(cons(sk5,skaf43(sk5,sk4)),X21)) ) | (~spl0_5 | ~spl0_7)),
  inference(subsumption_resolution,[],[f10809,f53])).
fof(f53,axiom,(
  ( ! [X0,X1] : (ssList(skaf42(X0,X1))) )),
  file('./p/SWC115-1.tptp',unknown)).
fof(f10809,plain,(
  ( ! [X21,X20] : (segmentP(app(sk4,X20),X21) | ~ssList(cons(sk5,skaf43(sk5,sk4))) | ~ssList(X21) | ~ssList(skaf42(sk4,sk5)) | ~ssList(X20) | ~segmentP(cons(sk5,skaf43(sk5,sk4)),X21)) ) | (~spl0_5 | ~spl0_7)),
  inference(superposition,[],[f172,f1722])).
fof(f1722,plain,(
  sk4 = app(skaf42(sk4,sk5),cons(sk5,skaf43(sk5,sk4))) | (~spl0_5 | ~spl0_7)),
  inference(subsumption_resolution,[],[f1721,f257])).
fof(f1721,plain,(
  ~ssItem(sk5) | sk4 = app(skaf42(sk4,sk5),cons(sk5,skaf43(sk5,sk4))) | ~spl0_7),
  inference(subsumption_resolution,[],[f1708,f203])).
fof(f1708,plain,(
  ~ssList(sk4) | ~ssItem(sk5) | sk4 = app(skaf42(sk4,sk5),cons(sk5,skaf43(sk5,sk4))) | ~spl0_7),
  inference(resolution,[],[f169,f268])).
fof(f268,plain,(
  memberP(sk4,sk5) | ~spl0_7),
  inference(avatar_component_clause,[],[f266])).
fof(f169,axiom,(
  ( ! [X0,X1] : (~memberP(X0,X1) | ~ssList(X0) | ~ssItem(X1) | app(skaf42(X0,X1),cons(X1,skaf43(X1,X0))) = X0) )),
  file('./p/SWC115-1.tptp',unknown)).
fof(f172,axiom,(
  ( ! [X2,X0,X3,X1] : (segmentP(app(app(X3,X0),X2),X1) | ~ssList(X0) | ~ssList(X1) | ~ssList(X3) | ~ssList(X2) | ~segmentP(X0,X1)) )),
  file('./p/SWC115-1.tptp',unknown)).
fof(f17685,plain,(
  ~spl0_5 | spl0_145),
  inference(avatar_contradiction_clause,[],[f17684])).
fof(f17684,plain,(
  $false | (~spl0_5 | spl0_145)),
  inference(subsumption_resolution,[],[f17683,f257])).
fof(f17683,plain,(
  ~ssItem(sk5) | spl0_145),
  inference(subsumption_resolution,[],[f17682,f52])).
fof(f17682,plain,(
  ~ssList(skaf43(sk5,sk4)) | ~ssItem(sk5) | spl0_145),
  inference(resolution,[],[f7593,f86])).
fof(f86,axiom,(
  ( ! [X0,X1] : (ssList(cons(X0,X1)) | ~ssList(X1) | ~ssItem(X0)) )),
  file('./p/SWC115-1.tptp',unknown)).
fof(f7593,plain,(
  ~ssList(cons(sk5,skaf43(sk5,sk4))) | spl0_145),
  inference(avatar_component_clause,[],[f7591])).
fof(f507,plain,(
  spl0_4 | spl0_9),
  inference(avatar_contradiction_clause,[],[f506])).
fof(f506,plain,(
  $false | (spl0_4 | spl0_9)),
  inference(subsumption_resolution,[],[f505,f8])).
fof(f505,plain,(
  ~ssList(nil) | (spl0_4 | spl0_9)),
  inference(subsumption_resolution,[],[f504,f202])).
fof(f504,plain,(
  ~ssList(sk3) | ~ssList(nil) | (spl0_4 | spl0_9)),
  inference(subsumption_resolution,[],[f501,f252])).
fof(f252,plain,(
  nil != sk3 | spl0_4),
  inference(avatar_component_clause,[],[f250])).
fof(f250,plain,(
  spl0_4 <=> nil = sk3),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f501,plain,(
  nil = sk3 | ~ssList(sk3) | ~ssList(nil) | spl0_9),
  inference(resolution,[],[f100,f280])).
fof(f280,plain,(
  ~neq(sk3,nil) | spl0_9),
  inference(avatar_component_clause,[],[f278])).
fof(f278,plain,(
  spl0_9 <=> neq(sk3,nil)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_9])])).
fof(f100,axiom,(
  ( ! [X0,X1] : (neq(X1,X0) | X0 = X1 | ~ssList(X1) | ~ssList(X0)) )),
  file('./p/SWC115-1.tptp',unknown)).
fof(f398,plain,(
  ~spl0_4 | ~spl0_5 | ~spl0_6),
  inference(avatar_contradiction_clause,[],[f397])).
fof(f397,plain,(
  $false | (~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(subsumption_resolution,[],[f396,f8])).
fof(f396,plain,(
  ~ssList(nil) | (~spl0_4 | ~spl0_5 | ~spl0_6)),
  inference(subsumption_resolution,[],[f395,f257])).
fof(f395,plain,(
  ~ssItem(sk5) | ~ssList(nil) | (~spl0_4 | ~spl0_6)),
  inference(trivial_inequality_removal,[],[f394])).
fof(f394,plain,(
  nil != nil | ~ssItem(sk5) | ~ssList(nil) | (~spl0_4 | ~spl0_6)),
  inference(superposition,[],[f98,f286])).
fof(f286,plain,(
  nil = cons(sk5,nil) | (~spl0_4 | ~spl0_6)),
  inference(backward_demodulation,[],[f251,f263])).
fof(f251,plain,(
  nil = sk3 | ~spl0_4),
  inference(avatar_component_clause,[],[f250])).
fof(f98,axiom,(
  ( ! [X0,X1] : (nil != cons(X0,X1) | ~ssItem(X0) | ~ssList(X1)) )),
  file('./p/SWC115-1.tptp',unknown)).
fof(f285,plain,(
  ~spl0_9 | ~spl0_10),
  inference(avatar_split_clause,[],[f205,f282,f278])).
fof(f205,plain,(
  ~segmentP(sk4,sk3) | ~neq(sk3,nil)),
  inference(definition_unfolding,[],[f201,f190,f191,f191])).
fof(f201,axiom,(
  ~segmentP(sk2,sk1) | ~neq(sk1,nil)),
  file('./p/SWC115-1.tptp',unknown)).
fof(f275,plain,(
  spl0_7 | spl0_4),
  inference(avatar_split_clause,[],[f199,f250,f266])).
fof(f199,axiom,(
  nil = sk3 | memberP(sk4,sk5)),
  file('./p/SWC115-1.tptp',unknown)).
fof(f274,plain,(
  spl0_6 | spl0_4),
  inference(avatar_split_clause,[],[f198,f250,f261])).
fof(f198,axiom,(
  nil = sk3 | sk3 = cons(sk5,nil)),
  file('./p/SWC115-1.tptp',unknown)).
fof(f264,plain,(
  spl0_6 | spl0_3),
  inference(avatar_split_clause,[],[f195,f246,f261])).
fof(f246,plain,(
  spl0_3 <=> nil = sk4),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f195,axiom,(
  nil = sk4 | sk3 = cons(sk5,nil)),
  file('./p/SWC115-1.tptp',unknown)).
fof(f259,plain,(
  spl0_5 | spl0_4),
  inference(avatar_split_clause,[],[f194,f250,f255])).
fof(f194,axiom,(
  nil = sk3 | ssItem(sk5)),
  file('./p/SWC115-1.tptp',unknown)).
fof(f258,plain,(
  spl0_5 | spl0_3),
  inference(avatar_split_clause,[],[f193,f246,f255])).
fof(f193,axiom,(
  nil = sk4 | ssItem(sk5)),
  file('./p/SWC115-1.tptp',unknown)).
fof(f253,plain,(
  ~spl0_3 | ~spl0_4),
  inference(avatar_split_clause,[],[f204,f250,f246])).
fof(f204,plain,(
  nil != sk3 | nil != sk4),
  inference(definition_unfolding,[],[f192,f191,f190])).
fof(f192,axiom,(
  nil != sk1 | nil != sk2),
  file('./p/SWC115-1.tptp',unknown)).
% SZS output end Proof for SWC115-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 264046
% Time elapsed: 20.717 s
% ------------------------------
% ------------------------------
