% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SWC153-1
% SZS output start Proof for SWC153-1
fof(f5979,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f255,f260,f269,f5497,f5705,f5740,f5978])).
fof(f5978,plain,(
  spl0_4),
  inference(avatar_contradiction_clause,[],[f5977])).
fof(f5977,plain,(
  $false | spl0_4),
  inference(subsumption_resolution,[],[f5976,f195])).
fof(f195,axiom,(
  ssItem(sk5)),
  file('./p/SWC153-1.tptp',unknown)).
fof(f5976,plain,(
  ~ssItem(sk5) | spl0_4),
  inference(subsumption_resolution,[],[f5975,f196])).
fof(f196,axiom,(
  ssItem(sk6)),
  file('./p/SWC153-1.tptp',unknown)).
fof(f5975,plain,(
  ~ssItem(sk6) | ~ssItem(sk5) | spl0_4),
  inference(subsumption_resolution,[],[f5974,f197])).
fof(f197,axiom,(
  ssList(sk7)),
  file('./p/SWC153-1.tptp',unknown)).
fof(f5974,plain,(
  ~ssList(sk7) | ~ssItem(sk6) | ~ssItem(sk5) | spl0_4),
  inference(subsumption_resolution,[],[f5973,f198])).
fof(f198,axiom,(
  ssList(sk8)),
  file('./p/SWC153-1.tptp',unknown)).
fof(f5973,plain,(
  ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk6) | ~ssItem(sk5) | spl0_4),
  inference(subsumption_resolution,[],[f5972,f199])).
fof(f199,axiom,(
  ssList(sk9)),
  file('./p/SWC153-1.tptp',unknown)).
fof(f5972,plain,(
  ~ssList(sk9) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk6) | ~ssItem(sk5) | spl0_4),
  inference(subsumption_resolution,[],[f5971,f254])).
fof(f254,plain,(
  ~leq(sk5,sk6) | spl0_4),
  inference(avatar_component_clause,[],[f252])).
fof(f252,plain,(
  spl0_4 <=> leq(sk5,sk6)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f5971,plain,(
  leq(sk5,sk6) | ~ssList(sk9) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk6) | ~ssItem(sk5)),
  inference(subsumption_resolution,[],[f5970,f201])).
fof(f201,axiom,(
  leq(sk6,sk5)),
  file('./p/SWC153-1.tptp',unknown)).
fof(f5970,plain,(
  ~leq(sk6,sk5) | leq(sk5,sk6) | ~ssList(sk9) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk6) | ~ssItem(sk5)),
  inference(trivial_inequality_removal,[],[f5969])).
fof(f5969,plain,(
  sk3 != sk3 | ~leq(sk6,sk5) | leq(sk5,sk6) | ~ssList(sk9) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk6) | ~ssItem(sk5)),
  inference(superposition,[],[f192,f207])).
fof(f207,plain,(
  sk3 = app(app(app(app(sk7,cons(sk5,nil)),sk8),cons(sk6,nil)),sk9)),
  inference(definition_unfolding,[],[f200,f191])).
fof(f191,axiom,(
  sk1 = sk3),
  file('./p/SWC153-1.tptp',unknown)).
fof(f200,axiom,(
  sk1 = app(app(app(app(sk7,cons(sk5,nil)),sk8),cons(sk6,nil)),sk9)),
  file('./p/SWC153-1.tptp',unknown)).
fof(f192,axiom,(
  ( ! [X6,X10,X8,X7,X9] : (sk3 != app(app(app(app(X8,cons(X6,nil)),X9),cons(X7,nil)),X10) | ~leq(X7,X6) | leq(X6,X7) | ~ssList(X10) | ~ssList(X9) | ~ssList(X8) | ~ssItem(X7) | ~ssItem(X6)) )),
  file('./p/SWC153-1.tptp',unknown)).
fof(f5740,plain,(
  ~spl0_4 | spl0_18),
  inference(avatar_split_clause,[],[f5739,f755,f252])).
fof(f755,plain,(
  spl0_18 <=> sk5 = sk6),
  introduced(avatar_definition,[new_symbols(naming,[spl0_18])])).
fof(f5739,plain,(
  sk5 = sk6 | ~leq(sk5,sk6)),
  inference(subsumption_resolution,[],[f5738,f195])).
fof(f5738,plain,(
  ~ssItem(sk5) | sk5 = sk6 | ~leq(sk5,sk6)),
  inference(subsumption_resolution,[],[f815,f196])).
fof(f815,plain,(
  ~ssItem(sk6) | ~ssItem(sk5) | sk5 = sk6 | ~leq(sk5,sk6)),
  inference(resolution,[],[f130,f201])).
fof(f130,axiom,(
  ( ! [X0,X1] : (~leq(X1,X0) | ~ssItem(X1) | ~ssItem(X0) | X0 = X1 | ~leq(X0,X1)) )),
  file('./p/SWC153-1.tptp',unknown)).
fof(f5705,plain,(
  ~spl0_3 | ~spl0_5 | spl0_6 | ~spl0_18),
  inference(avatar_contradiction_clause,[],[f5704])).
fof(f5704,plain,(
  $false | (~spl0_3 | ~spl0_5 | spl0_6 | ~spl0_18)),
  inference(subsumption_resolution,[],[f5703,f3188])).
fof(f3188,plain,(
  ~leq(sk10,sk5) | (spl0_6 | ~spl0_18)),
  inference(backward_demodulation,[],[f757,f264])).
fof(f264,plain,(
  ~leq(sk10,sk6) | spl0_6),
  inference(avatar_component_clause,[],[f262])).
fof(f262,plain,(
  spl0_6 <=> leq(sk10,sk6)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).
fof(f757,plain,(
  sk5 = sk6 | ~spl0_18),
  inference(avatar_component_clause,[],[f755])).
fof(f5703,plain,(
  leq(sk10,sk5) | (~spl0_3 | ~spl0_5 | ~spl0_18)),
  inference(subsumption_resolution,[],[f5700,f250])).
fof(f250,plain,(
  ssItem(sk10) | ~spl0_3),
  inference(avatar_component_clause,[],[f248])).
fof(f248,plain,(
  spl0_3 <=> ssItem(sk10)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f5700,plain,(
  ~ssItem(sk10) | leq(sk10,sk5) | (~spl0_5 | ~spl0_18)),
  inference(resolution,[],[f3611,f259])).
fof(f259,plain,(
  memberP(sk8,sk10) | ~spl0_5),
  inference(avatar_component_clause,[],[f257])).
fof(f257,plain,(
  spl0_5 <=> memberP(sk8,sk10)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f3611,plain,(
  ( ! [X0] : (~memberP(sk8,X0) | ~ssItem(X0) | leq(X0,sk5)) ) | ~spl0_18),
  inference(subsumption_resolution,[],[f3610,f195])).
fof(f3610,plain,(
  ( ! [X0] : (~memberP(sk8,X0) | ~ssItem(X0) | leq(X0,sk5) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(subsumption_resolution,[],[f3609,f197])).
fof(f3609,plain,(
  ( ! [X0] : (~memberP(sk8,X0) | ~ssItem(X0) | leq(X0,sk5) | ~ssList(sk7) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(subsumption_resolution,[],[f3608,f198])).
fof(f3608,plain,(
  ( ! [X0] : (~memberP(sk8,X0) | ~ssItem(X0) | leq(X0,sk5) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(subsumption_resolution,[],[f3607,f199])).
fof(f3607,plain,(
  ( ! [X0] : (~memberP(sk8,X0) | ~ssItem(X0) | leq(X0,sk5) | ~ssList(sk9) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(subsumption_resolution,[],[f3606,f3185])).
fof(f3185,plain,(
  leq(sk5,sk5) | ~spl0_18),
  inference(backward_demodulation,[],[f757,f201])).
fof(f3606,plain,(
  ( ! [X0] : (~memberP(sk8,X0) | ~ssItem(X0) | ~leq(sk5,sk5) | leq(X0,sk5) | ~ssList(sk9) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(trivial_inequality_removal,[],[f3605])).
fof(f3605,plain,(
  ( ! [X0] : (sk3 != sk3 | ~memberP(sk8,X0) | ~ssItem(X0) | ~leq(sk5,sk5) | leq(X0,sk5) | ~ssList(sk9) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(duplicate_literal_removal,[],[f3604])).
fof(f3604,plain,(
  ( ! [X0] : (sk3 != sk3 | ~memberP(sk8,X0) | ~ssItem(X0) | ~leq(sk5,sk5) | leq(X0,sk5) | ~ssList(sk9) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk5) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(superposition,[],[f194,f3186])).
fof(f3186,plain,(
  sk3 = app(app(app(app(sk7,cons(sk5,nil)),sk8),cons(sk5,nil)),sk9) | ~spl0_18),
  inference(backward_demodulation,[],[f757,f207])).
fof(f194,axiom,(
  ( ! [X6,X10,X8,X7,X11,X9] : (sk3 != app(app(app(app(X8,cons(X6,nil)),X9),cons(X7,nil)),X10) | ~memberP(X9,X11) | ~ssItem(X11) | ~leq(X7,X6) | leq(X11,X7) | ~ssList(X10) | ~ssList(X9) | ~ssList(X8) | ~ssItem(X7) | ~ssItem(X6)) )),
  file('./p/SWC153-1.tptp',unknown)).
fof(f5497,plain,(
  spl0_7 | ~spl0_3 | ~spl0_5 | ~spl0_18),
  inference(avatar_split_clause,[],[f5496,f755,f257,f248,f266])).
fof(f266,plain,(
  spl0_7 <=> leq(sk5,sk10)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_7])])).
fof(f5496,plain,(
  leq(sk5,sk10) | (~spl0_3 | ~spl0_5 | ~spl0_18)),
  inference(subsumption_resolution,[],[f5493,f250])).
fof(f5493,plain,(
  ~ssItem(sk10) | leq(sk5,sk10) | (~spl0_5 | ~spl0_18)),
  inference(resolution,[],[f3602,f259])).
fof(f3602,plain,(
  ( ! [X0] : (~memberP(sk8,X0) | ~ssItem(X0) | leq(sk5,X0)) ) | ~spl0_18),
  inference(subsumption_resolution,[],[f3601,f195])).
fof(f3601,plain,(
  ( ! [X0] : (~memberP(sk8,X0) | ~ssItem(X0) | leq(sk5,X0) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(subsumption_resolution,[],[f3600,f197])).
fof(f3600,plain,(
  ( ! [X0] : (~memberP(sk8,X0) | ~ssItem(X0) | leq(sk5,X0) | ~ssList(sk7) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(subsumption_resolution,[],[f3599,f198])).
fof(f3599,plain,(
  ( ! [X0] : (~memberP(sk8,X0) | ~ssItem(X0) | leq(sk5,X0) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(subsumption_resolution,[],[f3598,f199])).
fof(f3598,plain,(
  ( ! [X0] : (~memberP(sk8,X0) | ~ssItem(X0) | leq(sk5,X0) | ~ssList(sk9) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(subsumption_resolution,[],[f3597,f3185])).
fof(f3597,plain,(
  ( ! [X0] : (~memberP(sk8,X0) | ~ssItem(X0) | ~leq(sk5,sk5) | leq(sk5,X0) | ~ssList(sk9) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(trivial_inequality_removal,[],[f3596])).
fof(f3596,plain,(
  ( ! [X0] : (sk3 != sk3 | ~memberP(sk8,X0) | ~ssItem(X0) | ~leq(sk5,sk5) | leq(sk5,X0) | ~ssList(sk9) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(duplicate_literal_removal,[],[f3595])).
fof(f3595,plain,(
  ( ! [X0] : (sk3 != sk3 | ~memberP(sk8,X0) | ~ssItem(X0) | ~leq(sk5,sk5) | leq(sk5,X0) | ~ssList(sk9) | ~ssList(sk8) | ~ssList(sk7) | ~ssItem(sk5) | ~ssItem(sk5)) ) | ~spl0_18),
  inference(superposition,[],[f193,f3186])).
fof(f193,axiom,(
  ( ! [X6,X10,X8,X7,X11,X9] : (sk3 != app(app(app(app(X8,cons(X6,nil)),X9),cons(X7,nil)),X10) | ~memberP(X9,X11) | ~ssItem(X11) | ~leq(X7,X6) | leq(X6,X11) | ~ssList(X10) | ~ssList(X9) | ~ssList(X8) | ~ssItem(X7) | ~ssItem(X6)) )),
  file('./p/SWC153-1.tptp',unknown)).
fof(f269,plain,(
  ~spl0_6 | ~spl0_7 | ~spl0_4),
  inference(avatar_split_clause,[],[f204,f252,f266,f262])).
fof(f204,axiom,(
  ~leq(sk5,sk6) | ~leq(sk5,sk10) | ~leq(sk10,sk6)),
  file('./p/SWC153-1.tptp',unknown)).
fof(f260,plain,(
  spl0_5 | ~spl0_4),
  inference(avatar_split_clause,[],[f203,f252,f257])).
fof(f203,axiom,(
  ~leq(sk5,sk6) | memberP(sk8,sk10)),
  file('./p/SWC153-1.tptp',unknown)).
fof(f255,plain,(
  spl0_3 | ~spl0_4),
  inference(avatar_split_clause,[],[f202,f252,f248])).
fof(f202,axiom,(
  ~leq(sk5,sk6) | ssItem(sk10)),
  file('./p/SWC153-1.tptp',unknown)).
% SZS output end Proof for SWC153-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 9210
% Time elapsed: 0.050 s
% ------------------------------
% ------------------------------
