% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SWC408-1
% SZS output start Proof for SWC408-1
fof(f3892,plain,(
  $false),
  inference(subsumption_resolution,[],[f3891,f199])).
fof(f199,plain,(
  ~memberP(sk3,sk5)),
  inference(definition_unfolding,[],[f195,f192])).
fof(f192,axiom,(
  sk1 = sk3),
  file('./p/SWC408-1.tptp',unknown)).
fof(f195,axiom,(
  ~memberP(sk1,sk5)),
  file('./p/SWC408-1.tptp',unknown)).
fof(f3891,plain,(
  memberP(sk3,sk5)),
  inference(subsumption_resolution,[],[f3888,f193])).
fof(f193,axiom,(
  ssItem(sk5)),
  file('./p/SWC408-1.tptp',unknown)).
fof(f3888,plain,(
  ~ssItem(sk5) | memberP(sk3,sk5)),
  inference(resolution,[],[f1253,f198])).
fof(f198,plain,(
  memberP(sk4,sk5)),
  inference(definition_unfolding,[],[f194,f191])).
fof(f191,axiom,(
  sk2 = sk4),
  file('./p/SWC408-1.tptp',unknown)).
fof(f194,axiom,(
  memberP(sk2,sk5)),
  file('./p/SWC408-1.tptp',unknown)).
fof(f1253,plain,(
  ( ! [X58] : (~memberP(sk4,X58) | ~ssItem(X58) | memberP(sk3,X58)) )),
  inference(subsumption_resolution,[],[f1244,f197])).
fof(f197,plain,(
  ssList(sk4)),
  inference(definition_unfolding,[],[f187,f191])).
fof(f187,axiom,(
  ssList(sk2)),
  file('./p/SWC408-1.tptp',unknown)).
fof(f1244,plain,(
  ( ! [X58] : (memberP(sk3,X58) | ~ssItem(X58) | ~ssList(sk4) | ~memberP(sk4,X58)) )),
  inference(duplicate_literal_removal,[],[f1242])).
fof(f1242,plain,(
  ( ! [X58] : (memberP(sk3,X58) | ~ssItem(X58) | ~ssList(sk4) | ~ssList(sk4) | ~memberP(sk4,X58)) )),
  inference(superposition,[],[f141,f190])).
fof(f190,axiom,(
  sk3 = app(sk4,sk4)),
  file('./p/SWC408-1.tptp',unknown)).
fof(f141,axiom,(
  ( ! [X2,X0,X1] : (memberP(app(X2,X0),X1) | ~ssItem(X1) | ~ssList(X2) | ~ssList(X0) | ~memberP(X0,X1)) )),
  file('./p/SWC408-1.tptp',unknown)).
% SZS output end Proof for SWC408-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 8059
% Time elapsed: 0.032 s
% ------------------------------
% ------------------------------
