% Refutation found. Thanks to Tanya!
% SZS status Theorem for SWV191+1
% SZS output start Proof for SWV191+1
fof(f6787,plain,(
  $false),
  inference(subsumption_resolution,[],[f6775,f640])).
fof(f640,plain,(
  ~leq(n0,tptp_minus_1)),
  inference(superposition,[],[f561,f399])).
fof(f399,plain,(
  n0 = plus(tptp_minus_1,n1)),
  inference(definition_unfolding,[],[f298,f299])).
fof(f299,plain,(
  ( ! [X0] : (succ(X0) = plus(X0,n1)) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,axiom,(
  ! [X0] : succ(X0) = plus(X0,n1)),
  file('./p/SWV191+1.tptp',unknown)).
fof(f298,plain,(
  n0 = succ(tptp_minus_1)),
  inference(cnf_transformation,[],[f28])).
fof(f28,axiom,(
  n0 = succ(tptp_minus_1)),
  file('./p/SWV191+1.tptp',unknown)).
fof(f561,plain,(
  ( ! [X2] : (~leq(plus(X2,n1),X2)) )),
  inference(resolution,[],[f398,f221])).
fof(f221,plain,(
  ( ! [X0] : (~gt(X0,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : ~gt(X0,X0)),
  file('./p/SWV191+1.tptp',unknown)).
fof(f398,plain,(
  ( ! [X0,X1] : (gt(plus(X1,n1),X0) | ~leq(X0,X1)) )),
  inference(definition_unfolding,[],[f230,f299])).
fof(f230,plain,(
  ( ! [X0,X1] : (gt(succ(X1),X0) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : ((leq(X0,X1) | ~gt(succ(X1),X0)) & (gt(succ(X1),X0) | ~leq(X0,X1)))),
  inference(nnf_transformation,[],[f13])).
fof(f13,axiom,(
  ! [X0,X1] : (leq(X0,X1) <=> gt(succ(X1),X0))),
  file('./p/SWV191+1.tptp',unknown)).
fof(f6775,plain,(
  leq(n0,tptp_minus_1)),
  inference(resolution,[],[f694,f357])).
fof(f357,plain,(
  leq(n0,sK32)),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  (use != a_select3(z_defuse,sK31,sK32) & leq(sK32,tptp_minus_1) & leq(sK31,n2) & leq(n0,sK32) & leq(n0,sK31)) & use = a_select2(xinit_noise_defuse,n5) & use = a_select2(xinit_noise_defuse,n4) & use = a_select2(xinit_noise_defuse,n3) & use = a_select2(xinit_noise_defuse,n2) & use = a_select2(xinit_noise_defuse,n1) & use = a_select2(xinit_noise_defuse,n0) & use = a_select2(xinit_mean_defuse,n5) & use = a_select2(xinit_mean_defuse,n4) & use = a_select2(xinit_mean_defuse,n3) & use = a_select2(xinit_mean_defuse,n2) & use = a_select2(xinit_mean_defuse,n1) & use = a_select2(xinit_mean_defuse,n0) & use = a_select2(xinit_defuse,n5) & use = a_select2(xinit_defuse,n4) & use = a_select2(xinit_defuse,n3) & use = a_select3(u_defuse,n2,n0) & use = a_select3(u_defuse,n1,n0) & use = a_select3(u_defuse,n0,n0) & use = a_select2(sigma_defuse,n5) & use = a_select2(sigma_defuse,n4) & use = a_select2(sigma_defuse,n3) & use = a_select2(sigma_defuse,n2) & use = a_select2(sigma_defuse,n1) & use = a_select2(sigma_defuse,n0) & use = a_select2(rho_defuse,n2) & use = a_select2(rho_defuse,n1) & use = a_select2(rho_defuse,n0)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31,sK32])],[f150,f217])).
fof(f217,plain,(
  ? [X0,X1] : (use != a_select3(z_defuse,X0,X1) & leq(X1,tptp_minus_1) & leq(X0,n2) & leq(n0,X1) & leq(n0,X0)) => (use != a_select3(z_defuse,sK31,sK32) & leq(sK32,tptp_minus_1) & leq(sK31,n2) & leq(n0,sK32) & leq(n0,sK31))),
  introduced(choice_axiom,[])).
fof(f150,plain,(
  ? [X0,X1] : (use != a_select3(z_defuse,X0,X1) & leq(X1,tptp_minus_1) & leq(X0,n2) & leq(n0,X1) & leq(n0,X0)) & use = a_select2(xinit_noise_defuse,n5) & use = a_select2(xinit_noise_defuse,n4) & use = a_select2(xinit_noise_defuse,n3) & use = a_select2(xinit_noise_defuse,n2) & use = a_select2(xinit_noise_defuse,n1) & use = a_select2(xinit_noise_defuse,n0) & use = a_select2(xinit_mean_defuse,n5) & use = a_select2(xinit_mean_defuse,n4) & use = a_select2(xinit_mean_defuse,n3) & use = a_select2(xinit_mean_defuse,n2) & use = a_select2(xinit_mean_defuse,n1) & use = a_select2(xinit_mean_defuse,n0) & use = a_select2(xinit_defuse,n5) & use = a_select2(xinit_defuse,n4) & use = a_select2(xinit_defuse,n3) & use = a_select3(u_defuse,n2,n0) & use = a_select3(u_defuse,n1,n0) & use = a_select3(u_defuse,n0,n0) & use = a_select2(sigma_defuse,n5) & use = a_select2(sigma_defuse,n4) & use = a_select2(sigma_defuse,n3) & use = a_select2(sigma_defuse,n2) & use = a_select2(sigma_defuse,n1) & use = a_select2(sigma_defuse,n0) & use = a_select2(rho_defuse,n2) & use = a_select2(rho_defuse,n1) & use = a_select2(rho_defuse,n0)),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ? [X0,X1] : (use != a_select3(z_defuse,X0,X1) & (leq(X1,tptp_minus_1) & leq(X0,n2) & leq(n0,X1) & leq(n0,X0))) & (use = a_select2(xinit_noise_defuse,n5) & use = a_select2(xinit_noise_defuse,n4) & use = a_select2(xinit_noise_defuse,n3) & use = a_select2(xinit_noise_defuse,n2) & use = a_select2(xinit_noise_defuse,n1) & use = a_select2(xinit_noise_defuse,n0) & use = a_select2(xinit_mean_defuse,n5) & use = a_select2(xinit_mean_defuse,n4) & use = a_select2(xinit_mean_defuse,n3) & use = a_select2(xinit_mean_defuse,n2) & use = a_select2(xinit_mean_defuse,n1) & use = a_select2(xinit_mean_defuse,n0) & use = a_select2(xinit_defuse,n5) & use = a_select2(xinit_defuse,n4) & use = a_select2(xinit_defuse,n3) & use = a_select3(u_defuse,n2,n0) & use = a_select3(u_defuse,n1,n0) & use = a_select3(u_defuse,n0,n0) & use = a_select2(sigma_defuse,n5) & use = a_select2(sigma_defuse,n4) & use = a_select2(sigma_defuse,n3) & use = a_select2(sigma_defuse,n2) & use = a_select2(sigma_defuse,n1) & use = a_select2(sigma_defuse,n0) & use = a_select2(rho_defuse,n2) & use = a_select2(rho_defuse,n1) & use = a_select2(rho_defuse,n0))),
  inference(ennf_transformation,[],[f107])).
fof(f107,plain,(
  ~((use = a_select2(xinit_noise_defuse,n5) & use = a_select2(xinit_noise_defuse,n4) & use = a_select2(xinit_noise_defuse,n3) & use = a_select2(xinit_noise_defuse,n2) & use = a_select2(xinit_noise_defuse,n1) & use = a_select2(xinit_noise_defuse,n0) & use = a_select2(xinit_mean_defuse,n5) & use = a_select2(xinit_mean_defuse,n4) & use = a_select2(xinit_mean_defuse,n3) & use = a_select2(xinit_mean_defuse,n2) & use = a_select2(xinit_mean_defuse,n1) & use = a_select2(xinit_mean_defuse,n0) & use = a_select2(xinit_defuse,n5) & use = a_select2(xinit_defuse,n4) & use = a_select2(xinit_defuse,n3) & use = a_select3(u_defuse,n2,n0) & use = a_select3(u_defuse,n1,n0) & use = a_select3(u_defuse,n0,n0) & use = a_select2(sigma_defuse,n5) & use = a_select2(sigma_defuse,n4) & use = a_select2(sigma_defuse,n3) & use = a_select2(sigma_defuse,n2) & use = a_select2(sigma_defuse,n1) & use = a_select2(sigma_defuse,n0) & use = a_select2(rho_defuse,n2) & use = a_select2(rho_defuse,n1) & use = a_select2(rho_defuse,n0)) => ! [X0,X1] : ((leq(X1,tptp_minus_1) & leq(X0,n2) & leq(n0,X1) & leq(n0,X0)) => use = a_select3(z_defuse,X0,X1)))),
  inference(rectify,[],[f54])).
fof(f54,negated_conjecture,(
  ~((use = a_select2(xinit_noise_defuse,n5) & use = a_select2(xinit_noise_defuse,n4) & use = a_select2(xinit_noise_defuse,n3) & use = a_select2(xinit_noise_defuse,n2) & use = a_select2(xinit_noise_defuse,n1) & use = a_select2(xinit_noise_defuse,n0) & use = a_select2(xinit_mean_defuse,n5) & use = a_select2(xinit_mean_defuse,n4) & use = a_select2(xinit_mean_defuse,n3) & use = a_select2(xinit_mean_defuse,n2) & use = a_select2(xinit_mean_defuse,n1) & use = a_select2(xinit_mean_defuse,n0) & use = a_select2(xinit_defuse,n5) & use = a_select2(xinit_defuse,n4) & use = a_select2(xinit_defuse,n3) & use = a_select3(u_defuse,n2,n0) & use = a_select3(u_defuse,n1,n0) & use = a_select3(u_defuse,n0,n0) & use = a_select2(sigma_defuse,n5) & use = a_select2(sigma_defuse,n4) & use = a_select2(sigma_defuse,n3) & use = a_select2(sigma_defuse,n2) & use = a_select2(sigma_defuse,n1) & use = a_select2(sigma_defuse,n0) & use = a_select2(rho_defuse,n2) & use = a_select2(rho_defuse,n1) & use = a_select2(rho_defuse,n0)) => ! [X13,X17] : ((leq(X17,tptp_minus_1) & leq(X13,n2) & leq(n0,X17) & leq(n0,X13)) => use = a_select3(z_defuse,X13,X17)))),
  inference(negated_conjecture,[],[f53])).
fof(f53,conjecture,(
  (use = a_select2(xinit_noise_defuse,n5) & use = a_select2(xinit_noise_defuse,n4) & use = a_select2(xinit_noise_defuse,n3) & use = a_select2(xinit_noise_defuse,n2) & use = a_select2(xinit_noise_defuse,n1) & use = a_select2(xinit_noise_defuse,n0) & use = a_select2(xinit_mean_defuse,n5) & use = a_select2(xinit_mean_defuse,n4) & use = a_select2(xinit_mean_defuse,n3) & use = a_select2(xinit_mean_defuse,n2) & use = a_select2(xinit_mean_defuse,n1) & use = a_select2(xinit_mean_defuse,n0) & use = a_select2(xinit_defuse,n5) & use = a_select2(xinit_defuse,n4) & use = a_select2(xinit_defuse,n3) & use = a_select3(u_defuse,n2,n0) & use = a_select3(u_defuse,n1,n0) & use = a_select3(u_defuse,n0,n0) & use = a_select2(sigma_defuse,n5) & use = a_select2(sigma_defuse,n4) & use = a_select2(sigma_defuse,n3) & use = a_select2(sigma_defuse,n2) & use = a_select2(sigma_defuse,n1) & use = a_select2(sigma_defuse,n0) & use = a_select2(rho_defuse,n2) & use = a_select2(rho_defuse,n1) & use = a_select2(rho_defuse,n0)) => ! [X13,X17] : ((leq(X17,tptp_minus_1) & leq(X13,n2) & leq(n0,X17) & leq(n0,X13)) => use = a_select3(z_defuse,X13,X17))),
  file('./p/SWV191+1.tptp',unknown)).
fof(f694,plain,(
  ( ! [X88] : (~leq(X88,sK32) | leq(X88,tptp_minus_1)) )),
  inference(resolution,[],[f223,f359])).
fof(f359,plain,(
  leq(sK32,tptp_minus_1)),
  inference(cnf_transformation,[],[f218])).
fof(f223,plain,(
  ( ! [X2,X0,X1] : (~leq(X1,X2) | leq(X0,X2) | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2] : (leq(X0,X2) | ~leq(X1,X2) | ~leq(X0,X1))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2] : (leq(X0,X2) | (~leq(X1,X2) | ~leq(X0,X1)))),
  inference(ennf_transformation,[],[f5])).
fof(f5,axiom,(
  ! [X0,X1,X2] : ((leq(X1,X2) & leq(X0,X1)) => leq(X0,X2))),
  file('./p/SWV191+1.tptp',unknown)).
% SZS output end Proof for SWV191+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 8443
% Time elapsed: 0.111 s
% ------------------------------
% ------------------------------
