% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SWV241-2
% SZS output start Proof for SWV241-2
fof(f128,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f88,f127])).
fof(f127,plain,(
  ~spl0_1),
  inference(avatar_contradiction_clause,[],[f126])).
fof(f126,plain,(
  $false | ~spl0_1),
  inference(subsumption_resolution,[],[f122,f8])).
fof(f8,axiom,(
  ~c_in(v_Y,c_Message_Oanalz(v_H),tc_Message_Omsg)),
  file('./p/SWV241-2.tptp',unknown)).
fof(f122,plain,(
  c_in(v_Y,c_Message_Oanalz(v_H),tc_Message_Omsg) | ~spl0_1),
  inference(resolution,[],[f99,f6])).
fof(f6,axiom,(
  c_in(v_Y,c_Message_Oanalz(c_insert(v_X,v_H,tc_Message_Omsg)),tc_Message_Omsg)),
  file('./p/SWV241-2.tptp',unknown)).
fof(f99,plain,(
  ( ! [X0] : (~c_in(X0,c_Message_Oanalz(c_insert(v_X,v_H,tc_Message_Omsg)),tc_Message_Omsg) | c_in(X0,c_Message_Oanalz(v_H),tc_Message_Omsg)) ) | ~spl0_1),
  inference(resolution,[],[f81,f2])).
fof(f2,axiom,(
  ( ! [X2,X0,X1] : (~c_lessequals(X2,c_Message_Oanalz(X1),tc_set(tc_Message_Omsg)) | c_in(X0,c_Message_Oanalz(X1),tc_Message_Omsg) | ~c_in(X0,c_Message_Oanalz(X2),tc_Message_Omsg)) )),
  file('./p/SWV241-2.tptp',unknown)).
fof(f81,plain,(
  c_lessequals(c_insert(v_X,v_H,tc_Message_Omsg),c_Message_Oanalz(v_H),tc_set(tc_Message_Omsg)) | ~spl0_1),
  inference(avatar_component_clause,[],[f79])).
fof(f79,plain,(
  spl0_1 <=> c_lessequals(c_insert(v_X,v_H,tc_Message_Omsg),c_Message_Oanalz(v_H),tc_set(tc_Message_Omsg))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f88,plain,(
  spl0_1),
  inference(avatar_split_clause,[],[f87,f79])).
fof(f87,plain,(
  c_lessequals(c_insert(v_X,v_H,tc_Message_Omsg),c_Message_Oanalz(v_H),tc_set(tc_Message_Omsg))),
  inference(subsumption_resolution,[],[f76,f7])).
fof(f7,axiom,(
  c_in(v_X,c_Message_Oanalz(v_H),tc_Message_Omsg)),
  file('./p/SWV241-2.tptp',unknown)).
fof(f76,plain,(
  ~c_in(v_X,c_Message_Oanalz(v_H),tc_Message_Omsg) | c_lessequals(c_insert(v_X,v_H,tc_Message_Omsg),c_Message_Oanalz(v_H),tc_set(tc_Message_Omsg))),
  inference(superposition,[],[f5,f73])).
fof(f73,plain,(
  v_X = c_Main_OsubsetI__1(c_insert(v_X,v_H,tc_Message_Omsg),c_Message_Oanalz(v_H),tc_Message_Omsg)),
  inference(subsumption_resolution,[],[f69,f8])).
fof(f69,plain,(
  c_in(v_Y,c_Message_Oanalz(v_H),tc_Message_Omsg) | v_X = c_Main_OsubsetI__1(c_insert(v_X,v_H,tc_Message_Omsg),c_Message_Oanalz(v_H),tc_Message_Omsg)),
  inference(resolution,[],[f34,f6])).
fof(f34,plain,(
  ( ! [X2,X0,X1] : (~c_in(X2,c_Message_Oanalz(c_insert(X0,X1,tc_Message_Omsg)),tc_Message_Omsg) | c_in(X2,c_Message_Oanalz(X1),tc_Message_Omsg) | c_Main_OsubsetI__1(c_insert(X0,X1,tc_Message_Omsg),c_Message_Oanalz(X1),tc_Message_Omsg) = X0) )),
  inference(resolution,[],[f25,f2])).
fof(f25,plain,(
  ( ! [X4,X3] : (c_lessequals(c_insert(X3,X4,tc_Message_Omsg),c_Message_Oanalz(X4),tc_set(tc_Message_Omsg)) | c_Main_OsubsetI__1(c_insert(X3,X4,tc_Message_Omsg),c_Message_Oanalz(X4),tc_Message_Omsg) = X3) )),
  inference(duplicate_literal_removal,[],[f22])).
fof(f22,plain,(
  ( ! [X4,X3] : (c_Main_OsubsetI__1(c_insert(X3,X4,tc_Message_Omsg),c_Message_Oanalz(X4),tc_Message_Omsg) = X3 | c_lessequals(c_insert(X3,X4,tc_Message_Omsg),c_Message_Oanalz(X4),tc_set(tc_Message_Omsg)) | c_lessequals(c_insert(X3,X4,tc_Message_Omsg),c_Message_Oanalz(X4),tc_set(tc_Message_Omsg))) )),
  inference(resolution,[],[f13,f10])).
fof(f10,plain,(
  ( ! [X0,X1] : (~c_in(c_Main_OsubsetI__1(X0,c_Message_Oanalz(X1),tc_Message_Omsg),X1,tc_Message_Omsg) | c_lessequals(X0,c_Message_Oanalz(X1),tc_set(tc_Message_Omsg))) )),
  inference(resolution,[],[f5,f1])).
fof(f1,axiom,(
  ( ! [X0,X1] : (c_in(X0,c_Message_Oanalz(X1),tc_Message_Omsg) | ~c_in(X0,X1,tc_Message_Omsg)) )),
  file('./p/SWV241-2.tptp',unknown)).
fof(f13,plain,(
  ( ! [X2,X0,X3,X1] : (c_in(c_Main_OsubsetI__1(c_insert(X0,X1,X2),X3,X2),X1,X2) | c_Main_OsubsetI__1(c_insert(X0,X1,X2),X3,X2) = X0 | c_lessequals(c_insert(X0,X1,X2),X3,tc_set(X2))) )),
  inference(resolution,[],[f3,f4])).
fof(f4,axiom,(
  ( ! [X6,X7,X5] : (c_in(c_Main_OsubsetI__1(X5,X7,X6),X5,X6) | c_lessequals(X5,X7,tc_set(X6))) )),
  file('./p/SWV241-2.tptp',unknown)).
fof(f3,axiom,(
  ( ! [X6,X4,X5,X3] : (~c_in(X3,c_insert(X4,X5,X6),X6) | c_in(X3,X5,X6) | X3 = X4) )),
  file('./p/SWV241-2.tptp',unknown)).
fof(f5,axiom,(
  ( ! [X6,X7,X5] : (~c_in(c_Main_OsubsetI__1(X5,X7,X6),X7,X6) | c_lessequals(X5,X7,tc_set(X6))) )),
  file('./p/SWV241-2.tptp',unknown)).
% SZS output end Proof for SWV241-2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4989
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
