% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SWV347-2
% SZS output start Proof for SWV347-2
fof(f14,plain,(
  $false),
  inference(subsumption_resolution,[],[f13,f2])).
fof(f2,axiom,(
  ~c_in(v_A,c_Event_Obad,tc_Message_Oagent)),
  file('./p/SWV347-2.tptp',unknown)).
fof(f13,plain,(
  c_in(v_A,c_Event_Obad,tc_Message_Oagent)),
  inference(backward_demodulation,[],[f11,f3])).
fof(f3,axiom,(
  c_in(v_Aa,c_Event_Obad,tc_Message_Oagent)),
  file('./p/SWV347-2.tptp',unknown)).
fof(f11,plain,(
  v_A = v_Aa),
  inference(resolution,[],[f10,f5])).
fof(f5,axiom,(
  c_in(c_Event_Oevent_OGets(c_Message_Oagent_OServer,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_Ba),c_Message_Omsg_OCrypt(c_Public_OshrK(v_Ba),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_Aa),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(v_NAa),c_Message_Omsg_ONonce(v_NB)))))),c_List_Oset(v_evs3,tc_Event_Oevent),tc_Event_Oevent)),
  file('./p/SWV347-2.tptp',unknown)).
fof(f10,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~c_in(c_Event_Oevent_OGets(X1,c_Message_Omsg_OMPair(X2,c_Message_Omsg_OCrypt(c_Public_OshrK(X3),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(X0),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(X4),c_Message_Omsg_ONonce(v_NB)))))),c_List_Oset(v_evs3,tc_Event_Oevent),tc_Event_Oevent) | v_A = X0) )),
  inference(subsumption_resolution,[],[f9,f4])).
fof(f4,axiom,(
  c_in(v_evs3,c_Yahalom_Oyahalom,tc_List_Olist(tc_Event_Oevent))),
  file('./p/SWV347-2.tptp',unknown)).
fof(f9,plain,(
  ( ! [X4,X2,X0,X3,X1] : (v_A = X0 | ~c_in(c_Event_Oevent_OGets(X1,c_Message_Omsg_OMPair(X2,c_Message_Omsg_OCrypt(c_Public_OshrK(X3),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(X0),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(X4),c_Message_Omsg_ONonce(v_NB)))))),c_List_Oset(v_evs3,tc_Event_Oevent),tc_Event_Oevent) | ~c_in(v_evs3,c_Yahalom_Oyahalom,tc_List_Olist(tc_Event_Oevent))) )),
  inference(subsumption_resolution,[],[f8,f7])).
fof(f7,axiom,(
  ~c_in(c_Message_Omsg_ONonce(v_NB),c_Message_Oanalz(c_Event_Oknows(c_Message_Oagent_OSpy,v_evs3)),tc_Message_Omsg)),
  file('./p/SWV347-2.tptp',unknown)).
fof(f8,plain,(
  ( ! [X4,X2,X0,X3,X1] : (c_in(c_Message_Omsg_ONonce(v_NB),c_Message_Oanalz(c_Event_Oknows(c_Message_Oagent_OSpy,v_evs3)),tc_Message_Omsg) | v_A = X0 | ~c_in(c_Event_Oevent_OGets(X1,c_Message_Omsg_OMPair(X2,c_Message_Omsg_OCrypt(c_Public_OshrK(X3),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(X0),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(X4),c_Message_Omsg_ONonce(v_NB)))))),c_List_Oset(v_evs3,tc_Event_Oevent),tc_Event_Oevent) | ~c_in(v_evs3,c_Yahalom_Oyahalom,tc_List_Olist(tc_Event_Oevent))) )),
  inference(resolution,[],[f1,f6])).
fof(f6,axiom,(
  c_in(c_Event_Oevent_OSays(v_B,c_Message_Oagent_OServer,c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_B),c_Message_Omsg_OCrypt(c_Public_OshrK(v_B),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(v_A),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(v_NA),c_Message_Omsg_ONonce(v_NB)))))),c_List_Oset(v_evs3,tc_Event_Oevent),tc_Event_Oevent)),
  file('./p/SWV347-2.tptp',unknown)).
fof(f1,axiom,(
  ( ! [X6,X4,X2,X0,X12,X10,X8,X7,X5,X3,X1,X11,X9] : (~c_in(c_Event_Oevent_OSays(X7,X8,c_Message_Omsg_OMPair(X9,c_Message_Omsg_OCrypt(c_Public_OshrK(X10),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(X11),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(X12),X6))))),c_List_Oset(X0,tc_Event_Oevent),tc_Event_Oevent) | c_in(X6,c_Message_Oanalz(c_Event_Oknows(c_Message_Oagent_OSpy,X0)),tc_Message_Omsg) | X4 = X11 | ~c_in(c_Event_Oevent_OGets(X1,c_Message_Omsg_OMPair(X2,c_Message_Omsg_OCrypt(c_Public_OshrK(X3),c_Message_Omsg_OMPair(c_Message_Omsg_OAgent(X4),c_Message_Omsg_OMPair(c_Message_Omsg_ONonce(X5),X6))))),c_List_Oset(X0,tc_Event_Oevent),tc_Event_Oevent) | ~c_in(X0,c_Yahalom_Oyahalom,tc_List_Olist(tc_Event_Oevent))) )),
  file('./p/SWV347-2.tptp',unknown)).
% SZS output end Proof for SWV347-2
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.0000 s
% ------------------------------
% ------------------------------
