% Refutation found. Thanks to Tanya!
% SZS status Theorem for SWV464+1
% SZS output start Proof for SWV464+1
fof(f1717,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f333,f816,f1004,f1412,f1716])).
fof(f1716,plain,(
  ~spl12_24),
  inference(avatar_contradiction_clause,[],[f1715])).
fof(f1715,plain,(
  $false | ~spl12_24),
  inference(subsumption_resolution,[],[f1714,f654])).
fof(f654,plain,(
  elem(m_Ack(sK10,sK8),queue(host(sK10)))),
  inference(subsumption_resolution,[],[f649,f187])).
fof(f187,plain,(
  ( ! [X2,X0,X1] : (m_Ack(X0,X1) != m_Halt(X2)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : m_Ack(X0,X1) != m_Halt(X2)),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X4,X5,X6] : m_Ack(X4,X5) != m_Halt(X6)),
  file('./p/SWV464+1.tptp',unknown)).
fof(f649,plain,(
  m_Halt(sK5) = m_Ack(sK10,sK8) | elem(m_Ack(sK10,sK8),queue(host(sK10)))),
  inference(resolution,[],[f230,f301])).
fof(f301,plain,(
  elem(m_Ack(sK10,sK8),snoc(queue(host(sK10)),m_Halt(sK5)))),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  (norm = index(status,host(sK7)) & host(sK7) = index(ldr,host(sK7)) & setIn(sK7,alive) & host(sK9) = s(index(pendack,host(sK10))) & index(pendack,host(sK10)) = host(sK8) & elec_2 = index(status,host(sK10)) & elem(m_Ack(sK10,sK8),snoc(queue(host(sK10)),m_Halt(sK5))) & elem(m_Down(sK9),snoc(queue(host(sK10)),m_Halt(sK5))) & leq(nbr_proc,s(index(pendack,host(sK10)))) & setIn(sK10,alive) & host(sK5) != host(sK10) & s(index(pendack,host(sK5))) = host(sK10)) & ~leq(nbr_proc,index(pendack,host(sK5))) & index(pendack,host(sK5)) = host(sK6) & elec_2 = index(status,host(sK5)) & sK4 = index(elid,host(sK5)) & setIn(sK5,alive) & queue(host(sK5)) = cons(m_Ack(sK4,sK6),sK3) & ! [X8,X9,X10,X11] : (norm != index(status,host(X8)) | host(X8) != index(ldr,host(X8)) | ~setIn(X8,alive) | host(X10) != s(index(pendack,host(X11))) | host(X9) != index(pendack,host(X11)) | elec_2 != index(status,host(X11)) | ~leq(nbr_proc,s(index(pendack,host(X11)))) | ~elem(m_Ack(X11,X9),queue(host(X11))) | ~elem(m_Down(X10),queue(host(X11))) | ~setIn(X11,alive)) & ! [X12,X13,X14,X15] : (~elem(m_Down(X13),queue(host(X12))) | ~setIn(X12,alive) | elec_1 != index(status,host(X15)) | host(X13) != host(X15) | nbr_proc != host(X15) | ~elem(m_Down(X14),queue(host(X15))) | (host(X14) != sK11(X14,X15) & ~setIn(sK11(X14,X15),index(down,host(X15))) & leq(s(zero),sK11(X14,X15)) & ~leq(host(X15),sK11(X14,X15)))) & ! [X17,X18,X19] : (norm != index(status,host(X17)) | host(X17) != index(ldr,host(X17)) | ~setIn(X17,alive) | elec_2 != index(status,host(X19)) | ~elem(m_Halt(X19),queue(host(X18))) | ~setIn(X19,alive) | leq(index(pendack,host(X19)),host(X17))) & ! [X20,X21] : (~leq(index(pendack,host(X20)),index(pendack,host(X21))) | elec_2 != index(status,host(X21)) | elec_2 != index(status,host(X20)) | ~setIn(X21,alive) | ~setIn(X20,alive) | leq(host(X20),host(X21))) & ! [X22,X23,X24] : (~elem(m_Ack(X22,X23),queue(host(X22))) | elec_2 != index(status,host(X24)) | elec_2 != index(status,host(X22)) | host(X23) != host(X24) | ~setIn(X24,alive) | ~setIn(X22,alive)) & ! [X25,X26] : (leq(index(pendack,host(X26)),host(X25)) | elec_2 != index(status,host(X26)) | elec_2 != index(status,host(X25)) | ~setIn(X26,alive) | ~setIn(X25,alive) | leq(host(X25),host(X26))) & ! [X27,X28,X29] : (norm != index(status,host(X27)) | host(X27) != index(ldr,host(X27)) | ~setIn(X27,alive) | host(X28) != host(X27) | ~elem(m_Down(X28),queue(host(X29))) | ~setIn(X29,alive)) & ! [X30] : (index(elid,host(X30)) = X30 | ~setIn(X30,alive) | (elec_2 != index(status,host(X30)) & elec_1 != index(status,host(X30)))) & ! [X31,X32] : (~elem(m_Ack(X32,X31),queue(host(X32))) | elec_1 != index(status,host(X32)) | ~setIn(X32,alive)) & ! [X33,X34] : (leq(host(X33),index(pendack,host(X34))) | ~elem(m_Ack(X34,X33),queue(host(X34))) | ~setIn(X34,alive)) & ! [X35,X36] : (~setIn(X36,alive) | ~setIn(X35,alive) | host(X35) != host(X36) | X35 = X36) & ! [X37,X38,X39] : (~leq(host(X37),host(X39)) | ~elem(m_Ack(X39,X37),queue(host(X38)))) & ! [X40,X41] : (~leq(host(X40),host(X41)) | ~elem(m_Halt(X41),queue(host(X40)))) & ! [X42,X43] : (host(X43) != host(X42) | ~elem(m_Down(X43),queue(host(X42)))) & ! [X44,X45] : (~leq(host(X44),host(X45)) | ~elem(m_Ldr(X45),queue(host(X44))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4,sK5,sK6,sK7,sK8,sK9,sK10,sK11])],[f170,f173,f172,f171])).
fof(f171,plain,(
  ? [X0,X1,X2,X3] : (? [X4,X5,X6,X7] : (norm = index(status,host(X4)) & host(X4) = index(ldr,host(X4)) & setIn(X4,alive) & host(X6) = s(index(pendack,host(X7))) & host(X5) = index(pendack,host(X7)) & elec_2 = index(status,host(X7)) & elem(m_Ack(X7,X5),snoc(queue(host(X7)),m_Halt(X2))) & elem(m_Down(X6),snoc(queue(host(X7)),m_Halt(X2))) & leq(nbr_proc,s(index(pendack,host(X7)))) & setIn(X7,alive) & host(X2) != host(X7) & host(X7) = s(index(pendack,host(X2)))) & ~leq(nbr_proc,index(pendack,host(X2))) & host(X3) = index(pendack,host(X2)) & elec_2 = index(status,host(X2)) & index(elid,host(X2)) = X1 & setIn(X2,alive) & queue(host(X2)) = cons(m_Ack(X1,X3),X0) & ! [X8,X9,X10,X11] : (norm != index(status,host(X8)) | host(X8) != index(ldr,host(X8)) | ~setIn(X8,alive) | host(X10) != s(index(pendack,host(X11))) | host(X9) != index(pendack,host(X11)) | elec_2 != index(status,host(X11)) | ~leq(nbr_proc,s(index(pendack,host(X11)))) | ~elem(m_Ack(X11,X9),queue(host(X11))) | ~elem(m_Down(X10),queue(host(X11))) | ~setIn(X11,alive)) & ! [X12,X13,X14,X15] : (~elem(m_Down(X13),queue(host(X12))) | ~setIn(X12,alive) | elec_1 != index(status,host(X15)) | host(X13) != host(X15) | nbr_proc != host(X15) | ~elem(m_Down(X14),queue(host(X15))) | ? [X16] : (host(X14) != X16 & ~setIn(X16,index(down,host(X15))) & leq(s(zero),X16) & ~leq(host(X15),X16))) & ! [X17,X18,X19] : (norm != index(status,host(X17)) | host(X17) != index(ldr,host(X17)) | ~setIn(X17,alive) | elec_2 != index(status,host(X19)) | ~elem(m_Halt(X19),queue(host(X18))) | ~setIn(X19,alive) | leq(index(pendack,host(X19)),host(X17))) & ! [X20,X21] : (~leq(index(pendack,host(X20)),index(pendack,host(X21))) | elec_2 != index(status,host(X21)) | elec_2 != index(status,host(X20)) | ~setIn(X21,alive) | ~setIn(X20,alive) | leq(host(X20),host(X21))) & ! [X22,X23,X24] : (~elem(m_Ack(X22,X23),queue(host(X22))) | elec_2 != index(status,host(X24)) | elec_2 != index(status,host(X22)) | host(X23) != host(X24) | ~setIn(X24,alive) | ~setIn(X22,alive)) & ! [X25,X26] : (leq(index(pendack,host(X26)),host(X25)) | elec_2 != index(status,host(X26)) | elec_2 != index(status,host(X25)) | ~setIn(X26,alive) | ~setIn(X25,alive) | leq(host(X25),host(X26))) & ! [X27,X28,X29] : (norm != index(status,host(X27)) | host(X27) != index(ldr,host(X27)) | ~setIn(X27,alive) | host(X28) != host(X27) | ~elem(m_Down(X28),queue(host(X29))) | ~setIn(X29,alive)) & ! [X30] : (index(elid,host(X30)) = X30 | ~setIn(X30,alive) | (elec_2 != index(status,host(X30)) & elec_1 != index(status,host(X30)))) & ! [X31,X32] : (~elem(m_Ack(X32,X31),queue(host(X32))) | elec_1 != index(status,host(X32)) | ~setIn(X32,alive)) & ! [X33,X34] : (leq(host(X33),index(pendack,host(X34))) | ~elem(m_Ack(X34,X33),queue(host(X34))) | ~setIn(X34,alive)) & ! [X35,X36] : (~setIn(X36,alive) | ~setIn(X35,alive) | host(X35) != host(X36) | X35 = X36) & ! [X37,X38,X39] : (~leq(host(X37),host(X39)) | ~elem(m_Ack(X39,X37),queue(host(X38)))) & ! [X40,X41] : (~leq(host(X40),host(X41)) | ~elem(m_Halt(X41),queue(host(X40)))) & ! [X42,X43] : (host(X43) != host(X42) | ~elem(m_Down(X43),queue(host(X42)))) & ! [X44,X45] : (~leq(host(X44),host(X45)) | ~elem(m_Ldr(X45),queue(host(X44))))) => (? [X7,X6,X5,X4] : (norm = index(status,host(X4)) & host(X4) = index(ldr,host(X4)) & setIn(X4,alive) & host(X6) = s(index(pendack,host(X7))) & host(X5) = index(pendack,host(X7)) & elec_2 = index(status,host(X7)) & elem(m_Ack(X7,X5),snoc(queue(host(X7)),m_Halt(sK5))) & elem(m_Down(X6),snoc(queue(host(X7)),m_Halt(sK5))) & leq(nbr_proc,s(index(pendack,host(X7)))) & setIn(X7,alive) & host(X7) != host(sK5) & host(X7) = s(index(pendack,host(sK5)))) & ~leq(nbr_proc,index(pendack,host(sK5))) & index(pendack,host(sK5)) = host(sK6) & elec_2 = index(status,host(sK5)) & sK4 = index(elid,host(sK5)) & setIn(sK5,alive) & queue(host(sK5)) = cons(m_Ack(sK4,sK6),sK3) & ! [X8,X9,X10,X11] : (norm != index(status,host(X8)) | host(X8) != index(ldr,host(X8)) | ~setIn(X8,alive) | host(X10) != s(index(pendack,host(X11))) | host(X9) != index(pendack,host(X11)) | elec_2 != index(status,host(X11)) | ~leq(nbr_proc,s(index(pendack,host(X11)))) | ~elem(m_Ack(X11,X9),queue(host(X11))) | ~elem(m_Down(X10),queue(host(X11))) | ~setIn(X11,alive)) & ! [X12,X13,X14,X15] : (~elem(m_Down(X13),queue(host(X12))) | ~setIn(X12,alive) | elec_1 != index(status,host(X15)) | host(X13) != host(X15) | nbr_proc != host(X15) | ~elem(m_Down(X14),queue(host(X15))) | ? [X16] : (host(X14) != X16 & ~setIn(X16,index(down,host(X15))) & leq(s(zero),X16) & ~leq(host(X15),X16))) & ! [X17,X18,X19] : (norm != index(status,host(X17)) | host(X17) != index(ldr,host(X17)) | ~setIn(X17,alive) | elec_2 != index(status,host(X19)) | ~elem(m_Halt(X19),queue(host(X18))) | ~setIn(X19,alive) | leq(index(pendack,host(X19)),host(X17))) & ! [X20,X21] : (~leq(index(pendack,host(X20)),index(pendack,host(X21))) | elec_2 != index(status,host(X21)) | elec_2 != index(status,host(X20)) | ~setIn(X21,alive) | ~setIn(X20,alive) | leq(host(X20),host(X21))) & ! [X22,X23,X24] : (~elem(m_Ack(X22,X23),queue(host(X22))) | elec_2 != index(status,host(X24)) | elec_2 != index(status,host(X22)) | host(X23) != host(X24) | ~setIn(X24,alive) | ~setIn(X22,alive)) & ! [X25,X26] : (leq(index(pendack,host(X26)),host(X25)) | elec_2 != index(status,host(X26)) | elec_2 != index(status,host(X25)) | ~setIn(X26,alive) | ~setIn(X25,alive) | leq(host(X25),host(X26))) & ! [X27,X28,X29] : (norm != index(status,host(X27)) | host(X27) != index(ldr,host(X27)) | ~setIn(X27,alive) | host(X28) != host(X27) | ~elem(m_Down(X28),queue(host(X29))) | ~setIn(X29,alive)) & ! [X30] : (index(elid,host(X30)) = X30 | ~setIn(X30,alive) | (elec_2 != index(status,host(X30)) & elec_1 != index(status,host(X30)))) & ! [X31,X32] : (~elem(m_Ack(X32,X31),queue(host(X32))) | elec_1 != index(status,host(X32)) | ~setIn(X32,alive)) & ! [X33,X34] : (leq(host(X33),index(pendack,host(X34))) | ~elem(m_Ack(X34,X33),queue(host(X34))) | ~setIn(X34,alive)) & ! [X35,X36] : (~setIn(X36,alive) | ~setIn(X35,alive) | host(X35) != host(X36) | X35 = X36) & ! [X37,X38,X39] : (~leq(host(X37),host(X39)) | ~elem(m_Ack(X39,X37),queue(host(X38)))) & ! [X40,X41] : (~leq(host(X40),host(X41)) | ~elem(m_Halt(X41),queue(host(X40)))) & ! [X42,X43] : (host(X43) != host(X42) | ~elem(m_Down(X43),queue(host(X42)))) & ! [X44,X45] : (~leq(host(X44),host(X45)) | ~elem(m_Ldr(X45),queue(host(X44)))))),
  introduced(choice_axiom,[])).
fof(f172,plain,(
  ? [X7,X6,X5,X4] : (norm = index(status,host(X4)) & host(X4) = index(ldr,host(X4)) & setIn(X4,alive) & host(X6) = s(index(pendack,host(X7))) & host(X5) = index(pendack,host(X7)) & elec_2 = index(status,host(X7)) & elem(m_Ack(X7,X5),snoc(queue(host(X7)),m_Halt(sK5))) & elem(m_Down(X6),snoc(queue(host(X7)),m_Halt(sK5))) & leq(nbr_proc,s(index(pendack,host(X7)))) & setIn(X7,alive) & host(X7) != host(sK5) & host(X7) = s(index(pendack,host(sK5)))) => (norm = index(status,host(sK7)) & host(sK7) = index(ldr,host(sK7)) & setIn(sK7,alive) & host(sK9) = s(index(pendack,host(sK10))) & index(pendack,host(sK10)) = host(sK8) & elec_2 = index(status,host(sK10)) & elem(m_Ack(sK10,sK8),snoc(queue(host(sK10)),m_Halt(sK5))) & elem(m_Down(sK9),snoc(queue(host(sK10)),m_Halt(sK5))) & leq(nbr_proc,s(index(pendack,host(sK10)))) & setIn(sK10,alive) & host(sK5) != host(sK10) & s(index(pendack,host(sK5))) = host(sK10))),
  introduced(choice_axiom,[])).
fof(f173,plain,(
  ! [X15,X14] : (? [X16] : (host(X14) != X16 & ~setIn(X16,index(down,host(X15))) & leq(s(zero),X16) & ~leq(host(X15),X16)) => (host(X14) != sK11(X14,X15) & ~setIn(sK11(X14,X15),index(down,host(X15))) & leq(s(zero),sK11(X14,X15)) & ~leq(host(X15),sK11(X14,X15))))),
  introduced(choice_axiom,[])).
fof(f170,plain,(
  ? [X0,X1,X2,X3] : (? [X4,X5,X6,X7] : (norm = index(status,host(X4)) & host(X4) = index(ldr,host(X4)) & setIn(X4,alive) & host(X6) = s(index(pendack,host(X7))) & host(X5) = index(pendack,host(X7)) & elec_2 = index(status,host(X7)) & elem(m_Ack(X7,X5),snoc(queue(host(X7)),m_Halt(X2))) & elem(m_Down(X6),snoc(queue(host(X7)),m_Halt(X2))) & leq(nbr_proc,s(index(pendack,host(X7)))) & setIn(X7,alive) & host(X2) != host(X7) & host(X7) = s(index(pendack,host(X2)))) & ~leq(nbr_proc,index(pendack,host(X2))) & host(X3) = index(pendack,host(X2)) & elec_2 = index(status,host(X2)) & index(elid,host(X2)) = X1 & setIn(X2,alive) & queue(host(X2)) = cons(m_Ack(X1,X3),X0) & ! [X8,X9,X10,X11] : (norm != index(status,host(X8)) | host(X8) != index(ldr,host(X8)) | ~setIn(X8,alive) | host(X10) != s(index(pendack,host(X11))) | host(X9) != index(pendack,host(X11)) | elec_2 != index(status,host(X11)) | ~leq(nbr_proc,s(index(pendack,host(X11)))) | ~elem(m_Ack(X11,X9),queue(host(X11))) | ~elem(m_Down(X10),queue(host(X11))) | ~setIn(X11,alive)) & ! [X12,X13,X14,X15] : (~elem(m_Down(X13),queue(host(X12))) | ~setIn(X12,alive) | elec_1 != index(status,host(X15)) | host(X13) != host(X15) | nbr_proc != host(X15) | ~elem(m_Down(X14),queue(host(X15))) | ? [X16] : (host(X14) != X16 & ~setIn(X16,index(down,host(X15))) & leq(s(zero),X16) & ~leq(host(X15),X16))) & ! [X17,X18,X19] : (norm != index(status,host(X17)) | host(X17) != index(ldr,host(X17)) | ~setIn(X17,alive) | elec_2 != index(status,host(X19)) | ~elem(m_Halt(X19),queue(host(X18))) | ~setIn(X19,alive) | leq(index(pendack,host(X19)),host(X17))) & ! [X20,X21] : (~leq(index(pendack,host(X20)),index(pendack,host(X21))) | elec_2 != index(status,host(X21)) | elec_2 != index(status,host(X20)) | ~setIn(X21,alive) | ~setIn(X20,alive) | leq(host(X20),host(X21))) & ! [X22,X23,X24] : (~elem(m_Ack(X22,X23),queue(host(X22))) | elec_2 != index(status,host(X24)) | elec_2 != index(status,host(X22)) | host(X23) != host(X24) | ~setIn(X24,alive) | ~setIn(X22,alive)) & ! [X25,X26] : (leq(index(pendack,host(X26)),host(X25)) | elec_2 != index(status,host(X26)) | elec_2 != index(status,host(X25)) | ~setIn(X26,alive) | ~setIn(X25,alive) | leq(host(X25),host(X26))) & ! [X27,X28,X29] : (norm != index(status,host(X27)) | host(X27) != index(ldr,host(X27)) | ~setIn(X27,alive) | host(X28) != host(X27) | ~elem(m_Down(X28),queue(host(X29))) | ~setIn(X29,alive)) & ! [X30] : (index(elid,host(X30)) = X30 | ~setIn(X30,alive) | (elec_2 != index(status,host(X30)) & elec_1 != index(status,host(X30)))) & ! [X31,X32] : (~elem(m_Ack(X32,X31),queue(host(X32))) | elec_1 != index(status,host(X32)) | ~setIn(X32,alive)) & ! [X33,X34] : (leq(host(X33),index(pendack,host(X34))) | ~elem(m_Ack(X34,X33),queue(host(X34))) | ~setIn(X34,alive)) & ! [X35,X36] : (~setIn(X36,alive) | ~setIn(X35,alive) | host(X35) != host(X36) | X35 = X36) & ! [X37,X38,X39] : (~leq(host(X37),host(X39)) | ~elem(m_Ack(X39,X37),queue(host(X38)))) & ! [X40,X41] : (~leq(host(X40),host(X41)) | ~elem(m_Halt(X41),queue(host(X40)))) & ! [X42,X43] : (host(X43) != host(X42) | ~elem(m_Down(X43),queue(host(X42)))) & ! [X44,X45] : (~leq(host(X44),host(X45)) | ~elem(m_Ldr(X45),queue(host(X44)))))),
  inference(rectify,[],[f141])).
fof(f141,plain,(
  ? [X0,X1,X2,X3] : (? [X42,X43,X44,X45] : (norm = index(status,host(X42)) & host(X42) = index(ldr,host(X42)) & setIn(X42,alive) & s(index(pendack,host(X45))) = host(X44) & index(pendack,host(X45)) = host(X43) & elec_2 = index(status,host(X45)) & elem(m_Ack(X45,X43),snoc(queue(host(X45)),m_Halt(X2))) & elem(m_Down(X44),snoc(queue(host(X45)),m_Halt(X2))) & leq(nbr_proc,s(index(pendack,host(X45)))) & setIn(X45,alive) & host(X2) != host(X45) & s(index(pendack,host(X2))) = host(X45)) & ~leq(nbr_proc,index(pendack,host(X2))) & host(X3) = index(pendack,host(X2)) & elec_2 = index(status,host(X2)) & index(elid,host(X2)) = X1 & setIn(X2,alive) & queue(host(X2)) = cons(m_Ack(X1,X3),X0) & ! [X4,X5,X6,X7] : (norm != index(status,host(X4)) | host(X4) != index(ldr,host(X4)) | ~setIn(X4,alive) | host(X6) != s(index(pendack,host(X7))) | host(X5) != index(pendack,host(X7)) | elec_2 != index(status,host(X7)) | ~leq(nbr_proc,s(index(pendack,host(X7)))) | ~elem(m_Ack(X7,X5),queue(host(X7))) | ~elem(m_Down(X6),queue(host(X7))) | ~setIn(X7,alive)) & ! [X8,X9,X10,X11] : (~elem(m_Down(X9),queue(host(X8))) | ~setIn(X8,alive) | elec_1 != index(status,host(X11)) | host(X11) != host(X9) | nbr_proc != host(X11) | ~elem(m_Down(X10),queue(host(X11))) | ? [X12] : (host(X10) != X12 & ~setIn(X12,index(down,host(X11))) & leq(s(zero),X12) & ~leq(host(X11),X12))) & ! [X13,X14,X15] : (norm != index(status,host(X13)) | host(X13) != index(ldr,host(X13)) | ~setIn(X13,alive) | elec_2 != index(status,host(X15)) | ~elem(m_Halt(X15),queue(host(X14))) | ~setIn(X15,alive) | leq(index(pendack,host(X15)),host(X13))) & ! [X16,X17] : (~leq(index(pendack,host(X16)),index(pendack,host(X17))) | elec_2 != index(status,host(X17)) | elec_2 != index(status,host(X16)) | ~setIn(X17,alive) | ~setIn(X16,alive) | leq(host(X16),host(X17))) & ! [X18,X19,X20] : (~elem(m_Ack(X18,X19),queue(host(X18))) | elec_2 != index(status,host(X20)) | elec_2 != index(status,host(X18)) | host(X19) != host(X20) | ~setIn(X20,alive) | ~setIn(X18,alive)) & ! [X21,X22] : (leq(index(pendack,host(X22)),host(X21)) | elec_2 != index(status,host(X22)) | elec_2 != index(status,host(X21)) | ~setIn(X22,alive) | ~setIn(X21,alive) | leq(host(X21),host(X22))) & ! [X23,X24,X25] : (norm != index(status,host(X23)) | host(X23) != index(ldr,host(X23)) | ~setIn(X23,alive) | host(X23) != host(X24) | ~elem(m_Down(X24),queue(host(X25))) | ~setIn(X25,alive)) & ! [X26] : (index(elid,host(X26)) = X26 | ~setIn(X26,alive) | (elec_2 != index(status,host(X26)) & elec_1 != index(status,host(X26)))) & ! [X27,X28] : (~elem(m_Ack(X28,X27),queue(host(X28))) | elec_1 != index(status,host(X28)) | ~setIn(X28,alive)) & ! [X29,X30] : (leq(host(X29),index(pendack,host(X30))) | ~elem(m_Ack(X30,X29),queue(host(X30))) | ~setIn(X30,alive)) & ! [X31,X32] : (~setIn(X32,alive) | ~setIn(X31,alive) | host(X31) != host(X32) | X31 = X32) & ! [X33,X34,X35] : (~leq(host(X33),host(X35)) | ~elem(m_Ack(X35,X33),queue(host(X34)))) & ! [X36,X37] : (~leq(host(X36),host(X37)) | ~elem(m_Halt(X37),queue(host(X36)))) & ! [X38,X39] : (host(X38) != host(X39) | ~elem(m_Down(X39),queue(host(X38)))) & ! [X40,X41] : (~leq(host(X40),host(X41)) | ~elem(m_Ldr(X41),queue(host(X40)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ? [X0,X1,X2,X3] : ((((? [X42,X43,X44,X45] : ((((norm = index(status,host(X42)) & host(X42) = index(ldr,host(X42)) & setIn(X42,alive)) & (s(index(pendack,host(X45))) = host(X44) & index(pendack,host(X45)) = host(X43) & elec_2 = index(status,host(X45)) & elem(m_Ack(X45,X43),snoc(queue(host(X45)),m_Halt(X2))) & elem(m_Down(X44),snoc(queue(host(X45)),m_Halt(X2))) & leq(nbr_proc,s(index(pendack,host(X45)))) & setIn(X45,alive))) & host(X2) != host(X45)) & s(index(pendack,host(X2))) = host(X45)) & ~leq(nbr_proc,index(pendack,host(X2)))) & (host(X3) = index(pendack,host(X2)) & elec_2 = index(status,host(X2)) & index(elid,host(X2)) = X1)) & setIn(X2,alive)) & (queue(host(X2)) = cons(m_Ack(X1,X3),X0) & ! [X4,X5,X6,X7] : ((norm != index(status,host(X4)) | host(X4) != index(ldr,host(X4)) | ~setIn(X4,alive)) | (host(X6) != s(index(pendack,host(X7))) | host(X5) != index(pendack,host(X7)) | elec_2 != index(status,host(X7)) | ~leq(nbr_proc,s(index(pendack,host(X7)))) | ~elem(m_Ack(X7,X5),queue(host(X7))) | ~elem(m_Down(X6),queue(host(X7))) | ~setIn(X7,alive))) & ! [X8,X9,X10,X11] : ((~elem(m_Down(X9),queue(host(X8))) | ~setIn(X8,alive)) | (elec_1 != index(status,host(X11)) | host(X11) != host(X9) | nbr_proc != host(X11) | ~elem(m_Down(X10),queue(host(X11))) | ? [X12] : ((host(X10) != X12 & ~setIn(X12,index(down,host(X11)))) & (leq(s(zero),X12) & ~leq(host(X11),X12))))) & ! [X13,X14,X15] : ((norm != index(status,host(X13)) | host(X13) != index(ldr,host(X13)) | ~setIn(X13,alive)) | (elec_2 != index(status,host(X15)) | ~elem(m_Halt(X15),queue(host(X14))) | ~setIn(X15,alive) | leq(index(pendack,host(X15)),host(X13)))) & ! [X16,X17] : (~leq(index(pendack,host(X16)),index(pendack,host(X17))) | (elec_2 != index(status,host(X17)) | elec_2 != index(status,host(X16)) | ~setIn(X17,alive) | ~setIn(X16,alive) | leq(host(X16),host(X17)))) & ! [X18,X19,X20] : (~elem(m_Ack(X18,X19),queue(host(X18))) | (elec_2 != index(status,host(X20)) | elec_2 != index(status,host(X18)) | host(X19) != host(X20) | ~setIn(X20,alive) | ~setIn(X18,alive))) & ! [X21,X22] : (leq(index(pendack,host(X22)),host(X21)) | (elec_2 != index(status,host(X22)) | elec_2 != index(status,host(X21)) | ~setIn(X22,alive) | ~setIn(X21,alive) | leq(host(X21),host(X22)))) & ! [X23,X24,X25] : ((norm != index(status,host(X23)) | host(X23) != index(ldr,host(X23)) | ~setIn(X23,alive)) | (host(X23) != host(X24) | ~elem(m_Down(X24),queue(host(X25))) | ~setIn(X25,alive))) & ! [X26] : (index(elid,host(X26)) = X26 | (~setIn(X26,alive) | (elec_2 != index(status,host(X26)) & elec_1 != index(status,host(X26))))) & ! [X27,X28] : (~elem(m_Ack(X28,X27),queue(host(X28))) | (elec_1 != index(status,host(X28)) | ~setIn(X28,alive))) & ! [X29,X30] : (leq(host(X29),index(pendack,host(X30))) | (~elem(m_Ack(X30,X29),queue(host(X30))) | ~setIn(X30,alive))) & ! [X31,X32] : ((~setIn(X32,alive) | ~setIn(X31,alive)) | (host(X31) != host(X32) | X31 = X32)) & ! [X33,X34,X35] : (~leq(host(X33),host(X35)) | ~elem(m_Ack(X35,X33),queue(host(X34)))) & ! [X36,X37] : (~leq(host(X36),host(X37)) | ~elem(m_Halt(X37),queue(host(X36)))) & ! [X38,X39] : (host(X38) != host(X39) | ~elem(m_Down(X39),queue(host(X38)))) & ! [X40,X41] : (~leq(host(X40),host(X41)) | ~elem(m_Ldr(X41),queue(host(X40))))))),
  inference(ennf_transformation,[],[f124])).
fof(f124,plain,(
  ~! [X0,X1,X2,X3] : ((queue(host(X2)) = cons(m_Ack(X1,X3),X0) & ! [X4,X5,X6,X7] : ((host(X6) = s(index(pendack,host(X7))) & host(X5) = index(pendack,host(X7)) & elec_2 = index(status,host(X7)) & leq(nbr_proc,s(index(pendack,host(X7)))) & elem(m_Ack(X7,X5),queue(host(X7))) & elem(m_Down(X6),queue(host(X7))) & setIn(X7,alive)) => ~(norm = index(status,host(X4)) & host(X4) = index(ldr,host(X4)) & setIn(X4,alive))) & ! [X8,X9,X10,X11] : ((elec_1 = index(status,host(X11)) & host(X11) = host(X9) & nbr_proc = host(X11) & elem(m_Down(X10),queue(host(X11))) & ! [X12] : ((leq(s(zero),X12) & ~leq(host(X11),X12)) => (host(X10) = X12 | setIn(X12,index(down,host(X11)))))) => ~(elem(m_Down(X9),queue(host(X8))) & setIn(X8,alive))) & ! [X13,X14,X15] : ((elec_2 = index(status,host(X15)) & elem(m_Halt(X15),queue(host(X14))) & setIn(X15,alive) & ~leq(index(pendack,host(X15)),host(X13))) => ~(norm = index(status,host(X13)) & host(X13) = index(ldr,host(X13)) & setIn(X13,alive))) & ! [X16,X17] : ((elec_2 = index(status,host(X17)) & elec_2 = index(status,host(X16)) & setIn(X17,alive) & setIn(X16,alive) & ~leq(host(X16),host(X17))) => ~leq(index(pendack,host(X16)),index(pendack,host(X17)))) & ! [X18,X19,X20] : ((elec_2 = index(status,host(X20)) & elec_2 = index(status,host(X18)) & host(X19) = host(X20) & setIn(X20,alive) & setIn(X18,alive)) => ~elem(m_Ack(X18,X19),queue(host(X18)))) & ! [X21,X22] : ((elec_2 = index(status,host(X22)) & elec_2 = index(status,host(X21)) & setIn(X22,alive) & setIn(X21,alive) & ~leq(host(X21),host(X22))) => leq(index(pendack,host(X22)),host(X21))) & ! [X23,X24,X25] : ((host(X23) = host(X24) & elem(m_Down(X24),queue(host(X25))) & setIn(X25,alive)) => ~(norm = index(status,host(X23)) & host(X23) = index(ldr,host(X23)) & setIn(X23,alive))) & ! [X26] : ((setIn(X26,alive) & (elec_2 = index(status,host(X26)) | elec_1 = index(status,host(X26)))) => index(elid,host(X26)) = X26) & ! [X27,X28] : ((elec_1 = index(status,host(X28)) & setIn(X28,alive)) => ~elem(m_Ack(X28,X27),queue(host(X28)))) & ! [X29,X30] : ((elem(m_Ack(X30,X29),queue(host(X30))) & setIn(X30,alive)) => leq(host(X29),index(pendack,host(X30)))) & ! [X31,X32] : ((host(X31) = host(X32) & X31 != X32) => (~setIn(X32,alive) | ~setIn(X31,alive))) & ! [X33,X34,X35] : (elem(m_Ack(X35,X33),queue(host(X34))) => ~leq(host(X33),host(X35))) & ! [X36,X37] : (elem(m_Halt(X37),queue(host(X36))) => ~leq(host(X36),host(X37))) & ! [X38,X39] : (elem(m_Down(X39),queue(host(X38))) => host(X38) != host(X39)) & ! [X40,X41] : (elem(m_Ldr(X41),queue(host(X40))) => ~leq(host(X40),host(X41)))) => (setIn(X2,alive) => ((host(X3) = index(pendack,host(X2)) & elec_2 = index(status,host(X2)) & index(elid,host(X2)) = X1) => (~leq(nbr_proc,index(pendack,host(X2))) => ! [X42,X43,X44,X45] : (s(index(pendack,host(X2))) = host(X45) => (host(X2) != host(X45) => ((s(index(pendack,host(X45))) = host(X44) & index(pendack,host(X45)) = host(X43) & elec_2 = index(status,host(X45)) & elem(m_Ack(X45,X43),snoc(queue(host(X45)),m_Halt(X2))) & elem(m_Down(X44),snoc(queue(host(X45)),m_Halt(X2))) & leq(nbr_proc,s(index(pendack,host(X45)))) & setIn(X45,alive)) => ~(norm = index(status,host(X42)) & host(X42) = index(ldr,host(X42)) & setIn(X42,alive)))))))))),
  inference(rectify,[],[f68])).
fof(f68,negated_conjecture,(
  ~! [X11,X12,X4,X5] : ((queue(host(X4)) = cons(m_Ack(X12,X5),X11) & ! [X6,X15,X14,X13] : ((host(X14) = s(index(pendack,host(X13))) & index(pendack,host(X13)) = host(X15) & elec_2 = index(status,host(X13)) & leq(nbr_proc,s(index(pendack,host(X13)))) & elem(m_Ack(X13,X15),queue(host(X13))) & elem(m_Down(X14),queue(host(X13))) & setIn(X13,alive)) => ~(norm = index(status,host(X6)) & host(X6) = index(ldr,host(X6)) & setIn(X6,alive))) & ! [X6,X15,X14,X13] : ((elec_1 = index(status,host(X13)) & host(X13) = host(X15) & nbr_proc = host(X13) & elem(m_Down(X14),queue(host(X13))) & ! [X16] : ((leq(s(zero),X16) & ~leq(host(X13),X16)) => (host(X14) = X16 | setIn(X16,index(down,host(X13)))))) => ~(elem(m_Down(X15),queue(host(X6))) & setIn(X6,alive))) & ! [X6,X14,X13] : ((elec_2 = index(status,host(X13)) & elem(m_Halt(X13),queue(host(X14))) & setIn(X13,alive) & ~leq(index(pendack,host(X13)),host(X6))) => ~(norm = index(status,host(X6)) & host(X6) = index(ldr,host(X6)) & setIn(X6,alive))) & ! [X6,X13] : ((elec_2 = index(status,host(X13)) & elec_2 = index(status,host(X6)) & setIn(X13,alive) & setIn(X6,alive) & ~leq(host(X6),host(X13))) => ~leq(index(pendack,host(X6)),index(pendack,host(X13)))) & ! [X6,X14,X13] : ((elec_2 = index(status,host(X13)) & elec_2 = index(status,host(X6)) & host(X13) = host(X14) & setIn(X13,alive) & setIn(X6,alive)) => ~elem(m_Ack(X6,X14),queue(host(X6)))) & ! [X6,X13] : ((elec_2 = index(status,host(X13)) & elec_2 = index(status,host(X6)) & setIn(X13,alive) & setIn(X6,alive) & ~leq(host(X6),host(X13))) => leq(index(pendack,host(X13)),host(X6))) & ! [X6,X14,X13] : ((host(X6) = host(X14) & elem(m_Down(X14),queue(host(X13))) & setIn(X13,alive)) => ~(norm = index(status,host(X6)) & host(X6) = index(ldr,host(X6)) & setIn(X6,alive))) & ! [X6] : ((setIn(X6,alive) & (elec_2 = index(status,host(X6)) | elec_1 = index(status,host(X6)))) => index(elid,host(X6)) = X6) & ! [X6,X13] : ((elec_1 = index(status,host(X13)) & setIn(X13,alive)) => ~elem(m_Ack(X13,X6),queue(host(X13)))) & ! [X6,X13] : ((elem(m_Ack(X13,X6),queue(host(X13))) & setIn(X13,alive)) => leq(host(X6),index(pendack,host(X13)))) & ! [X6,X13] : ((host(X6) = host(X13) & X6 != X13) => (~setIn(X13,alive) | ~setIn(X6,alive))) & ! [X6,X14,X13] : (elem(m_Ack(X13,X6),queue(host(X14))) => ~leq(host(X6),host(X13))) & ! [X6,X13] : (elem(m_Halt(X13),queue(host(X6))) => ~leq(host(X6),host(X13))) & ! [X6,X13] : (elem(m_Down(X13),queue(host(X6))) => host(X6) != host(X13)) & ! [X6,X13] : (elem(m_Ldr(X13),queue(host(X6))) => ~leq(host(X6),host(X13)))) => (setIn(X4,alive) => ((host(X5) = index(pendack,host(X4)) & elec_2 = index(status,host(X4)) & index(elid,host(X4)) = X12) => (~leq(nbr_proc,index(pendack,host(X4))) => ! [X6,X17,X18,X19] : (s(index(pendack,host(X4))) = host(X19) => (host(X4) != host(X19) => ((s(index(pendack,host(X19))) = host(X18) & index(pendack,host(X19)) = host(X17) & elec_2 = index(status,host(X19)) & elem(m_Ack(X19,X17),snoc(queue(host(X19)),m_Halt(X4))) & elem(m_Down(X18),snoc(queue(host(X19)),m_Halt(X4))) & leq(nbr_proc,s(index(pendack,host(X19)))) & setIn(X19,alive)) => ~(norm = index(status,host(X6)) & host(X6) = index(ldr,host(X6)) & setIn(X6,alive)))))))))),
  inference(negated_conjecture,[],[f67])).
fof(f67,conjecture,(
  ! [X11,X12,X4,X5] : ((queue(host(X4)) = cons(m_Ack(X12,X5),X11) & ! [X6,X15,X14,X13] : ((host(X14) = s(index(pendack,host(X13))) & index(pendack,host(X13)) = host(X15) & elec_2 = index(status,host(X13)) & leq(nbr_proc,s(index(pendack,host(X13)))) & elem(m_Ack(X13,X15),queue(host(X13))) & elem(m_Down(X14),queue(host(X13))) & setIn(X13,alive)) => ~(norm = index(status,host(X6)) & host(X6) = index(ldr,host(X6)) & setIn(X6,alive))) & ! [X6,X15,X14,X13] : ((elec_1 = index(status,host(X13)) & host(X13) = host(X15) & nbr_proc = host(X13) & elem(m_Down(X14),queue(host(X13))) & ! [X16] : ((leq(s(zero),X16) & ~leq(host(X13),X16)) => (host(X14) = X16 | setIn(X16,index(down,host(X13)))))) => ~(elem(m_Down(X15),queue(host(X6))) & setIn(X6,alive))) & ! [X6,X14,X13] : ((elec_2 = index(status,host(X13)) & elem(m_Halt(X13),queue(host(X14))) & setIn(X13,alive) & ~leq(index(pendack,host(X13)),host(X6))) => ~(norm = index(status,host(X6)) & host(X6) = index(ldr,host(X6)) & setIn(X6,alive))) & ! [X6,X13] : ((elec_2 = index(status,host(X13)) & elec_2 = index(status,host(X6)) & setIn(X13,alive) & setIn(X6,alive) & ~leq(host(X6),host(X13))) => ~leq(index(pendack,host(X6)),index(pendack,host(X13)))) & ! [X6,X14,X13] : ((elec_2 = index(status,host(X13)) & elec_2 = index(status,host(X6)) & host(X13) = host(X14) & setIn(X13,alive) & setIn(X6,alive)) => ~elem(m_Ack(X6,X14),queue(host(X6)))) & ! [X6,X13] : ((elec_2 = index(status,host(X13)) & elec_2 = index(status,host(X6)) & setIn(X13,alive) & setIn(X6,alive) & ~leq(host(X6),host(X13))) => leq(index(pendack,host(X13)),host(X6))) & ! [X6,X14,X13] : ((host(X6) = host(X14) & elem(m_Down(X14),queue(host(X13))) & setIn(X13,alive)) => ~(norm = index(status,host(X6)) & host(X6) = index(ldr,host(X6)) & setIn(X6,alive))) & ! [X6] : ((setIn(X6,alive) & (elec_2 = index(status,host(X6)) | elec_1 = index(status,host(X6)))) => index(elid,host(X6)) = X6) & ! [X6,X13] : ((elec_1 = index(status,host(X13)) & setIn(X13,alive)) => ~elem(m_Ack(X13,X6),queue(host(X13)))) & ! [X6,X13] : ((elem(m_Ack(X13,X6),queue(host(X13))) & setIn(X13,alive)) => leq(host(X6),index(pendack,host(X13)))) & ! [X6,X13] : ((host(X6) = host(X13) & X6 != X13) => (~setIn(X13,alive) | ~setIn(X6,alive))) & ! [X6,X14,X13] : (elem(m_Ack(X13,X6),queue(host(X14))) => ~leq(host(X6),host(X13))) & ! [X6,X13] : (elem(m_Halt(X13),queue(host(X6))) => ~leq(host(X6),host(X13))) & ! [X6,X13] : (elem(m_Down(X13),queue(host(X6))) => host(X6) != host(X13)) & ! [X6,X13] : (elem(m_Ldr(X13),queue(host(X6))) => ~leq(host(X6),host(X13)))) => (setIn(X4,alive) => ((host(X5) = index(pendack,host(X4)) & elec_2 = index(status,host(X4)) & index(elid,host(X4)) = X12) => (~leq(nbr_proc,index(pendack,host(X4))) => ! [X6,X17,X18,X19] : (s(index(pendack,host(X4))) = host(X19) => (host(X4) != host(X19) => ((s(index(pendack,host(X19))) = host(X18) & index(pendack,host(X19)) = host(X17) & elec_2 = index(status,host(X19)) & elem(m_Ack(X19,X17),snoc(queue(host(X19)),m_Halt(X4))) & elem(m_Down(X18),snoc(queue(host(X19)),m_Halt(X4))) & leq(nbr_proc,s(index(pendack,host(X19)))) & setIn(X19,alive)) => ~(norm = index(status,host(X6)) & host(X6) = index(ldr,host(X6)) & setIn(X6,alive)))))))))),
  file('./p/SWV464+1.tptp',unknown)).
fof(f230,plain,(
  ( ! [X2,X0,X1] : (~elem(X0,snoc(X2,X1)) | X0 = X1 | elem(X0,X2)) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((elem(X0,snoc(X2,X1)) | (~elem(X0,X2) & X0 != X1)) & (elem(X0,X2) | X0 = X1 | ~elem(X0,snoc(X2,X1))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : ((elem(X0,snoc(X2,X1)) | (~elem(X0,X2) & X0 != X1)) & ((elem(X0,X2) | X0 = X1) | ~elem(X0,snoc(X2,X1))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2] : (elem(X0,snoc(X2,X1)) <=> (elem(X0,X2) | X0 = X1))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X4,X5,X3] : (elem(X4,snoc(X3,X5)) <=> (elem(X4,X3) | X4 = X5))),
  file('./p/SWV464+1.tptp',unknown)).
fof(f1714,plain,(
  ~elem(m_Ack(sK10,sK8),queue(host(sK10))) | ~spl12_24),
  inference(equality_resolution,[],[f1003])).
fof(f1003,plain,(
  ( ! [X1] : (host(X1) != host(sK8) | ~elem(m_Ack(sK10,X1),queue(host(sK10)))) ) | ~spl12_24),
  inference(avatar_component_clause,[],[f1002])).
fof(f1002,plain,(
  spl12_24 <=> ! [X1] : (host(X1) != host(sK8) | ~elem(m_Ack(sK10,X1),queue(host(sK10))))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_24])])).
fof(f1412,plain,(
  ~spl12_23),
  inference(avatar_contradiction_clause,[],[f1411])).
fof(f1411,plain,(
  $false | ~spl12_23),
  inference(subsumption_resolution,[],[f1410,f653])).
fof(f653,plain,(
  elem(m_Down(sK9),queue(host(sK10)))),
  inference(subsumption_resolution,[],[f648,f193])).
fof(f193,plain,(
  ( ! [X0,X1] : (m_Halt(X1) != m_Down(X0)) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : m_Halt(X1) != m_Down(X0)),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X4,X5] : m_Halt(X5) != m_Down(X4)),
  file('./p/SWV464+1.tptp',unknown)).
fof(f648,plain,(
  m_Halt(sK5) = m_Down(sK9) | elem(m_Down(sK9),queue(host(sK10)))),
  inference(resolution,[],[f230,f300])).
fof(f300,plain,(
  elem(m_Down(sK9),snoc(queue(host(sK10)),m_Halt(sK5)))),
  inference(cnf_transformation,[],[f174])).
fof(f1410,plain,(
  ~elem(m_Down(sK9),queue(host(sK10))) | ~spl12_23),
  inference(trivial_inequality_removal,[],[f1409])).
fof(f1409,plain,(
  nbr_proc != nbr_proc | ~elem(m_Down(sK9),queue(host(sK10))) | ~spl12_23),
  inference(superposition,[],[f1000,f484])).
fof(f484,plain,(
  nbr_proc = host(sK9)),
  inference(subsumption_resolution,[],[f465,f180])).
fof(f180,plain,(
  ( ! [X0] : (leq(host(X0),nbr_proc)) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : leq(host(X0),nbr_proc)),
  inference(rectify,[],[f5])).
fof(f5,axiom,(
  ! [X2] : leq(host(X2),nbr_proc)),
  file('./p/SWV464+1.tptp',unknown)).
fof(f465,plain,(
  nbr_proc = host(sK9) | ~leq(host(sK9),nbr_proc)),
  inference(resolution,[],[f261,f324])).
fof(f324,plain,(
  leq(nbr_proc,host(sK9))),
  inference(forward_demodulation,[],[f323,f322])).
fof(f322,plain,(
  host(sK9) = s(host(sK8))),
  inference(backward_demodulation,[],[f303,f304])).
fof(f304,plain,(
  host(sK9) = s(index(pendack,host(sK10)))),
  inference(cnf_transformation,[],[f174])).
fof(f303,plain,(
  index(pendack,host(sK10)) = host(sK8)),
  inference(cnf_transformation,[],[f174])).
fof(f323,plain,(
  leq(nbr_proc,s(host(sK8)))),
  inference(forward_demodulation,[],[f299,f303])).
fof(f299,plain,(
  leq(nbr_proc,s(index(pendack,host(sK10))))),
  inference(cnf_transformation,[],[f174])).
fof(f261,plain,(
  ( ! [X0,X1] : (~leq(X1,X0) | X0 = X1 | ~leq(X0,X1)) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : (((leq(X1,X0) & leq(X0,X1)) | X0 != X1) & (X0 = X1 | ~leq(X1,X0) | ~leq(X0,X1)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1] : (((leq(X1,X0) & leq(X0,X1)) | X0 != X1) & (X0 = X1 | (~leq(X1,X0) | ~leq(X0,X1))))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : ((leq(X1,X0) & leq(X0,X1)) <=> X0 = X1)),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X4,X5] : ((leq(X5,X4) & leq(X4,X5)) <=> X4 = X5)),
  file('./p/SWV464+1.tptp',unknown)).
fof(f1000,plain,(
  ( ! [X0] : (host(X0) != nbr_proc | ~elem(m_Down(X0),queue(host(sK10)))) ) | ~spl12_23),
  inference(avatar_component_clause,[],[f999])).
fof(f999,plain,(
  spl12_23 <=> ! [X0] : (host(X0) != nbr_proc | ~elem(m_Down(X0),queue(host(sK10))))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_23])])).
fof(f1004,plain,(
  spl12_23 | spl12_24 | ~spl12_1),
  inference(avatar_split_clause,[],[f997,f328,f1002,f999])).
fof(f328,plain,(
  spl12_1 <=> ! [X9,X11,X10] : (host(X10) != s(index(pendack,host(X11))) | ~setIn(X11,alive) | ~elem(m_Down(X10),queue(host(X11))) | ~elem(m_Ack(X11,X9),queue(host(X11))) | ~leq(nbr_proc,s(index(pendack,host(X11)))) | elec_2 != index(status,host(X11)) | host(X9) != index(pendack,host(X11)))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_1])])).
fof(f997,plain,(
  ( ! [X0,X1] : (host(X1) != host(sK8) | host(X0) != nbr_proc | ~elem(m_Down(X0),queue(host(sK10))) | ~elem(m_Ack(sK10,X1),queue(host(sK10)))) ) | ~spl12_1),
  inference(forward_demodulation,[],[f996,f303])).
fof(f996,plain,(
  ( ! [X0,X1] : (host(X0) != nbr_proc | ~elem(m_Down(X0),queue(host(sK10))) | ~elem(m_Ack(sK10,X1),queue(host(sK10))) | host(X1) != index(pendack,host(sK10))) ) | ~spl12_1),
  inference(subsumption_resolution,[],[f995,f259])).
fof(f259,plain,(
  ( ! [X0] : (leq(X0,X0)) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0] : leq(X0,X0)),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X4] : leq(X4,X4)),
  file('./p/SWV464+1.tptp',unknown)).
fof(f995,plain,(
  ( ! [X0,X1] : (~leq(nbr_proc,nbr_proc) | host(X0) != nbr_proc | ~elem(m_Down(X0),queue(host(sK10))) | ~elem(m_Ack(sK10,X1),queue(host(sK10))) | host(X1) != index(pendack,host(sK10))) ) | ~spl12_1),
  inference(forward_demodulation,[],[f994,f485])).
fof(f485,plain,(
  nbr_proc = s(host(sK8))),
  inference(backward_demodulation,[],[f484,f322])).
fof(f994,plain,(
  ( ! [X0,X1] : (~leq(nbr_proc,s(host(sK8))) | host(X0) != nbr_proc | ~elem(m_Down(X0),queue(host(sK10))) | ~elem(m_Ack(sK10,X1),queue(host(sK10))) | host(X1) != index(pendack,host(sK10))) ) | ~spl12_1),
  inference(forward_demodulation,[],[f993,f303])).
fof(f993,plain,(
  ( ! [X0,X1] : (host(X0) != nbr_proc | ~elem(m_Down(X0),queue(host(sK10))) | ~elem(m_Ack(sK10,X1),queue(host(sK10))) | ~leq(nbr_proc,s(index(pendack,host(sK10)))) | host(X1) != index(pendack,host(sK10))) ) | ~spl12_1),
  inference(forward_demodulation,[],[f992,f485])).
fof(f992,plain,(
  ( ! [X0,X1] : (host(X0) != s(host(sK8)) | ~elem(m_Down(X0),queue(host(sK10))) | ~elem(m_Ack(sK10,X1),queue(host(sK10))) | ~leq(nbr_proc,s(index(pendack,host(sK10)))) | host(X1) != index(pendack,host(sK10))) ) | ~spl12_1),
  inference(subsumption_resolution,[],[f991,f302])).
fof(f302,plain,(
  elec_2 = index(status,host(sK10))),
  inference(cnf_transformation,[],[f174])).
fof(f991,plain,(
  ( ! [X0,X1] : (host(X0) != s(host(sK8)) | ~elem(m_Down(X0),queue(host(sK10))) | ~elem(m_Ack(sK10,X1),queue(host(sK10))) | ~leq(nbr_proc,s(index(pendack,host(sK10)))) | elec_2 != index(status,host(sK10)) | host(X1) != index(pendack,host(sK10))) ) | ~spl12_1),
  inference(subsumption_resolution,[],[f989,f298])).
fof(f298,plain,(
  setIn(sK10,alive)),
  inference(cnf_transformation,[],[f174])).
fof(f989,plain,(
  ( ! [X0,X1] : (host(X0) != s(host(sK8)) | ~setIn(sK10,alive) | ~elem(m_Down(X0),queue(host(sK10))) | ~elem(m_Ack(sK10,X1),queue(host(sK10))) | ~leq(nbr_proc,s(index(pendack,host(sK10)))) | elec_2 != index(status,host(sK10)) | host(X1) != index(pendack,host(sK10))) ) | ~spl12_1),
  inference(superposition,[],[f329,f303])).
fof(f329,plain,(
  ( ! [X10,X11,X9] : (host(X10) != s(index(pendack,host(X11))) | ~setIn(X11,alive) | ~elem(m_Down(X10),queue(host(X11))) | ~elem(m_Ack(X11,X9),queue(host(X11))) | ~leq(nbr_proc,s(index(pendack,host(X11)))) | elec_2 != index(status,host(X11)) | host(X9) != index(pendack,host(X11))) ) | ~spl12_1),
  inference(avatar_component_clause,[],[f328])).
fof(f816,plain,(
  ~spl12_2),
  inference(avatar_contradiction_clause,[],[f815])).
fof(f815,plain,(
  $false | ~spl12_2),
  inference(subsumption_resolution,[],[f814,f307])).
fof(f307,plain,(
  norm = index(status,host(sK7))),
  inference(cnf_transformation,[],[f174])).
fof(f814,plain,(
  norm != index(status,host(sK7)) | ~spl12_2),
  inference(subsumption_resolution,[],[f813,f305])).
fof(f305,plain,(
  setIn(sK7,alive)),
  inference(cnf_transformation,[],[f174])).
fof(f813,plain,(
  ~setIn(sK7,alive) | norm != index(status,host(sK7)) | ~spl12_2),
  inference(trivial_inequality_removal,[],[f812])).
fof(f812,plain,(
  host(sK7) != host(sK7) | ~setIn(sK7,alive) | norm != index(status,host(sK7)) | ~spl12_2),
  inference(superposition,[],[f332,f306])).
fof(f306,plain,(
  host(sK7) = index(ldr,host(sK7))),
  inference(cnf_transformation,[],[f174])).
fof(f332,plain,(
  ( ! [X8] : (host(X8) != index(ldr,host(X8)) | ~setIn(X8,alive) | norm != index(status,host(X8))) ) | ~spl12_2),
  inference(avatar_component_clause,[],[f331])).
fof(f331,plain,(
  spl12_2 <=> ! [X8] : (norm != index(status,host(X8)) | ~setIn(X8,alive) | host(X8) != index(ldr,host(X8)))),
  introduced(avatar_definition,[new_symbols(naming,[spl12_2])])).
fof(f333,plain,(
  spl12_1 | spl12_2),
  inference(avatar_split_clause,[],[f289,f331,f328])).
fof(f289,plain,(
  ( ! [X10,X8,X11,X9] : (norm != index(status,host(X8)) | host(X8) != index(ldr,host(X8)) | ~setIn(X8,alive) | host(X10) != s(index(pendack,host(X11))) | host(X9) != index(pendack,host(X11)) | elec_2 != index(status,host(X11)) | ~leq(nbr_proc,s(index(pendack,host(X11)))) | ~elem(m_Ack(X11,X9),queue(host(X11))) | ~elem(m_Down(X10),queue(host(X11))) | ~setIn(X11,alive)) )),
  inference(cnf_transformation,[],[f174])).
% SZS output end Proof for SWV464+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 6012
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
