% Refutation found. Thanks to Tanya!
% SZS status Theorem for SWW098+1
% SZS output start Proof for SWW098+1
fof(f1185311,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f463,f497,f1381,f1721,f1757,f1895,f2067,f2077,f2082,f2175,f2630,f2731,f2847,f2874,f4033,f4069,f4073,f4125,f4126,f4127,f4152,f4159,f4192,f5268,f5721,f5951,f6277,f6415,f6451,f6482,f8676,f8744,f10036,f12287,f12375,f13265,f13376,f13437,f13492,f15499,f15752,f15787,f15793,f15937,f15939,f15948,f15956,f16508,f21231,f21239,f30225,f30237,f30272,f35065,f37601,f39742,f43093,f43460,f43463,f43919,f43944,f44284,f44479,f46882,f46895,f50439,f51382,f57475,f57496,f57899,f58114,f58381,f58501,f59227,f59558,f60750,f61094,f61550,f61555,f61987,f65286,f65288,f65354,f66048,f69268,f69296,f69299,f96081,f107177,f108193,f108213,f110848,f111108,f113446,f113516,f113599,f113904,f114364,f114386,f120288,f135599,f214468,f214790,f215329,f269763,f269905,f269906,f272497,f273630,f273894,f292761,f292828,f311501,f311889,f312035,f313436,f313633,f314301,f314373,f315098,f315168,f317922,f317931,f318192,f319142,f321563,f321638,f322005,f322054,f323409,f323610,f323614,f323777,f324470,f326598,f377345,f377399,f378072,f390588,f391032,f391251,f756939,f767766,f767930,f770111,f770246,f770395,f771645,f771836,f772161,f772634,f773262,f774271,f774731,f774964,f775638,f776439,f776805,f776971,f778948,f788726,f789478,f789608,f795615,f831493,f847380,f847534,f847838,f847840,f847900,f847901,f922184,f1042697,f1043101,f1044266,f1044268,f1044395,f1044596,f1178427,f1178496,f1179296,f1183667])).
fof(f1183667,plain,(
  ~spl9_52 | ~spl9_260 | ~spl9_271 | ~spl9_406 | ~spl9_918 | ~spl9_1083 | ~spl9_1364),
  inference(avatar_contradiction_clause,[],[f1183666])).
fof(f1183666,plain,(
  $false | (~spl9_52 | ~spl9_260 | ~spl9_271 | ~spl9_406 | ~spl9_918 | ~spl9_1083 | ~spl9_1364)),
  inference(subsumption_resolution,[],[f1183665,f208])).
fof(f208,plain,(
  true != false1),
  inference(definition_unfolding,[],[f134,f189])).
fof(f189,plain,(
  false = false1),
  inference(cnf_transformation,[],[f38])).
fof(f38,axiom,(
  false = false1),
  file('./p/SWW098+1.tptp',unknown)).
fof(f134,plain,(
  false != true),
  inference(cnf_transformation,[],[f2])).
fof(f2,axiom,(
  false != err & true != err & false != true),
  file('./p/SWW098+1.tptp',unknown)).
fof(f1183665,plain,(
  true = false1 | (~spl9_52 | ~spl9_260 | ~spl9_271 | ~spl9_406 | ~spl9_918 | ~spl9_1083 | ~spl9_1364)),
  inference(forward_demodulation,[],[f1183664,f767193])).
fof(f767193,plain,(
  true = not1(false1) | ~spl9_1364),
  inference(avatar_component_clause,[],[f767192])).
fof(f767192,plain,(
  spl9_1364 <=> true = not1(false1)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_1364])])).
fof(f1183664,plain,(
  false1 = not1(false1) | (~spl9_52 | ~spl9_260 | ~spl9_271 | ~spl9_406 | ~spl9_918 | ~spl9_1083)),
  inference(forward_demodulation,[],[f1183663,f34988])).
fof(f34988,plain,(
  false1 = or1(false1,false1) | ~spl9_406),
  inference(avatar_component_clause,[],[f34986])).
fof(f34986,plain,(
  spl9_406 <=> false1 = or1(false1,false1)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_406])])).
fof(f1183663,plain,(
  not1(false1) = or1(false1,false1) | (~spl9_52 | ~spl9_260 | ~spl9_271 | ~spl9_918 | ~spl9_1083)),
  inference(forward_demodulation,[],[f1183493,f1183555])).
fof(f1183555,plain,(
  false1 = lazy_impl(prop(sK2(false1,false1)),impl(not1(sK2(false1,false1)),impl(not1(sK2(false1,false1)),sK2(false1,false1)))) | (~spl9_260 | ~spl9_271 | ~spl9_918 | ~spl9_1083)),
  inference(forward_demodulation,[],[f391027,f1179459])).
fof(f1179459,plain,(
  impl(false1,sK2(false1,false1)) = not1(sK2(false1,false1)) | (~spl9_260 | ~spl9_271 | ~spl9_918)),
  inference(backward_demodulation,[],[f15749,f1056329])).
fof(f1056329,plain,(
  impl(false1,sK2(false1,sK8)) = not1(sK2(false1,sK8)) | (~spl9_260 | ~spl9_918)),
  inference(forward_demodulation,[],[f1056318,f1056316])).
fof(f1056316,plain,(
  lazy_impl(true,sK2(false1,sK8)) = not1(sK2(false1,sK8)) | (~spl9_260 | ~spl9_918)),
  inference(resolution,[],[f1056267,f238])).
fof(f238,plain,(
  ( ! [X0] : (bool(X0) | lazy_impl(true,X0) = not1(X0)) )),
  inference(definition_unfolding,[],[f193,f160])).
fof(f160,plain,(
  ( ! [X0] : (phi(X0) = lazy_impl(true,X0)) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : phi(X0) = lazy_impl(true,X0)),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X3] : phi(X3) = lazy_impl(true,X3)),
  file('./p/SWW098+1.tptp',unknown)).
fof(f193,plain,(
  ( ! [X0] : (phi(X0) = not1(X0) | bool(X0)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : (phi(X0) = not1(X0) | bool(X0))),
  inference(ennf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : (~bool(X0) => phi(X0) = not1(X0))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X2] : (~bool(X2) => phi(X2) = not1(X2))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f1056267,plain,(
  ~bool(sK2(false1,sK8)) | (~spl9_260 | ~spl9_918)),
  inference(trivial_inequality_removal,[],[f1056246])).
fof(f1056246,plain,(
  false1 != false1 | ~bool(sK2(false1,sK8)) | (~spl9_260 | ~spl9_918)),
  inference(superposition,[],[f216,f1044784])).
fof(f1044784,plain,(
  false1 = prop(sK2(false1,sK8)) | (~spl9_260 | ~spl9_918)),
  inference(forward_demodulation,[],[f15259,f314920])).
fof(f314920,plain,(
  false1 = sK7 | ~spl9_918),
  inference(avatar_component_clause,[],[f314918])).
fof(f314918,plain,(
  spl9_918 <=> false1 = sK7),
  introduced(avatar_definition,[new_symbols(naming,[spl9_918])])).
fof(f15259,plain,(
  sK7 = prop(sK2(sK7,sK8)) | ~spl9_260),
  inference(avatar_component_clause,[],[f15258])).
fof(f15258,plain,(
  spl9_260 <=> sK7 = prop(sK2(sK7,sK8))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_260])])).
fof(f216,plain,(
  ( ! [X0] : (prop(X0) != false1 | ~bool(X0)) )),
  inference(definition_unfolding,[],[f152,f189])).
fof(f152,plain,(
  ( ! [X0] : (~bool(X0) | false != prop(X0)) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : ((false = prop(X0) | bool(X0)) & (~bool(X0) | false != prop(X0)))),
  inference(nnf_transformation,[],[f8])).
fof(f8,axiom,(
  ! [X0] : (false = prop(X0) <=> ~bool(X0))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f1056318,plain,(
  impl(false1,sK2(false1,sK8)) = lazy_impl(true,sK2(false1,sK8)) | (~spl9_260 | ~spl9_918)),
  inference(resolution,[],[f1056267,f525])).
fof(f525,plain,(
  ( ! [X1] : (bool(X1) | lazy_impl(true,X1) = impl(false1,X1)) )),
  inference(resolution,[],[f218,f243])).
fof(f243,plain,(
  bool(false1)),
  inference(equality_resolution,[],[f205])).
fof(f205,plain,(
  ( ! [X0] : (bool(X0) | false1 != X0) )),
  inference(definition_unfolding,[],[f132,f189])).
fof(f132,plain,(
  ( ! [X0] : (bool(X0) | false != X0) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : ((bool(X0) | (true != X0 & false != X0)) & (true = X0 | false = X0 | ~bool(X0)))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ((bool(X0) | (true != X0 & false != X0)) & ((true = X0 | false = X0) | ~bool(X0)))),
  inference(nnf_transformation,[],[f1])).
fof(f1,axiom,(
  ! [X0] : (bool(X0) <=> (true = X0 | false = X0))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f218,plain,(
  ( ! [X0,X1] : (~bool(X0) | bool(X1) | impl(X0,X1) = lazy_impl(true,X1)) )),
  inference(definition_unfolding,[],[f155,f160])).
fof(f155,plain,(
  ( ! [X0,X1] : (impl(X0,X1) = phi(X1) | bool(X1) | ~bool(X0)) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : (impl(X0,X1) = phi(X1) | bool(X1) | ~bool(X0))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : (impl(X0,X1) = phi(X1) | (bool(X1) | ~bool(X0)))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ((~bool(X1) & bool(X0)) => impl(X0,X1) = phi(X1))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X2,X3] : ((~bool(X3) & bool(X2)) => impl(X2,X3) = phi(X3))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f15749,plain,(
  false1 = sK8 | ~spl9_271),
  inference(avatar_component_clause,[],[f15747])).
fof(f15747,plain,(
  spl9_271 <=> false1 = sK8),
  introduced(avatar_definition,[new_symbols(naming,[spl9_271])])).
fof(f391027,plain,(
  false1 = lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1)))) | ~spl9_1083),
  inference(avatar_component_clause,[],[f391025])).
fof(f391025,plain,(
  spl9_1083 <=> false1 = lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_1083])])).
fof(f1183493,plain,(
  or1(false1,false1) = not1(lazy_impl(prop(sK2(false1,false1)),impl(not1(sK2(false1,false1)),impl(not1(sK2(false1,false1)),sK2(false1,false1))))) | (~spl9_52 | ~spl9_260 | ~spl9_271 | ~spl9_918)),
  inference(backward_demodulation,[],[f1179459,f1183445])).
fof(f1183445,plain,(
  or1(false1,false1) = not1(lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1))))) | (~spl9_52 | ~spl9_271 | ~spl9_918)),
  inference(forward_demodulation,[],[f1183444,f314920])).
fof(f1183444,plain,(
  or1(sK7,false1) = not1(lazy_impl(prop(sK2(sK7,false1)),impl(impl(sK7,sK2(sK7,false1)),impl(impl(false1,sK2(sK7,false1)),sK2(sK7,false1))))) | (~spl9_52 | ~spl9_271)),
  inference(forward_demodulation,[],[f1728,f15749])).
fof(f1728,plain,(
  or1(sK7,sK8) = not1(lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | ~spl9_52),
  inference(avatar_component_clause,[],[f1727])).
fof(f1727,plain,(
  spl9_52 <=> or1(sK7,sK8) = not1(lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_52])])).
fof(f1179296,plain,(
  spl9_271 | ~spl9_87 | ~spl9_458 | ~spl9_606),
  inference(avatar_split_clause,[],[f1179295,f87093,f50398,f2862,f15747])).
fof(f2862,plain,(
  spl9_87 <=> true = sK2(true,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_87])])).
fof(f50398,plain,(
  spl9_458 <=> true = impl(false1,true)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_458])])).
fof(f87093,plain,(
  spl9_606 <=> sK8 = not1(lazy_impl(true,impl(not1(sK2(true,sK8)),true)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_606])])).
fof(f1179295,plain,(
  false1 = sK8 | (~spl9_87 | ~spl9_458 | ~spl9_606)),
  inference(forward_demodulation,[],[f1179294,f240])).
fof(f240,plain,(
  false1 = not1(true)),
  inference(definition_unfolding,[],[f195,f189])).
fof(f195,plain,(
  false = not1(true)),
  inference(cnf_transformation,[],[f43])).
fof(f43,axiom,(
  false = not1(true)),
  file('./p/SWW098+1.tptp',unknown)).
fof(f1179294,plain,(
  not1(true) = sK8 | (~spl9_87 | ~spl9_458 | ~spl9_606)),
  inference(forward_demodulation,[],[f1179293,f267])).
fof(f267,plain,(
  true = lazy_impl(true,true)),
  inference(resolution,[],[f214,f137])).
fof(f137,plain,(
  d(true)),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  d(err) & d(false) & d(true)),
  file('./p/SWW098+1.tptp',unknown)).
fof(f214,plain,(
  ( ! [X0] : (~d(X0) | lazy_impl(true,X0) = X0) )),
  inference(definition_unfolding,[],[f147,f160])).
fof(f147,plain,(
  ( ! [X0] : (~d(X0) | phi(X0) = X0) )),
  inference(cnf_transformation,[],[f6])).
fof(f6,axiom,(
  ! [X0] : ((err = phi(X0) & ~d(X0)) | (phi(X0) = X0 & d(X0)))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f1179293,plain,(
  sK8 = not1(lazy_impl(true,true)) | (~spl9_87 | ~spl9_458 | ~spl9_606)),
  inference(forward_demodulation,[],[f1179292,f50400])).
fof(f50400,plain,(
  true = impl(false1,true) | ~spl9_458),
  inference(avatar_component_clause,[],[f50398])).
fof(f1179292,plain,(
  sK8 = not1(lazy_impl(true,impl(false1,true))) | (~spl9_87 | ~spl9_606)),
  inference(forward_demodulation,[],[f1179291,f240])).
fof(f1179291,plain,(
  sK8 = not1(lazy_impl(true,impl(not1(true),true))) | (~spl9_87 | ~spl9_606)),
  inference(forward_demodulation,[],[f87094,f2864])).
fof(f2864,plain,(
  true = sK2(true,sK8) | ~spl9_87),
  inference(avatar_component_clause,[],[f2862])).
fof(f87094,plain,(
  sK8 = not1(lazy_impl(true,impl(not1(sK2(true,sK8)),true))) | ~spl9_606),
  inference(avatar_component_clause,[],[f87093])).
fof(f1178496,plain,(
  spl9_47 | ~spl9_901),
  inference(avatar_split_clause,[],[f1178495,f311836,f1576])).
fof(f1576,plain,(
  spl9_47 <=> false1 = prop(sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_47])])).
fof(f311836,plain,(
  spl9_901 <=> err = or1(true,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_901])])).
fof(f1178495,plain,(
  false1 = prop(sK8) | ~spl9_901),
  inference(subsumption_resolution,[],[f1044970,f135])).
fof(f135,plain,(
  true != err),
  inference(cnf_transformation,[],[f2])).
fof(f1044970,plain,(
  true = err | false1 = prop(sK8) | ~spl9_901),
  inference(superposition,[],[f311837,f265])).
fof(f265,plain,(
  ( ! [X0] : (true = or1(true,X0) | prop(X0) = false1) )),
  inference(resolution,[],[f176,f215])).
fof(f215,plain,(
  ( ! [X0] : (bool(X0) | prop(X0) = false1) )),
  inference(definition_unfolding,[],[f153,f189])).
fof(f153,plain,(
  ( ! [X0] : (false = prop(X0) | bool(X0)) )),
  inference(cnf_transformation,[],[f114])).
fof(f176,plain,(
  ( ! [X0] : (~bool(X0) | true = or1(true,X0)) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : (true = or1(true,X0) | ~bool(X0))),
  inference(ennf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0] : (bool(X0) => true = or1(true,X0))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X3] : (bool(X3) => true = or1(true,X3))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f311837,plain,(
  err = or1(true,sK8) | ~spl9_901),
  inference(avatar_component_clause,[],[f311836])).
fof(f1178427,plain,(
  spl9_49 | ~spl9_114 | ~spl9_167 | ~spl9_234 | ~spl9_260 | ~spl9_266 | ~spl9_918),
  inference(avatar_contradiction_clause,[],[f1178426])).
fof(f1178426,plain,(
  $false | (spl9_49 | ~spl9_114 | ~spl9_167 | ~spl9_234 | ~spl9_260 | ~spl9_266 | ~spl9_918)),
  inference(subsumption_resolution,[],[f1178425,f11776])).
fof(f11776,plain,(
  forallprefers(err,true) | ~spl9_234),
  inference(avatar_component_clause,[],[f11775])).
fof(f11775,plain,(
  spl9_234 <=> forallprefers(err,true)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_234])])).
fof(f1178425,plain,(
  ~forallprefers(err,true) | (spl9_49 | ~spl9_114 | ~spl9_167 | ~spl9_260 | ~spl9_266 | ~spl9_918)),
  inference(forward_demodulation,[],[f1178424,f268])).
fof(f268,plain,(
  err = lazy_impl(true,err)),
  inference(resolution,[],[f214,f139])).
fof(f139,plain,(
  d(err)),
  inference(cnf_transformation,[],[f3])).
fof(f1178424,plain,(
  ~forallprefers(lazy_impl(true,err),true) | (spl9_49 | ~spl9_114 | ~spl9_167 | ~spl9_260 | ~spl9_266 | ~spl9_918)),
  inference(forward_demodulation,[],[f1178423,f8121])).
fof(f8121,plain,(
  ( ! [X2] : (true = prop(sK1(false1,X2))) ) | ~spl9_167),
  inference(avatar_component_clause,[],[f8120])).
fof(f8120,plain,(
  spl9_167 <=> ! [X2] : true = prop(sK1(false1,X2))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_167])])).
fof(f1178423,plain,(
  ( ! [X11] : (~forallprefers(lazy_impl(prop(sK1(false1,X11)),err),true)) ) | (spl9_49 | ~spl9_114 | ~spl9_167 | ~spl9_260 | ~spl9_266 | ~spl9_918)),
  inference(forward_demodulation,[],[f1178236,f15576])).
fof(f15576,plain,(
  err = impl(true,err) | ~spl9_114),
  inference(forward_demodulation,[],[f14290,f268])).
fof(f14290,plain,(
  lazy_impl(true,err) = impl(true,err) | ~spl9_114),
  inference(resolution,[],[f524,f6849])).
fof(f6849,plain,(
  ~bool(err) | ~spl9_114),
  inference(trivial_inequality_removal,[],[f6834])).
fof(f6834,plain,(
  false1 != false1 | ~bool(err) | ~spl9_114),
  inference(superposition,[],[f216,f4968])).
fof(f4968,plain,(
  false1 = prop(err) | ~spl9_114),
  inference(avatar_component_clause,[],[f4966])).
fof(f4966,plain,(
  spl9_114 <=> false1 = prop(err)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_114])])).
fof(f524,plain,(
  ( ! [X0] : (bool(X0) | impl(true,X0) = lazy_impl(true,X0)) )),
  inference(resolution,[],[f218,f242])).
fof(f242,plain,(
  bool(true)),
  inference(equality_resolution,[],[f133])).
fof(f133,plain,(
  ( ! [X0] : (bool(X0) | true != X0) )),
  inference(cnf_transformation,[],[f112])).
fof(f1178236,plain,(
  ( ! [X11] : (~forallprefers(lazy_impl(prop(sK1(false1,X11)),impl(true,err)),true)) ) | (spl9_49 | ~spl9_167 | ~spl9_260 | ~spl9_266 | ~spl9_918)),
  inference(superposition,[],[f1056270,f12787])).
fof(f12787,plain,(
  ( ! [X7] : (true = impl(false1,sK1(false1,X7))) ) | ~spl9_167),
  inference(resolution,[],[f12764,f219])).
fof(f219,plain,(
  ( ! [X0] : (~bool(X0) | true = impl(false1,X0)) )),
  inference(definition_unfolding,[],[f156,f189])).
fof(f156,plain,(
  ( ! [X0] : (true = impl(false,X0) | ~bool(X0)) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : (true = impl(false,X0) | ~bool(X0))),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (bool(X0) => true = impl(false,X0))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X3] : (bool(X3) => true = impl(false,X3))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f12764,plain,(
  ( ! [X6] : (bool(sK1(false1,X6))) ) | ~spl9_167),
  inference(trivial_inequality_removal,[],[f12735])).
fof(f12735,plain,(
  ( ! [X6] : (true != true | bool(sK1(false1,X6))) ) | ~spl9_167),
  inference(superposition,[],[f150,f8121])).
fof(f150,plain,(
  ( ! [X0] : (true != prop(X0) | bool(X0)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : ((true = prop(X0) | ~bool(X0)) & (bool(X0) | true != prop(X0)))),
  inference(nnf_transformation,[],[f7])).
fof(f7,axiom,(
  ! [X0] : (true = prop(X0) <=> bool(X0))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f1056270,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),impl(impl(false1,X0),err)),true)) ) | (spl9_49 | ~spl9_260 | ~spl9_266 | ~spl9_918)),
  inference(forward_demodulation,[],[f1056269,f1864])).
fof(f1864,plain,(
  ( ! [X14] : (err = impl(err,X14)) ) | spl9_49),
  inference(subsumption_resolution,[],[f1814,f1714])).
fof(f1714,plain,(
  true != prop(err) | spl9_49),
  inference(avatar_component_clause,[],[f1713])).
fof(f1713,plain,(
  spl9_49 <=> true = prop(err)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_49])])).
fof(f1814,plain,(
  ( ! [X14] : (err = impl(err,X14) | true = prop(err)) )),
  inference(superposition,[],[f291,f268])).
fof(f291,plain,(
  ( ! [X0,X1] : (impl(X0,X1) = lazy_impl(true,X0) | true = prop(X0)) )),
  inference(resolution,[],[f217,f151])).
fof(f151,plain,(
  ( ! [X0] : (~bool(X0) | true = prop(X0)) )),
  inference(cnf_transformation,[],[f113])).
fof(f217,plain,(
  ( ! [X0,X1] : (bool(X0) | impl(X0,X1) = lazy_impl(true,X0)) )),
  inference(definition_unfolding,[],[f154,f160])).
fof(f154,plain,(
  ( ! [X0,X1] : (phi(X0) = impl(X0,X1) | bool(X0)) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : (phi(X0) = impl(X0,X1) | bool(X0))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : (~bool(X0) => phi(X0) = impl(X0,X1))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X2,X3] : (~bool(X2) => impl(X2,X3) = phi(X2))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f1056269,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),impl(impl(false1,X0),impl(err,X0))),true)) ) | (~spl9_260 | ~spl9_266 | ~spl9_918)),
  inference(forward_demodulation,[],[f1056268,f15603])).
fof(f15603,plain,(
  ( ! [X101] : (err = impl(sK8,X101)) ) | ~spl9_266),
  inference(avatar_component_clause,[],[f15602])).
fof(f15602,plain,(
  spl9_266 <=> ! [X101] : err = impl(sK8,X101)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_266])])).
fof(f1056268,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),impl(impl(false1,X0),impl(impl(sK8,X0),X0))),true)) ) | (~spl9_260 | ~spl9_918)),
  inference(forward_demodulation,[],[f1056244,f221])).
fof(f221,plain,(
  ( ! [X0] : (true = lazy_impl(false1,X0)) )),
  inference(definition_unfolding,[],[f159,f189])).
fof(f159,plain,(
  ( ! [X0] : (true = lazy_impl(false,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : true = lazy_impl(false,X0)),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X3] : true = lazy_impl(false,X3)),
  file('./p/SWW098+1.tptp',unknown)).
fof(f1056244,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),impl(impl(false1,X0),impl(impl(sK8,X0),X0))),lazy_impl(false1,impl(impl(false1,sK2(false1,sK8)),impl(impl(sK8,sK2(false1,sK8)),sK2(false1,sK8)))))) ) | (~spl9_260 | ~spl9_918)),
  inference(superposition,[],[f233,f1044784])).
fof(f233,plain,(
  ( ! [X0,X3,X1] : (~forallprefers(lazy_impl(prop(X3),impl(impl(X0,X3),impl(impl(X1,X3),X3))),lazy_impl(prop(sK2(X0,X1)),impl(impl(X0,sK2(X0,X1)),impl(impl(X1,sK2(X0,X1)),sK2(X0,X1)))))) )),
  inference(definition_unfolding,[],[f180,f178,f178])).
fof(f178,plain,(
  ( ! [X2,X0,X1] : (f3(X0,X1,X2) = lazy_impl(prop(X2),impl(impl(X0,X2),impl(impl(X1,X2),X2)))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : f3(X0,X1,X2) = lazy_impl(prop(X2),impl(impl(X0,X2),impl(impl(X1,X2),X2)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X4,X5,X6] : f3(X4,X5,X6) = lazy_impl(prop(X6),impl(impl(X4,X6),impl(impl(X5,X6),X6)))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f180,plain,(
  ( ! [X0,X3,X1] : (~forallprefers(f3(X0,X1,X3),f3(X0,X1,sK2(X0,X1)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : (! [X3] : ~forallprefers(f3(X0,X1,X3),f3(X0,X1,sK2(X0,X1))) & or2(X0,X1) = phi(f3(X0,X1,sK2(X0,X1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f104,f119])).
fof(f119,plain,(
  ! [X1,X0] : (? [X2] : (! [X3] : ~forallprefers(f3(X0,X1,X3),f3(X0,X1,X2)) & or2(X0,X1) = phi(f3(X0,X1,X2))) => (! [X3] : ~forallprefers(f3(X0,X1,X3),f3(X0,X1,sK2(X0,X1))) & or2(X0,X1) = phi(f3(X0,X1,sK2(X0,X1)))))),
  introduced(choice_axiom,[])).
fof(f104,plain,(
  ! [X0,X1] : ? [X2] : (! [X3] : ~forallprefers(f3(X0,X1,X3),f3(X0,X1,X2)) & or2(X0,X1) = phi(f3(X0,X1,X2)))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : ? [X2] : (~? [X3] : forallprefers(f3(X0,X1,X3),f3(X0,X1,X2)) & or2(X0,X1) = phi(f3(X0,X1,X2)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X4,X5] : ? [X6] : (~? [X7] : forallprefers(f3(X4,X5,X7),f3(X4,X5,X6)) & or2(X4,X5) = phi(f3(X4,X5,X6)))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f1044596,plain,(
  spl9_111 | spl9_690 | spl9_85),
  inference(avatar_split_clause,[],[f784867,f2853,f113518,f4108])).
fof(f4108,plain,(
  spl9_111 <=> err = lazy_impl(true,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_111])])).
fof(f113518,plain,(
  spl9_690 <=> ! [X33] : sK8 = impl(sK8,X33)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_690])])).
fof(f2853,plain,(
  spl9_85 <=> true = or1(true,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_85])])).
fof(f784867,plain,(
  ( ! [X6] : (sK8 = impl(sK8,X6) | err = lazy_impl(true,sK8)) ) | spl9_85),
  inference(trivial_inequality_removal,[],[f784744])).
fof(f784744,plain,(
  ( ! [X6] : (true != true | sK8 = impl(sK8,X6) | err = lazy_impl(true,sK8)) ) | spl9_85),
  inference(superposition,[],[f2855,f2408])).
fof(f2408,plain,(
  ( ! [X2,X3] : (impl(X2,X3) = X2 | true = or1(true,X2) | err = lazy_impl(true,X2)) )),
  inference(superposition,[],[f294,f212])).
fof(f212,plain,(
  ( ! [X0] : (err = lazy_impl(true,X0) | lazy_impl(true,X0) = X0) )),
  inference(definition_unfolding,[],[f149,f160,f160])).
fof(f149,plain,(
  ( ! [X0] : (err = phi(X0) | phi(X0) = X0) )),
  inference(cnf_transformation,[],[f6])).
fof(f294,plain,(
  ( ! [X6,X7] : (lazy_impl(true,X6) = impl(X6,X7) | true = or1(true,X6)) )),
  inference(resolution,[],[f217,f176])).
fof(f2855,plain,(
  true != or1(true,sK8) | spl9_85),
  inference(avatar_component_clause,[],[f2853])).
fof(f1044395,plain,(
  spl9_113 | spl9_266 | spl9_85),
  inference(avatar_split_clause,[],[f784868,f2853,f15602,f4117])).
fof(f4117,plain,(
  spl9_113 <=> sK8 = lazy_impl(true,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_113])])).
fof(f784868,plain,(
  ( ! [X5] : (err = impl(sK8,X5) | sK8 = lazy_impl(true,sK8)) ) | spl9_85),
  inference(trivial_inequality_removal,[],[f784743])).
fof(f784743,plain,(
  ( ! [X5] : (true != true | err = impl(sK8,X5) | sK8 = lazy_impl(true,sK8)) ) | spl9_85),
  inference(superposition,[],[f2855,f2407])).
fof(f2407,plain,(
  ( ! [X0,X1] : (err = impl(X0,X1) | true = or1(true,X0) | lazy_impl(true,X0) = X0) )),
  inference(superposition,[],[f294,f212])).
fof(f1044268,plain,(
  spl9_13 | spl9_260 | ~spl9_918),
  inference(avatar_contradiction_clause,[],[f1044267])).
fof(f1044267,plain,(
  $false | (spl9_13 | spl9_260 | ~spl9_918)),
  inference(subsumption_resolution,[],[f1044262,f922374])).
fof(f922374,plain,(
  true != prop(sK2(false1,sK8)) | (spl9_13 | ~spl9_918)),
  inference(forward_demodulation,[],[f637,f314920])).
fof(f637,plain,(
  true != prop(sK2(sK7,sK8)) | spl9_13),
  inference(avatar_component_clause,[],[f636])).
fof(f636,plain,(
  spl9_13 <=> true = prop(sK2(sK7,sK8))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_13])])).
fof(f1044262,plain,(
  true = prop(sK2(false1,sK8)) | (spl9_260 | ~spl9_918)),
  inference(trivial_inequality_removal,[],[f1044261])).
fof(f1044261,plain,(
  false1 != false1 | true = prop(sK2(false1,sK8)) | (spl9_260 | ~spl9_918)),
  inference(superposition,[],[f1043127,f252])).
fof(f252,plain,(
  ( ! [X0] : (prop(X0) = false1 | true = prop(X0)) )),
  inference(resolution,[],[f215,f151])).
fof(f1043127,plain,(
  false1 != prop(sK2(false1,sK8)) | (spl9_260 | ~spl9_918)),
  inference(forward_demodulation,[],[f15260,f314920])).
fof(f15260,plain,(
  sK7 != prop(sK2(sK7,sK8)) | spl9_260),
  inference(avatar_component_clause,[],[f15258])).
fof(f1044266,plain,(
  spl9_32 | spl9_182 | spl9_260 | ~spl9_918),
  inference(avatar_contradiction_clause,[],[f1044265])).
fof(f1044265,plain,(
  $false | (spl9_32 | spl9_182 | spl9_260 | ~spl9_918)),
  inference(subsumption_resolution,[],[f1044264,f922364])).
fof(f922364,plain,(
  false1 != sK2(false1,sK8) | (spl9_182 | ~spl9_918)),
  inference(forward_demodulation,[],[f8568,f314920])).
fof(f8568,plain,(
  false1 != sK2(sK7,sK8) | spl9_182),
  inference(avatar_component_clause,[],[f8567])).
fof(f8567,plain,(
  spl9_182 <=> false1 = sK2(sK7,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_182])])).
fof(f1044264,plain,(
  false1 = sK2(false1,sK8) | (spl9_32 | spl9_260 | ~spl9_918)),
  inference(subsumption_resolution,[],[f1044263,f922372])).
fof(f922372,plain,(
  true != sK2(false1,sK8) | (spl9_32 | ~spl9_918)),
  inference(forward_demodulation,[],[f1182,f314920])).
fof(f1182,plain,(
  true != sK2(sK7,sK8) | spl9_32),
  inference(avatar_component_clause,[],[f1181])).
fof(f1181,plain,(
  spl9_32 <=> true = sK2(sK7,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_32])])).
fof(f1044263,plain,(
  true = sK2(false1,sK8) | false1 = sK2(false1,sK8) | (spl9_260 | ~spl9_918)),
  inference(trivial_inequality_removal,[],[f1044260])).
fof(f1044260,plain,(
  false1 != false1 | true = sK2(false1,sK8) | false1 = sK2(false1,sK8) | (spl9_260 | ~spl9_918)),
  inference(superposition,[],[f1043127,f282])).
fof(f282,plain,(
  ( ! [X0] : (prop(X0) = false1 | true = X0 | false1 = X0) )),
  inference(resolution,[],[f206,f215])).
fof(f206,plain,(
  ( ! [X0] : (~bool(X0) | false1 = X0 | true = X0) )),
  inference(definition_unfolding,[],[f131,f189])).
fof(f131,plain,(
  ( ! [X0] : (true = X0 | false = X0 | ~bool(X0)) )),
  inference(cnf_transformation,[],[f112])).
fof(f1043101,plain,(
  ~spl9_47 | ~spl9_95 | ~spl9_112 | ~spl9_113 | ~spl9_269 | ~spl9_690),
  inference(avatar_contradiction_clause,[],[f1043100])).
fof(f1043100,plain,(
  $false | (~spl9_47 | ~spl9_95 | ~spl9_112 | ~spl9_113 | ~spl9_269 | ~spl9_690)),
  inference(subsumption_resolution,[],[f968391,f1042833])).
fof(f1042833,plain,(
  bool(sK8) | (~spl9_95 | ~spl9_112 | ~spl9_113 | ~spl9_269 | ~spl9_690)),
  inference(forward_demodulation,[],[f1042832,f4119])).
fof(f4119,plain,(
  sK8 = lazy_impl(true,sK8) | ~spl9_113),
  inference(avatar_component_clause,[],[f4117])).
fof(f1042832,plain,(
  bool(lazy_impl(true,sK8)) | (~spl9_95 | ~spl9_112 | ~spl9_269 | ~spl9_690)),
  inference(forward_demodulation,[],[f1042831,f779197])).
fof(f779197,plain,(
  sK8 = impl(false1,sK8) | (~spl9_95 | ~spl9_112 | ~spl9_690)),
  inference(forward_demodulation,[],[f779181,f113519])).
fof(f113519,plain,(
  ( ! [X33] : (sK8 = impl(sK8,X33)) ) | ~spl9_690),
  inference(avatar_component_clause,[],[f113518])).
fof(f779181,plain,(
  sK8 = impl(false1,impl(sK8,false1)) | (~spl9_95 | ~spl9_112)),
  inference(backward_demodulation,[],[f4114,f3164])).
fof(f3164,plain,(
  not1(sK8) = impl(false1,impl(not1(sK8),false1)) | ~spl9_95),
  inference(avatar_component_clause,[],[f3163])).
fof(f3163,plain,(
  spl9_95 <=> not1(sK8) = impl(false1,impl(not1(sK8),false1))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_95])])).
fof(f4114,plain,(
  sK8 = not1(sK8) | ~spl9_112),
  inference(avatar_component_clause,[],[f4112])).
fof(f4112,plain,(
  spl9_112 <=> sK8 = not1(sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_112])])).
fof(f1042831,plain,(
  bool(lazy_impl(true,impl(false1,sK8))) | (~spl9_269 | ~spl9_690)),
  inference(forward_demodulation,[],[f1042830,f113519])).
fof(f1042830,plain,(
  bool(lazy_impl(true,impl(false1,impl(sK8,false1)))) | (~spl9_269 | ~spl9_690)),
  inference(forward_demodulation,[],[f15625,f113519])).
fof(f15625,plain,(
  bool(lazy_impl(true,impl(false1,impl(impl(sK8,false1),false1)))) | ~spl9_269),
  inference(avatar_component_clause,[],[f15624])).
fof(f15624,plain,(
  spl9_269 <=> bool(lazy_impl(true,impl(false1,impl(impl(sK8,false1),false1))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_269])])).
fof(f968391,plain,(
  ~bool(sK8) | ~spl9_47),
  inference(trivial_inequality_removal,[],[f968390])).
fof(f968390,plain,(
  false1 != false1 | ~bool(sK8) | ~spl9_47),
  inference(superposition,[],[f216,f1578])).
fof(f1578,plain,(
  false1 = prop(sK8) | ~spl9_47),
  inference(avatar_component_clause,[],[f1576])).
fof(f1042697,plain,(
  ~spl9_112 | ~spl9_113 | ~spl9_260 | ~spl9_265 | ~spl9_425 | ~spl9_458 | ~spl9_690 | spl9_712 | ~spl9_918),
  inference(avatar_contradiction_clause,[],[f1042696])).
fof(f1042696,plain,(
  $false | (~spl9_112 | ~spl9_113 | ~spl9_260 | ~spl9_265 | ~spl9_425 | ~spl9_458 | ~spl9_690 | spl9_712 | ~spl9_918)),
  inference(subsumption_resolution,[],[f1042695,f776267])).
fof(f776267,plain,(
  d(sK8) | (~spl9_113 | ~spl9_425)),
  inference(backward_demodulation,[],[f4119,f117300])).
fof(f117300,plain,(
  d(lazy_impl(true,sK8)) | ~spl9_425),
  inference(superposition,[],[f41243,f248])).
fof(f248,plain,(
  true = prop(true)),
  inference(resolution,[],[f151,f242])).
fof(f41243,plain,(
  ( ! [X33] : (d(lazy_impl(prop(X33),sK8))) ) | ~spl9_425),
  inference(avatar_component_clause,[],[f41242])).
fof(f41242,plain,(
  spl9_425 <=> ! [X33] : d(lazy_impl(prop(X33),sK8))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_425])])).
fof(f1042695,plain,(
  ~d(sK8) | (~spl9_112 | ~spl9_113 | ~spl9_260 | ~spl9_265 | ~spl9_458 | ~spl9_690 | spl9_712 | ~spl9_918)),
  inference(subsumption_resolution,[],[f1042694,f137])).
fof(f1042694,plain,(
  ~d(true) | ~d(sK8) | (~spl9_112 | ~spl9_113 | ~spl9_260 | ~spl9_265 | ~spl9_458 | ~spl9_690 | spl9_712 | ~spl9_918)),
  inference(subsumption_resolution,[],[f1042693,f779183])).
fof(f779183,plain,(
  ~bool(sK8) | (~spl9_112 | spl9_712)),
  inference(backward_demodulation,[],[f4114,f113809])).
fof(f113809,plain,(
  ~bool(not1(sK8)) | spl9_712),
  inference(avatar_component_clause,[],[f113807])).
fof(f113807,plain,(
  spl9_712 <=> bool(not1(sK8))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_712])])).
fof(f1042693,plain,(
  bool(sK8) | ~d(true) | ~d(sK8) | (~spl9_113 | ~spl9_260 | ~spl9_265 | ~spl9_458 | ~spl9_690 | ~spl9_918)),
  inference(subsumption_resolution,[],[f1042691,f242])).
fof(f1042691,plain,(
  ~bool(true) | bool(sK8) | ~d(true) | ~d(sK8) | (~spl9_113 | ~spl9_260 | ~spl9_265 | ~spl9_458 | ~spl9_690 | ~spl9_918)),
  inference(resolution,[],[f1042509,f141])).
fof(f141,plain,(
  ( ! [X0,X1] : (forallprefers(X0,X1) | ~bool(X1) | bool(X0) | ~d(X1) | ~d(X0)) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : (forallprefers(X0,X1) | ((true != X1 | false != X0) & (~bool(X1) | bool(X0) | ~d(X1) | ~d(X0)) & (~d(X1) | d(X0))))),
  inference(ennf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : (((true = X1 & false = X0) | (bool(X1) & ~bool(X0) & d(X1) & d(X0)) | (d(X1) & ~d(X0))) => forallprefers(X0,X1))),
  inference(unused_predicate_definition_removal,[],[f4])).
fof(f4,axiom,(
  ! [X0,X1] : (forallprefers(X0,X1) <=> ((true = X1 & false = X0) | (bool(X1) & ~bool(X0) & d(X1) & d(X0)) | (d(X1) & ~d(X0))))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f1042509,plain,(
  ~forallprefers(sK8,true) | (~spl9_113 | ~spl9_260 | ~spl9_265 | ~spl9_458 | ~spl9_690 | ~spl9_918)),
  inference(forward_demodulation,[],[f1042508,f4119])).
fof(f1042508,plain,(
  ~forallprefers(lazy_impl(true,sK8),true) | (~spl9_260 | ~spl9_265 | ~spl9_458 | ~spl9_690 | ~spl9_918)),
  inference(forward_demodulation,[],[f1042507,f15600])).
fof(f15600,plain,(
  sK8 = impl(true,sK8) | ~spl9_265),
  inference(avatar_component_clause,[],[f15598])).
fof(f15598,plain,(
  spl9_265 <=> sK8 = impl(true,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_265])])).
fof(f1042507,plain,(
  ~forallprefers(lazy_impl(true,impl(true,sK8)),true) | (~spl9_260 | ~spl9_458 | ~spl9_690 | ~spl9_918)),
  inference(forward_demodulation,[],[f1042335,f50400])).
fof(f1042335,plain,(
  ~forallprefers(lazy_impl(true,impl(impl(false1,true),sK8)),true) | (~spl9_260 | ~spl9_690 | ~spl9_918)),
  inference(superposition,[],[f935951,f248])).
fof(f935951,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),impl(impl(false1,X0),sK8)),true)) ) | (~spl9_260 | ~spl9_690 | ~spl9_918)),
  inference(forward_demodulation,[],[f935950,f113519])).
fof(f935950,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),impl(impl(false1,X0),impl(sK8,X0))),true)) ) | (~spl9_260 | ~spl9_690 | ~spl9_918)),
  inference(forward_demodulation,[],[f935949,f113519])).
fof(f935949,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),impl(impl(false1,X0),impl(impl(sK8,X0),X0))),true)) ) | (~spl9_260 | ~spl9_918)),
  inference(forward_demodulation,[],[f935926,f221])).
fof(f935926,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),impl(impl(false1,X0),impl(impl(sK8,X0),X0))),lazy_impl(false1,impl(impl(false1,sK2(false1,sK8)),impl(impl(sK8,sK2(false1,sK8)),sK2(false1,sK8)))))) ) | (~spl9_260 | ~spl9_918)),
  inference(superposition,[],[f233,f922112])).
fof(f922112,plain,(
  false1 = prop(sK2(false1,sK8)) | (~spl9_260 | ~spl9_918)),
  inference(forward_demodulation,[],[f15259,f314920])).
fof(f922184,plain,(
  ~spl9_42 | spl9_52 | ~spl9_260 | ~spl9_918),
  inference(avatar_split_clause,[],[f922183,f314918,f15258,f1727,f1241])).
fof(f1241,plain,(
  spl9_42 <=> false1 = or1(false1,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_42])])).
fof(f922183,plain,(
  false1 != or1(false1,sK8) | (spl9_52 | ~spl9_260 | ~spl9_918)),
  inference(forward_demodulation,[],[f922182,f240])).
fof(f922182,plain,(
  not1(true) != or1(false1,sK8) | (spl9_52 | ~spl9_260 | ~spl9_918)),
  inference(forward_demodulation,[],[f922181,f221])).
fof(f922181,plain,(
  or1(false1,sK8) != not1(lazy_impl(false1,impl(impl(false1,sK2(false1,sK8)),impl(impl(sK8,sK2(false1,sK8)),sK2(false1,sK8))))) | (spl9_52 | ~spl9_260 | ~spl9_918)),
  inference(forward_demodulation,[],[f922180,f922112])).
fof(f922180,plain,(
  or1(false1,sK8) != not1(lazy_impl(prop(sK2(false1,sK8)),impl(impl(false1,sK2(false1,sK8)),impl(impl(sK8,sK2(false1,sK8)),sK2(false1,sK8))))) | (spl9_52 | ~spl9_918)),
  inference(forward_demodulation,[],[f1729,f314920])).
fof(f1729,plain,(
  or1(sK7,sK8) != not1(lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | spl9_52),
  inference(avatar_component_clause,[],[f1727])).
fof(f847901,plain,(
  spl9_456 | spl9_57 | spl9_70),
  inference(avatar_split_clause,[],[f797229,f2172,f1888,f46638])).
fof(f46638,plain,(
  spl9_456 <=> sK7 = not1(sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_456])])).
fof(f1888,plain,(
  spl9_57 <=> true = prop(sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_57])])).
fof(f2172,plain,(
  spl9_70 <=> err = lazy_impl(true,sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_70])])).
fof(f797229,plain,(
  true = prop(sK7) | sK7 = not1(sK7) | spl9_70),
  inference(trivial_inequality_removal,[],[f797056])).
fof(f797056,plain,(
  err != err | true = prop(sK7) | sK7 = not1(sK7) | spl9_70),
  inference(superposition,[],[f2174,f1696])).
fof(f1696,plain,(
  ( ! [X1] : (err = lazy_impl(true,X1) | true = prop(X1) | not1(X1) = X1) )),
  inference(superposition,[],[f283,f212])).
fof(f283,plain,(
  ( ! [X0] : (lazy_impl(true,X0) = not1(X0) | true = prop(X0)) )),
  inference(resolution,[],[f238,f151])).
fof(f2174,plain,(
  err != lazy_impl(true,sK7) | spl9_70),
  inference(avatar_component_clause,[],[f2172])).
fof(f847900,plain,(
  spl9_456 | spl9_73 | spl9_70),
  inference(avatar_split_clause,[],[f797228,f2172,f2198,f46638])).
fof(f2198,plain,(
  spl9_73 <=> true = impl(false1,sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_73])])).
fof(f797228,plain,(
  true = impl(false1,sK7) | sK7 = not1(sK7) | spl9_70),
  inference(trivial_inequality_removal,[],[f797058])).
fof(f797058,plain,(
  err != err | true = impl(false1,sK7) | sK7 = not1(sK7) | spl9_70),
  inference(superposition,[],[f2174,f1786])).
fof(f1786,plain,(
  ( ! [X3] : (err = lazy_impl(true,X3) | true = impl(false1,X3) | not1(X3) = X3) )),
  inference(superposition,[],[f288,f212])).
fof(f288,plain,(
  ( ! [X5] : (lazy_impl(true,X5) = not1(X5) | true = impl(false1,X5)) )),
  inference(resolution,[],[f238,f219])).
fof(f847840,plain,(
  spl9_57 | spl9_262 | spl9_70),
  inference(avatar_split_clause,[],[f797227,f2172,f15326,f1888])).
fof(f15326,plain,(
  spl9_262 <=> ! [X7] : sK7 = impl(sK7,X7)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_262])])).
fof(f797227,plain,(
  ( ! [X16] : (sK7 = impl(sK7,X16) | true = prop(sK7)) ) | spl9_70),
  inference(trivial_inequality_removal,[],[f797060])).
fof(f797060,plain,(
  ( ! [X16] : (err != err | sK7 = impl(sK7,X16) | true = prop(sK7)) ) | spl9_70),
  inference(superposition,[],[f2174,f1811])).
fof(f1811,plain,(
  ( ! [X10,X9] : (impl(X9,X10) = X9 | err = lazy_impl(true,X9) | true = prop(X9)) )),
  inference(superposition,[],[f291,f212])).
fof(f847838,plain,(
  spl9_918 | spl9_262 | spl9_34 | spl9_70),
  inference(avatar_split_clause,[],[f847837,f2172,f1190,f15326,f314918])).
fof(f1190,plain,(
  spl9_34 <=> true = sK7),
  introduced(avatar_definition,[new_symbols(naming,[spl9_34])])).
fof(f847837,plain,(
  ( ! [X25] : (sK7 = impl(sK7,X25) | false1 = sK7) ) | (spl9_34 | spl9_70)),
  inference(subsumption_resolution,[],[f797222,f1191])).
fof(f1191,plain,(
  true != sK7 | spl9_34),
  inference(avatar_component_clause,[],[f1190])).
fof(f797222,plain,(
  ( ! [X25] : (sK7 = impl(sK7,X25) | true = sK7 | false1 = sK7) ) | spl9_70),
  inference(trivial_inequality_removal,[],[f797071])).
fof(f797071,plain,(
  ( ! [X25] : (err != err | sK7 = impl(sK7,X25) | true = sK7 | false1 = sK7) ) | spl9_70),
  inference(superposition,[],[f2174,f5139])).
fof(f5139,plain,(
  ( ! [X2,X3] : (impl(X2,X3) = X2 | err = lazy_impl(true,X2) | true = X2 | false1 = X2) )),
  inference(superposition,[],[f295,f212])).
fof(f295,plain,(
  ( ! [X8,X9] : (impl(X8,X9) = lazy_impl(true,X8) | false1 = X8 | true = X8) )),
  inference(resolution,[],[f217,f206])).
fof(f847534,plain,(
  spl9_918 | ~spl9_73 | spl9_34 | spl9_69),
  inference(avatar_split_clause,[],[f847533,f2167,f1190,f2198,f314918])).
fof(f2167,plain,(
  spl9_69 <=> true = lazy_impl(true,sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_69])])).
fof(f847533,plain,(
  true != impl(false1,sK7) | false1 = sK7 | (spl9_34 | spl9_69)),
  inference(subsumption_resolution,[],[f785958,f1191])).
fof(f785958,plain,(
  true != impl(false1,sK7) | false1 = sK7 | true = sK7 | spl9_69),
  inference(superposition,[],[f2169,f1601])).
fof(f1601,plain,(
  ( ! [X4] : (lazy_impl(true,X4) = impl(false1,X4) | false1 = X4 | true = X4) )),
  inference(resolution,[],[f525,f206])).
fof(f2169,plain,(
  true != lazy_impl(true,sK7) | spl9_69),
  inference(avatar_component_clause,[],[f2167])).
fof(f847380,plain,(
  spl9_1486 | ~spl9_203 | ~spl9_256),
  inference(avatar_split_clause,[],[f847379,f13487,f10009,f788650])).
fof(f788650,plain,(
  spl9_1486 <=> ! [X25] : true = prop(sK2(sK7,X25))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_1486])])).
fof(f10009,plain,(
  spl9_203 <=> ! [X26] : d(lazy_impl(prop(X26),sK7))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_203])])).
fof(f13487,plain,(
  spl9_256 <=> ! [X33] : (~bool(lazy_impl(prop(sK2(sK7,X33)),sK7)) | ~d(lazy_impl(prop(sK2(sK7,X33)),sK7)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_256])])).
fof(f847379,plain,(
  ( ! [X1] : (true = prop(sK2(sK7,X1))) ) | (~spl9_203 | ~spl9_256)),
  inference(subsumption_resolution,[],[f847378,f242])).
fof(f847378,plain,(
  ( ! [X1] : (~bool(true) | true = prop(sK2(sK7,X1))) ) | (~spl9_203 | ~spl9_256)),
  inference(forward_demodulation,[],[f847322,f221])).
fof(f847322,plain,(
  ( ! [X1] : (~bool(lazy_impl(false1,sK7)) | true = prop(sK2(sK7,X1))) ) | (~spl9_203 | ~spl9_256)),
  inference(superposition,[],[f774481,f252])).
fof(f774481,plain,(
  ( ! [X33] : (~bool(lazy_impl(prop(sK2(sK7,X33)),sK7))) ) | (~spl9_203 | ~spl9_256)),
  inference(subsumption_resolution,[],[f13488,f10010])).
fof(f10010,plain,(
  ( ! [X26] : (d(lazy_impl(prop(X26),sK7))) ) | ~spl9_203),
  inference(avatar_component_clause,[],[f10009])).
fof(f13488,plain,(
  ( ! [X33] : (~bool(lazy_impl(prop(sK2(sK7,X33)),sK7)) | ~d(lazy_impl(prop(sK2(sK7,X33)),sK7))) ) | ~spl9_256),
  inference(avatar_component_clause,[],[f13487])).
fof(f831493,plain,(
  spl9_845 | ~spl9_458 | ~spl9_1002),
  inference(avatar_split_clause,[],[f831383,f333413,f50398,f235756])).
fof(f235756,plain,(
  spl9_845 <=> ! [X66] : true = impl(false1,X66)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_845])])).
fof(f333413,plain,(
  spl9_1002 <=> ! [X155,X154] : impl(false1,X155) = impl(false1,X154)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_1002])])).
fof(f831383,plain,(
  ( ! [X11] : (true = impl(false1,X11)) ) | (~spl9_458 | ~spl9_1002)),
  inference(superposition,[],[f333414,f50400])).
fof(f333414,plain,(
  ( ! [X154,X155] : (impl(false1,X155) = impl(false1,X154)) ) | ~spl9_1002),
  inference(avatar_component_clause,[],[f333413])).
fof(f795615,plain,(
  spl9_1486 | spl9_395 | spl9_49 | ~spl9_71 | ~spl9_906),
  inference(avatar_split_clause,[],[f795614,f314292,f2188,f1713,f27929,f788650])).
fof(f27929,plain,(
  spl9_395 <=> ! [X1] : ~forallprefers(lazy_impl(prop(X1),err),true)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_395])])).
fof(f2188,plain,(
  spl9_71 <=> err = not1(sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_71])])).
fof(f314292,plain,(
  spl9_906 <=> err = impl(false1,sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_906])])).
fof(f795614,plain,(
  ( ! [X26,X25] : (~forallprefers(lazy_impl(prop(X25),err),true) | true = prop(sK2(sK7,X26))) ) | (spl9_49 | ~spl9_71 | ~spl9_906)),
  inference(forward_demodulation,[],[f795510,f1864])).
fof(f795510,plain,(
  ( ! [X26,X25] : (~forallprefers(lazy_impl(prop(X25),impl(err,impl(impl(X26,X25),X25))),true) | true = prop(sK2(sK7,X26))) ) | (~spl9_71 | ~spl9_906)),
  inference(superposition,[],[f1501,f788750])).
fof(f788750,plain,(
  ( ! [X2] : (err = impl(sK7,X2)) ) | (~spl9_71 | ~spl9_906)),
  inference(forward_demodulation,[],[f788160,f2189])).
fof(f2189,plain,(
  err = not1(sK7) | ~spl9_71),
  inference(avatar_component_clause,[],[f2188])).
fof(f788160,plain,(
  ( ! [X2] : (not1(sK7) = impl(sK7,X2)) ) | ~spl9_906),
  inference(subsumption_resolution,[],[f788107,f135])).
fof(f788107,plain,(
  ( ! [X2] : (true = err | not1(sK7) = impl(sK7,X2)) ) | ~spl9_906),
  inference(superposition,[],[f2567,f314293])).
fof(f314293,plain,(
  err = impl(false1,sK7) | ~spl9_906),
  inference(avatar_component_clause,[],[f314292])).
fof(f2567,plain,(
  ( ! [X10,X11] : (impl(X10,X11) = not1(X10) | true = impl(false1,X10)) )),
  inference(duplicate_literal_removal,[],[f2497])).
fof(f2497,plain,(
  ( ! [X10,X11] : (impl(X10,X11) = not1(X10) | true = impl(false1,X10) | true = impl(false1,X10)) )),
  inference(superposition,[],[f296,f288])).
fof(f296,plain,(
  ( ! [X10,X11] : (impl(X10,X11) = lazy_impl(true,X10) | true = impl(false1,X10)) )),
  inference(resolution,[],[f217,f219])).
fof(f1501,plain,(
  ( ! [X2,X0,X1] : (~forallprefers(lazy_impl(prop(X2),impl(impl(X0,X2),impl(impl(X1,X2),X2))),true) | true = prop(sK2(X0,X1))) )),
  inference(forward_demodulation,[],[f1483,f221])).
fof(f1483,plain,(
  ( ! [X2,X0,X1] : (~forallprefers(lazy_impl(prop(X2),impl(impl(X0,X2),impl(impl(X1,X2),X2))),lazy_impl(false1,impl(impl(X0,sK2(X0,X1)),impl(impl(X1,sK2(X0,X1)),sK2(X0,X1))))) | true = prop(sK2(X0,X1))) )),
  inference(superposition,[],[f233,f252])).
fof(f789608,plain,(
  ~spl9_48 | ~spl9_114 | spl9_901),
  inference(avatar_contradiction_clause,[],[f789607])).
fof(f789607,plain,(
  $false | (~spl9_48 | ~spl9_114 | spl9_901)),
  inference(subsumption_resolution,[],[f789540,f6942])).
fof(f6942,plain,(
  err = or1(true,err) | ~spl9_114),
  inference(forward_demodulation,[],[f6922,f268])).
fof(f6922,plain,(
  lazy_impl(true,err) = or1(true,err) | ~spl9_114),
  inference(resolution,[],[f6849,f546])).
fof(f546,plain,(
  ( ! [X0] : (bool(X0) | lazy_impl(true,X0) = or1(true,X0)) )),
  inference(resolution,[],[f231,f242])).
fof(f231,plain,(
  ( ! [X0,X1] : (~bool(X0) | bool(X1) | or1(X0,X1) = lazy_impl(true,X1)) )),
  inference(definition_unfolding,[],[f175,f160])).
fof(f175,plain,(
  ( ! [X0,X1] : (phi(X1) = or1(X0,X1) | bool(X1) | ~bool(X0)) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : (phi(X1) = or1(X0,X1) | bool(X1) | ~bool(X0))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : (phi(X1) = or1(X0,X1) | (bool(X1) | ~bool(X0)))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : ((~bool(X1) & bool(X0)) => phi(X1) = or1(X0,X1))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X2,X3] : ((~bool(X3) & bool(X2)) => phi(X3) = or1(X2,X3))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f789540,plain,(
  err != or1(true,err) | (~spl9_48 | spl9_901)),
  inference(backward_demodulation,[],[f1582,f311838])).
fof(f311838,plain,(
  err != or1(true,sK8) | spl9_901),
  inference(avatar_component_clause,[],[f311836])).
fof(f1582,plain,(
  err = sK8 | ~spl9_48),
  inference(avatar_component_clause,[],[f1581])).
fof(f1581,plain,(
  spl9_48 <=> err = sK8),
  introduced(avatar_definition,[new_symbols(naming,[spl9_48])])).
fof(f789478,plain,(
  spl9_48 | ~spl9_111 | ~spl9_113),
  inference(avatar_split_clause,[],[f789476,f4117,f4108,f1581])).
fof(f789476,plain,(
  err = sK8 | (~spl9_111 | ~spl9_113)),
  inference(backward_demodulation,[],[f4110,f4119])).
fof(f4110,plain,(
  err = lazy_impl(true,sK8) | ~spl9_111),
  inference(avatar_component_clause,[],[f4108])).
fof(f788726,plain,(
  spl9_69 | ~spl9_173 | ~spl9_1486),
  inference(avatar_contradiction_clause,[],[f788725])).
fof(f788725,plain,(
  $false | (spl9_69 | ~spl9_173 | ~spl9_1486)),
  inference(subsumption_resolution,[],[f788705,f2169])).
fof(f788705,plain,(
  true = lazy_impl(true,sK7) | (~spl9_173 | ~spl9_1486)),
  inference(backward_demodulation,[],[f788651,f8417])).
fof(f8417,plain,(
  true = lazy_impl(prop(sK2(sK7,sK8)),sK7) | ~spl9_173),
  inference(avatar_component_clause,[],[f8415])).
fof(f8415,plain,(
  spl9_173 <=> true = lazy_impl(prop(sK2(sK7,sK8)),sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_173])])).
fof(f788651,plain,(
  ( ! [X25] : (true = prop(sK2(sK7,X25))) ) | ~spl9_1486),
  inference(avatar_component_clause,[],[f788650])).
fof(f778948,plain,(
  ~spl9_13 | ~spl9_18 | ~spl9_32 | spl9_49 | spl9_66 | ~spl9_71),
  inference(avatar_contradiction_clause,[],[f778947])).
fof(f778947,plain,(
  $false | (~spl9_13 | ~spl9_18 | ~spl9_32 | spl9_49 | spl9_66 | ~spl9_71)),
  inference(subsumption_resolution,[],[f778946,f659])).
fof(f659,plain,(
  err = or1(sK7,sK8) | ~spl9_18),
  inference(avatar_component_clause,[],[f658])).
fof(f658,plain,(
  spl9_18 <=> err = or1(sK7,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_18])])).
fof(f778946,plain,(
  err != or1(sK7,sK8) | (~spl9_13 | ~spl9_32 | spl9_49 | spl9_66 | ~spl9_71)),
  inference(forward_demodulation,[],[f778945,f268])).
fof(f778945,plain,(
  or1(sK7,sK8) != lazy_impl(true,err) | (~spl9_13 | ~spl9_32 | spl9_49 | spl9_66 | ~spl9_71)),
  inference(forward_demodulation,[],[f778944,f268])).
fof(f778944,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,err)) | (~spl9_13 | ~spl9_32 | spl9_49 | spl9_66 | ~spl9_71)),
  inference(forward_demodulation,[],[f778943,f1864])).
fof(f778943,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(err,impl(impl(sK8,true),true)))) | (~spl9_13 | ~spl9_32 | spl9_66 | ~spl9_71)),
  inference(forward_demodulation,[],[f778942,f2189])).
fof(f778942,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(not1(sK7),impl(impl(sK8,true),true)))) | (~spl9_13 | ~spl9_32 | spl9_66)),
  inference(forward_demodulation,[],[f775925,f1183])).
fof(f1183,plain,(
  true = sK2(sK7,sK8) | ~spl9_32),
  inference(avatar_component_clause,[],[f1181])).
fof(f775925,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(not1(sK7),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | (~spl9_13 | spl9_66)),
  inference(backward_demodulation,[],[f638,f2076])).
fof(f2076,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(not1(sK7),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | spl9_66),
  inference(avatar_component_clause,[],[f2074])).
fof(f2074,plain,(
  spl9_66 <=> or1(sK7,sK8) = lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(not1(sK7),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_66])])).
fof(f638,plain,(
  true = prop(sK2(sK7,sK8)) | ~spl9_13),
  inference(avatar_component_clause,[],[f636])).
fof(f776971,plain,(
  ~spl9_13 | ~spl9_39 | ~spl9_262 | ~spl9_271 | ~spl9_949),
  inference(avatar_contradiction_clause,[],[f776970])).
fof(f776970,plain,(
  $false | (~spl9_13 | ~spl9_39 | ~spl9_262 | ~spl9_271 | ~spl9_949)),
  inference(subsumption_resolution,[],[f776969,f776176])).
fof(f776176,plain,(
  sK7 = or1(sK7,false1) | (~spl9_39 | ~spl9_262 | ~spl9_949)),
  inference(forward_demodulation,[],[f776175,f318190])).
fof(f318190,plain,(
  sK7 = lazy_impl(true,sK7) | ~spl9_949),
  inference(avatar_component_clause,[],[f318188])).
fof(f318188,plain,(
  spl9_949 <=> sK7 = lazy_impl(true,sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_949])])).
fof(f776175,plain,(
  or1(sK7,false1) = lazy_impl(true,sK7) | (~spl9_39 | ~spl9_262 | ~spl9_949)),
  inference(forward_demodulation,[],[f776174,f318190])).
fof(f776174,plain,(
  or1(sK7,false1) = lazy_impl(true,lazy_impl(true,sK7)) | (~spl9_39 | ~spl9_262)),
  inference(forward_demodulation,[],[f776100,f15327])).
fof(f15327,plain,(
  ( ! [X7] : (sK7 = impl(sK7,X7)) ) | ~spl9_262),
  inference(avatar_component_clause,[],[f15326])).
fof(f776100,plain,(
  or1(sK7,false1) = lazy_impl(true,lazy_impl(true,impl(sK7,sK2(sK7,false1)))) | (~spl9_39 | ~spl9_262)),
  inference(backward_demodulation,[],[f15327,f1219])).
fof(f1219,plain,(
  or1(sK7,false1) = lazy_impl(true,lazy_impl(true,impl(impl(sK7,sK2(sK7,false1)),sK2(sK7,false1)))) | ~spl9_39),
  inference(avatar_component_clause,[],[f1218])).
fof(f1218,plain,(
  spl9_39 <=> or1(sK7,false1) = lazy_impl(true,lazy_impl(true,impl(impl(sK7,sK2(sK7,false1)),sK2(sK7,false1))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_39])])).
fof(f776969,plain,(
  sK7 != or1(sK7,false1) | (~spl9_13 | ~spl9_262 | ~spl9_271 | ~spl9_949)),
  inference(forward_demodulation,[],[f776968,f318190])).
fof(f776968,plain,(
  or1(sK7,false1) != lazy_impl(true,sK7) | (~spl9_13 | ~spl9_262 | ~spl9_271 | ~spl9_949)),
  inference(forward_demodulation,[],[f776967,f318190])).
fof(f776967,plain,(
  or1(sK7,false1) != lazy_impl(true,lazy_impl(true,sK7)) | (~spl9_13 | ~spl9_262 | ~spl9_271)),
  inference(forward_demodulation,[],[f776966,f15327])).
fof(f776966,plain,(
  or1(sK7,false1) != lazy_impl(true,lazy_impl(true,impl(sK7,impl(impl(false1,sK2(sK7,false1)),sK2(sK7,false1))))) | (~spl9_13 | ~spl9_262 | ~spl9_271)),
  inference(forward_demodulation,[],[f776965,f15327])).
fof(f776965,plain,(
  or1(sK7,false1) != lazy_impl(true,lazy_impl(true,impl(impl(sK7,sK2(sK7,false1)),impl(impl(false1,sK2(sK7,false1)),sK2(sK7,false1))))) | (~spl9_13 | ~spl9_271)),
  inference(forward_demodulation,[],[f389908,f15749])).
fof(f389908,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | ~spl9_13),
  inference(backward_demodulation,[],[f638,f241])).
fof(f241,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))))),
  inference(definition_unfolding,[],[f197,f200])).
fof(f200,plain,(
  ( ! [X0,X1] : (or2(X0,X1) = lazy_impl(true,lazy_impl(prop(sK2(X0,X1)),impl(impl(X0,sK2(X0,X1)),impl(impl(X1,sK2(X0,X1)),sK2(X0,X1)))))) )),
  inference(definition_unfolding,[],[f179,f160,f178])).
fof(f179,plain,(
  ( ! [X0,X1] : (or2(X0,X1) = phi(f3(X0,X1,sK2(X0,X1)))) )),
  inference(cnf_transformation,[],[f120])).
fof(f197,plain,(
  or1(sK7,sK8) != or2(sK7,sK8)),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  or1(sK7,sK8) != or2(sK7,sK8)),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f110,f129])).
fof(f129,plain,(
  ? [X0,X1] : or1(X0,X1) != or2(X0,X1) => or1(sK7,sK8) != or2(sK7,sK8)),
  introduced(choice_axiom,[])).
fof(f110,plain,(
  ? [X0,X1] : or1(X0,X1) != or2(X0,X1)),
  inference(ennf_transformation,[],[f80])).
fof(f80,plain,(
  ~! [X0,X1] : or1(X0,X1) = or2(X0,X1)),
  inference(rectify,[],[f46])).
fof(f46,negated_conjecture,(
  ~! [X4,X5] : or2(X4,X5) = or1(X4,X5)),
  inference(negated_conjecture,[],[f45])).
fof(f45,conjecture,(
  ! [X4,X5] : or2(X4,X5) = or1(X4,X5)),
  file('./p/SWW098+1.tptp',unknown)).
fof(f776805,plain,(
  ~spl9_271 | ~spl9_689),
  inference(avatar_contradiction_clause,[],[f776804])).
fof(f776804,plain,(
  $false | (~spl9_271 | ~spl9_689)),
  inference(subsumption_resolution,[],[f776803,f208])).
fof(f776803,plain,(
  true = false1 | (~spl9_271 | ~spl9_689)),
  inference(forward_demodulation,[],[f776685,f221])).
fof(f776685,plain,(
  ( ! [X35] : (false1 = lazy_impl(false1,X35)) ) | (~spl9_271 | ~spl9_689)),
  inference(backward_demodulation,[],[f15749,f113514])).
fof(f113514,plain,(
  ( ! [X35] : (sK8 = lazy_impl(sK8,X35)) ) | ~spl9_689),
  inference(avatar_component_clause,[],[f113513])).
fof(f113513,plain,(
  spl9_689 <=> ! [X35] : sK8 = lazy_impl(sK8,X35)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_689])])).
fof(f776439,plain,(
  ~spl9_13 | ~spl9_33 | spl9_66 | ~spl9_108 | ~spl9_262 | ~spl9_456 | ~spl9_690 | ~spl9_949),
  inference(avatar_contradiction_clause,[],[f776438])).
fof(f776438,plain,(
  $false | (~spl9_13 | ~spl9_33 | spl9_66 | ~spl9_108 | ~spl9_262 | ~spl9_456 | ~spl9_690 | ~spl9_949)),
  inference(subsumption_resolution,[],[f776154,f776408])).
fof(f776408,plain,(
  sK7 = or1(sK7,sK8) | (~spl9_33 | ~spl9_262 | ~spl9_690 | ~spl9_949)),
  inference(forward_demodulation,[],[f776407,f318190])).
fof(f776407,plain,(
  or1(sK7,sK8) = lazy_impl(true,sK7) | (~spl9_33 | ~spl9_262 | ~spl9_690 | ~spl9_949)),
  inference(forward_demodulation,[],[f776406,f318190])).
fof(f776406,plain,(
  or1(sK7,sK8) = lazy_impl(true,lazy_impl(true,sK7)) | (~spl9_33 | ~spl9_262 | ~spl9_690)),
  inference(forward_demodulation,[],[f776405,f15327])).
fof(f776405,plain,(
  or1(sK7,sK8) = lazy_impl(true,lazy_impl(true,impl(sK7,sK8))) | (~spl9_33 | ~spl9_262 | ~spl9_690)),
  inference(forward_demodulation,[],[f776404,f15327])).
fof(f776404,plain,(
  or1(sK7,sK8) = lazy_impl(true,lazy_impl(true,impl(impl(sK7,false1),sK8))) | (~spl9_33 | ~spl9_690)),
  inference(forward_demodulation,[],[f776403,f113519])).
fof(f776403,plain,(
  or1(sK7,sK8) = lazy_impl(true,lazy_impl(true,impl(impl(sK7,false1),impl(sK8,false1)))) | (~spl9_33 | ~spl9_690)),
  inference(forward_demodulation,[],[f1186,f113519])).
fof(f1186,plain,(
  or1(sK7,sK8) = lazy_impl(true,lazy_impl(true,impl(impl(sK7,false1),impl(impl(sK8,false1),false1)))) | ~spl9_33),
  inference(avatar_component_clause,[],[f1185])).
fof(f1185,plain,(
  spl9_33 <=> or1(sK7,sK8) = lazy_impl(true,lazy_impl(true,impl(impl(sK7,false1),impl(impl(sK8,false1),false1))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_33])])).
fof(f776154,plain,(
  sK7 != or1(sK7,sK8) | (~spl9_13 | spl9_66 | ~spl9_108 | ~spl9_262 | ~spl9_456 | ~spl9_949)),
  inference(forward_demodulation,[],[f776153,f318190])).
fof(f776153,plain,(
  or1(sK7,sK8) != lazy_impl(true,sK7) | (~spl9_13 | spl9_66 | ~spl9_108 | ~spl9_262 | ~spl9_456 | ~spl9_949)),
  inference(forward_demodulation,[],[f776087,f318190])).
fof(f776087,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,sK7)) | (~spl9_13 | spl9_66 | ~spl9_108 | ~spl9_262 | ~spl9_456)),
  inference(backward_demodulation,[],[f15327,f776067])).
fof(f776067,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(sK7,impl(not1(sK8),sK2(sK7,sK8))))) | (~spl9_13 | spl9_66 | ~spl9_108 | ~spl9_456)),
  inference(backward_demodulation,[],[f46640,f775988])).
fof(f775988,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(not1(sK7),impl(not1(sK8),sK2(sK7,sK8))))) | (~spl9_13 | spl9_66 | ~spl9_108)),
  inference(backward_demodulation,[],[f4094,f775925])).
fof(f4094,plain,(
  ( ! [X6] : (not1(sK8) = impl(sK8,X6)) ) | ~spl9_108),
  inference(avatar_component_clause,[],[f4093])).
fof(f4093,plain,(
  spl9_108 <=> ! [X6] : not1(sK8) = impl(sK8,X6)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_108])])).
fof(f46640,plain,(
  sK7 = not1(sK7) | ~spl9_456),
  inference(avatar_component_clause,[],[f46638])).
fof(f775638,plain,(
  ~spl9_18 | spl9_49 | spl9_66 | ~spl9_71 | ~spl9_182 | ~spl9_1198),
  inference(avatar_contradiction_clause,[],[f775637])).
fof(f775637,plain,(
  $false | (~spl9_18 | spl9_49 | spl9_66 | ~spl9_71 | ~spl9_182 | ~spl9_1198)),
  inference(subsumption_resolution,[],[f775636,f659])).
fof(f775636,plain,(
  err != or1(sK7,sK8) | (spl9_49 | spl9_66 | ~spl9_71 | ~spl9_182 | ~spl9_1198)),
  inference(forward_demodulation,[],[f775635,f268])).
fof(f775635,plain,(
  or1(sK7,sK8) != lazy_impl(true,err) | (spl9_49 | spl9_66 | ~spl9_71 | ~spl9_182 | ~spl9_1198)),
  inference(forward_demodulation,[],[f775634,f268])).
fof(f775634,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,err)) | (spl9_49 | spl9_66 | ~spl9_71 | ~spl9_182 | ~spl9_1198)),
  inference(forward_demodulation,[],[f775633,f765460])).
fof(f765460,plain,(
  true = prop(false1) | ~spl9_1198),
  inference(avatar_component_clause,[],[f765458])).
fof(f765458,plain,(
  spl9_1198 <=> true = prop(false1)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_1198])])).
fof(f775633,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(false1),err)) | (spl9_49 | spl9_66 | ~spl9_71 | ~spl9_182)),
  inference(forward_demodulation,[],[f775632,f1864])).
fof(f775632,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(false1),impl(err,impl(impl(sK8,false1),false1)))) | (spl9_66 | ~spl9_71 | ~spl9_182)),
  inference(forward_demodulation,[],[f774498,f2189])).
fof(f774498,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(false1),impl(not1(sK7),impl(impl(sK8,false1),false1)))) | (spl9_66 | ~spl9_182)),
  inference(forward_demodulation,[],[f2076,f8569])).
fof(f8569,plain,(
  false1 = sK2(sK7,sK8) | ~spl9_182),
  inference(avatar_component_clause,[],[f8567])).
fof(f774964,plain,(
  ~spl9_114 | spl9_906 | ~spl9_934),
  inference(avatar_contradiction_clause,[],[f774963])).
fof(f774963,plain,(
  $false | (~spl9_114 | spl9_906 | ~spl9_934)),
  inference(subsumption_resolution,[],[f774915,f15641])).
fof(f15641,plain,(
  err = impl(false1,err) | ~spl9_114),
  inference(forward_demodulation,[],[f6919,f268])).
fof(f6919,plain,(
  lazy_impl(true,err) = impl(false1,err) | ~spl9_114),
  inference(resolution,[],[f6849,f525])).
fof(f774915,plain,(
  err != impl(false1,err) | (spl9_906 | ~spl9_934)),
  inference(backward_demodulation,[],[f317928,f314294])).
fof(f314294,plain,(
  err != impl(false1,sK7) | spl9_906),
  inference(avatar_component_clause,[],[f314292])).
fof(f317928,plain,(
  err = sK7 | ~spl9_934),
  inference(avatar_component_clause,[],[f317926])).
fof(f317926,plain,(
  spl9_934 <=> err = sK7),
  introduced(avatar_definition,[new_symbols(naming,[spl9_934])])).
fof(f774731,plain,(
  ~spl9_53 | spl9_59 | ~spl9_690),
  inference(avatar_contradiction_clause,[],[f774730])).
fof(f774730,plain,(
  $false | (~spl9_53 | spl9_59 | ~spl9_690)),
  inference(subsumption_resolution,[],[f774729,f1898])).
fof(f1898,plain,(
  true != prop(sK8) | spl9_59),
  inference(avatar_component_clause,[],[f1897])).
fof(f1897,plain,(
  spl9_59 <=> true = prop(sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_59])])).
fof(f774729,plain,(
  true = prop(sK8) | (~spl9_53 | ~spl9_690)),
  inference(forward_demodulation,[],[f1872,f113519])).
fof(f1872,plain,(
  true = prop(impl(sK8,sK2(sK7,sK8))) | ~spl9_53),
  inference(avatar_component_clause,[],[f1870])).
fof(f1870,plain,(
  spl9_53 <=> true = prop(impl(sK8,sK2(sK7,sK8)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_53])])).
fof(f774271,plain,(
  spl9_79 | spl9_70),
  inference(avatar_split_clause,[],[f390663,f2172,f2598])).
fof(f2598,plain,(
  spl9_79 <=> sK7 = and1(true,sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_79])])).
fof(f390663,plain,(
  sK7 = and1(true,sK7) | spl9_70),
  inference(subsumption_resolution,[],[f368706,f292726])).
fof(f292726,plain,(
  ( ! [X0,X1] : (and1(true,X0) = X0 | err = impl(X0,X1)) )),
  inference(subsumption_resolution,[],[f2329,f75215])).
fof(f75215,plain,(
  ( ! [X0] : (lazy_impl(true,X0) != X0 | and1(true,X0) = X0) )),
  inference(equality_factoring,[],[f1610])).
fof(f1610,plain,(
  ( ! [X2] : (lazy_impl(true,X2) = and1(true,X2) | and1(true,X2) = X2) )),
  inference(resolution,[],[f535,f164])).
fof(f164,plain,(
  ( ! [X0] : (~bool(X0) | and1(true,X0) = X0) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : (and1(true,X0) = X0 | ~bool(X0))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : (bool(X0) => and1(true,X0) = X0)),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X3] : (bool(X3) => and1(true,X3) = X3)),
  file('./p/SWW098+1.tptp',unknown)).
fof(f535,plain,(
  ( ! [X0] : (bool(X0) | lazy_impl(true,X0) = and1(true,X0)) )),
  inference(resolution,[],[f223,f242])).
fof(f223,plain,(
  ( ! [X0,X1] : (~bool(X0) | bool(X1) | and1(X0,X1) = lazy_impl(true,X1)) )),
  inference(definition_unfolding,[],[f162,f160])).
fof(f162,plain,(
  ( ! [X0,X1] : (phi(X1) = and1(X0,X1) | bool(X1) | ~bool(X0)) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1] : (phi(X1) = and1(X0,X1) | bool(X1) | ~bool(X0))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : (phi(X1) = and1(X0,X1) | (bool(X1) | ~bool(X0)))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ((~bool(X1) & bool(X0)) => phi(X1) = and1(X0,X1))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X2,X3] : ((~bool(X3) & bool(X2)) => phi(X3) = and1(X2,X3))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f2329,plain,(
  ( ! [X0,X1] : (and1(true,X0) = X0 | err = impl(X0,X1) | lazy_impl(true,X0) = X0) )),
  inference(superposition,[],[f293,f212])).
fof(f293,plain,(
  ( ! [X4,X5] : (lazy_impl(true,X4) = impl(X4,X5) | and1(true,X4) = X4) )),
  inference(resolution,[],[f217,f164])).
fof(f368706,plain,(
  ( ! [X2] : (err != impl(sK7,X2) | sK7 = and1(true,sK7)) ) | spl9_70),
  inference(superposition,[],[f2174,f293])).
fof(f773262,plain,(
  ~spl9_107 | spl9_64 | ~spl9_73 | ~spl9_92 | ~spl9_95 | ~spl9_97 | ~spl9_98 | ~spl9_108 | ~spl9_112 | ~spl9_182 | ~spl9_918 | ~spl9_1198),
  inference(avatar_split_clause,[],[f773222,f765458,f314918,f8567,f4112,f4093,f3177,f3172,f3163,f3150,f2198,f2064,f4089])).
fof(f4089,plain,(
  spl9_107 <=> sK8 = or1(false1,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_107])])).
fof(f2064,plain,(
  spl9_64 <=> or1(sK7,sK8) = lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),not1(impl(sK8,sK2(sK7,sK8))))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_64])])).
fof(f3150,plain,(
  spl9_92 <=> impl(false1,impl(not1(sK8),false1)) = lazy_impl(true,impl(false1,impl(not1(sK8),false1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_92])])).
fof(f3172,plain,(
  spl9_97 <=> not1(sK8) = not1(impl(false1,impl(not1(sK8),false1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_97])])).
fof(f3177,plain,(
  spl9_98 <=> impl(false1,impl(not1(sK8),false1)) = impl(true,impl(false1,impl(not1(sK8),false1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_98])])).
fof(f773222,plain,(
  sK8 != or1(false1,sK8) | (spl9_64 | ~spl9_73 | ~spl9_92 | ~spl9_95 | ~spl9_97 | ~spl9_98 | ~spl9_108 | ~spl9_112 | ~spl9_182 | ~spl9_918 | ~spl9_1198)),
  inference(backward_demodulation,[],[f4114,f772987])).
fof(f772987,plain,(
  or1(false1,sK8) != not1(sK8) | (spl9_64 | ~spl9_73 | ~spl9_92 | ~spl9_95 | ~spl9_97 | ~spl9_98 | ~spl9_108 | ~spl9_182 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772986,f771918])).
fof(f771918,plain,(
  not1(sK8) = lazy_impl(true,not1(sK8)) | (~spl9_92 | ~spl9_95)),
  inference(backward_demodulation,[],[f3164,f3152])).
fof(f3152,plain,(
  impl(false1,impl(not1(sK8),false1)) = lazy_impl(true,impl(false1,impl(not1(sK8),false1))) | ~spl9_92),
  inference(avatar_component_clause,[],[f3150])).
fof(f772986,plain,(
  or1(false1,sK8) != lazy_impl(true,not1(sK8)) | (spl9_64 | ~spl9_73 | ~spl9_92 | ~spl9_95 | ~spl9_97 | ~spl9_98 | ~spl9_108 | ~spl9_182 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772985,f771918])).
fof(f772985,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(true,not1(sK8))) | (spl9_64 | ~spl9_73 | ~spl9_95 | ~spl9_97 | ~spl9_98 | ~spl9_108 | ~spl9_182 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772984,f771917])).
fof(f771917,plain,(
  not1(sK8) = impl(true,not1(sK8)) | (~spl9_95 | ~spl9_98)),
  inference(backward_demodulation,[],[f3164,f3179])).
fof(f3179,plain,(
  impl(false1,impl(not1(sK8),false1)) = impl(true,impl(false1,impl(not1(sK8),false1))) | ~spl9_98),
  inference(avatar_component_clause,[],[f3177])).
fof(f772984,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(true,impl(true,not1(sK8)))) | (spl9_64 | ~spl9_73 | ~spl9_95 | ~spl9_97 | ~spl9_108 | ~spl9_182 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772971,f771927])).
fof(f771927,plain,(
  not1(sK8) = not1(not1(sK8)) | (~spl9_95 | ~spl9_97)),
  inference(forward_demodulation,[],[f3173,f3164])).
fof(f3173,plain,(
  not1(sK8) = not1(impl(false1,impl(not1(sK8),false1))) | ~spl9_97),
  inference(avatar_component_clause,[],[f3172])).
fof(f772971,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(true,impl(true,not1(not1(sK8))))) | (spl9_64 | ~spl9_73 | ~spl9_108 | ~spl9_182 | ~spl9_918 | ~spl9_1198)),
  inference(backward_demodulation,[],[f4094,f772965])).
fof(f772965,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(true,impl(true,not1(impl(sK8,false1))))) | (spl9_64 | ~spl9_73 | ~spl9_182 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772964,f765460])).
fof(f772964,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(prop(false1),impl(true,not1(impl(sK8,false1))))) | (spl9_64 | ~spl9_73 | ~spl9_182 | ~spl9_918)),
  inference(forward_demodulation,[],[f772963,f390071])).
fof(f390071,plain,(
  true = impl(false1,false1) | (~spl9_73 | ~spl9_918)),
  inference(backward_demodulation,[],[f314920,f2200])).
fof(f2200,plain,(
  true = impl(false1,sK7) | ~spl9_73),
  inference(avatar_component_clause,[],[f2198])).
fof(f772963,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(prop(false1),impl(impl(false1,false1),not1(impl(sK8,false1))))) | (spl9_64 | ~spl9_182 | ~spl9_918)),
  inference(forward_demodulation,[],[f772962,f772454])).
fof(f772454,plain,(
  false1 = sK2(false1,sK8) | (~spl9_182 | ~spl9_918)),
  inference(forward_demodulation,[],[f8569,f314920])).
fof(f772962,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(prop(sK2(false1,sK8)),impl(impl(false1,sK2(false1,sK8)),not1(impl(sK8,sK2(false1,sK8)))))) | (spl9_64 | ~spl9_918)),
  inference(forward_demodulation,[],[f2066,f314920])).
fof(f2066,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),not1(impl(sK8,sK2(sK7,sK8)))))) | spl9_64),
  inference(avatar_component_clause,[],[f2064])).
fof(f772634,plain,(
  ~spl9_35 | spl9_49 | ~spl9_50 | ~spl9_73 | ~spl9_114 | spl9_116 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198),
  inference(avatar_contradiction_clause,[],[f772633])).
fof(f772633,plain,(
  $false | (~spl9_35 | spl9_49 | ~spl9_50 | ~spl9_73 | ~spl9_114 | spl9_116 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198)),
  inference(subsumption_resolution,[],[f772632,f5310])).
fof(f5310,plain,(
  true != impl(false1,err) | spl9_116),
  inference(avatar_component_clause,[],[f5309])).
fof(f5309,plain,(
  spl9_116 <=> true = impl(false1,err)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_116])])).
fof(f772632,plain,(
  true = impl(false1,err) | (~spl9_35 | spl9_49 | ~spl9_50 | ~spl9_73 | ~spl9_114 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772631,f268])).
fof(f772631,plain,(
  true = impl(false1,lazy_impl(true,err)) | (~spl9_35 | spl9_49 | ~spl9_50 | ~spl9_73 | ~spl9_114 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772630,f765460])).
fof(f772630,plain,(
  true = impl(false1,lazy_impl(prop(false1),err)) | (~spl9_35 | spl9_49 | ~spl9_50 | ~spl9_73 | ~spl9_114 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772629,f15576])).
fof(f772629,plain,(
  true = impl(false1,lazy_impl(prop(false1),impl(true,err))) | (~spl9_35 | spl9_49 | ~spl9_50 | ~spl9_73 | ~spl9_114 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772628,f390071])).
fof(f772628,plain,(
  true = impl(false1,lazy_impl(prop(false1),impl(impl(false1,false1),err))) | (~spl9_35 | spl9_49 | ~spl9_50 | ~spl9_73 | ~spl9_114 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772627,f1864])).
fof(f772627,plain,(
  true = impl(false1,lazy_impl(prop(false1),impl(impl(false1,false1),impl(err,false1)))) | (~spl9_35 | spl9_49 | ~spl9_50 | ~spl9_73 | ~spl9_114 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772626,f15603])).
fof(f772626,plain,(
  true = impl(false1,lazy_impl(prop(false1),impl(impl(false1,false1),impl(impl(sK8,false1),false1)))) | (~spl9_35 | spl9_49 | ~spl9_50 | ~spl9_73 | ~spl9_114 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772625,f772454])).
fof(f772625,plain,(
  true = impl(false1,lazy_impl(prop(sK2(false1,sK8)),impl(impl(false1,sK2(false1,sK8)),impl(impl(sK8,sK2(false1,sK8)),sK2(false1,sK8))))) | (~spl9_35 | spl9_49 | ~spl9_50 | ~spl9_73 | ~spl9_114 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772624,f314920])).
fof(f772624,plain,(
  true = impl(false1,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | (~spl9_35 | spl9_49 | ~spl9_50 | ~spl9_73 | ~spl9_114 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198)),
  inference(subsumption_resolution,[],[f772623,f1195])).
fof(f1195,plain,(
  err = or1(false1,sK8) | ~spl9_35),
  inference(avatar_component_clause,[],[f1194])).
fof(f1194,plain,(
  spl9_35 <=> err = or1(false1,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_35])])).
fof(f772623,plain,(
  err != or1(false1,sK8) | true = impl(false1,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | (spl9_49 | ~spl9_50 | ~spl9_73 | ~spl9_114 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772622,f1719])).
fof(f1719,plain,(
  err = not1(err) | ~spl9_50),
  inference(avatar_component_clause,[],[f1717])).
fof(f1717,plain,(
  spl9_50 <=> err = not1(err)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_50])])).
fof(f772622,plain,(
  or1(false1,sK8) != not1(err) | true = impl(false1,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | (spl9_49 | ~spl9_73 | ~spl9_114 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772621,f268])).
fof(f772621,plain,(
  or1(false1,sK8) != not1(lazy_impl(true,err)) | true = impl(false1,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | (spl9_49 | ~spl9_73 | ~spl9_114 | ~spl9_182 | ~spl9_266 | ~spl9_918 | ~spl9_1198)),
  inference(forward_demodulation,[],[f772620,f765460])).
fof(f772620,plain,(
  or1(false1,sK8) != not1(lazy_impl(prop(false1),err)) | true = impl(false1,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | (spl9_49 | ~spl9_73 | ~spl9_114 | ~spl9_182 | ~spl9_266 | ~spl9_918)),
  inference(forward_demodulation,[],[f772619,f15576])).
fof(f772619,plain,(
  or1(false1,sK8) != not1(lazy_impl(prop(false1),impl(true,err))) | true = impl(false1,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | (spl9_49 | ~spl9_73 | ~spl9_182 | ~spl9_266 | ~spl9_918)),
  inference(forward_demodulation,[],[f772618,f390071])).
fof(f772618,plain,(
  or1(false1,sK8) != not1(lazy_impl(prop(false1),impl(impl(false1,false1),err))) | true = impl(false1,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | (spl9_49 | ~spl9_182 | ~spl9_266 | ~spl9_918)),
  inference(forward_demodulation,[],[f772237,f772454])).
fof(f772237,plain,(
  or1(false1,sK8) != not1(lazy_impl(prop(sK2(false1,sK8)),impl(impl(false1,sK2(false1,sK8)),err))) | true = impl(false1,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | (spl9_49 | ~spl9_266 | ~spl9_918)),
  inference(forward_demodulation,[],[f772236,f1864])).
fof(f772236,plain,(
  or1(false1,sK8) != not1(lazy_impl(prop(sK2(false1,sK8)),impl(impl(false1,sK2(false1,sK8)),impl(err,sK2(false1,sK8))))) | true = impl(false1,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | (~spl9_266 | ~spl9_918)),
  inference(forward_demodulation,[],[f770005,f15603])).
fof(f770005,plain,(
  or1(false1,sK8) != not1(lazy_impl(prop(sK2(false1,sK8)),impl(impl(false1,sK2(false1,sK8)),impl(impl(sK8,sK2(false1,sK8)),sK2(false1,sK8))))) | true = impl(false1,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | ~spl9_918),
  inference(forward_demodulation,[],[f1797,f314920])).
fof(f1797,plain,(
  or1(sK7,sK8) != not1(lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | true = impl(false1,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))))),
  inference(superposition,[],[f241,f288])).
fof(f772161,plain,(
  ~spl9_35 | ~spl9_13 | ~spl9_32 | spl9_49 | spl9_67 | ~spl9_93 | ~spl9_114 | ~spl9_458 | ~spl9_918),
  inference(avatar_split_clause,[],[f772160,f314918,f50398,f4966,f3154,f2079,f1713,f1181,f636,f1194])).
fof(f2079,plain,(
  spl9_67 <=> or1(sK7,sK8) = lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(not1(sK8),sK2(sK7,sK8)))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_67])])).
fof(f3154,plain,(
  spl9_93 <=> err = not1(sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_93])])).
fof(f772160,plain,(
  err != or1(false1,sK8) | (~spl9_13 | ~spl9_32 | spl9_49 | spl9_67 | ~spl9_93 | ~spl9_114 | ~spl9_458 | ~spl9_918)),
  inference(forward_demodulation,[],[f772159,f268])).
fof(f772159,plain,(
  lazy_impl(true,err) != or1(false1,sK8) | (~spl9_13 | ~spl9_32 | spl9_49 | spl9_67 | ~spl9_93 | ~spl9_114 | ~spl9_458 | ~spl9_918)),
  inference(forward_demodulation,[],[f772158,f268])).
fof(f772158,plain,(
  lazy_impl(true,lazy_impl(true,err)) != or1(false1,sK8) | (~spl9_13 | ~spl9_32 | spl9_49 | spl9_67 | ~spl9_93 | ~spl9_114 | ~spl9_458 | ~spl9_918)),
  inference(forward_demodulation,[],[f772157,f15576])).
fof(f772157,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(true,impl(true,err))) | (~spl9_13 | ~spl9_32 | spl9_49 | spl9_67 | ~spl9_93 | ~spl9_458 | ~spl9_918)),
  inference(forward_demodulation,[],[f772135,f1864])).
fof(f772135,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(true,impl(true,impl(err,true)))) | (~spl9_13 | ~spl9_32 | spl9_67 | ~spl9_93 | ~spl9_458 | ~spl9_918)),
  inference(backward_demodulation,[],[f3155,f769876])).
fof(f769876,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(true,impl(true,impl(not1(sK8),true)))) | (~spl9_13 | ~spl9_32 | spl9_67 | ~spl9_458 | ~spl9_918)),
  inference(forward_demodulation,[],[f769875,f50400])).
fof(f769875,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(true,impl(impl(false1,true),impl(not1(sK8),true)))) | (~spl9_13 | ~spl9_32 | spl9_67 | ~spl9_918)),
  inference(forward_demodulation,[],[f769874,f769792])).
fof(f769792,plain,(
  true = sK2(false1,sK8) | (~spl9_32 | ~spl9_918)),
  inference(forward_demodulation,[],[f1183,f314920])).
fof(f769874,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(true,impl(impl(false1,sK2(false1,sK8)),impl(not1(sK8),sK2(false1,sK8))))) | (~spl9_13 | spl9_67 | ~spl9_918)),
  inference(forward_demodulation,[],[f389918,f314920])).
fof(f389918,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(impl(sK7,sK2(sK7,sK8)),impl(not1(sK8),sK2(sK7,sK8))))) | (~spl9_13 | spl9_67)),
  inference(backward_demodulation,[],[f638,f2081])).
fof(f2081,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(not1(sK8),sK2(sK7,sK8))))) | spl9_67),
  inference(avatar_component_clause,[],[f2079])).
fof(f3155,plain,(
  err = not1(sK8) | ~spl9_93),
  inference(avatar_component_clause,[],[f3154])).
fof(f771836,plain,(
  ~spl9_32 | spl9_64 | ~spl9_107 | ~spl9_112 | ~spl9_113 | ~spl9_265 | ~spl9_458 | ~spl9_690 | ~spl9_918),
  inference(avatar_contradiction_clause,[],[f771835])).
fof(f771835,plain,(
  $false | (~spl9_32 | spl9_64 | ~spl9_107 | ~spl9_112 | ~spl9_113 | ~spl9_265 | ~spl9_458 | ~spl9_690 | ~spl9_918)),
  inference(subsumption_resolution,[],[f771834,f4091])).
fof(f4091,plain,(
  sK8 = or1(false1,sK8) | ~spl9_107),
  inference(avatar_component_clause,[],[f4089])).
fof(f771834,plain,(
  sK8 != or1(false1,sK8) | (~spl9_32 | spl9_64 | ~spl9_112 | ~spl9_113 | ~spl9_265 | ~spl9_458 | ~spl9_690 | ~spl9_918)),
  inference(forward_demodulation,[],[f771833,f4119])).
fof(f771833,plain,(
  or1(false1,sK8) != lazy_impl(true,sK8) | (~spl9_32 | spl9_64 | ~spl9_112 | ~spl9_113 | ~spl9_265 | ~spl9_458 | ~spl9_690 | ~spl9_918)),
  inference(forward_demodulation,[],[f771832,f4119])).
fof(f771832,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(true,sK8)) | (~spl9_32 | spl9_64 | ~spl9_112 | ~spl9_265 | ~spl9_458 | ~spl9_690 | ~spl9_918)),
  inference(forward_demodulation,[],[f771831,f248])).
fof(f771831,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(prop(true),sK8)) | (~spl9_32 | spl9_64 | ~spl9_112 | ~spl9_265 | ~spl9_458 | ~spl9_690 | ~spl9_918)),
  inference(forward_demodulation,[],[f771830,f15600])).
fof(f771830,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(prop(true),impl(true,sK8))) | (~spl9_32 | spl9_64 | ~spl9_112 | ~spl9_458 | ~spl9_690 | ~spl9_918)),
  inference(forward_demodulation,[],[f771829,f50400])).
fof(f771829,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(prop(true),impl(impl(false1,true),sK8))) | (~spl9_32 | spl9_64 | ~spl9_112 | ~spl9_690 | ~spl9_918)),
  inference(forward_demodulation,[],[f771828,f769792])).
fof(f771828,plain,(
  or1(false1,sK8) != lazy_impl(true,lazy_impl(prop(sK2(false1,sK8)),impl(impl(false1,sK2(false1,sK8)),sK8))) | (spl9_64 | ~spl9_112 | ~spl9_690 | ~spl9_918)),
  inference(forward_demodulation,[],[f771827,f314920])).
fof(f771827,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),sK8))) | (spl9_64 | ~spl9_112 | ~spl9_690)),
  inference(forward_demodulation,[],[f771826,f4114])).
fof(f771826,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),not1(sK8)))) | (spl9_64 | ~spl9_690)),
  inference(forward_demodulation,[],[f2066,f113519])).
fof(f771645,plain,(
  ~spl9_1084 | spl9_1403),
  inference(avatar_contradiction_clause,[],[f771644])).
fof(f771644,plain,(
  $false | (~spl9_1084 | spl9_1403)),
  inference(subsumption_resolution,[],[f771643,f245])).
fof(f245,plain,(
  forallprefers(false1,true)),
  inference(equality_resolution,[],[f244])).
fof(f244,plain,(
  ( ! [X0] : (forallprefers(X0,true) | false1 != X0) )),
  inference(equality_resolution,[],[f210])).
fof(f210,plain,(
  ( ! [X0,X1] : (forallprefers(X0,X1) | true != X1 | false1 != X0) )),
  inference(definition_unfolding,[],[f142,f189])).
fof(f142,plain,(
  ( ! [X0,X1] : (forallprefers(X0,X1) | true != X1 | false != X0) )),
  inference(cnf_transformation,[],[f83])).
fof(f771643,plain,(
  ~forallprefers(false1,true) | (~spl9_1084 | spl9_1403)),
  inference(forward_demodulation,[],[f767923,f391031])).
fof(f391031,plain,(
  true = lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1)))) | ~spl9_1084),
  inference(avatar_component_clause,[],[f391029])).
fof(f391029,plain,(
  spl9_1084 <=> true = lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_1084])])).
fof(f767923,plain,(
  ~forallprefers(false1,lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1))))) | spl9_1403),
  inference(avatar_component_clause,[],[f767921])).
fof(f767921,plain,(
  spl9_1403 <=> forallprefers(false1,lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1)))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_1403])])).
fof(f770395,plain,(
  spl9_690 | ~spl9_108 | ~spl9_112),
  inference(avatar_split_clause,[],[f770379,f4112,f4093,f113518])).
fof(f770379,plain,(
  ( ! [X6] : (sK8 = impl(sK8,X6)) ) | (~spl9_108 | ~spl9_112)),
  inference(backward_demodulation,[],[f4114,f4094])).
fof(f770246,plain,(
  ~spl9_87 | ~spl9_458 | spl9_608),
  inference(avatar_contradiction_clause,[],[f770245])).
fof(f770245,plain,(
  $false | (~spl9_87 | ~spl9_458 | spl9_608)),
  inference(subsumption_resolution,[],[f770244,f267])).
fof(f770244,plain,(
  true != lazy_impl(true,true) | (~spl9_87 | ~spl9_458 | spl9_608)),
  inference(forward_demodulation,[],[f770243,f50400])).
fof(f770243,plain,(
  true != lazy_impl(true,impl(false1,true)) | (~spl9_87 | spl9_608)),
  inference(forward_demodulation,[],[f770242,f240])).
fof(f770242,plain,(
  true != lazy_impl(true,impl(not1(true),true)) | (~spl9_87 | spl9_608)),
  inference(forward_demodulation,[],[f87104,f2864])).
fof(f87104,plain,(
  true != lazy_impl(true,impl(not1(sK2(true,sK8)),true)) | spl9_608),
  inference(avatar_component_clause,[],[f87103])).
fof(f87103,plain,(
  spl9_608 <=> true = lazy_impl(true,impl(not1(sK2(true,sK8)),true))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_608])])).
fof(f770111,plain,(
  spl9_1364),
  inference(avatar_split_clause,[],[f239,f767192])).
fof(f239,plain,(
  true = not1(false1)),
  inference(definition_unfolding,[],[f194,f189])).
fof(f194,plain,(
  true = not1(false)),
  inference(cnf_transformation,[],[f42])).
fof(f42,axiom,(
  true = not1(false)),
  file('./p/SWW098+1.tptp',unknown)).
fof(f767930,plain,(
  spl9_1002 | ~spl9_1403 | ~spl9_846 | ~spl9_847),
  inference(avatar_split_clause,[],[f767929,f252859,f252855,f767921,f333413])).
fof(f252855,plain,(
  spl9_846 <=> false1 = lazy_impl(true,false1)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_846])])).
fof(f252859,plain,(
  spl9_847 <=> false1 = impl(true,false1)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_847])])).
fof(f767929,plain,(
  ( ! [X196,X197] : (~forallprefers(false1,lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1))))) | impl(false1,X197) = impl(false1,X196)) ) | (~spl9_846 | ~spl9_847)),
  inference(forward_demodulation,[],[f767928,f252856])).
fof(f252856,plain,(
  false1 = lazy_impl(true,false1) | ~spl9_846),
  inference(avatar_component_clause,[],[f252855])).
fof(f767928,plain,(
  ( ! [X196,X197] : (~forallprefers(lazy_impl(true,false1),lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1))))) | impl(false1,X197) = impl(false1,X196)) ) | ~spl9_847),
  inference(forward_demodulation,[],[f756054,f252861])).
fof(f252861,plain,(
  false1 = impl(true,false1) | ~spl9_847),
  inference(avatar_component_clause,[],[f252859])).
fof(f756054,plain,(
  ( ! [X196,X197] : (~forallprefers(lazy_impl(true,impl(true,false1)),lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1))))) | impl(false1,X197) = impl(false1,X196)) )),
  inference(superposition,[],[f1500,f2566])).
fof(f2566,plain,(
  ( ! [X26,X24,X25] : (impl(X24,X25) = impl(X24,X26) | true = impl(false1,X24)) )),
  inference(duplicate_literal_removal,[],[f2502])).
fof(f2502,plain,(
  ( ! [X26,X24,X25] : (impl(X24,X25) = impl(X24,X26) | true = impl(false1,X24) | true = impl(false1,X24)) )),
  inference(superposition,[],[f296,f296])).
fof(f1500,plain,(
  ( ! [X3] : (~forallprefers(lazy_impl(true,impl(impl(X3,false1),false1)),lazy_impl(prop(sK2(X3,false1)),impl(impl(X3,sK2(X3,false1)),impl(impl(false1,sK2(X3,false1)),sK2(X3,false1)))))) )),
  inference(forward_demodulation,[],[f1499,f249])).
fof(f249,plain,(
  true = prop(false1)),
  inference(resolution,[],[f151,f243])).
fof(f1499,plain,(
  ( ! [X3] : (~forallprefers(lazy_impl(prop(false1),impl(impl(X3,false1),false1)),lazy_impl(prop(sK2(X3,false1)),impl(impl(X3,sK2(X3,false1)),impl(impl(false1,sK2(X3,false1)),sK2(X3,false1)))))) )),
  inference(forward_demodulation,[],[f1482,f258])).
fof(f258,plain,(
  false1 = impl(true,false1)),
  inference(resolution,[],[f157,f243])).
fof(f157,plain,(
  ( ! [X0] : (~bool(X0) | impl(true,X0) = X0) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0] : (impl(true,X0) = X0 | ~bool(X0))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (bool(X0) => impl(true,X0) = X0)),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X3] : (bool(X3) => impl(true,X3) = X3)),
  file('./p/SWW098+1.tptp',unknown)).
fof(f1482,plain,(
  ( ! [X3] : (~forallprefers(lazy_impl(prop(false1),impl(impl(X3,false1),impl(true,false1))),lazy_impl(prop(sK2(X3,false1)),impl(impl(X3,sK2(X3,false1)),impl(impl(false1,sK2(X3,false1)),sK2(X3,false1)))))) )),
  inference(superposition,[],[f233,f271])).
fof(f271,plain,(
  true = impl(false1,false1)),
  inference(resolution,[],[f219,f243])).
fof(f767766,plain,(
  spl9_1198),
  inference(avatar_split_clause,[],[f249,f765458])).
fof(f756939,plain,(
  spl9_1002 | ~spl9_32 | ~spl9_271 | ~spl9_458 | ~spl9_846 | ~spl9_847 | ~spl9_918),
  inference(avatar_split_clause,[],[f756938,f314918,f252859,f252855,f50398,f15747,f1181,f333413])).
fof(f756938,plain,(
  ( ! [X196,X197] : (impl(false1,X197) = impl(false1,X196)) ) | (~spl9_32 | ~spl9_271 | ~spl9_458 | ~spl9_846 | ~spl9_847 | ~spl9_918)),
  inference(subsumption_resolution,[],[f756937,f245])).
fof(f756937,plain,(
  ( ! [X196,X197] : (~forallprefers(false1,true) | impl(false1,X197) = impl(false1,X196)) ) | (~spl9_32 | ~spl9_271 | ~spl9_458 | ~spl9_846 | ~spl9_847 | ~spl9_918)),
  inference(forward_demodulation,[],[f756936,f252856])).
fof(f756936,plain,(
  ( ! [X196,X197] : (~forallprefers(lazy_impl(true,false1),true) | impl(false1,X197) = impl(false1,X196)) ) | (~spl9_32 | ~spl9_271 | ~spl9_458 | ~spl9_847 | ~spl9_918)),
  inference(forward_demodulation,[],[f756935,f252861])).
fof(f756935,plain,(
  ( ! [X196,X197] : (~forallprefers(lazy_impl(true,impl(true,false1)),true) | impl(false1,X197) = impl(false1,X196)) ) | (~spl9_32 | ~spl9_271 | ~spl9_458 | ~spl9_918)),
  inference(forward_demodulation,[],[f756934,f267])).
fof(f756934,plain,(
  ( ! [X196,X197] : (~forallprefers(lazy_impl(true,impl(true,false1)),lazy_impl(true,true)) | impl(false1,X197) = impl(false1,X196)) ) | (~spl9_32 | ~spl9_271 | ~spl9_458 | ~spl9_918)),
  inference(forward_demodulation,[],[f756933,f248])).
fof(f756933,plain,(
  ( ! [X196,X197] : (~forallprefers(lazy_impl(true,impl(true,false1)),lazy_impl(prop(true),true)) | impl(false1,X197) = impl(false1,X196)) ) | (~spl9_32 | ~spl9_271 | ~spl9_458 | ~spl9_918)),
  inference(forward_demodulation,[],[f756932,f257])).
fof(f257,plain,(
  true = impl(true,true)),
  inference(resolution,[],[f157,f242])).
fof(f756932,plain,(
  ( ! [X196,X197] : (~forallprefers(lazy_impl(true,impl(true,false1)),lazy_impl(prop(true),impl(true,true))) | impl(false1,X197) = impl(false1,X196)) ) | (~spl9_32 | ~spl9_271 | ~spl9_458 | ~spl9_918)),
  inference(forward_demodulation,[],[f756931,f257])).
fof(f756931,plain,(
  ( ! [X196,X197] : (~forallprefers(lazy_impl(true,impl(true,false1)),lazy_impl(prop(true),impl(true,impl(true,true)))) | impl(false1,X197) = impl(false1,X196)) ) | (~spl9_32 | ~spl9_271 | ~spl9_458 | ~spl9_918)),
  inference(forward_demodulation,[],[f756930,f50400])).
fof(f756930,plain,(
  ( ! [X196,X197] : (~forallprefers(lazy_impl(true,impl(true,false1)),lazy_impl(prop(true),impl(impl(false1,true),impl(impl(false1,true),true)))) | impl(false1,X197) = impl(false1,X196)) ) | (~spl9_32 | ~spl9_271 | ~spl9_918)),
  inference(forward_demodulation,[],[f756054,f391457])).
fof(f391457,plain,(
  true = sK2(false1,false1) | (~spl9_32 | ~spl9_271 | ~spl9_918)),
  inference(forward_demodulation,[],[f391456,f314920])).
fof(f391456,plain,(
  true = sK2(sK7,false1) | (~spl9_32 | ~spl9_271)),
  inference(forward_demodulation,[],[f1183,f15749])).
fof(f391251,plain,(
  ~spl9_13 | ~spl9_73 | ~spl9_182 | ~spl9_271 | ~spl9_406 | ~spl9_846 | ~spl9_847 | ~spl9_918),
  inference(avatar_contradiction_clause,[],[f391250])).
fof(f391250,plain,(
  $false | (~spl9_13 | ~spl9_73 | ~spl9_182 | ~spl9_271 | ~spl9_406 | ~spl9_846 | ~spl9_847 | ~spl9_918)),
  inference(subsumption_resolution,[],[f391249,f252856])).
fof(f391249,plain,(
  false1 != lazy_impl(true,false1) | (~spl9_13 | ~spl9_73 | ~spl9_182 | ~spl9_271 | ~spl9_406 | ~spl9_846 | ~spl9_847 | ~spl9_918)),
  inference(forward_demodulation,[],[f391248,f252856])).
fof(f391248,plain,(
  false1 != lazy_impl(true,lazy_impl(true,false1)) | (~spl9_13 | ~spl9_73 | ~spl9_182 | ~spl9_271 | ~spl9_406 | ~spl9_847 | ~spl9_918)),
  inference(forward_demodulation,[],[f391247,f252861])).
fof(f391247,plain,(
  false1 != lazy_impl(true,lazy_impl(true,impl(true,false1))) | (~spl9_13 | ~spl9_73 | ~spl9_182 | ~spl9_271 | ~spl9_406 | ~spl9_847 | ~spl9_918)),
  inference(forward_demodulation,[],[f391246,f252861])).
fof(f391246,plain,(
  false1 != lazy_impl(true,lazy_impl(true,impl(true,impl(true,false1)))) | (~spl9_13 | ~spl9_73 | ~spl9_182 | ~spl9_271 | ~spl9_406 | ~spl9_918)),
  inference(forward_demodulation,[],[f391163,f390071])).
fof(f391163,plain,(
  false1 != lazy_impl(true,lazy_impl(true,impl(impl(false1,false1),impl(impl(false1,false1),false1)))) | (~spl9_13 | ~spl9_182 | ~spl9_271 | ~spl9_406 | ~spl9_918)),
  inference(backward_demodulation,[],[f391149,f391087])).
fof(f391087,plain,(
  false1 != lazy_impl(true,lazy_impl(true,impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1))))) | (~spl9_13 | ~spl9_271 | ~spl9_406 | ~spl9_918)),
  inference(forward_demodulation,[],[f391086,f34988])).
fof(f391086,plain,(
  or1(false1,false1) != lazy_impl(true,lazy_impl(true,impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1))))) | (~spl9_13 | ~spl9_271 | ~spl9_918)),
  inference(forward_demodulation,[],[f391085,f314920])).
fof(f391085,plain,(
  or1(sK7,false1) != lazy_impl(true,lazy_impl(true,impl(impl(sK7,sK2(sK7,false1)),impl(impl(false1,sK2(sK7,false1)),sK2(sK7,false1))))) | (~spl9_13 | ~spl9_271)),
  inference(forward_demodulation,[],[f389908,f15749])).
fof(f391149,plain,(
  false1 = sK2(false1,false1) | (~spl9_182 | ~spl9_271 | ~spl9_918)),
  inference(forward_demodulation,[],[f391148,f314920])).
fof(f391148,plain,(
  false1 = sK2(sK7,false1) | (~spl9_182 | ~spl9_271)),
  inference(forward_demodulation,[],[f8569,f15749])).
fof(f391032,plain,(
  spl9_1083 | spl9_1084 | ~spl9_52 | ~spl9_271 | ~spl9_406 | ~spl9_918),
  inference(avatar_split_clause,[],[f391023,f314918,f34986,f15747,f1727,f391029,f391025])).
fof(f391023,plain,(
  true = lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1)))) | false1 = lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1)))) | (~spl9_52 | ~spl9_271 | ~spl9_406 | ~spl9_918)),
  inference(forward_demodulation,[],[f391022,f314920])).
fof(f391022,plain,(
  true = lazy_impl(prop(sK2(sK7,false1)),impl(impl(sK7,sK2(sK7,false1)),impl(impl(false1,sK2(sK7,false1)),sK2(sK7,false1)))) | false1 = lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1)))) | (~spl9_52 | ~spl9_271 | ~spl9_406 | ~spl9_918)),
  inference(forward_demodulation,[],[f391021,f15749])).
fof(f391021,plain,(
  false1 = lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1)))) | true = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))) | (~spl9_52 | ~spl9_271 | ~spl9_406 | ~spl9_918)),
  inference(forward_demodulation,[],[f391020,f314920])).
fof(f391020,plain,(
  false1 = lazy_impl(prop(sK2(sK7,false1)),impl(impl(sK7,sK2(sK7,false1)),impl(impl(false1,sK2(sK7,false1)),sK2(sK7,false1)))) | true = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))) | (~spl9_52 | ~spl9_271 | ~spl9_406 | ~spl9_918)),
  inference(forward_demodulation,[],[f391019,f15749])).
fof(f391019,plain,(
  false1 = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))) | true = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))) | (~spl9_52 | ~spl9_271 | ~spl9_406 | ~spl9_918)),
  inference(subsumption_resolution,[],[f391018,f34988])).
fof(f391018,plain,(
  false1 != or1(false1,false1) | false1 = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))) | true = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))) | (~spl9_52 | ~spl9_271 | ~spl9_406 | ~spl9_918)),
  inference(forward_demodulation,[],[f390629,f390966])).
fof(f390966,plain,(
  false1 = not1(lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1))))) | (~spl9_52 | ~spl9_271 | ~spl9_406 | ~spl9_918)),
  inference(forward_demodulation,[],[f390935,f34988])).
fof(f390935,plain,(
  or1(false1,false1) = not1(lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1))))) | (~spl9_52 | ~spl9_271 | ~spl9_918)),
  inference(forward_demodulation,[],[f390934,f314920])).
fof(f390934,plain,(
  or1(sK7,false1) = not1(lazy_impl(prop(sK2(sK7,false1)),impl(impl(sK7,sK2(sK7,false1)),impl(impl(false1,sK2(sK7,false1)),sK2(sK7,false1))))) | (~spl9_52 | ~spl9_271)),
  inference(forward_demodulation,[],[f1728,f15749])).
fof(f390629,plain,(
  or1(false1,false1) != not1(lazy_impl(prop(sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),impl(impl(false1,sK2(false1,false1)),sK2(false1,false1))))) | false1 = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))) | true = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))) | (~spl9_271 | ~spl9_918)),
  inference(forward_demodulation,[],[f390628,f314920])).
fof(f390628,plain,(
  or1(sK7,false1) != not1(lazy_impl(prop(sK2(sK7,false1)),impl(impl(sK7,sK2(sK7,false1)),impl(impl(false1,sK2(sK7,false1)),sK2(sK7,false1))))) | false1 = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))) | true = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))) | ~spl9_271),
  inference(forward_demodulation,[],[f2223,f15749])).
fof(f2223,plain,(
  or1(sK7,sK8) != not1(lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | false1 = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))) | true = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))),
  inference(superposition,[],[f241,f287])).
fof(f287,plain,(
  ( ! [X4] : (lazy_impl(true,X4) = not1(X4) | false1 = X4 | true = X4) )),
  inference(resolution,[],[f238,f206])).
fof(f390588,plain,(
  spl9_42 | ~spl9_271 | ~spl9_406),
  inference(avatar_contradiction_clause,[],[f390587])).
fof(f390587,plain,(
  $false | (spl9_42 | ~spl9_271 | ~spl9_406)),
  inference(subsumption_resolution,[],[f390586,f34988])).
fof(f390586,plain,(
  false1 != or1(false1,false1) | (spl9_42 | ~spl9_271)),
  inference(forward_demodulation,[],[f1243,f15749])).
fof(f1243,plain,(
  false1 != or1(false1,sK8) | spl9_42),
  inference(avatar_component_clause,[],[f1241])).
fof(f378072,plain,(
  ~spl9_456 | ~spl9_918),
  inference(avatar_contradiction_clause,[],[f378071])).
fof(f378071,plain,(
  $false | (~spl9_456 | ~spl9_918)),
  inference(subsumption_resolution,[],[f378070,f208])).
fof(f378070,plain,(
  true = false1 | (~spl9_456 | ~spl9_918)),
  inference(forward_demodulation,[],[f377883,f239])).
fof(f377883,plain,(
  false1 = not1(false1) | (~spl9_456 | ~spl9_918)),
  inference(backward_demodulation,[],[f314920,f46640])).
fof(f377399,plain,(
  spl9_32 | ~spl9_14 | spl9_182),
  inference(avatar_split_clause,[],[f377398,f8567,f640,f1181])).
fof(f640,plain,(
  spl9_14 <=> true = or1(sK7,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_14])])).
fof(f377398,plain,(
  true != or1(sK7,sK8) | true = sK2(sK7,sK8) | spl9_182),
  inference(subsumption_resolution,[],[f1693,f8568])).
fof(f1693,plain,(
  true != or1(sK7,sK8) | true = sK2(sK7,sK8) | false1 = sK2(sK7,sK8)),
  inference(forward_demodulation,[],[f1692,f267])).
fof(f1692,plain,(
  or1(sK7,sK8) != lazy_impl(true,true) | true = sK2(sK7,sK8) | false1 = sK2(sK7,sK8)),
  inference(forward_demodulation,[],[f1674,f221])).
fof(f1674,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(false1,impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | true = sK2(sK7,sK8) | false1 = sK2(sK7,sK8)),
  inference(superposition,[],[f241,f282])).
fof(f377345,plain,(
  ~spl9_112 | ~spl9_590 | spl9_606),
  inference(avatar_split_clause,[],[f377344,f87093,f87010,f4112])).
fof(f87010,plain,(
  spl9_590 <=> sK8 = lazy_impl(true,impl(not1(sK2(true,sK8)),true))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_590])])).
fof(f377344,plain,(
  sK8 != not1(sK8) | (~spl9_590 | spl9_606)),
  inference(forward_demodulation,[],[f87095,f87011])).
fof(f87011,plain,(
  sK8 = lazy_impl(true,impl(not1(sK2(true,sK8)),true)) | ~spl9_590),
  inference(avatar_component_clause,[],[f87010])).
fof(f87095,plain,(
  sK8 != not1(lazy_impl(true,impl(not1(sK2(true,sK8)),true))) | spl9_606),
  inference(avatar_component_clause,[],[f87093])).
fof(f326598,plain,(
  spl9_906 | ~spl9_70 | spl9_79),
  inference(avatar_split_clause,[],[f326587,f2598,f2172,f314292])).
fof(f326587,plain,(
  err = impl(false1,sK7) | (~spl9_70 | spl9_79)),
  inference(forward_demodulation,[],[f326573,f2173])).
fof(f2173,plain,(
  err = lazy_impl(true,sK7) | ~spl9_70),
  inference(avatar_component_clause,[],[f2172])).
fof(f326573,plain,(
  lazy_impl(true,sK7) = impl(false1,sK7) | spl9_79),
  inference(resolution,[],[f326530,f525])).
fof(f326530,plain,(
  ~bool(sK7) | spl9_79),
  inference(trivial_inequality_removal,[],[f326522])).
fof(f326522,plain,(
  false1 != false1 | ~bool(sK7) | spl9_79),
  inference(superposition,[],[f216,f326315])).
fof(f326315,plain,(
  false1 = prop(sK7) | spl9_79),
  inference(trivial_inequality_removal,[],[f326307])).
fof(f326307,plain,(
  sK7 != sK7 | false1 = prop(sK7) | spl9_79),
  inference(superposition,[],[f2599,f262])).
fof(f262,plain,(
  ( ! [X0] : (and1(true,X0) = X0 | prop(X0) = false1) )),
  inference(resolution,[],[f164,f215])).
fof(f2599,plain,(
  sK7 != and1(true,sK7) | spl9_79),
  inference(avatar_component_clause,[],[f2598])).
fof(f324470,plain,(
  spl9_71 | spl9_34 | ~spl9_70 | spl9_918),
  inference(avatar_split_clause,[],[f324469,f314918,f2172,f1190,f2188])).
fof(f324469,plain,(
  err = not1(sK7) | (spl9_34 | ~spl9_70 | spl9_918)),
  inference(subsumption_resolution,[],[f324468,f1191])).
fof(f324468,plain,(
  err = not1(sK7) | true = sK7 | (~spl9_70 | spl9_918)),
  inference(subsumption_resolution,[],[f324330,f314919])).
fof(f314919,plain,(
  false1 != sK7 | spl9_918),
  inference(avatar_component_clause,[],[f314918])).
fof(f324330,plain,(
  err = not1(sK7) | false1 = sK7 | true = sK7 | ~spl9_70),
  inference(superposition,[],[f2173,f287])).
fof(f323777,plain,(
  spl9_34 | ~spl9_73 | ~spl9_907),
  inference(avatar_contradiction_clause,[],[f323776])).
fof(f323776,plain,(
  $false | (spl9_34 | ~spl9_73 | ~spl9_907)),
  inference(subsumption_resolution,[],[f323775,f1191])).
fof(f323775,plain,(
  true = sK7 | (~spl9_73 | ~spl9_907)),
  inference(backward_demodulation,[],[f2200,f314299])).
fof(f314299,plain,(
  sK7 = impl(false1,sK7) | ~spl9_907),
  inference(avatar_component_clause,[],[f314297])).
fof(f314297,plain,(
  spl9_907 <=> sK7 = impl(false1,sK7)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_907])])).
fof(f323614,plain,(
  spl9_73 | spl9_232 | ~spl9_846),
  inference(avatar_contradiction_clause,[],[f323613])).
fof(f323613,plain,(
  $false | (spl9_73 | spl9_232 | ~spl9_846)),
  inference(subsumption_resolution,[],[f323612,f243])).
fof(f323612,plain,(
  ~bool(false1) | (spl9_73 | spl9_232 | ~spl9_846)),
  inference(forward_demodulation,[],[f323611,f252856])).
fof(f323611,plain,(
  ~bool(lazy_impl(true,false1)) | (spl9_73 | spl9_232)),
  inference(forward_demodulation,[],[f11756,f323008])).
fof(f323008,plain,(
  false1 = prop(sK7) | spl9_73),
  inference(trivial_inequality_removal,[],[f323007])).
fof(f323007,plain,(
  true != true | false1 = prop(sK7) | spl9_73),
  inference(superposition,[],[f2199,f272])).
fof(f272,plain,(
  ( ! [X0] : (true = impl(false1,X0) | prop(X0) = false1) )),
  inference(resolution,[],[f219,f215])).
fof(f2199,plain,(
  true != impl(false1,sK7) | spl9_73),
  inference(avatar_component_clause,[],[f2198])).
fof(f11756,plain,(
  ~bool(lazy_impl(true,prop(sK7))) | spl9_232),
  inference(avatar_component_clause,[],[f11755])).
fof(f11755,plain,(
  spl9_232 <=> bool(lazy_impl(true,prop(sK7)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_232])])).
fof(f323610,plain,(
  spl9_73 | spl9_420 | ~spl9_846),
  inference(avatar_contradiction_clause,[],[f323609])).
fof(f323609,plain,(
  $false | (spl9_73 | spl9_420 | ~spl9_846)),
  inference(subsumption_resolution,[],[f323608,f252856])).
fof(f323608,plain,(
  false1 != lazy_impl(true,false1) | (spl9_73 | spl9_420)),
  inference(forward_demodulation,[],[f39733,f323008])).
fof(f39733,plain,(
  false1 != lazy_impl(true,prop(sK7)) | spl9_420),
  inference(avatar_component_clause,[],[f39732])).
fof(f39732,plain,(
  spl9_420 <=> false1 = lazy_impl(true,prop(sK7))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_420])])).
fof(f323409,plain,(
  ~spl9_233 | ~spl9_262 | ~spl9_420 | spl9_916),
  inference(avatar_contradiction_clause,[],[f323408])).
fof(f323408,plain,(
  $false | (~spl9_233 | ~spl9_262 | ~spl9_420 | spl9_916)),
  inference(subsumption_resolution,[],[f322210,f323191])).
fof(f323191,plain,(
  ( ! [X4] : (lazy_impl(true,sK7) = or1(sK7,X4)) ) | (~spl9_233 | ~spl9_420)),
  inference(resolution,[],[f322465,f230])).
fof(f230,plain,(
  ( ! [X0,X1] : (bool(X0) | lazy_impl(true,X0) = or1(X0,X1)) )),
  inference(definition_unfolding,[],[f174,f160])).
fof(f174,plain,(
  ( ! [X0,X1] : (phi(X0) = or1(X0,X1) | bool(X0)) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : (phi(X0) = or1(X0,X1) | bool(X0))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : (~bool(X0) => phi(X0) = or1(X0,X1))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X2,X3] : (~bool(X2) => phi(X2) = or1(X2,X3))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f322465,plain,(
  ~bool(sK7) | (~spl9_233 | ~spl9_420)),
  inference(trivial_inequality_removal,[],[f322463])).
fof(f322463,plain,(
  false1 != false1 | ~bool(sK7) | (~spl9_233 | ~spl9_420)),
  inference(superposition,[],[f216,f321758])).
fof(f321758,plain,(
  false1 = prop(sK7) | (~spl9_233 | ~spl9_420)),
  inference(backward_demodulation,[],[f39734,f11773])).
fof(f11773,plain,(
  prop(sK7) = lazy_impl(true,prop(sK7)) | ~spl9_233),
  inference(avatar_component_clause,[],[f11771])).
fof(f11771,plain,(
  spl9_233 <=> prop(sK7) = lazy_impl(true,prop(sK7))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_233])])).
fof(f39734,plain,(
  false1 = lazy_impl(true,prop(sK7)) | ~spl9_420),
  inference(avatar_component_clause,[],[f39732])).
fof(f322210,plain,(
  lazy_impl(true,sK7) != or1(sK7,true) | (~spl9_262 | spl9_916)),
  inference(forward_demodulation,[],[f322209,f15327])).
fof(f322209,plain,(
  or1(sK7,true) != lazy_impl(true,impl(sK7,true)) | (~spl9_262 | spl9_916)),
  inference(forward_demodulation,[],[f314905,f15327])).
fof(f314905,plain,(
  or1(sK7,true) != lazy_impl(true,impl(impl(sK7,true),true)) | spl9_916),
  inference(avatar_component_clause,[],[f314903])).
fof(f314903,plain,(
  spl9_916 <=> or1(sK7,true) = lazy_impl(true,impl(impl(sK7,true),true))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_916])])).
fof(f322054,plain,(
  ~spl9_13 | ~spl9_37 | ~spl9_262 | ~spl9_916 | ~spl9_949 | spl9_964),
  inference(avatar_contradiction_clause,[],[f322053])).
fof(f322053,plain,(
  $false | (~spl9_13 | ~spl9_37 | ~spl9_262 | ~spl9_916 | ~spl9_949 | spl9_964)),
  inference(subsumption_resolution,[],[f322052,f318190])).
fof(f322052,plain,(
  sK7 != lazy_impl(true,sK7) | (~spl9_13 | ~spl9_37 | ~spl9_262 | ~spl9_916 | ~spl9_949 | spl9_964)),
  inference(forward_demodulation,[],[f322051,f318190])).
fof(f322051,plain,(
  sK7 != lazy_impl(true,lazy_impl(true,sK7)) | (~spl9_13 | ~spl9_37 | ~spl9_262 | ~spl9_916 | ~spl9_949 | spl9_964)),
  inference(forward_demodulation,[],[f322050,f321959])).
fof(f321959,plain,(
  sK7 = or1(sK7,true) | (~spl9_262 | ~spl9_916 | ~spl9_949)),
  inference(forward_demodulation,[],[f321958,f318190])).
fof(f321958,plain,(
  lazy_impl(true,sK7) = or1(sK7,true) | (~spl9_262 | ~spl9_916)),
  inference(forward_demodulation,[],[f321957,f15327])).
fof(f321957,plain,(
  or1(sK7,true) = lazy_impl(true,impl(sK7,true)) | (~spl9_262 | ~spl9_916)),
  inference(forward_demodulation,[],[f314904,f15327])).
fof(f314904,plain,(
  or1(sK7,true) = lazy_impl(true,impl(impl(sK7,true),true)) | ~spl9_916),
  inference(avatar_component_clause,[],[f314903])).
fof(f322050,plain,(
  lazy_impl(true,lazy_impl(true,sK7)) != or1(sK7,true) | (~spl9_13 | ~spl9_37 | ~spl9_262 | ~spl9_949 | spl9_964)),
  inference(forward_demodulation,[],[f322049,f313820])).
fof(f313820,plain,(
  true = prop(sK2(sK7,true)) | (~spl9_13 | ~spl9_37)),
  inference(forward_demodulation,[],[f638,f1209])).
fof(f1209,plain,(
  true = sK8 | ~spl9_37),
  inference(avatar_component_clause,[],[f1207])).
fof(f1207,plain,(
  spl9_37 <=> true = sK8),
  introduced(avatar_definition,[new_symbols(naming,[spl9_37])])).
fof(f322049,plain,(
  or1(sK7,true) != lazy_impl(true,lazy_impl(prop(sK2(sK7,true)),sK7)) | (~spl9_262 | ~spl9_949 | spl9_964)),
  inference(forward_demodulation,[],[f322048,f15327])).
fof(f322048,plain,(
  or1(sK7,true) != lazy_impl(true,lazy_impl(prop(sK2(sK7,true)),impl(sK7,impl(impl(true,sK2(sK7,true)),sK2(sK7,true))))) | (~spl9_949 | spl9_964)),
  inference(forward_demodulation,[],[f321568,f318190])).
fof(f321568,plain,(
  or1(sK7,true) != lazy_impl(true,lazy_impl(prop(sK2(sK7,true)),impl(lazy_impl(true,sK7),impl(impl(true,sK2(sK7,true)),sK2(sK7,true))))) | spl9_964),
  inference(avatar_component_clause,[],[f321566])).
fof(f321566,plain,(
  spl9_964 <=> or1(sK7,true) = lazy_impl(true,lazy_impl(prop(sK2(sK7,true)),impl(lazy_impl(true,sK7),impl(impl(true,sK2(sK7,true)),sK2(sK7,true)))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_964])])).
fof(f322005,plain,(
  ~spl9_13 | ~spl9_37 | spl9_66 | ~spl9_262 | ~spl9_456 | ~spl9_916 | ~spl9_949),
  inference(avatar_contradiction_clause,[],[f322004])).
fof(f322004,plain,(
  $false | (~spl9_13 | ~spl9_37 | spl9_66 | ~spl9_262 | ~spl9_456 | ~spl9_916 | ~spl9_949)),
  inference(subsumption_resolution,[],[f322003,f318190])).
fof(f322003,plain,(
  sK7 != lazy_impl(true,sK7) | (~spl9_13 | ~spl9_37 | spl9_66 | ~spl9_262 | ~spl9_456 | ~spl9_916 | ~spl9_949)),
  inference(forward_demodulation,[],[f322002,f318190])).
fof(f322002,plain,(
  sK7 != lazy_impl(true,lazy_impl(true,sK7)) | (~spl9_13 | ~spl9_37 | spl9_66 | ~spl9_262 | ~spl9_456 | ~spl9_916 | ~spl9_949)),
  inference(forward_demodulation,[],[f321779,f321959])).
fof(f321779,plain,(
  lazy_impl(true,lazy_impl(true,sK7)) != or1(sK7,true) | (~spl9_13 | ~spl9_37 | spl9_66 | ~spl9_262 | ~spl9_456)),
  inference(backward_demodulation,[],[f15327,f321755])).
fof(f321755,plain,(
  or1(sK7,true) != lazy_impl(true,lazy_impl(true,impl(sK7,impl(impl(true,sK2(sK7,true)),sK2(sK7,true))))) | (~spl9_13 | ~spl9_37 | spl9_66 | ~spl9_456)),
  inference(backward_demodulation,[],[f46640,f321690])).
fof(f321690,plain,(
  or1(sK7,true) != lazy_impl(true,lazy_impl(true,impl(not1(sK7),impl(impl(true,sK2(sK7,true)),sK2(sK7,true))))) | (~spl9_13 | ~spl9_37 | spl9_66)),
  inference(backward_demodulation,[],[f313820,f317940])).
fof(f317940,plain,(
  or1(sK7,true) != lazy_impl(true,lazy_impl(prop(sK2(sK7,true)),impl(not1(sK7),impl(impl(true,sK2(sK7,true)),sK2(sK7,true))))) | (~spl9_37 | spl9_66)),
  inference(forward_demodulation,[],[f2076,f1209])).
fof(f321638,plain,(
  ~spl9_964 | ~spl9_37 | spl9_58),
  inference(avatar_split_clause,[],[f314022,f1892,f1207,f321566])).
fof(f1892,plain,(
  spl9_58 <=> or1(sK7,sK8) = lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(lazy_impl(true,sK7),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8)))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_58])])).
fof(f314022,plain,(
  or1(sK7,true) != lazy_impl(true,lazy_impl(prop(sK2(sK7,true)),impl(lazy_impl(true,sK7),impl(impl(true,sK2(sK7,true)),sK2(sK7,true))))) | (~spl9_37 | spl9_58)),
  inference(forward_demodulation,[],[f1894,f1209])).
fof(f1894,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(lazy_impl(true,sK7),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | spl9_58),
  inference(avatar_component_clause,[],[f1892])).
fof(f321563,plain,(
  spl9_14 | ~spl9_37 | ~spl9_520 | ~spl9_918),
  inference(avatar_contradiction_clause,[],[f321562])).
fof(f321562,plain,(
  $false | (spl9_14 | ~spl9_37 | ~spl9_520 | ~spl9_918)),
  inference(subsumption_resolution,[],[f321561,f57346])).
fof(f57346,plain,(
  true = or1(false1,true) | ~spl9_520),
  inference(avatar_component_clause,[],[f57344])).
fof(f57344,plain,(
  spl9_520 <=> true = or1(false1,true)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_520])])).
fof(f321561,plain,(
  true != or1(false1,true) | (spl9_14 | ~spl9_37 | ~spl9_918)),
  inference(forward_demodulation,[],[f321560,f314920])).
fof(f321560,plain,(
  true != or1(sK7,true) | (spl9_14 | ~spl9_37)),
  inference(forward_demodulation,[],[f642,f1209])).
fof(f642,plain,(
  true != or1(sK7,sK8) | spl9_14),
  inference(avatar_component_clause,[],[f640])).
fof(f319142,plain,(
  ~spl9_37 | ~spl9_182 | ~spl9_520 | ~spl9_847 | ~spl9_918),
  inference(avatar_contradiction_clause,[],[f319141])).
fof(f319141,plain,(
  $false | (~spl9_37 | ~spl9_182 | ~spl9_520 | ~spl9_847 | ~spl9_918)),
  inference(subsumption_resolution,[],[f319140,f267])).
fof(f319140,plain,(
  true != lazy_impl(true,true) | (~spl9_37 | ~spl9_182 | ~spl9_520 | ~spl9_847 | ~spl9_918)),
  inference(forward_demodulation,[],[f319139,f57346])).
fof(f319139,plain,(
  lazy_impl(true,true) != or1(false1,true) | (~spl9_37 | ~spl9_182 | ~spl9_847 | ~spl9_918)),
  inference(forward_demodulation,[],[f319138,f267])).
fof(f319138,plain,(
  or1(false1,true) != lazy_impl(true,lazy_impl(true,true)) | (~spl9_37 | ~spl9_182 | ~spl9_847 | ~spl9_918)),
  inference(forward_demodulation,[],[f319137,f249])).
fof(f319137,plain,(
  or1(false1,true) != lazy_impl(true,lazy_impl(prop(false1),true)) | (~spl9_37 | ~spl9_182 | ~spl9_847 | ~spl9_918)),
  inference(forward_demodulation,[],[f319136,f257])).
fof(f319136,plain,(
  or1(false1,true) != lazy_impl(true,lazy_impl(prop(false1),impl(true,true))) | (~spl9_37 | ~spl9_182 | ~spl9_847 | ~spl9_918)),
  inference(forward_demodulation,[],[f319135,f271])).
fof(f319135,plain,(
  or1(false1,true) != lazy_impl(true,lazy_impl(prop(false1),impl(impl(false1,false1),impl(false1,false1)))) | (~spl9_37 | ~spl9_182 | ~spl9_847 | ~spl9_918)),
  inference(forward_demodulation,[],[f319134,f252861])).
fof(f319134,plain,(
  or1(false1,true) != lazy_impl(true,lazy_impl(prop(false1),impl(impl(false1,false1),impl(impl(true,false1),false1)))) | (~spl9_37 | ~spl9_182 | ~spl9_918)),
  inference(forward_demodulation,[],[f319133,f318574])).
fof(f318574,plain,(
  false1 = sK2(false1,true) | (~spl9_37 | ~spl9_182 | ~spl9_918)),
  inference(forward_demodulation,[],[f318573,f314920])).
fof(f318573,plain,(
  false1 = sK2(sK7,true) | (~spl9_37 | ~spl9_182)),
  inference(forward_demodulation,[],[f8569,f1209])).
fof(f319133,plain,(
  or1(false1,true) != lazy_impl(true,lazy_impl(prop(sK2(false1,true)),impl(impl(false1,sK2(false1,true)),impl(impl(true,sK2(false1,true)),sK2(false1,true))))) | (~spl9_37 | ~spl9_918)),
  inference(forward_demodulation,[],[f314249,f314920])).
fof(f314249,plain,(
  or1(sK7,true) != lazy_impl(true,lazy_impl(prop(sK2(sK7,true)),impl(impl(sK7,sK2(sK7,true)),impl(impl(true,sK2(sK7,true)),sK2(sK7,true))))) | ~spl9_37),
  inference(forward_demodulation,[],[f241,f1209])).
fof(f318192,plain,(
  spl9_57 | spl9_949 | ~spl9_79),
  inference(avatar_split_clause,[],[f46523,f2598,f318188,f1888])).
fof(f46523,plain,(
  sK7 = lazy_impl(true,sK7) | true = prop(sK7) | ~spl9_79),
  inference(superposition,[],[f1608,f2600])).
fof(f2600,plain,(
  sK7 = and1(true,sK7) | ~spl9_79),
  inference(avatar_component_clause,[],[f2598])).
fof(f1608,plain,(
  ( ! [X0] : (lazy_impl(true,X0) = and1(true,X0) | true = prop(X0)) )),
  inference(resolution,[],[f535,f151])).
fof(f317931,plain,(
  spl9_57 | spl9_934 | ~spl9_70 | ~spl9_79),
  inference(avatar_split_clause,[],[f317930,f2598,f2172,f317926,f1888])).
fof(f317930,plain,(
  err = sK7 | true = prop(sK7) | (~spl9_70 | ~spl9_79)),
  inference(forward_demodulation,[],[f46543,f2173])).
fof(f46543,plain,(
  sK7 = lazy_impl(true,sK7) | true = prop(sK7) | ~spl9_79),
  inference(superposition,[],[f2600,f1608])).
fof(f317922,plain,(
  spl9_34 | ~spl9_57 | spl9_918),
  inference(avatar_contradiction_clause,[],[f317921])).
fof(f317921,plain,(
  $false | (spl9_34 | ~spl9_57 | spl9_918)),
  inference(subsumption_resolution,[],[f317920,f1191])).
fof(f317920,plain,(
  true = sK7 | (~spl9_57 | spl9_918)),
  inference(subsumption_resolution,[],[f317919,f314919])).
fof(f317919,plain,(
  false1 = sK7 | true = sK7 | ~spl9_57),
  inference(resolution,[],[f206,f316925])).
fof(f316925,plain,(
  bool(sK7) | ~spl9_57),
  inference(trivial_inequality_removal,[],[f316924])).
fof(f316924,plain,(
  true != true | bool(sK7) | ~spl9_57),
  inference(superposition,[],[f150,f1890])).
fof(f1890,plain,(
  true = prop(sK7) | ~spl9_57),
  inference(avatar_component_clause,[],[f1888])).
fof(f315168,plain,(
  ~spl9_14 | ~spl9_32 | ~spl9_37 | ~spl9_458 | ~spl9_918),
  inference(avatar_contradiction_clause,[],[f315167])).
fof(f315167,plain,(
  $false | (~spl9_14 | ~spl9_32 | ~spl9_37 | ~spl9_458 | ~spl9_918)),
  inference(subsumption_resolution,[],[f315166,f267])).
fof(f315166,plain,(
  true != lazy_impl(true,true) | (~spl9_14 | ~spl9_32 | ~spl9_37 | ~spl9_458 | ~spl9_918)),
  inference(forward_demodulation,[],[f315165,f267])).
fof(f315165,plain,(
  true != lazy_impl(true,lazy_impl(true,true)) | (~spl9_14 | ~spl9_32 | ~spl9_37 | ~spl9_458 | ~spl9_918)),
  inference(forward_demodulation,[],[f315164,f257])).
fof(f315164,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(true,true))) | (~spl9_14 | ~spl9_32 | ~spl9_37 | ~spl9_458 | ~spl9_918)),
  inference(forward_demodulation,[],[f315092,f50400])).
fof(f315092,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(impl(false1,true),true))) | (~spl9_14 | ~spl9_32 | ~spl9_37 | ~spl9_918)),
  inference(backward_demodulation,[],[f314920,f314963])).
fof(f314963,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(impl(sK7,true),true))) | (~spl9_14 | ~spl9_32 | ~spl9_37)),
  inference(backward_demodulation,[],[f313822,f314253])).
fof(f314253,plain,(
  or1(sK7,true) != lazy_impl(true,lazy_impl(true,impl(impl(sK7,true),true))) | (~spl9_32 | ~spl9_37)),
  inference(forward_demodulation,[],[f314252,f248])).
fof(f314252,plain,(
  or1(sK7,true) != lazy_impl(true,lazy_impl(prop(true),impl(impl(sK7,true),true))) | (~spl9_32 | ~spl9_37)),
  inference(forward_demodulation,[],[f314251,f257])).
fof(f314251,plain,(
  or1(sK7,true) != lazy_impl(true,lazy_impl(prop(true),impl(impl(sK7,true),impl(true,true)))) | (~spl9_32 | ~spl9_37)),
  inference(forward_demodulation,[],[f314250,f257])).
fof(f314250,plain,(
  or1(sK7,true) != lazy_impl(true,lazy_impl(prop(true),impl(impl(sK7,true),impl(impl(true,true),true)))) | (~spl9_32 | ~spl9_37)),
  inference(forward_demodulation,[],[f314249,f313727])).
fof(f313727,plain,(
  true = sK2(sK7,true) | (~spl9_32 | ~spl9_37)),
  inference(forward_demodulation,[],[f1183,f1209])).
fof(f313822,plain,(
  true = or1(sK7,true) | (~spl9_14 | ~spl9_37)),
  inference(forward_demodulation,[],[f641,f1209])).
fof(f641,plain,(
  true = or1(sK7,sK8) | ~spl9_14),
  inference(avatar_component_clause,[],[f640])).
fof(f315098,plain,(
  ~spl9_70 | spl9_848 | ~spl9_918),
  inference(avatar_contradiction_clause,[],[f315097])).
fof(f315097,plain,(
  $false | (~spl9_70 | spl9_848 | ~spl9_918)),
  inference(subsumption_resolution,[],[f315058,f252890])).
fof(f252890,plain,(
  err != lazy_impl(true,false1) | spl9_848),
  inference(avatar_component_clause,[],[f252889])).
fof(f252889,plain,(
  spl9_848 <=> err = lazy_impl(true,false1)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_848])])).
fof(f315058,plain,(
  err = lazy_impl(true,false1) | (~spl9_70 | ~spl9_918)),
  inference(backward_demodulation,[],[f314920,f2173])).
fof(f314373,plain,(
  spl9_529 | ~spl9_37 | ~spl9_87),
  inference(avatar_split_clause,[],[f313724,f2862,f1207,f60548])).
fof(f60548,plain,(
  spl9_529 <=> true = sK2(true,true)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_529])])).
fof(f313724,plain,(
  true = sK2(true,true) | (~spl9_37 | ~spl9_87)),
  inference(forward_demodulation,[],[f2864,f1209])).
fof(f314301,plain,(
  spl9_57 | spl9_907 | spl9_70),
  inference(avatar_split_clause,[],[f46417,f2172,f314297,f1888])).
fof(f46417,plain,(
  sK7 = impl(false1,sK7) | true = prop(sK7) | spl9_70),
  inference(superposition,[],[f1597,f7094])).
fof(f7094,plain,(
  sK7 = lazy_impl(true,sK7) | spl9_70),
  inference(trivial_inequality_removal,[],[f7072])).
fof(f7072,plain,(
  err != err | sK7 = lazy_impl(true,sK7) | spl9_70),
  inference(superposition,[],[f2174,f212])).
fof(f1597,plain,(
  ( ! [X0] : (lazy_impl(true,X0) = impl(false1,X0) | true = prop(X0)) )),
  inference(resolution,[],[f525,f151])).
fof(f313633,plain,(
  ~spl9_37 | spl9_83 | ~spl9_84 | ~spl9_85 | ~spl9_529),
  inference(avatar_contradiction_clause,[],[f313632])).
fof(f313632,plain,(
  $false | (~spl9_37 | spl9_83 | ~spl9_84 | ~spl9_85 | ~spl9_529)),
  inference(subsumption_resolution,[],[f313631,f267])).
fof(f313631,plain,(
  true != lazy_impl(true,true) | (~spl9_37 | spl9_83 | ~spl9_84 | ~spl9_85 | ~spl9_529)),
  inference(forward_demodulation,[],[f313630,f267])).
fof(f313630,plain,(
  true != lazy_impl(true,lazy_impl(true,true)) | (~spl9_37 | spl9_83 | ~spl9_84 | ~spl9_85 | ~spl9_529)),
  inference(forward_demodulation,[],[f313629,f257])).
fof(f313629,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(true,true))) | (~spl9_37 | spl9_83 | ~spl9_84 | ~spl9_85 | ~spl9_529)),
  inference(forward_demodulation,[],[f313628,f257])).
fof(f313628,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(true,impl(true,true)))) | (~spl9_37 | spl9_83 | ~spl9_84 | ~spl9_85 | ~spl9_529)),
  inference(forward_demodulation,[],[f313548,f257])).
fof(f313548,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(true,impl(impl(true,true),true)))) | (~spl9_37 | spl9_83 | ~spl9_84 | ~spl9_85 | ~spl9_529)),
  inference(backward_demodulation,[],[f60550,f312324])).
fof(f312324,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(sK2(true,true),impl(impl(true,sK2(true,true)),sK2(true,true))))) | (~spl9_37 | spl9_83 | ~spl9_84 | ~spl9_85)),
  inference(forward_demodulation,[],[f273478,f1209])).
fof(f273478,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(sK2(true,sK8),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8))))) | (spl9_83 | ~spl9_84 | ~spl9_85)),
  inference(backward_demodulation,[],[f2851,f273452])).
fof(f273452,plain,(
  true != lazy_impl(true,lazy_impl(prop(sK2(true,sK8)),impl(sK2(true,sK8),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8))))) | (spl9_83 | ~spl9_85)),
  inference(backward_demodulation,[],[f2854,f2730])).
fof(f2730,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(sK2(true,sK8)),impl(sK2(true,sK8),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8))))) | spl9_83),
  inference(avatar_component_clause,[],[f2728])).
fof(f2728,plain,(
  spl9_83 <=> or1(true,sK8) = lazy_impl(true,lazy_impl(prop(sK2(true,sK8)),impl(sK2(true,sK8),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8)))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_83])])).
fof(f2854,plain,(
  true = or1(true,sK8) | ~spl9_85),
  inference(avatar_component_clause,[],[f2853])).
fof(f2851,plain,(
  true = prop(sK2(true,sK8)) | ~spl9_84),
  inference(avatar_component_clause,[],[f2849])).
fof(f2849,plain,(
  spl9_84 <=> true = prop(sK2(true,sK8))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_84])])).
fof(f60550,plain,(
  true = sK2(true,true) | ~spl9_529),
  inference(avatar_component_clause,[],[f60548])).
fof(f313436,plain,(
  ~spl9_114 | ~spl9_845),
  inference(avatar_contradiction_clause,[],[f313435])).
fof(f313435,plain,(
  $false | (~spl9_114 | ~spl9_845)),
  inference(subsumption_resolution,[],[f312922,f135])).
fof(f312922,plain,(
  true = err | (~spl9_114 | ~spl9_845)),
  inference(backward_demodulation,[],[f235757,f15641])).
fof(f235757,plain,(
  ( ! [X66] : (true = impl(false1,X66)) ) | ~spl9_845),
  inference(avatar_component_clause,[],[f235756])).
fof(f312035,plain,(
  ~spl9_34 | ~spl9_84 | ~spl9_85 | ~spl9_512),
  inference(avatar_contradiction_clause,[],[f312034])).
fof(f312034,plain,(
  $false | (~spl9_34 | ~spl9_84 | ~spl9_85 | ~spl9_512)),
  inference(subsumption_resolution,[],[f312030,f267])).
fof(f312030,plain,(
  true != lazy_impl(true,true) | (~spl9_34 | ~spl9_84 | ~spl9_85 | ~spl9_512)),
  inference(backward_demodulation,[],[f312024,f273479])).
fof(f273479,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8))))) | (~spl9_34 | ~spl9_84 | ~spl9_85)),
  inference(backward_demodulation,[],[f2851,f273453])).
fof(f273453,plain,(
  true != lazy_impl(true,lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8))))) | (~spl9_34 | ~spl9_85)),
  inference(backward_demodulation,[],[f2854,f57231])).
fof(f57231,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8))))) | ~spl9_34),
  inference(forward_demodulation,[],[f241,f1192])).
fof(f1192,plain,(
  true = sK7 | ~spl9_34),
  inference(avatar_component_clause,[],[f1190])).
fof(f312024,plain,(
  true = lazy_impl(true,impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8)))) | (~spl9_84 | ~spl9_512)),
  inference(forward_demodulation,[],[f57259,f2851])).
fof(f57259,plain,(
  true = lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8)))) | ~spl9_512),
  inference(avatar_component_clause,[],[f57257])).
fof(f57257,plain,(
  spl9_512 <=> true = lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_512])])).
fof(f311889,plain,(
  ~spl9_37 | spl9_590 | ~spl9_608),
  inference(avatar_split_clause,[],[f311880,f87103,f87010,f1207])).
fof(f311880,plain,(
  true != sK8 | (spl9_590 | ~spl9_608)),
  inference(backward_demodulation,[],[f87105,f87012])).
fof(f87012,plain,(
  sK8 != lazy_impl(true,impl(not1(sK2(true,sK8)),true)) | spl9_590),
  inference(avatar_component_clause,[],[f87010])).
fof(f87105,plain,(
  true = lazy_impl(true,impl(not1(sK2(true,sK8)),true)) | ~spl9_608),
  inference(avatar_component_clause,[],[f87103])).
fof(f311501,plain,(
  ~spl9_34 | ~spl9_84 | ~spl9_85 | ~spl9_87 | ~spl9_271 | ~spl9_458),
  inference(avatar_contradiction_clause,[],[f311500])).
fof(f311500,plain,(
  $false | (~spl9_34 | ~spl9_84 | ~spl9_85 | ~spl9_87 | ~spl9_271 | ~spl9_458)),
  inference(subsumption_resolution,[],[f311499,f267])).
fof(f311499,plain,(
  true != lazy_impl(true,true) | (~spl9_34 | ~spl9_84 | ~spl9_85 | ~spl9_87 | ~spl9_271 | ~spl9_458)),
  inference(forward_demodulation,[],[f311498,f267])).
fof(f311498,plain,(
  true != lazy_impl(true,lazy_impl(true,true)) | (~spl9_34 | ~spl9_84 | ~spl9_85 | ~spl9_87 | ~spl9_271 | ~spl9_458)),
  inference(forward_demodulation,[],[f311497,f257])).
fof(f311497,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(true,true))) | (~spl9_34 | ~spl9_84 | ~spl9_85 | ~spl9_87 | ~spl9_271 | ~spl9_458)),
  inference(forward_demodulation,[],[f311473,f257])).
fof(f311473,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(true,impl(true,true)))) | (~spl9_34 | ~spl9_84 | ~spl9_85 | ~spl9_87 | ~spl9_271 | ~spl9_458)),
  inference(backward_demodulation,[],[f50400,f311348])).
fof(f311348,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(true,impl(impl(false1,true),true)))) | (~spl9_34 | ~spl9_84 | ~spl9_85 | ~spl9_87 | ~spl9_271)),
  inference(backward_demodulation,[],[f15749,f311290])).
fof(f311290,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(true,impl(impl(sK8,true),true)))) | (~spl9_34 | ~spl9_84 | ~spl9_85 | ~spl9_87)),
  inference(forward_demodulation,[],[f311278,f257])).
fof(f311278,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(impl(true,true),impl(impl(sK8,true),true)))) | (~spl9_34 | ~spl9_84 | ~spl9_85 | ~spl9_87)),
  inference(backward_demodulation,[],[f2864,f273479])).
fof(f292828,plain,(
  spl9_87 | spl9_86 | ~spl9_34 | ~spl9_85),
  inference(avatar_split_clause,[],[f292827,f2853,f1190,f2858,f2862])).
fof(f2858,plain,(
  spl9_86 <=> false1 = sK2(true,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_86])])).
fof(f292827,plain,(
  false1 = sK2(true,sK8) | true = sK2(true,sK8) | (~spl9_34 | ~spl9_85)),
  inference(forward_demodulation,[],[f292826,f1192])).
fof(f292826,plain,(
  true = sK2(true,sK8) | false1 = sK2(sK7,sK8) | (~spl9_34 | ~spl9_85)),
  inference(subsumption_resolution,[],[f61545,f2854])).
fof(f61545,plain,(
  true = sK2(true,sK8) | true != or1(true,sK8) | false1 = sK2(sK7,sK8) | ~spl9_34),
  inference(forward_demodulation,[],[f58572,f1192])).
fof(f58572,plain,(
  true != or1(true,sK8) | true = sK2(sK7,sK8) | false1 = sK2(sK7,sK8) | ~spl9_34),
  inference(forward_demodulation,[],[f1693,f1192])).
fof(f292761,plain,(
  spl9_512 | ~spl9_85 | ~spl9_508),
  inference(avatar_split_clause,[],[f292760,f57237,f2853,f57257])).
fof(f57237,plain,(
  spl9_508 <=> or1(true,sK8) = lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_508])])).
fof(f292760,plain,(
  true = lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8)))) | (~spl9_85 | ~spl9_508)),
  inference(forward_demodulation,[],[f57238,f2854])).
fof(f57238,plain,(
  or1(true,sK8) = lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8)))) | ~spl9_508),
  inference(avatar_component_clause,[],[f57237])).
fof(f273894,plain,(
  ~spl9_37 | spl9_113),
  inference(avatar_contradiction_clause,[],[f273893])).
fof(f273893,plain,(
  $false | (~spl9_37 | spl9_113)),
  inference(subsumption_resolution,[],[f273784,f267])).
fof(f273784,plain,(
  true != lazy_impl(true,true) | (~spl9_37 | spl9_113)),
  inference(backward_demodulation,[],[f1209,f4118])).
fof(f4118,plain,(
  sK8 != lazy_impl(true,sK8) | spl9_113),
  inference(avatar_component_clause,[],[f4117])).
fof(f273630,plain,(
  ~spl9_34 | ~spl9_86 | ~spl9_114 | spl9_133 | ~spl9_152 | ~spl9_847),
  inference(avatar_contradiction_clause,[],[f273629])).
fof(f273629,plain,(
  $false | (~spl9_34 | ~spl9_86 | ~spl9_114 | spl9_133 | ~spl9_152 | ~spl9_847)),
  inference(subsumption_resolution,[],[f273628,f58694])).
fof(f58694,plain,(
  true != lazy_impl(true,err) | (~spl9_114 | spl9_133)),
  inference(forward_demodulation,[],[f5879,f15641])).
fof(f5879,plain,(
  true != lazy_impl(true,impl(false1,err)) | spl9_133),
  inference(avatar_component_clause,[],[f5878])).
fof(f5878,plain,(
  spl9_133 <=> true = lazy_impl(true,impl(false1,err))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_133])])).
fof(f273628,plain,(
  true = lazy_impl(true,err) | (~spl9_34 | ~spl9_86 | ~spl9_114 | ~spl9_152 | ~spl9_847)),
  inference(forward_demodulation,[],[f273627,f249])).
fof(f273627,plain,(
  true = lazy_impl(prop(false1),err) | (~spl9_34 | ~spl9_86 | ~spl9_114 | ~spl9_152 | ~spl9_847)),
  inference(forward_demodulation,[],[f273626,f15641])).
fof(f273626,plain,(
  true = lazy_impl(prop(false1),impl(false1,err)) | (~spl9_34 | ~spl9_86 | ~spl9_152 | ~spl9_847)),
  inference(forward_demodulation,[],[f273625,f252861])).
fof(f273625,plain,(
  true = lazy_impl(prop(false1),impl(impl(true,false1),err)) | (~spl9_34 | ~spl9_86 | ~spl9_152)),
  inference(forward_demodulation,[],[f65521,f2860])).
fof(f2860,plain,(
  false1 = sK2(true,sK8) | ~spl9_86),
  inference(avatar_component_clause,[],[f2858])).
fof(f65521,plain,(
  true = lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),err)) | (~spl9_34 | ~spl9_152)),
  inference(forward_demodulation,[],[f7251,f1192])).
fof(f7251,plain,(
  true = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),err)) | ~spl9_152),
  inference(avatar_component_clause,[],[f7249])).
fof(f7249,plain,(
  spl9_152 <=> true = lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),err))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_152])])).
fof(f272497,plain,(
  ~spl9_34 | ~spl9_182 | ~spl9_271 | ~spl9_847),
  inference(avatar_contradiction_clause,[],[f272496])).
fof(f272496,plain,(
  $false | (~spl9_34 | ~spl9_182 | ~spl9_271 | ~spl9_847)),
  inference(subsumption_resolution,[],[f272495,f267])).
fof(f272495,plain,(
  true != lazy_impl(true,true) | (~spl9_34 | ~spl9_182 | ~spl9_271 | ~spl9_847)),
  inference(forward_demodulation,[],[f272494,f267])).
fof(f272494,plain,(
  true != lazy_impl(true,lazy_impl(true,true)) | (~spl9_34 | ~spl9_182 | ~spl9_271 | ~spl9_847)),
  inference(forward_demodulation,[],[f272397,f271])).
fof(f272397,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(false1,false1))) | (~spl9_34 | ~spl9_182 | ~spl9_271 | ~spl9_847)),
  inference(backward_demodulation,[],[f252861,f269838])).
fof(f269838,plain,(
  true != lazy_impl(true,lazy_impl(true,impl(impl(true,false1),impl(true,false1)))) | (~spl9_34 | ~spl9_182 | ~spl9_271)),
  inference(forward_demodulation,[],[f269837,f264])).
fof(f264,plain,(
  true = or1(true,false1)),
  inference(resolution,[],[f176,f243])).
fof(f269837,plain,(
  or1(true,false1) != lazy_impl(true,lazy_impl(true,impl(impl(true,false1),impl(true,false1)))) | (~spl9_34 | ~spl9_182 | ~spl9_271)),
  inference(forward_demodulation,[],[f269836,f249])).
fof(f269836,plain,(
  or1(true,false1) != lazy_impl(true,lazy_impl(prop(false1),impl(impl(true,false1),impl(true,false1)))) | (~spl9_34 | ~spl9_182 | ~spl9_271)),
  inference(forward_demodulation,[],[f269835,f271])).
fof(f269835,plain,(
  or1(true,false1) != lazy_impl(true,lazy_impl(prop(false1),impl(impl(true,false1),impl(impl(false1,false1),false1)))) | (~spl9_34 | ~spl9_182 | ~spl9_271)),
  inference(forward_demodulation,[],[f269834,f252748])).
fof(f252748,plain,(
  false1 = sK2(true,false1) | (~spl9_34 | ~spl9_182 | ~spl9_271)),
  inference(forward_demodulation,[],[f252747,f1192])).
fof(f252747,plain,(
  false1 = sK2(sK7,false1) | (~spl9_182 | ~spl9_271)),
  inference(forward_demodulation,[],[f8569,f15749])).
fof(f269834,plain,(
  or1(true,false1) != lazy_impl(true,lazy_impl(prop(sK2(true,false1)),impl(impl(true,sK2(true,false1)),impl(impl(false1,sK2(true,false1)),sK2(true,false1))))) | (~spl9_34 | ~spl9_271)),
  inference(forward_demodulation,[],[f57231,f15749])).
fof(f269906,plain,(
  ~spl9_848 | spl9_6 | ~spl9_283),
  inference(avatar_split_clause,[],[f34841,f16383,f457,f252889])).
fof(f457,plain,(
  spl9_6 <=> err = lazy_impl(true,sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_6])])).
fof(f16383,plain,(
  spl9_283 <=> false1 = sK6),
  introduced(avatar_definition,[new_symbols(naming,[spl9_283])])).
fof(f34841,plain,(
  err != lazy_impl(true,false1) | (spl9_6 | ~spl9_283)),
  inference(backward_demodulation,[],[f16385,f458])).
fof(f458,plain,(
  err != lazy_impl(true,sK6) | spl9_6),
  inference(avatar_component_clause,[],[f457])).
fof(f16385,plain,(
  false1 = sK6 | ~spl9_283),
  inference(avatar_component_clause,[],[f16383])).
fof(f269905,plain,(
  spl9_846 | ~spl9_8 | ~spl9_283),
  inference(avatar_split_clause,[],[f34842,f16383,f465,f252855])).
fof(f465,plain,(
  spl9_8 <=> sK6 = lazy_impl(true,sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_8])])).
fof(f34842,plain,(
  false1 = lazy_impl(true,false1) | (~spl9_8 | ~spl9_283)),
  inference(forward_demodulation,[],[f467,f16385])).
fof(f467,plain,(
  sK6 = lazy_impl(true,sK6) | ~spl9_8),
  inference(avatar_component_clause,[],[f465])).
fof(f269763,plain,(
  spl9_847 | ~spl9_1 | spl9_2 | ~spl9_283),
  inference(avatar_split_clause,[],[f34886,f16383,f339,f335,f252859])).
fof(f335,plain,(
  spl9_1 <=> true = prop(sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_1])])).
fof(f339,plain,(
  spl9_2 <=> true = false2),
  introduced(avatar_definition,[new_symbols(naming,[spl9_2])])).
fof(f34886,plain,(
  false1 = impl(true,false1) | (~spl9_1 | spl9_2 | ~spl9_283)),
  inference(forward_demodulation,[],[f34885,f16385])).
fof(f34885,plain,(
  sK6 = impl(true,sK6) | (~spl9_1 | spl9_2 | ~spl9_283)),
  inference(subsumption_resolution,[],[f34884,f340])).
fof(f340,plain,(
  true != false2 | spl9_2),
  inference(avatar_component_clause,[],[f339])).
fof(f34884,plain,(
  true = false2 | sK6 = impl(true,sK6) | (~spl9_1 | ~spl9_283)),
  inference(forward_demodulation,[],[f34883,f267])).
fof(f34883,plain,(
  false2 = lazy_impl(true,true) | sK6 = impl(true,sK6) | (~spl9_1 | ~spl9_283)),
  inference(forward_demodulation,[],[f34882,f239])).
fof(f34882,plain,(
  false2 = lazy_impl(true,not1(false1)) | sK6 = impl(true,sK6) | (~spl9_1 | ~spl9_283)),
  inference(forward_demodulation,[],[f16237,f16385])).
fof(f16237,plain,(
  false2 = lazy_impl(true,not1(sK6)) | sK6 = impl(true,sK6) | ~spl9_1),
  inference(superposition,[],[f579,f284])).
fof(f284,plain,(
  ( ! [X1] : (lazy_impl(true,X1) = not1(X1) | impl(true,X1) = X1) )),
  inference(resolution,[],[f238,f157])).
fof(f579,plain,(
  false2 = lazy_impl(true,lazy_impl(true,sK6)) | ~spl9_1),
  inference(forward_demodulation,[],[f237,f337])).
fof(f337,plain,(
  true = prop(sK6) | ~spl9_1),
  inference(avatar_component_clause,[],[f335])).
fof(f237,plain,(
  false2 = lazy_impl(true,lazy_impl(prop(sK6),sK6))),
  inference(definition_unfolding,[],[f191,f160,f190])).
fof(f190,plain,(
  ( ! [X0] : (f7(X0) = lazy_impl(prop(X0),X0)) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : f7(X0) = lazy_impl(prop(X0),X0)),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X4] : f7(X4) = lazy_impl(prop(X4),X4)),
  file('./p/SWW098+1.tptp',unknown)).
fof(f191,plain,(
  false2 = phi(f7(sK6))),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X1] : ~forallprefers(f7(X1),f7(sK6)) & false2 = phi(f7(sK6))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f108,f127])).
fof(f127,plain,(
  ? [X0] : (! [X1] : ~forallprefers(f7(X1),f7(X0)) & false2 = phi(f7(X0))) => (! [X1] : ~forallprefers(f7(X1),f7(sK6)) & false2 = phi(f7(sK6)))),
  introduced(choice_axiom,[])).
fof(f108,plain,(
  ? [X0] : (! [X1] : ~forallprefers(f7(X1),f7(X0)) & false2 = phi(f7(X0)))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ? [X0] : (~? [X1] : forallprefers(f7(X1),f7(X0)) & false2 = phi(f7(X0)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ? [X4] : (~? [X9] : forallprefers(f7(X9),f7(X4)) & false2 = phi(f7(X4)))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f215329,plain,(
  spl9_6 | ~spl9_111 | ~spl9_271 | ~spl9_283),
  inference(avatar_contradiction_clause,[],[f215328])).
fof(f215328,plain,(
  $false | (spl9_6 | ~spl9_111 | ~spl9_271 | ~spl9_283)),
  inference(subsumption_resolution,[],[f215008,f34841])).
fof(f215008,plain,(
  err = lazy_impl(true,false1) | (~spl9_111 | ~spl9_271)),
  inference(backward_demodulation,[],[f15749,f4110])).
fof(f214790,plain,(
  spl9_59 | spl9_93 | ~spl9_266),
  inference(avatar_split_clause,[],[f117029,f15602,f3154,f1897])).
fof(f117029,plain,(
  err = not1(sK8) | true = prop(sK8) | ~spl9_266),
  inference(superposition,[],[f15603,f1862])).
fof(f1862,plain,(
  ( ! [X8,X7] : (not1(X7) = impl(X7,X8) | true = prop(X7)) )),
  inference(duplicate_literal_removal,[],[f1810])).
fof(f1810,plain,(
  ( ! [X8,X7] : (not1(X7) = impl(X7,X8) | true = prop(X7) | true = prop(X7)) )),
  inference(superposition,[],[f291,f283])).
fof(f214468,plain,(
  spl9_49 | ~spl9_114 | ~spl9_266 | ~spl9_704),
  inference(avatar_contradiction_clause,[],[f214467])).
fof(f214467,plain,(
  $false | (spl9_49 | ~spl9_114 | ~spl9_266 | ~spl9_704)),
  inference(subsumption_resolution,[],[f214466,f6849])).
fof(f214466,plain,(
  bool(err) | (spl9_49 | ~spl9_114 | ~spl9_266 | ~spl9_704)),
  inference(forward_demodulation,[],[f214465,f268])).
fof(f214465,plain,(
  bool(lazy_impl(true,err)) | (spl9_49 | ~spl9_114 | ~spl9_266 | ~spl9_704)),
  inference(forward_demodulation,[],[f214464,f248])).
fof(f214464,plain,(
  bool(lazy_impl(prop(true),err)) | (spl9_49 | ~spl9_114 | ~spl9_266 | ~spl9_704)),
  inference(forward_demodulation,[],[f214181,f15576])).
fof(f214181,plain,(
  bool(lazy_impl(prop(true),impl(true,err))) | (spl9_49 | ~spl9_266 | ~spl9_704)),
  inference(superposition,[],[f114414,f257])).
fof(f114414,plain,(
  ( ! [X0] : (bool(lazy_impl(prop(X0),impl(impl(true,X0),err)))) ) | (spl9_49 | ~spl9_266 | ~spl9_704)),
  inference(forward_demodulation,[],[f114413,f1864])).
fof(f114413,plain,(
  ( ! [X0] : (bool(lazy_impl(prop(X0),impl(impl(true,X0),impl(err,X0))))) ) | (~spl9_266 | ~spl9_704)),
  inference(forward_demodulation,[],[f113598,f15603])).
fof(f113598,plain,(
  ( ! [X0] : (bool(lazy_impl(prop(X0),impl(impl(true,X0),impl(impl(sK8,X0),X0))))) ) | ~spl9_704),
  inference(avatar_component_clause,[],[f113597])).
fof(f113597,plain,(
  spl9_704 <=> ! [X0] : bool(lazy_impl(prop(X0),impl(impl(true,X0),impl(impl(sK8,X0),X0))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_704])])).
fof(f135599,plain,(
  ~spl9_234 | ~spl9_395),
  inference(avatar_contradiction_clause,[],[f135598])).
fof(f135598,plain,(
  $false | (~spl9_234 | ~spl9_395)),
  inference(subsumption_resolution,[],[f135597,f11776])).
fof(f135597,plain,(
  ~forallprefers(err,true) | ~spl9_395),
  inference(forward_demodulation,[],[f135520,f268])).
fof(f135520,plain,(
  ~forallprefers(lazy_impl(true,err),true) | ~spl9_395),
  inference(superposition,[],[f27930,f249])).
fof(f27930,plain,(
  ( ! [X1] : (~forallprefers(lazy_impl(prop(X1),err),true)) ) | ~spl9_395),
  inference(avatar_component_clause,[],[f27929])).
fof(f120288,plain,(
  ~spl9_114 | spl9_234),
  inference(avatar_contradiction_clause,[],[f120287])).
fof(f120287,plain,(
  $false | (~spl9_114 | spl9_234)),
  inference(subsumption_resolution,[],[f120286,f139])).
fof(f120286,plain,(
  ~d(err) | (~spl9_114 | spl9_234)),
  inference(subsumption_resolution,[],[f120285,f137])).
fof(f120285,plain,(
  ~d(true) | ~d(err) | (~spl9_114 | spl9_234)),
  inference(subsumption_resolution,[],[f120284,f6849])).
fof(f120284,plain,(
  bool(err) | ~d(true) | ~d(err) | spl9_234),
  inference(subsumption_resolution,[],[f120282,f242])).
fof(f120282,plain,(
  ~bool(true) | bool(err) | ~d(true) | ~d(err) | spl9_234),
  inference(resolution,[],[f11777,f141])).
fof(f11777,plain,(
  ~forallprefers(err,true) | spl9_234),
  inference(avatar_component_clause,[],[f11775])).
fof(f114386,plain,(
  spl9_49 | ~spl9_82 | ~spl9_266 | spl9_703),
  inference(avatar_contradiction_clause,[],[f114385])).
fof(f114385,plain,(
  $false | (spl9_49 | ~spl9_82 | ~spl9_266 | spl9_703)),
  inference(subsumption_resolution,[],[f114384,f242])).
fof(f114384,plain,(
  ~bool(true) | (spl9_49 | ~spl9_82 | ~spl9_266 | spl9_703)),
  inference(forward_demodulation,[],[f114383,f221])).
fof(f114383,plain,(
  ~bool(lazy_impl(false1,impl(not1(sK2(true,sK8)),err))) | (spl9_49 | ~spl9_82 | ~spl9_266 | spl9_703)),
  inference(forward_demodulation,[],[f114382,f2726])).
fof(f2726,plain,(
  false1 = prop(sK2(true,sK8)) | ~spl9_82),
  inference(avatar_component_clause,[],[f2724])).
fof(f2724,plain,(
  spl9_82 <=> false1 = prop(sK2(true,sK8))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_82])])).
fof(f114382,plain,(
  ~bool(lazy_impl(prop(sK2(true,sK8)),impl(not1(sK2(true,sK8)),err))) | (spl9_49 | ~spl9_82 | ~spl9_266 | spl9_703)),
  inference(forward_demodulation,[],[f114381,f65742])).
fof(f65742,plain,(
  impl(true,sK2(true,sK8)) = not1(sK2(true,sK8)) | ~spl9_82),
  inference(forward_demodulation,[],[f65718,f65717])).
fof(f65717,plain,(
  lazy_impl(true,sK2(true,sK8)) = not1(sK2(true,sK8)) | ~spl9_82),
  inference(resolution,[],[f65684,f238])).
fof(f65684,plain,(
  ~bool(sK2(true,sK8)) | ~spl9_82),
  inference(trivial_inequality_removal,[],[f65675])).
fof(f65675,plain,(
  false1 != false1 | ~bool(sK2(true,sK8)) | ~spl9_82),
  inference(superposition,[],[f216,f2726])).
fof(f65718,plain,(
  impl(true,sK2(true,sK8)) = lazy_impl(true,sK2(true,sK8)) | ~spl9_82),
  inference(resolution,[],[f65684,f524])).
fof(f114381,plain,(
  ~bool(lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),err))) | (spl9_49 | ~spl9_266 | spl9_703)),
  inference(forward_demodulation,[],[f114380,f1864])).
fof(f114380,plain,(
  ~bool(lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(err,sK2(true,sK8))))) | (~spl9_266 | spl9_703)),
  inference(forward_demodulation,[],[f113595,f15603])).
fof(f113595,plain,(
  ~bool(lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8))))) | spl9_703),
  inference(avatar_component_clause,[],[f113593])).
fof(f113593,plain,(
  spl9_703 <=> bool(lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8)))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_703])])).
fof(f114364,plain,(
  ~spl9_682 | ~spl9_34 | spl9_55 | ~spl9_82),
  inference(avatar_split_clause,[],[f114119,f2724,f1879,f1190,f110855])).
fof(f110855,plain,(
  spl9_682 <=> true = prop(not1(sK2(true,sK8)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_682])])).
fof(f1879,plain,(
  spl9_55 <=> true = prop(impl(sK7,sK2(sK7,sK8)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_55])])).
fof(f114119,plain,(
  true != prop(not1(sK2(true,sK8))) | (~spl9_34 | spl9_55 | ~spl9_82)),
  inference(forward_demodulation,[],[f114118,f65742])).
fof(f114118,plain,(
  true != prop(impl(true,sK2(true,sK8))) | (~spl9_34 | spl9_55)),
  inference(forward_demodulation,[],[f1880,f1192])).
fof(f1880,plain,(
  true != prop(impl(sK7,sK2(sK7,sK8))) | spl9_55),
  inference(avatar_component_clause,[],[f1879])).
fof(f113904,plain,(
  ~spl9_712 | spl9_109 | ~spl9_425 | ~spl9_542),
  inference(avatar_split_clause,[],[f113903,f69294,f41242,f4097,f113807])).
fof(f4097,plain,(
  spl9_109 <=> false1 = and1(false1,sK8)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_109])])).
fof(f69294,plain,(
  spl9_542 <=> ! [X80] : (~bool(lazy_impl(prop(sK1(sK8,X80)),sK8)) | ~d(lazy_impl(prop(sK1(sK8,X80)),sK8)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_542])])).
fof(f113903,plain,(
  ~bool(not1(sK8)) | (spl9_109 | ~spl9_425 | ~spl9_542)),
  inference(forward_demodulation,[],[f69525,f64964])).
fof(f64964,plain,(
  lazy_impl(true,sK8) = not1(sK8) | spl9_109),
  inference(trivial_inequality_removal,[],[f64953])).
fof(f64953,plain,(
  false1 != false1 | lazy_impl(true,sK8) = not1(sK8) | spl9_109),
  inference(superposition,[],[f4098,f289])).
fof(f289,plain,(
  ( ! [X6] : (false1 = and1(false1,X6) | lazy_impl(true,X6) = not1(X6)) )),
  inference(resolution,[],[f238,f224])).
fof(f224,plain,(
  ( ! [X0] : (~bool(X0) | false1 = and1(false1,X0)) )),
  inference(definition_unfolding,[],[f163,f189,f189])).
fof(f163,plain,(
  ( ! [X0] : (false = and1(false,X0) | ~bool(X0)) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : (false = and1(false,X0) | ~bool(X0))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : (bool(X0) => false = and1(false,X0))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X3] : (bool(X3) => false = and1(false,X3))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f4098,plain,(
  false1 != and1(false1,sK8) | spl9_109),
  inference(avatar_component_clause,[],[f4097])).
fof(f69525,plain,(
  ~bool(lazy_impl(true,sK8)) | (~spl9_425 | ~spl9_542)),
  inference(backward_demodulation,[],[f69523,f69300])).
fof(f69300,plain,(
  ( ! [X80] : (~bool(lazy_impl(prop(sK1(sK8,X80)),sK8))) ) | (~spl9_425 | ~spl9_542)),
  inference(subsumption_resolution,[],[f69295,f41243])).
fof(f69295,plain,(
  ( ! [X80] : (~bool(lazy_impl(prop(sK1(sK8,X80)),sK8)) | ~d(lazy_impl(prop(sK1(sK8,X80)),sK8))) ) | ~spl9_542),
  inference(avatar_component_clause,[],[f69294])).
fof(f69523,plain,(
  ( ! [X1] : (true = prop(sK1(sK8,X1))) ) | (~spl9_425 | ~spl9_542)),
  inference(subsumption_resolution,[],[f69522,f242])).
fof(f69522,plain,(
  ( ! [X1] : (~bool(true) | true = prop(sK1(sK8,X1))) ) | (~spl9_425 | ~spl9_542)),
  inference(forward_demodulation,[],[f69477,f221])).
fof(f69477,plain,(
  ( ! [X1] : (~bool(lazy_impl(false1,sK8)) | true = prop(sK1(sK8,X1))) ) | (~spl9_425 | ~spl9_542)),
  inference(superposition,[],[f69300,f252])).
fof(f113599,plain,(
  ~spl9_703 | spl9_704 | ~spl9_82),
  inference(avatar_split_clause,[],[f68337,f2724,f113597,f113593])).
fof(f68337,plain,(
  ( ! [X0] : (bool(lazy_impl(prop(X0),impl(impl(true,X0),impl(impl(sK8,X0),X0)))) | ~bool(lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8)))))) ) | ~spl9_82),
  inference(subsumption_resolution,[],[f68336,f137])).
fof(f68336,plain,(
  ( ! [X0] : (~d(true) | bool(lazy_impl(prop(X0),impl(impl(true,X0),impl(impl(sK8,X0),X0)))) | ~bool(lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8)))))) ) | ~spl9_82),
  inference(forward_demodulation,[],[f68283,f221])).
fof(f68283,plain,(
  ( ! [X0] : (~d(lazy_impl(false1,impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8))))) | bool(lazy_impl(prop(X0),impl(impl(true,X0),impl(impl(sK8,X0),X0)))) | ~bool(lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8)))))) ) | ~spl9_82),
  inference(superposition,[],[f21241,f2726])).
fof(f21241,plain,(
  ( ! [X2,X0,X1] : (~d(lazy_impl(prop(sK2(X0,X1)),impl(impl(X0,sK2(X0,X1)),impl(impl(X1,sK2(X0,X1)),sK2(X0,X1))))) | bool(lazy_impl(prop(X2),impl(impl(X0,X2),impl(impl(X1,X2),X2)))) | ~bool(lazy_impl(prop(sK2(X0,X1)),impl(impl(X0,sK2(X0,X1)),impl(impl(X1,sK2(X0,X1)),sK2(X0,X1)))))) )),
  inference(subsumption_resolution,[],[f1469,f1470])).
fof(f1470,plain,(
  ( ! [X4,X5,X3] : (~d(lazy_impl(prop(sK2(X3,X4)),impl(impl(X3,sK2(X3,X4)),impl(impl(X4,sK2(X3,X4)),sK2(X3,X4))))) | d(lazy_impl(prop(X5),impl(impl(X3,X5),impl(impl(X4,X5),X5))))) )),
  inference(resolution,[],[f233,f140])).
fof(f140,plain,(
  ( ! [X0,X1] : (forallprefers(X0,X1) | ~d(X1) | d(X0)) )),
  inference(cnf_transformation,[],[f83])).
fof(f1469,plain,(
  ( ! [X2,X0,X1] : (~d(lazy_impl(prop(sK2(X0,X1)),impl(impl(X0,sK2(X0,X1)),impl(impl(X1,sK2(X0,X1)),sK2(X0,X1))))) | bool(lazy_impl(prop(X2),impl(impl(X0,X2),impl(impl(X1,X2),X2)))) | ~bool(lazy_impl(prop(sK2(X0,X1)),impl(impl(X0,sK2(X0,X1)),impl(impl(X1,sK2(X0,X1)),sK2(X0,X1))))) | ~d(lazy_impl(prop(X2),impl(impl(X0,X2),impl(impl(X1,X2),X2))))) )),
  inference(resolution,[],[f233,f141])).
fof(f113516,plain,(
  spl9_111 | spl9_689 | spl9_85),
  inference(avatar_split_clause,[],[f109776,f2853,f113513,f4108])).
fof(f109776,plain,(
  ( ! [X44] : (sK8 = lazy_impl(sK8,X44) | err = lazy_impl(true,sK8)) ) | spl9_85),
  inference(trivial_inequality_removal,[],[f109766])).
fof(f109766,plain,(
  ( ! [X44] : (true != true | sK8 = lazy_impl(sK8,X44) | err = lazy_impl(true,sK8)) ) | spl9_85),
  inference(superposition,[],[f2855,f3426])).
fof(f3426,plain,(
  ( ! [X2,X3] : (lazy_impl(X2,X3) = X2 | true = or1(true,X2) | err = lazy_impl(true,X2)) )),
  inference(superposition,[],[f302,f212])).
fof(f302,plain,(
  ( ! [X6,X7] : (lazy_impl(true,X6) = lazy_impl(X6,X7) | true = or1(true,X6)) )),
  inference(resolution,[],[f220,f176])).
fof(f220,plain,(
  ( ! [X0,X1] : (bool(X0) | lazy_impl(X0,X1) = lazy_impl(true,X0)) )),
  inference(definition_unfolding,[],[f158,f160])).
fof(f158,plain,(
  ( ! [X0,X1] : (phi(X0) = lazy_impl(X0,X1) | bool(X0)) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : (phi(X0) = lazy_impl(X0,X1) | bool(X0))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : (~bool(X0) => phi(X0) = lazy_impl(X0,X1))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X2,X3] : (~bool(X2) => phi(X2) = lazy_impl(X2,X3))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f113446,plain,(
  ~spl9_82 | spl9_84 | spl9_86 | spl9_87 | ~spl9_114 | spl9_463 | ~spl9_540),
  inference(avatar_contradiction_clause,[],[f113445])).
fof(f113445,plain,(
  $false | (~spl9_82 | spl9_84 | spl9_86 | spl9_87 | ~spl9_114 | spl9_463 | ~spl9_540)),
  inference(subsumption_resolution,[],[f113444,f50620])).
fof(f50620,plain,(
  err != false1 | spl9_463),
  inference(avatar_component_clause,[],[f50619])).
fof(f50619,plain,(
  spl9_463 <=> err = false1),
  introduced(avatar_definition,[new_symbols(naming,[spl9_463])])).
fof(f113444,plain,(
  err = false1 | (~spl9_82 | spl9_84 | spl9_86 | spl9_87 | ~spl9_114 | spl9_463 | ~spl9_540)),
  inference(forward_demodulation,[],[f113443,f113228])).
fof(f113228,plain,(
  err = sK2(true,sK8) | (~spl9_82 | spl9_84 | spl9_86 | spl9_463 | ~spl9_540)),
  inference(forward_demodulation,[],[f113227,f113222])).
fof(f113222,plain,(
  ( ! [X18] : (err = lazy_impl(sK2(true,sK8),X18)) ) | (~spl9_82 | spl9_84 | spl9_86 | ~spl9_540)),
  inference(subsumption_resolution,[],[f113221,f2850])).
fof(f2850,plain,(
  true != prop(sK2(true,sK8)) | spl9_84),
  inference(avatar_component_clause,[],[f2849])).
fof(f113221,plain,(
  ( ! [X18] : (err = lazy_impl(sK2(true,sK8),X18) | true = prop(sK2(true,sK8))) ) | (~spl9_82 | spl9_86 | ~spl9_540)),
  inference(subsumption_resolution,[],[f113026,f2859])).
fof(f2859,plain,(
  false1 != sK2(true,sK8) | spl9_86),
  inference(avatar_component_clause,[],[f2858])).
fof(f113026,plain,(
  ( ! [X18] : (false1 = sK2(true,sK8) | err = lazy_impl(sK2(true,sK8),X18) | true = prop(sK2(true,sK8))) ) | (~spl9_82 | ~spl9_540)),
  inference(superposition,[],[f111111,f1953])).
fof(f1953,plain,(
  ( ! [X0,X1] : (lazy_impl(true,X0) = X0 | err = lazy_impl(X0,X1) | true = prop(X0)) )),
  inference(superposition,[],[f299,f212])).
fof(f299,plain,(
  ( ! [X0,X1] : (lazy_impl(X0,X1) = lazy_impl(true,X0) | true = prop(X0)) )),
  inference(resolution,[],[f220,f151])).
fof(f111111,plain,(
  false1 = lazy_impl(true,sK2(true,sK8)) | (~spl9_82 | ~spl9_540)),
  inference(backward_demodulation,[],[f66041,f65717])).
fof(f66041,plain,(
  false1 = not1(sK2(true,sK8)) | ~spl9_540),
  inference(avatar_component_clause,[],[f66039])).
fof(f66039,plain,(
  spl9_540 <=> false1 = not1(sK2(true,sK8))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_540])])).
fof(f113227,plain,(
  ( ! [X19] : (sK2(true,sK8) = lazy_impl(sK2(true,sK8),X19)) ) | (~spl9_82 | spl9_84 | spl9_463 | ~spl9_540)),
  inference(subsumption_resolution,[],[f113226,f2850])).
fof(f113226,plain,(
  ( ! [X19] : (sK2(true,sK8) = lazy_impl(sK2(true,sK8),X19) | true = prop(sK2(true,sK8))) ) | (~spl9_82 | spl9_463 | ~spl9_540)),
  inference(subsumption_resolution,[],[f113027,f50620])).
fof(f113027,plain,(
  ( ! [X19] : (err = false1 | sK2(true,sK8) = lazy_impl(sK2(true,sK8),X19) | true = prop(sK2(true,sK8))) ) | (~spl9_82 | ~spl9_540)),
  inference(superposition,[],[f111111,f1954])).
fof(f1954,plain,(
  ( ! [X2,X3] : (lazy_impl(X2,X3) = X2 | err = lazy_impl(true,X2) | true = prop(X2)) )),
  inference(superposition,[],[f299,f212])).
fof(f113443,plain,(
  false1 = sK2(true,sK8) | (~spl9_82 | spl9_84 | spl9_86 | spl9_87 | ~spl9_114 | spl9_463 | ~spl9_540)),
  inference(subsumption_resolution,[],[f113442,f50620])).
fof(f113442,plain,(
  err = false1 | false1 = sK2(true,sK8) | (~spl9_82 | spl9_84 | spl9_86 | spl9_87 | ~spl9_114 | spl9_463 | ~spl9_540)),
  inference(forward_demodulation,[],[f113441,f15641])).
fof(f113441,plain,(
  false1 = impl(false1,err) | false1 = sK2(true,sK8) | (~spl9_82 | spl9_84 | spl9_86 | spl9_87 | spl9_463 | ~spl9_540)),
  inference(forward_demodulation,[],[f113440,f113228])).
fof(f113440,plain,(
  false1 = impl(false1,sK2(true,sK8)) | false1 = sK2(true,sK8) | (~spl9_82 | spl9_87 | ~spl9_540)),
  inference(subsumption_resolution,[],[f113107,f2863])).
fof(f2863,plain,(
  true != sK2(true,sK8) | spl9_87),
  inference(avatar_component_clause,[],[f2862])).
fof(f113107,plain,(
  false1 = impl(false1,sK2(true,sK8)) | false1 = sK2(true,sK8) | true = sK2(true,sK8) | (~spl9_82 | ~spl9_540)),
  inference(superposition,[],[f1601,f111111])).
fof(f111108,plain,(
  spl9_682 | spl9_618 | spl9_541 | ~spl9_594),
  inference(avatar_split_clause,[],[f107965,f87029,f66043,f95458,f110855])).
fof(f95458,plain,(
  spl9_618 <=> err = lazy_impl(true,not1(sK2(true,sK8)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_618])])).
fof(f66043,plain,(
  spl9_541 <=> true = not1(sK2(true,sK8))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_541])])).
fof(f87029,plain,(
  spl9_594 <=> true = impl(not1(sK2(true,sK8)),true)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_594])])).
fof(f107965,plain,(
  true = not1(sK2(true,sK8)) | err = lazy_impl(true,not1(sK2(true,sK8))) | true = prop(not1(sK2(true,sK8))) | ~spl9_594),
  inference(superposition,[],[f87031,f1811])).
fof(f87031,plain,(
  true = impl(not1(sK2(true,sK8)),true) | ~spl9_594),
  inference(avatar_component_clause,[],[f87029])).
fof(f110848,plain,(
  ~spl9_82 | spl9_87 | ~spl9_541),
  inference(avatar_contradiction_clause,[],[f110847])).
fof(f110847,plain,(
  $false | (~spl9_82 | spl9_87 | ~spl9_541)),
  inference(subsumption_resolution,[],[f110846,f2863])).
fof(f110846,plain,(
  true = sK2(true,sK8) | (~spl9_82 | spl9_87 | ~spl9_541)),
  inference(forward_demodulation,[],[f110845,f108225])).
fof(f108225,plain,(
  true = impl(true,sK2(true,sK8)) | (~spl9_82 | ~spl9_541)),
  inference(backward_demodulation,[],[f66045,f65742])).
fof(f66045,plain,(
  true = not1(sK2(true,sK8)) | ~spl9_541),
  inference(avatar_component_clause,[],[f66043])).
fof(f110845,plain,(
  sK2(true,sK8) = impl(true,sK2(true,sK8)) | (~spl9_82 | spl9_87 | ~spl9_541)),
  inference(subsumption_resolution,[],[f110844,f2863])).
fof(f110844,plain,(
  true = sK2(true,sK8) | sK2(true,sK8) = impl(true,sK2(true,sK8)) | (~spl9_82 | ~spl9_541)),
  inference(forward_demodulation,[],[f110843,f108223])).
fof(f108223,plain,(
  ( ! [X1] : (true = lazy_impl(sK2(true,sK8),X1)) ) | (~spl9_82 | ~spl9_541)),
  inference(backward_demodulation,[],[f66045,f65740])).
fof(f65740,plain,(
  ( ! [X1] : (lazy_impl(sK2(true,sK8),X1) = not1(sK2(true,sK8))) ) | ~spl9_82),
  inference(backward_demodulation,[],[f65717,f65713])).
fof(f65713,plain,(
  ( ! [X1] : (lazy_impl(true,sK2(true,sK8)) = lazy_impl(sK2(true,sK8),X1)) ) | ~spl9_82),
  inference(resolution,[],[f65684,f220])).
fof(f110843,plain,(
  ( ! [X25] : (sK2(true,sK8) = lazy_impl(sK2(true,sK8),X25) | sK2(true,sK8) = impl(true,sK2(true,sK8))) ) | (~spl9_82 | ~spl9_541)),
  inference(subsumption_resolution,[],[f110666,f135])).
fof(f110666,plain,(
  ( ! [X25] : (true = err | sK2(true,sK8) = lazy_impl(sK2(true,sK8),X25) | sK2(true,sK8) = impl(true,sK2(true,sK8))) ) | (~spl9_82 | ~spl9_541)),
  inference(superposition,[],[f3240,f108219])).
fof(f108219,plain,(
  true = lazy_impl(true,sK2(true,sK8)) | (~spl9_82 | ~spl9_541)),
  inference(backward_demodulation,[],[f66045,f65717])).
fof(f3240,plain,(
  ( ! [X2,X3] : (lazy_impl(X2,X3) = X2 | err = lazy_impl(true,X2) | impl(true,X2) = X2) )),
  inference(superposition,[],[f300,f212])).
fof(f300,plain,(
  ( ! [X2,X3] : (lazy_impl(X2,X3) = lazy_impl(true,X2) | impl(true,X2) = X2) )),
  inference(resolution,[],[f220,f157])).
fof(f108213,plain,(
  spl9_681 | ~spl9_34 | spl9_55 | ~spl9_82 | ~spl9_594 | ~spl9_618),
  inference(avatar_split_clause,[],[f108212,f95458,f87029,f2724,f1879,f1190,f108086])).
fof(f108086,plain,(
  spl9_681 <=> err = not1(sK2(true,sK8))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_681])])).
fof(f108212,plain,(
  err = not1(sK2(true,sK8)) | (~spl9_34 | spl9_55 | ~spl9_82 | ~spl9_594 | ~spl9_618)),
  inference(forward_demodulation,[],[f108045,f95460])).
fof(f95460,plain,(
  err = lazy_impl(true,not1(sK2(true,sK8))) | ~spl9_618),
  inference(avatar_component_clause,[],[f95458])).
fof(f108045,plain,(
  not1(sK2(true,sK8)) = lazy_impl(true,not1(sK2(true,sK8))) | (~spl9_34 | spl9_55 | ~spl9_82 | ~spl9_594)),
  inference(subsumption_resolution,[],[f108044,f66168])).
fof(f66168,plain,(
  true != prop(not1(sK2(true,sK8))) | (~spl9_34 | spl9_55 | ~spl9_82)),
  inference(forward_demodulation,[],[f66167,f65742])).
fof(f66167,plain,(
  true != prop(impl(true,sK2(true,sK8))) | (~spl9_34 | spl9_55)),
  inference(forward_demodulation,[],[f1880,f1192])).
fof(f108044,plain,(
  not1(sK2(true,sK8)) = lazy_impl(true,not1(sK2(true,sK8))) | true = prop(not1(sK2(true,sK8))) | ~spl9_594),
  inference(subsumption_resolution,[],[f107966,f135])).
fof(f107966,plain,(
  true = err | not1(sK2(true,sK8)) = lazy_impl(true,not1(sK2(true,sK8))) | true = prop(not1(sK2(true,sK8))) | ~spl9_594),
  inference(superposition,[],[f87031,f1812])).
fof(f1812,plain,(
  ( ! [X12,X11] : (lazy_impl(true,X11) = X11 | err = impl(X11,X12) | true = prop(X11)) )),
  inference(superposition,[],[f291,f212])).
fof(f108193,plain,(
  spl9_116 | ~spl9_550 | ~spl9_681),
  inference(avatar_contradiction_clause,[],[f108192])).
fof(f108192,plain,(
  $false | (spl9_116 | ~spl9_550 | ~spl9_681)),
  inference(subsumption_resolution,[],[f108178,f5310])).
fof(f108178,plain,(
  true = impl(false1,err) | (~spl9_550 | ~spl9_681)),
  inference(backward_demodulation,[],[f108088,f86401])).
fof(f86401,plain,(
  true = impl(false1,not1(sK2(true,sK8))) | ~spl9_550),
  inference(avatar_component_clause,[],[f86399])).
fof(f86399,plain,(
  spl9_550 <=> true = impl(false1,not1(sK2(true,sK8)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_550])])).
fof(f108088,plain,(
  err = not1(sK2(true,sK8)) | ~spl9_681),
  inference(avatar_component_clause,[],[f108086])).
fof(f107177,plain,(
  ~spl9_82 | ~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | ~spl9_265 | spl9_269 | spl9_550),
  inference(avatar_contradiction_clause,[],[f107176])).
fof(f107176,plain,(
  $false | (~spl9_82 | ~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | ~spl9_265 | spl9_269 | spl9_550)),
  inference(subsumption_resolution,[],[f107175,f61806])).
fof(f61806,plain,(
  ~bool(sK8) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(forward_demodulation,[],[f61805,f4119])).
fof(f61805,plain,(
  ~bool(lazy_impl(true,sK8)) | (~spl9_95 | ~spl9_108 | ~spl9_112 | spl9_269)),
  inference(forward_demodulation,[],[f61804,f61613])).
fof(f61613,plain,(
  sK8 = impl(false1,sK8) | (~spl9_95 | ~spl9_108 | ~spl9_112)),
  inference(forward_demodulation,[],[f61612,f61557])).
fof(f61557,plain,(
  ( ! [X6] : (sK8 = impl(sK8,X6)) ) | (~spl9_108 | ~spl9_112)),
  inference(backward_demodulation,[],[f4114,f4094])).
fof(f61612,plain,(
  sK8 = impl(false1,impl(sK8,false1)) | (~spl9_95 | ~spl9_112)),
  inference(forward_demodulation,[],[f3164,f4114])).
fof(f61804,plain,(
  ~bool(lazy_impl(true,impl(false1,sK8))) | (~spl9_108 | ~spl9_112 | spl9_269)),
  inference(forward_demodulation,[],[f61803,f61557])).
fof(f61803,plain,(
  ~bool(lazy_impl(true,impl(false1,impl(sK8,false1)))) | (~spl9_108 | ~spl9_112 | spl9_269)),
  inference(forward_demodulation,[],[f15626,f61557])).
fof(f15626,plain,(
  ~bool(lazy_impl(true,impl(false1,impl(impl(sK8,false1),false1)))) | spl9_269),
  inference(avatar_component_clause,[],[f15624])).
fof(f107175,plain,(
  bool(sK8) | (~spl9_82 | ~spl9_108 | ~spl9_112 | ~spl9_113 | ~spl9_265 | spl9_550)),
  inference(forward_demodulation,[],[f107174,f4119])).
fof(f107174,plain,(
  bool(lazy_impl(true,sK8)) | (~spl9_82 | ~spl9_108 | ~spl9_112 | ~spl9_265 | spl9_550)),
  inference(forward_demodulation,[],[f107173,f248])).
fof(f107173,plain,(
  bool(lazy_impl(prop(true),sK8)) | (~spl9_82 | ~spl9_108 | ~spl9_112 | ~spl9_265 | spl9_550)),
  inference(forward_demodulation,[],[f106988,f15600])).
fof(f106988,plain,(
  bool(lazy_impl(prop(true),impl(true,sK8))) | (~spl9_82 | ~spl9_108 | ~spl9_112 | spl9_550)),
  inference(superposition,[],[f102072,f257])).
fof(f102072,plain,(
  ( ! [X0] : (bool(lazy_impl(prop(X0),impl(impl(true,X0),sK8)))) ) | (~spl9_82 | ~spl9_108 | ~spl9_112 | spl9_550)),
  inference(subsumption_resolution,[],[f102071,f242])).
fof(f102071,plain,(
  ( ! [X0] : (~bool(true) | bool(lazy_impl(prop(X0),impl(impl(true,X0),sK8)))) ) | (~spl9_82 | ~spl9_108 | ~spl9_112 | spl9_550)),
  inference(forward_demodulation,[],[f102070,f221])).
fof(f102070,plain,(
  ( ! [X0] : (~bool(lazy_impl(false1,not1(not1(sK2(true,sK8))))) | bool(lazy_impl(prop(X0),impl(impl(true,X0),sK8)))) ) | (~spl9_82 | ~spl9_108 | ~spl9_112 | spl9_550)),
  inference(forward_demodulation,[],[f102069,f2726])).
fof(f102069,plain,(
  ( ! [X0] : (~bool(lazy_impl(prop(sK2(true,sK8)),not1(not1(sK2(true,sK8))))) | bool(lazy_impl(prop(X0),impl(impl(true,X0),sK8)))) ) | (~spl9_82 | ~spl9_108 | ~spl9_112 | spl9_550)),
  inference(forward_demodulation,[],[f102068,f96618])).
fof(f96618,plain,(
  ( ! [X2] : (not1(not1(sK2(true,sK8))) = impl(not1(sK2(true,sK8)),X2)) ) | spl9_550),
  inference(trivial_inequality_removal,[],[f96577])).
fof(f96577,plain,(
  ( ! [X2] : (true != true | not1(not1(sK2(true,sK8))) = impl(not1(sK2(true,sK8)),X2)) ) | spl9_550),
  inference(superposition,[],[f86400,f2567])).
fof(f86400,plain,(
  true != impl(false1,not1(sK2(true,sK8))) | spl9_550),
  inference(avatar_component_clause,[],[f86399])).
fof(f102068,plain,(
  ( ! [X0] : (~bool(lazy_impl(prop(sK2(true,sK8)),impl(not1(sK2(true,sK8)),sK8))) | bool(lazy_impl(prop(X0),impl(impl(true,X0),sK8)))) ) | (~spl9_82 | ~spl9_108 | ~spl9_112)),
  inference(forward_demodulation,[],[f102067,f65742])).
fof(f102067,plain,(
  ( ! [X0] : (~bool(lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),sK8))) | bool(lazy_impl(prop(X0),impl(impl(true,X0),sK8)))) ) | (~spl9_82 | ~spl9_108 | ~spl9_112)),
  inference(forward_demodulation,[],[f102066,f61557])).
fof(f102066,plain,(
  ( ! [X0] : (~bool(lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(sK8,sK2(true,sK8))))) | bool(lazy_impl(prop(X0),impl(impl(true,X0),sK8)))) ) | (~spl9_82 | ~spl9_108 | ~spl9_112)),
  inference(forward_demodulation,[],[f68339,f61557])).
fof(f68339,plain,(
  ( ! [X0] : (bool(lazy_impl(prop(X0),impl(impl(true,X0),sK8))) | ~bool(lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8)))))) ) | (~spl9_82 | ~spl9_108 | ~spl9_112)),
  inference(forward_demodulation,[],[f68338,f61557])).
fof(f68338,plain,(
  ( ! [X0] : (bool(lazy_impl(prop(X0),impl(impl(true,X0),impl(sK8,X0)))) | ~bool(lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8)))))) ) | (~spl9_82 | ~spl9_108 | ~spl9_112)),
  inference(forward_demodulation,[],[f68337,f61557])).
fof(f96081,plain,(
  ~spl9_550 | ~spl9_34 | spl9_55 | ~spl9_82 | spl9_594),
  inference(avatar_split_clause,[],[f96080,f87029,f2724,f1879,f1190,f86399])).
fof(f96080,plain,(
  true != impl(false1,not1(sK2(true,sK8))) | (~spl9_34 | spl9_55 | ~spl9_82 | spl9_594)),
  inference(subsumption_resolution,[],[f96031,f66168])).
fof(f96031,plain,(
  true != impl(false1,not1(sK2(true,sK8))) | true = prop(not1(sK2(true,sK8))) | (~spl9_34 | spl9_55 | ~spl9_82 | spl9_594)),
  inference(superposition,[],[f91259,f1597])).
fof(f91259,plain,(
  true != lazy_impl(true,not1(sK2(true,sK8))) | (~spl9_34 | spl9_55 | ~spl9_82 | spl9_594)),
  inference(subsumption_resolution,[],[f91226,f66168])).
fof(f91226,plain,(
  true != lazy_impl(true,not1(sK2(true,sK8))) | true = prop(not1(sK2(true,sK8))) | (~spl9_34 | spl9_55 | ~spl9_82 | spl9_594)),
  inference(superposition,[],[f90719,f291])).
fof(f90719,plain,(
  ( ! [X0] : (true != impl(not1(sK2(true,sK8)),X0)) ) | (~spl9_34 | spl9_55 | ~spl9_82 | spl9_594)),
  inference(subsumption_resolution,[],[f90701,f66168])).
fof(f90701,plain,(
  ( ! [X0] : (true != impl(not1(sK2(true,sK8)),X0) | true = prop(not1(sK2(true,sK8)))) ) | spl9_594),
  inference(superposition,[],[f87030,f1863])).
fof(f1863,plain,(
  ( ! [X2,X0,X1] : (impl(X0,X1) = impl(X0,X2) | true = prop(X0)) )),
  inference(duplicate_literal_removal,[],[f1807])).
fof(f1807,plain,(
  ( ! [X2,X0,X1] : (impl(X0,X1) = impl(X0,X2) | true = prop(X0) | true = prop(X0)) )),
  inference(superposition,[],[f291,f291])).
fof(f87030,plain,(
  true != impl(not1(sK2(true,sK8)),true) | spl9_594),
  inference(avatar_component_clause,[],[f87029])).
fof(f69299,plain,(
  spl9_431 | spl9_425 | ~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269),
  inference(avatar_split_clause,[],[f69298,f15624,f4117,f4112,f4093,f3163,f41242,f43113])).
fof(f43113,plain,(
  spl9_431 <=> ! [X82] : ~d(lazy_impl(prop(sK1(sK8,X82)),sK8))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_431])])).
fof(f69298,plain,(
  ( ! [X80,X81] : (d(lazy_impl(prop(X81),sK8)) | ~d(lazy_impl(prop(sK1(sK8,X80)),sK8))) ) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(forward_demodulation,[],[f69297,f61557])).
fof(f69297,plain,(
  ( ! [X80,X81] : (d(lazy_impl(prop(X81),impl(sK8,X81))) | ~d(lazy_impl(prop(sK1(sK8,X80)),sK8))) ) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(forward_demodulation,[],[f66326,f61824])).
fof(f61824,plain,(
  ( ! [X1] : (sK8 = lazy_impl(sK8,X1)) ) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(forward_demodulation,[],[f61818,f4119])).
fof(f61818,plain,(
  ( ! [X1] : (lazy_impl(true,sK8) = lazy_impl(sK8,X1)) ) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(resolution,[],[f61806,f220])).
fof(f66326,plain,(
  ( ! [X80,X81] : (~d(lazy_impl(prop(sK1(sK8,X80)),sK8)) | d(lazy_impl(prop(X81),impl(lazy_impl(sK8,impl(X80,X81)),X81)))) ) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(forward_demodulation,[],[f66294,f61557])).
fof(f66294,plain,(
  ( ! [X80,X81] : (~d(lazy_impl(prop(sK1(sK8,X80)),impl(sK8,sK1(sK8,X80)))) | d(lazy_impl(prop(X81),impl(lazy_impl(sK8,impl(X80,X81)),X81)))) ) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(superposition,[],[f1414,f61824])).
fof(f1414,plain,(
  ( ! [X4,X5,X3] : (~d(lazy_impl(prop(sK1(X3,X4)),impl(lazy_impl(X3,impl(X4,sK1(X3,X4))),sK1(X3,X4)))) | d(lazy_impl(prop(X5),impl(lazy_impl(X3,impl(X4,X5)),X5)))) )),
  inference(resolution,[],[f229,f140])).
fof(f229,plain,(
  ( ! [X0,X3,X1] : (~forallprefers(lazy_impl(prop(X3),impl(lazy_impl(X0,impl(X1,X3)),X3)),lazy_impl(prop(sK1(X0,X1)),impl(lazy_impl(X0,impl(X1,sK1(X0,X1))),sK1(X0,X1))))) )),
  inference(definition_unfolding,[],[f173,f171,f171])).
fof(f171,plain,(
  ( ! [X2,X0,X1] : (f2(X0,X1,X2) = lazy_impl(prop(X2),impl(lazy_impl(X0,impl(X1,X2)),X2))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : f2(X0,X1,X2) = lazy_impl(prop(X2),impl(lazy_impl(X0,impl(X1,X2)),X2))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X4,X5,X6] : f2(X4,X5,X6) = lazy_impl(prop(X6),impl(lazy_impl(X4,impl(X5,X6)),X6))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f173,plain,(
  ( ! [X0,X3,X1] : (~forallprefers(f2(X0,X1,X3),f2(X0,X1,sK1(X0,X1)))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : (! [X3] : ~forallprefers(f2(X0,X1,X3),f2(X0,X1,sK1(X0,X1))) & lazy_and2(X0,X1) = phi(f2(X0,X1,sK1(X0,X1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f98,f117])).
fof(f117,plain,(
  ! [X1,X0] : (? [X2] : (! [X3] : ~forallprefers(f2(X0,X1,X3),f2(X0,X1,X2)) & lazy_and2(X0,X1) = phi(f2(X0,X1,X2))) => (! [X3] : ~forallprefers(f2(X0,X1,X3),f2(X0,X1,sK1(X0,X1))) & lazy_and2(X0,X1) = phi(f2(X0,X1,sK1(X0,X1)))))),
  introduced(choice_axiom,[])).
fof(f98,plain,(
  ! [X0,X1] : ? [X2] : (! [X3] : ~forallprefers(f2(X0,X1,X3),f2(X0,X1,X2)) & lazy_and2(X0,X1) = phi(f2(X0,X1,X2)))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : ? [X2] : (~? [X3] : forallprefers(f2(X0,X1,X3),f2(X0,X1,X2)) & lazy_and2(X0,X1) = phi(f2(X0,X1,X2)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X4,X5] : ? [X6] : (~? [X7] : forallprefers(f2(X4,X5,X7),f2(X4,X5,X6)) & lazy_and2(X4,X5) = phi(f2(X4,X5,X6)))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f69296,plain,(
  spl9_427 | spl9_542 | ~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269),
  inference(avatar_split_clause,[],[f69292,f15624,f4117,f4112,f4093,f3163,f69294,f43041])).
fof(f43041,plain,(
  spl9_427 <=> ! [X33] : bool(lazy_impl(prop(X33),sK8))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_427])])).
fof(f69292,plain,(
  ( ! [X80,X81] : (~bool(lazy_impl(prop(sK1(sK8,X80)),sK8)) | bool(lazy_impl(prop(X81),sK8)) | ~d(lazy_impl(prop(sK1(sK8,X80)),sK8))) ) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(forward_demodulation,[],[f69291,f61557])).
fof(f69291,plain,(
  ( ! [X80,X81] : (~bool(lazy_impl(prop(sK1(sK8,X80)),impl(sK8,sK1(sK8,X80)))) | bool(lazy_impl(prop(X81),sK8)) | ~d(lazy_impl(prop(sK1(sK8,X80)),sK8))) ) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(forward_demodulation,[],[f69290,f61824])).
fof(f69290,plain,(
  ( ! [X80,X81] : (bool(lazy_impl(prop(X81),sK8)) | ~d(lazy_impl(prop(sK1(sK8,X80)),sK8)) | ~bool(lazy_impl(prop(sK1(sK8,X80)),impl(lazy_impl(sK8,impl(X80,sK1(sK8,X80))),sK1(sK8,X80))))) ) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(forward_demodulation,[],[f69289,f61557])).
fof(f69289,plain,(
  ( ! [X80,X81] : (bool(lazy_impl(prop(X81),impl(sK8,X81))) | ~d(lazy_impl(prop(sK1(sK8,X80)),sK8)) | ~bool(lazy_impl(prop(sK1(sK8,X80)),impl(lazy_impl(sK8,impl(X80,sK1(sK8,X80))),sK1(sK8,X80))))) ) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(forward_demodulation,[],[f68152,f61824])).
fof(f68152,plain,(
  ( ! [X80,X81] : (~d(lazy_impl(prop(sK1(sK8,X80)),sK8)) | bool(lazy_impl(prop(X81),impl(lazy_impl(sK8,impl(X80,X81)),X81))) | ~bool(lazy_impl(prop(sK1(sK8,X80)),impl(lazy_impl(sK8,impl(X80,sK1(sK8,X80))),sK1(sK8,X80))))) ) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(forward_demodulation,[],[f68101,f61557])).
fof(f68101,plain,(
  ( ! [X80,X81] : (~d(lazy_impl(prop(sK1(sK8,X80)),impl(sK8,sK1(sK8,X80)))) | bool(lazy_impl(prop(X81),impl(lazy_impl(sK8,impl(X80,X81)),X81))) | ~bool(lazy_impl(prop(sK1(sK8,X80)),impl(lazy_impl(sK8,impl(X80,sK1(sK8,X80))),sK1(sK8,X80))))) ) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(superposition,[],[f21242,f61824])).
fof(f21242,plain,(
  ( ! [X2,X0,X1] : (~d(lazy_impl(prop(sK1(X0,X1)),impl(lazy_impl(X0,impl(X1,sK1(X0,X1))),sK1(X0,X1)))) | bool(lazy_impl(prop(X2),impl(lazy_impl(X0,impl(X1,X2)),X2))) | ~bool(lazy_impl(prop(sK1(X0,X1)),impl(lazy_impl(X0,impl(X1,sK1(X0,X1))),sK1(X0,X1))))) )),
  inference(subsumption_resolution,[],[f1413,f1414])).
fof(f1413,plain,(
  ( ! [X2,X0,X1] : (~d(lazy_impl(prop(sK1(X0,X1)),impl(lazy_impl(X0,impl(X1,sK1(X0,X1))),sK1(X0,X1)))) | bool(lazy_impl(prop(X2),impl(lazy_impl(X0,impl(X1,X2)),X2))) | ~bool(lazy_impl(prop(sK1(X0,X1)),impl(lazy_impl(X0,impl(X1,sK1(X0,X1))),sK1(X0,X1)))) | ~d(lazy_impl(prop(X2),impl(lazy_impl(X0,impl(X1,X2)),X2)))) )),
  inference(resolution,[],[f229,f141])).
fof(f69268,plain,(
  spl9_111 | ~spl9_113 | ~spl9_431),
  inference(avatar_contradiction_clause,[],[f69267])).
fof(f69267,plain,(
  $false | (spl9_111 | ~spl9_113 | ~spl9_431)),
  inference(subsumption_resolution,[],[f69266,f4109])).
fof(f4109,plain,(
  err != lazy_impl(true,sK8) | spl9_111),
  inference(avatar_component_clause,[],[f4108])).
fof(f69266,plain,(
  err = lazy_impl(true,sK8) | (~spl9_113 | ~spl9_431)),
  inference(forward_demodulation,[],[f69263,f4119])).
fof(f69263,plain,(
  err = lazy_impl(true,lazy_impl(true,sK8)) | ~spl9_431),
  inference(backward_demodulation,[],[f69262,f69247])).
fof(f69247,plain,(
  ( ! [X0] : (err = lazy_impl(true,lazy_impl(prop(sK1(sK8,X0)),sK8))) ) | ~spl9_431),
  inference(resolution,[],[f43114,f213])).
fof(f213,plain,(
  ( ! [X0] : (d(X0) | err = lazy_impl(true,X0)) )),
  inference(definition_unfolding,[],[f148,f160])).
fof(f148,plain,(
  ( ! [X0] : (err = phi(X0) | d(X0)) )),
  inference(cnf_transformation,[],[f6])).
fof(f43114,plain,(
  ( ! [X82] : (~d(lazy_impl(prop(sK1(sK8,X82)),sK8))) ) | ~spl9_431),
  inference(avatar_component_clause,[],[f43113])).
fof(f69262,plain,(
  ( ! [X1] : (true = prop(sK1(sK8,X1))) ) | ~spl9_431),
  inference(subsumption_resolution,[],[f69261,f137])).
fof(f69261,plain,(
  ( ! [X1] : (~d(true) | true = prop(sK1(sK8,X1))) ) | ~spl9_431),
  inference(forward_demodulation,[],[f69249,f221])).
fof(f69249,plain,(
  ( ! [X1] : (~d(lazy_impl(false1,sK8)) | true = prop(sK1(sK8,X1))) ) | ~spl9_431),
  inference(superposition,[],[f43114,f252])).
fof(f66048,plain,(
  spl9_540 | spl9_541 | ~spl9_34 | ~spl9_55 | ~spl9_82),
  inference(avatar_split_clause,[],[f66047,f2724,f1879,f1190,f66043,f66039])).
fof(f66047,plain,(
  true = not1(sK2(true,sK8)) | false1 = not1(sK2(true,sK8)) | (~spl9_34 | ~spl9_55 | ~spl9_82)),
  inference(subsumption_resolution,[],[f66029,f208])).
fof(f66029,plain,(
  true = false1 | true = not1(sK2(true,sK8)) | false1 = not1(sK2(true,sK8)) | (~spl9_34 | ~spl9_55 | ~spl9_82)),
  inference(superposition,[],[f282,f65743])).
fof(f65743,plain,(
  true = prop(not1(sK2(true,sK8))) | (~spl9_34 | ~spl9_55 | ~spl9_82)),
  inference(backward_demodulation,[],[f65742,f57473])).
fof(f57473,plain,(
  true = prop(impl(true,sK2(true,sK8))) | (~spl9_34 | ~spl9_55)),
  inference(forward_demodulation,[],[f1881,f1192])).
fof(f1881,plain,(
  true = prop(impl(sK7,sK2(sK7,sK8))) | ~spl9_55),
  inference(avatar_component_clause,[],[f1879])).
fof(f65354,plain,(
  ~spl9_34 | ~spl9_82 | spl9_152),
  inference(avatar_contradiction_clause,[],[f65353])).
fof(f65353,plain,(
  $false | (~spl9_34 | ~spl9_82 | spl9_152)),
  inference(subsumption_resolution,[],[f65318,f221])).
fof(f65318,plain,(
  true != lazy_impl(false1,impl(impl(true,sK2(true,sK8)),err)) | (~spl9_34 | ~spl9_82 | spl9_152)),
  inference(backward_demodulation,[],[f2726,f58560])).
fof(f58560,plain,(
  true != lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),err)) | (~spl9_34 | spl9_152)),
  inference(forward_demodulation,[],[f7250,f1192])).
fof(f7250,plain,(
  true != lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),err)) | spl9_152),
  inference(avatar_component_clause,[],[f7249])).
fof(f65288,plain,(
  spl9_82 | spl9_84),
  inference(avatar_contradiction_clause,[],[f65287])).
fof(f65287,plain,(
  $false | (spl9_82 | spl9_84)),
  inference(subsumption_resolution,[],[f65282,f2850])).
fof(f65282,plain,(
  true = prop(sK2(true,sK8)) | spl9_82),
  inference(trivial_inequality_removal,[],[f65281])).
fof(f65281,plain,(
  false1 != false1 | true = prop(sK2(true,sK8)) | spl9_82),
  inference(superposition,[],[f2725,f252])).
fof(f2725,plain,(
  false1 != prop(sK2(true,sK8)) | spl9_82),
  inference(avatar_component_clause,[],[f2724])).
fof(f65286,plain,(
  spl9_82 | spl9_86 | spl9_87),
  inference(avatar_contradiction_clause,[],[f65285])).
fof(f65285,plain,(
  $false | (spl9_82 | spl9_86 | spl9_87)),
  inference(subsumption_resolution,[],[f65284,f2859])).
fof(f65284,plain,(
  false1 = sK2(true,sK8) | (spl9_82 | spl9_87)),
  inference(subsumption_resolution,[],[f65283,f2863])).
fof(f65283,plain,(
  true = sK2(true,sK8) | false1 = sK2(true,sK8) | spl9_82),
  inference(trivial_inequality_removal,[],[f65280])).
fof(f65280,plain,(
  false1 != false1 | true = sK2(true,sK8) | false1 = sK2(true,sK8) | spl9_82),
  inference(superposition,[],[f2725,f282])).
fof(f61987,plain,(
  ~spl9_95 | ~spl9_108 | ~spl9_109 | ~spl9_112 | ~spl9_113 | spl9_269 | spl9_271),
  inference(avatar_contradiction_clause,[],[f61986])).
fof(f61986,plain,(
  $false | (~spl9_95 | ~spl9_108 | ~spl9_109 | ~spl9_112 | ~spl9_113 | spl9_269 | spl9_271)),
  inference(subsumption_resolution,[],[f61985,f15748])).
fof(f15748,plain,(
  false1 != sK8 | spl9_271),
  inference(avatar_component_clause,[],[f15747])).
fof(f61985,plain,(
  false1 = sK8 | (~spl9_95 | ~spl9_108 | ~spl9_109 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(backward_demodulation,[],[f61984,f4119])).
fof(f61984,plain,(
  false1 = lazy_impl(true,sK8) | (~spl9_95 | ~spl9_108 | ~spl9_109 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(forward_demodulation,[],[f61982,f4099])).
fof(f4099,plain,(
  false1 = and1(false1,sK8) | ~spl9_109),
  inference(avatar_component_clause,[],[f4097])).
fof(f61982,plain,(
  lazy_impl(true,sK8) = and1(false1,sK8) | (~spl9_95 | ~spl9_108 | ~spl9_112 | ~spl9_113 | spl9_269)),
  inference(resolution,[],[f536,f61806])).
fof(f536,plain,(
  ( ! [X1] : (bool(X1) | lazy_impl(true,X1) = and1(false1,X1)) )),
  inference(resolution,[],[f223,f243])).
fof(f61555,plain,(
  ~spl9_86 | ~spl9_34 | spl9_182),
  inference(avatar_split_clause,[],[f61554,f8567,f1190,f2858])).
fof(f61554,plain,(
  false1 != sK2(true,sK8) | (~spl9_34 | spl9_182)),
  inference(forward_demodulation,[],[f8568,f1192])).
fof(f61550,plain,(
  ~spl9_85 | spl9_84 | ~spl9_34),
  inference(avatar_split_clause,[],[f61549,f1190,f2849,f2853])).
fof(f61549,plain,(
  true = prop(sK2(true,sK8)) | true != or1(true,sK8) | ~spl9_34),
  inference(forward_demodulation,[],[f58574,f1192])).
fof(f58574,plain,(
  true != or1(true,sK8) | true = prop(sK2(sK7,sK8)) | ~spl9_34),
  inference(forward_demodulation,[],[f1468,f1192])).
fof(f1468,plain,(
  true != or1(sK7,sK8) | true = prop(sK2(sK7,sK8))),
  inference(forward_demodulation,[],[f1467,f267])).
fof(f1467,plain,(
  or1(sK7,sK8) != lazy_impl(true,true) | true = prop(sK2(sK7,sK8))),
  inference(forward_demodulation,[],[f1451,f221])).
fof(f1451,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(false1,impl(impl(sK7,sK2(sK7,sK8)),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | true = prop(sK2(sK7,sK8))),
  inference(superposition,[],[f241,f252])).
fof(f61094,plain,(
  ~spl9_1 | spl9_2 | ~spl9_37 | ~spl9_283 | ~spl9_458 | spl9_508 | ~spl9_530),
  inference(avatar_contradiction_clause,[],[f61093])).
fof(f61093,plain,(
  $false | (~spl9_1 | spl9_2 | ~spl9_37 | ~spl9_283 | ~spl9_458 | spl9_508 | ~spl9_530)),
  inference(subsumption_resolution,[],[f61092,f263])).
fof(f263,plain,(
  true = or1(true,true)),
  inference(resolution,[],[f176,f242])).
fof(f61092,plain,(
  true != or1(true,true) | (~spl9_1 | spl9_2 | ~spl9_37 | ~spl9_283 | ~spl9_458 | spl9_508 | ~spl9_530)),
  inference(forward_demodulation,[],[f61091,f267])).
fof(f61091,plain,(
  or1(true,true) != lazy_impl(true,true) | (~spl9_1 | spl9_2 | ~spl9_37 | ~spl9_283 | ~spl9_458 | spl9_508 | ~spl9_530)),
  inference(forward_demodulation,[],[f61090,f249])).
fof(f61090,plain,(
  or1(true,true) != lazy_impl(prop(false1),true) | (~spl9_1 | spl9_2 | ~spl9_37 | ~spl9_283 | ~spl9_458 | spl9_508 | ~spl9_530)),
  inference(forward_demodulation,[],[f61089,f50400])).
fof(f61089,plain,(
  or1(true,true) != lazy_impl(prop(false1),impl(false1,true)) | (~spl9_1 | spl9_2 | ~spl9_37 | ~spl9_283 | spl9_508 | ~spl9_530)),
  inference(forward_demodulation,[],[f61088,f271])).
fof(f61088,plain,(
  or1(true,true) != lazy_impl(prop(false1),impl(false1,impl(false1,false1))) | (~spl9_1 | spl9_2 | ~spl9_37 | ~spl9_283 | spl9_508 | ~spl9_530)),
  inference(forward_demodulation,[],[f61087,f34886])).
fof(f61087,plain,(
  or1(true,true) != lazy_impl(prop(false1),impl(impl(true,false1),impl(impl(true,false1),false1))) | (~spl9_37 | spl9_508 | ~spl9_530)),
  inference(forward_demodulation,[],[f61086,f60554])).
fof(f60554,plain,(
  false1 = sK2(true,true) | ~spl9_530),
  inference(avatar_component_clause,[],[f60552])).
fof(f60552,plain,(
  spl9_530 <=> false1 = sK2(true,true)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_530])])).
fof(f61086,plain,(
  or1(true,true) != lazy_impl(prop(sK2(true,true)),impl(impl(true,sK2(true,true)),impl(impl(true,sK2(true,true)),sK2(true,true)))) | (~spl9_37 | spl9_508)),
  inference(forward_demodulation,[],[f57239,f1209])).
fof(f57239,plain,(
  or1(true,sK8) != lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8)))) | spl9_508),
  inference(avatar_component_clause,[],[f57237])).
fof(f60750,plain,(
  spl9_530 | ~spl9_37 | ~spl9_86),
  inference(avatar_split_clause,[],[f60749,f2858,f1207,f60552])).
fof(f60749,plain,(
  false1 = sK2(true,true) | (~spl9_37 | ~spl9_86)),
  inference(forward_demodulation,[],[f2860,f1209])).
fof(f59558,plain,(
  spl9_37 | spl9_271 | ~spl9_59),
  inference(avatar_split_clause,[],[f59249,f1897,f15747,f1207])).
fof(f59249,plain,(
  false1 = sK8 | true = sK8 | ~spl9_59),
  inference(resolution,[],[f59240,f206])).
fof(f59240,plain,(
  bool(sK8) | ~spl9_59),
  inference(trivial_inequality_removal,[],[f59230])).
fof(f59230,plain,(
  true != true | bool(sK8) | ~spl9_59),
  inference(superposition,[],[f150,f1899])).
fof(f1899,plain,(
  true = prop(sK8) | ~spl9_59),
  inference(avatar_component_clause,[],[f1897])).
fof(f59227,plain,(
  spl9_59 | spl9_47),
  inference(avatar_split_clause,[],[f59212,f1576,f1897])).
fof(f59212,plain,(
  true = prop(sK8) | spl9_47),
  inference(trivial_inequality_removal,[],[f59209])).
fof(f59209,plain,(
  false1 != false1 | true = prop(sK8) | spl9_47),
  inference(superposition,[],[f1577,f252])).
fof(f1577,plain,(
  false1 != prop(sK8) | spl9_47),
  inference(avatar_component_clause,[],[f1576])).
fof(f58501,plain,(
  ~spl9_114 | ~spl9_133),
  inference(avatar_contradiction_clause,[],[f58500])).
fof(f58500,plain,(
  $false | (~spl9_114 | ~spl9_133)),
  inference(subsumption_resolution,[],[f58499,f135])).
fof(f58499,plain,(
  true = err | (~spl9_114 | ~spl9_133)),
  inference(forward_demodulation,[],[f268,f58156])).
fof(f58156,plain,(
  true = lazy_impl(true,err) | (~spl9_114 | ~spl9_133)),
  inference(forward_demodulation,[],[f5880,f15641])).
fof(f5880,plain,(
  true = lazy_impl(true,impl(false1,err)) | ~spl9_133),
  inference(avatar_component_clause,[],[f5878])).
fof(f58381,plain,(
  ~spl9_18 | ~spl9_34 | spl9_49 | ~spl9_87 | ~spl9_93 | ~spl9_94 | ~spl9_114 | ~spl9_266),
  inference(avatar_contradiction_clause,[],[f58380])).
fof(f58380,plain,(
  $false | (~spl9_18 | ~spl9_34 | spl9_49 | ~spl9_87 | ~spl9_93 | ~spl9_94 | ~spl9_114 | ~spl9_266)),
  inference(subsumption_resolution,[],[f58379,f58240])).
fof(f58240,plain,(
  err = lazy_impl(true,err) | (spl9_49 | ~spl9_93 | ~spl9_94 | ~spl9_114)),
  inference(forward_demodulation,[],[f58239,f15641])).
fof(f58239,plain,(
  err = lazy_impl(true,impl(false1,err)) | (spl9_49 | ~spl9_93 | ~spl9_94)),
  inference(forward_demodulation,[],[f58238,f1864])).
fof(f58238,plain,(
  err = lazy_impl(true,impl(false1,impl(err,false1))) | (~spl9_93 | ~spl9_94)),
  inference(forward_demodulation,[],[f3161,f3155])).
fof(f3161,plain,(
  err = lazy_impl(true,impl(false1,impl(not1(sK8),false1))) | ~spl9_94),
  inference(avatar_component_clause,[],[f3159])).
fof(f3159,plain,(
  spl9_94 <=> err = lazy_impl(true,impl(false1,impl(not1(sK8),false1)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_94])])).
fof(f58379,plain,(
  err != lazy_impl(true,err) | (~spl9_18 | ~spl9_34 | spl9_49 | ~spl9_87 | ~spl9_93 | ~spl9_94 | ~spl9_114 | ~spl9_266)),
  inference(forward_demodulation,[],[f58378,f58240])).
fof(f58378,plain,(
  err != lazy_impl(true,lazy_impl(true,err)) | (~spl9_18 | ~spl9_34 | spl9_49 | ~spl9_87 | ~spl9_114 | ~spl9_266)),
  inference(forward_demodulation,[],[f58377,f58222])).
fof(f58222,plain,(
  err = or1(true,sK8) | (~spl9_18 | ~spl9_34)),
  inference(forward_demodulation,[],[f659,f1192])).
fof(f58377,plain,(
  lazy_impl(true,lazy_impl(true,err)) != or1(true,sK8) | (~spl9_34 | spl9_49 | ~spl9_87 | ~spl9_114 | ~spl9_266)),
  inference(forward_demodulation,[],[f58376,f248])).
fof(f58376,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(true),err)) | (~spl9_34 | spl9_49 | ~spl9_87 | ~spl9_114 | ~spl9_266)),
  inference(forward_demodulation,[],[f58375,f15576])).
fof(f58375,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(true),impl(true,err))) | (~spl9_34 | spl9_49 | ~spl9_87 | ~spl9_266)),
  inference(forward_demodulation,[],[f58374,f257])).
fof(f58374,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(true),impl(impl(true,true),err))) | (~spl9_34 | spl9_49 | ~spl9_87 | ~spl9_266)),
  inference(forward_demodulation,[],[f58373,f1864])).
fof(f58373,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(true),impl(impl(true,true),impl(err,true)))) | (~spl9_34 | ~spl9_87 | ~spl9_266)),
  inference(forward_demodulation,[],[f58372,f15603])).
fof(f58372,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(true),impl(impl(true,true),impl(impl(sK8,true),true)))) | (~spl9_34 | ~spl9_87)),
  inference(forward_demodulation,[],[f57231,f2864])).
fof(f58114,plain,(
  spl9_113 | ~spl9_47 | spl9_59 | ~spl9_112),
  inference(avatar_split_clause,[],[f58113,f4112,f1897,f1576,f4117])).
fof(f58113,plain,(
  sK8 = lazy_impl(true,sK8) | (~spl9_47 | spl9_59 | ~spl9_112)),
  inference(subsumption_resolution,[],[f36130,f1898])).
fof(f36130,plain,(
  sK8 = lazy_impl(true,sK8) | true = prop(sK8) | (~spl9_47 | ~spl9_112)),
  inference(superposition,[],[f299,f15913])).
fof(f15913,plain,(
  ( ! [X1] : (sK8 = lazy_impl(sK8,X1)) ) | (~spl9_47 | ~spl9_112)),
  inference(backward_demodulation,[],[f4114,f3081])).
fof(f3081,plain,(
  ( ! [X1] : (not1(sK8) = lazy_impl(sK8,X1)) ) | ~spl9_47),
  inference(backward_demodulation,[],[f3071,f3067])).
fof(f3067,plain,(
  ( ! [X1] : (lazy_impl(true,sK8) = lazy_impl(sK8,X1)) ) | ~spl9_47),
  inference(resolution,[],[f2970,f220])).
fof(f2970,plain,(
  ~bool(sK8) | ~spl9_47),
  inference(trivial_inequality_removal,[],[f2955])).
fof(f2955,plain,(
  false1 != false1 | ~bool(sK8) | ~spl9_47),
  inference(superposition,[],[f216,f1578])).
fof(f3071,plain,(
  lazy_impl(true,sK8) = not1(sK8) | ~spl9_47),
  inference(resolution,[],[f2970,f238])).
fof(f57899,plain,(
  ~spl9_47 | spl9_83 | ~spl9_84 | ~spl9_87 | ~spl9_108 | ~spl9_112 | ~spl9_113 | ~spl9_265),
  inference(avatar_contradiction_clause,[],[f57898])).
fof(f57898,plain,(
  $false | (~spl9_47 | spl9_83 | ~spl9_84 | ~spl9_87 | ~spl9_108 | ~spl9_112 | ~spl9_113 | ~spl9_265)),
  inference(subsumption_resolution,[],[f57897,f4119])).
fof(f57897,plain,(
  sK8 != lazy_impl(true,sK8) | (~spl9_47 | spl9_83 | ~spl9_84 | ~spl9_87 | ~spl9_108 | ~spl9_112 | ~spl9_113 | ~spl9_265)),
  inference(forward_demodulation,[],[f57896,f4119])).
fof(f57896,plain,(
  sK8 != lazy_impl(true,lazy_impl(true,sK8)) | (~spl9_47 | spl9_83 | ~spl9_84 | ~spl9_87 | ~spl9_108 | ~spl9_112 | ~spl9_265)),
  inference(forward_demodulation,[],[f57844,f15600])).
fof(f57844,plain,(
  sK8 != lazy_impl(true,lazy_impl(true,impl(true,sK8))) | (~spl9_47 | spl9_83 | ~spl9_84 | ~spl9_87 | ~spl9_108 | ~spl9_112)),
  inference(backward_demodulation,[],[f2864,f57568])).
fof(f57568,plain,(
  sK8 != lazy_impl(true,lazy_impl(true,impl(sK2(true,sK8),sK8))) | (~spl9_47 | spl9_83 | ~spl9_84 | ~spl9_108 | ~spl9_112)),
  inference(backward_demodulation,[],[f2851,f57533])).
fof(f57533,plain,(
  sK8 != lazy_impl(true,lazy_impl(prop(sK2(true,sK8)),impl(sK2(true,sK8),sK8))) | (~spl9_47 | spl9_83 | ~spl9_108 | ~spl9_112)),
  inference(backward_demodulation,[],[f15921,f57508])).
fof(f57508,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(sK2(true,sK8)),impl(sK2(true,sK8),sK8))) | (spl9_83 | ~spl9_108 | ~spl9_112)),
  inference(forward_demodulation,[],[f57503,f15924])).
fof(f15924,plain,(
  ( ! [X6] : (sK8 = impl(sK8,X6)) ) | (~spl9_108 | ~spl9_112)),
  inference(backward_demodulation,[],[f4114,f4094])).
fof(f57503,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(sK2(true,sK8)),impl(sK2(true,sK8),impl(sK8,sK2(true,sK8))))) | (spl9_83 | ~spl9_108 | ~spl9_112)),
  inference(backward_demodulation,[],[f15924,f2730])).
fof(f15921,plain,(
  sK8 = or1(true,sK8) | (~spl9_47 | ~spl9_112)),
  inference(backward_demodulation,[],[f4114,f3097])).
fof(f3097,plain,(
  not1(sK8) = or1(true,sK8) | ~spl9_47),
  inference(forward_demodulation,[],[f3076,f3071])).
fof(f3076,plain,(
  lazy_impl(true,sK8) = or1(true,sK8) | ~spl9_47),
  inference(resolution,[],[f2970,f546])).
fof(f57496,plain,(
  spl9_458),
  inference(avatar_split_clause,[],[f270,f50398])).
fof(f270,plain,(
  true = impl(false1,true)),
  inference(resolution,[],[f219,f242])).
fof(f57475,plain,(
  spl9_520),
  inference(avatar_split_clause,[],[f276,f57344])).
fof(f276,plain,(
  true = or1(false1,true)),
  inference(resolution,[],[f232,f242])).
fof(f232,plain,(
  ( ! [X0] : (~bool(X0) | or1(false1,X0) = X0) )),
  inference(definition_unfolding,[],[f177,f189])).
fof(f177,plain,(
  ( ! [X0] : (or1(false,X0) = X0 | ~bool(X0)) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : (or1(false,X0) = X0 | ~bool(X0))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : (bool(X0) => or1(false,X0) = X0)),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X3] : (bool(X3) => or1(false,X3) = X3)),
  file('./p/SWW098+1.tptp',unknown)).
fof(f51382,plain,(
  ~spl9_463),
  inference(avatar_split_clause,[],[f207,f50619])).
fof(f207,plain,(
  err != false1),
  inference(definition_unfolding,[],[f136,f189])).
fof(f136,plain,(
  false != err),
  inference(cnf_transformation,[],[f2])).
fof(f50439,plain,(
  ~spl9_82 | ~spl9_84),
  inference(avatar_contradiction_clause,[],[f50438])).
fof(f50438,plain,(
  $false | (~spl9_82 | ~spl9_84)),
  inference(subsumption_resolution,[],[f50437,f208])).
fof(f50437,plain,(
  true = false1 | (~spl9_82 | ~spl9_84)),
  inference(backward_demodulation,[],[f2851,f2726])).
fof(f46895,plain,(
  spl9_34 | ~spl9_69 | spl9_70),
  inference(avatar_split_clause,[],[f46893,f2172,f2167,f1190])).
fof(f46893,plain,(
  true = sK7 | (~spl9_69 | spl9_70)),
  inference(backward_demodulation,[],[f2168,f7094])).
fof(f2168,plain,(
  true = lazy_impl(true,sK7) | ~spl9_69),
  inference(avatar_component_clause,[],[f2167])).
fof(f46882,plain,(
  ~spl9_82 | ~spl9_87),
  inference(avatar_contradiction_clause,[],[f46881])).
fof(f46881,plain,(
  $false | (~spl9_82 | ~spl9_87)),
  inference(subsumption_resolution,[],[f46880,f208])).
fof(f46880,plain,(
  true = false1 | (~spl9_82 | ~spl9_87)),
  inference(forward_demodulation,[],[f46879,f248])).
fof(f46879,plain,(
  false1 = prop(true) | (~spl9_82 | ~spl9_87)),
  inference(forward_demodulation,[],[f2726,f2864])).
fof(f44479,plain,(
  spl9_70 | ~spl9_204),
  inference(avatar_contradiction_clause,[],[f44478])).
fof(f44478,plain,(
  $false | (spl9_70 | ~spl9_204)),
  inference(subsumption_resolution,[],[f44477,f2174])).
fof(f44477,plain,(
  err = lazy_impl(true,sK7) | (spl9_70 | ~spl9_204)),
  inference(forward_demodulation,[],[f44467,f7094])).
fof(f44467,plain,(
  err = lazy_impl(true,lazy_impl(true,sK7)) | ~spl9_204),
  inference(backward_demodulation,[],[f44452,f44436])).
fof(f44436,plain,(
  ( ! [X0] : (err = lazy_impl(true,lazy_impl(prop(sK2(sK7,X0)),sK7))) ) | ~spl9_204),
  inference(resolution,[],[f10035,f213])).
fof(f10035,plain,(
  ( ! [X40] : (~d(lazy_impl(prop(sK2(sK7,X40)),sK7))) ) | ~spl9_204),
  inference(avatar_component_clause,[],[f10034])).
fof(f10034,plain,(
  spl9_204 <=> ! [X40] : ~d(lazy_impl(prop(sK2(sK7,X40)),sK7))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_204])])).
fof(f44452,plain,(
  ( ! [X1] : (true = prop(sK2(sK7,X1))) ) | ~spl9_204),
  inference(subsumption_resolution,[],[f44451,f137])).
fof(f44451,plain,(
  ( ! [X1] : (~d(true) | true = prop(sK2(sK7,X1))) ) | ~spl9_204),
  inference(forward_demodulation,[],[f44438,f221])).
fof(f44438,plain,(
  ( ! [X1] : (~d(lazy_impl(false1,sK7)) | true = prop(sK2(sK7,X1))) ) | ~spl9_204),
  inference(superposition,[],[f10035,f252])).
fof(f44284,plain,(
  spl9_70 | ~spl9_233 | ~spl9_420 | ~spl9_432),
  inference(avatar_contradiction_clause,[],[f44283])).
fof(f44283,plain,(
  $false | (spl9_70 | ~spl9_233 | ~spl9_420 | ~spl9_432)),
  inference(subsumption_resolution,[],[f44282,f39967])).
fof(f39967,plain,(
  ~bool(sK7) | (~spl9_233 | ~spl9_420)),
  inference(trivial_inequality_removal,[],[f39958])).
fof(f39958,plain,(
  false1 != false1 | ~bool(sK7) | (~spl9_233 | ~spl9_420)),
  inference(superposition,[],[f216,f39909])).
fof(f39909,plain,(
  false1 = prop(sK7) | (~spl9_233 | ~spl9_420)),
  inference(backward_demodulation,[],[f39734,f11773])).
fof(f44282,plain,(
  bool(sK7) | (spl9_70 | ~spl9_432)),
  inference(forward_demodulation,[],[f44258,f7094])).
fof(f44258,plain,(
  bool(lazy_impl(true,sK7)) | ~spl9_432),
  inference(superposition,[],[f43943,f249])).
fof(f43943,plain,(
  ( ! [X31] : (bool(lazy_impl(prop(X31),sK7))) ) | ~spl9_432),
  inference(avatar_component_clause,[],[f43942])).
fof(f43942,plain,(
  spl9_432 <=> ! [X31] : bool(lazy_impl(prop(X31),sK7))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_432])])).
fof(f43944,plain,(
  spl9_202 | spl9_432 | ~spl9_257 | ~spl9_262),
  inference(avatar_split_clause,[],[f43940,f15326,f13490,f43942,f10006])).
fof(f10006,plain,(
  spl9_202 <=> ! [X25] : ~d(lazy_impl(prop(sK0(sK7,X25)),sK7))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_202])])).
fof(f13490,plain,(
  spl9_257 <=> ! [X34] : (~d(lazy_impl(prop(X34),sK7)) | bool(lazy_impl(prop(X34),sK7)))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_257])])).
fof(f43940,plain,(
  ( ! [X30,X31] : (bool(lazy_impl(prop(X31),sK7)) | ~d(lazy_impl(prop(sK0(sK7,X30)),sK7))) ) | (~spl9_257 | ~spl9_262)),
  inference(subsumption_resolution,[],[f43939,f13491])).
fof(f13491,plain,(
  ( ! [X34] : (~d(lazy_impl(prop(X34),sK7)) | bool(lazy_impl(prop(X34),sK7))) ) | ~spl9_257),
  inference(avatar_component_clause,[],[f13490])).
fof(f43939,plain,(
  ( ! [X30,X31] : (~bool(lazy_impl(prop(sK0(sK7,X30)),sK7)) | bool(lazy_impl(prop(X31),sK7)) | ~d(lazy_impl(prop(sK0(sK7,X30)),sK7))) ) | ~spl9_262),
  inference(forward_demodulation,[],[f43938,f15327])).
fof(f43938,plain,(
  ( ! [X30,X31] : (~bool(lazy_impl(prop(sK0(sK7,X30)),impl(sK7,sK0(sK7,X30)))) | bool(lazy_impl(prop(X31),sK7)) | ~d(lazy_impl(prop(sK0(sK7,X30)),sK7))) ) | ~spl9_262),
  inference(forward_demodulation,[],[f43937,f15327])).
fof(f43937,plain,(
  ( ! [X30,X31] : (bool(lazy_impl(prop(X31),sK7)) | ~d(lazy_impl(prop(sK0(sK7,X30)),sK7)) | ~bool(lazy_impl(prop(sK0(sK7,X30)),impl(impl(sK7,impl(X30,sK0(sK7,X30))),sK0(sK7,X30))))) ) | ~spl9_262),
  inference(forward_demodulation,[],[f43936,f15327])).
fof(f43936,plain,(
  ( ! [X30,X31] : (bool(lazy_impl(prop(X31),impl(sK7,X31))) | ~d(lazy_impl(prop(sK0(sK7,X30)),sK7)) | ~bool(lazy_impl(prop(sK0(sK7,X30)),impl(impl(sK7,impl(X30,sK0(sK7,X30))),sK0(sK7,X30))))) ) | ~spl9_262),
  inference(forward_demodulation,[],[f43033,f15327])).
fof(f43033,plain,(
  ( ! [X30,X31] : (~d(lazy_impl(prop(sK0(sK7,X30)),sK7)) | bool(lazy_impl(prop(X31),impl(impl(sK7,impl(X30,X31)),X31))) | ~bool(lazy_impl(prop(sK0(sK7,X30)),impl(impl(sK7,impl(X30,sK0(sK7,X30))),sK0(sK7,X30))))) ) | ~spl9_262),
  inference(forward_demodulation,[],[f42991,f15327])).
fof(f42991,plain,(
  ( ! [X30,X31] : (~d(lazy_impl(prop(sK0(sK7,X30)),impl(sK7,sK0(sK7,X30)))) | bool(lazy_impl(prop(X31),impl(impl(sK7,impl(X30,X31)),X31))) | ~bool(lazy_impl(prop(sK0(sK7,X30)),impl(impl(sK7,impl(X30,sK0(sK7,X30))),sK0(sK7,X30))))) ) | ~spl9_262),
  inference(superposition,[],[f1255,f15327])).
fof(f1255,plain,(
  ( ! [X2,X0,X1] : (~d(lazy_impl(prop(sK0(X0,X1)),impl(impl(X0,impl(X1,sK0(X0,X1))),sK0(X0,X1)))) | bool(lazy_impl(prop(X2),impl(impl(X0,impl(X1,X2)),X2))) | ~bool(lazy_impl(prop(sK0(X0,X1)),impl(impl(X0,impl(X1,sK0(X0,X1))),sK0(X0,X1))))) )),
  inference(subsumption_resolution,[],[f700,f701])).
fof(f701,plain,(
  ( ! [X4,X5,X3] : (~d(lazy_impl(prop(sK0(X3,X4)),impl(impl(X3,impl(X4,sK0(X3,X4))),sK0(X3,X4)))) | d(lazy_impl(prop(X5),impl(impl(X3,impl(X4,X5)),X5)))) )),
  inference(resolution,[],[f225,f140])).
fof(f225,plain,(
  ( ! [X0,X3,X1] : (~forallprefers(lazy_impl(prop(X3),impl(impl(X0,impl(X1,X3)),X3)),lazy_impl(prop(sK0(X0,X1)),impl(impl(X0,impl(X1,sK0(X0,X1))),sK0(X0,X1))))) )),
  inference(definition_unfolding,[],[f167,f165,f165])).
fof(f165,plain,(
  ( ! [X2,X0,X1] : (f1(X0,X1,X2) = lazy_impl(prop(X2),impl(impl(X0,impl(X1,X2)),X2))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : f1(X0,X1,X2) = lazy_impl(prop(X2),impl(impl(X0,impl(X1,X2)),X2))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X4,X5,X6] : f1(X4,X5,X6) = lazy_impl(prop(X6),impl(impl(X4,impl(X5,X6)),X6))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f167,plain,(
  ( ! [X0,X3,X1] : (~forallprefers(f1(X0,X1,X3),f1(X0,X1,sK0(X0,X1)))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : (! [X3] : ~forallprefers(f1(X0,X1,X3),f1(X0,X1,sK0(X0,X1))) & and2(X0,X1) = phi(f1(X0,X1,sK0(X0,X1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f96,f115])).
fof(f115,plain,(
  ! [X1,X0] : (? [X2] : (! [X3] : ~forallprefers(f1(X0,X1,X3),f1(X0,X1,X2)) & and2(X0,X1) = phi(f1(X0,X1,X2))) => (! [X3] : ~forallprefers(f1(X0,X1,X3),f1(X0,X1,sK0(X0,X1))) & and2(X0,X1) = phi(f1(X0,X1,sK0(X0,X1)))))),
  introduced(choice_axiom,[])).
fof(f96,plain,(
  ! [X0,X1] : ? [X2] : (! [X3] : ~forallprefers(f1(X0,X1,X3),f1(X0,X1,X2)) & and2(X0,X1) = phi(f1(X0,X1,X2)))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ? [X2] : (~? [X3] : forallprefers(f1(X0,X1,X3),f1(X0,X1,X2)) & and2(X0,X1) = phi(f1(X0,X1,X2)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X4,X5] : ? [X6] : (~? [X7] : forallprefers(f1(X4,X5,X7),f1(X4,X5,X6)) & and2(X4,X5) = phi(f1(X4,X5,X6)))),
  file('./p/SWW098+1.tptp',unknown)).
fof(f700,plain,(
  ( ! [X2,X0,X1] : (~bool(lazy_impl(prop(sK0(X0,X1)),impl(impl(X0,impl(X1,sK0(X0,X1))),sK0(X0,X1)))) | bool(lazy_impl(prop(X2),impl(impl(X0,impl(X1,X2)),X2))) | ~d(lazy_impl(prop(sK0(X0,X1)),impl(impl(X0,impl(X1,sK0(X0,X1))),sK0(X0,X1)))) | ~d(lazy_impl(prop(X2),impl(impl(X0,impl(X1,X2)),X2)))) )),
  inference(resolution,[],[f225,f141])).
fof(f43919,plain,(
  spl9_70 | ~spl9_202),
  inference(avatar_contradiction_clause,[],[f43918])).
fof(f43918,plain,(
  $false | (spl9_70 | ~spl9_202)),
  inference(subsumption_resolution,[],[f43917,f2174])).
fof(f43917,plain,(
  err = lazy_impl(true,sK7) | (spl9_70 | ~spl9_202)),
  inference(forward_demodulation,[],[f43914,f7094])).
fof(f43914,plain,(
  err = lazy_impl(true,lazy_impl(true,sK7)) | ~spl9_202),
  inference(backward_demodulation,[],[f43913,f43899])).
fof(f43899,plain,(
  ( ! [X0] : (err = lazy_impl(true,lazy_impl(prop(sK0(sK7,X0)),sK7))) ) | ~spl9_202),
  inference(resolution,[],[f10007,f213])).
fof(f10007,plain,(
  ( ! [X25] : (~d(lazy_impl(prop(sK0(sK7,X25)),sK7))) ) | ~spl9_202),
  inference(avatar_component_clause,[],[f10006])).
fof(f43913,plain,(
  ( ! [X1] : (true = prop(sK0(sK7,X1))) ) | ~spl9_202),
  inference(subsumption_resolution,[],[f43912,f137])).
fof(f43912,plain,(
  ( ! [X1] : (~d(true) | true = prop(sK0(sK7,X1))) ) | ~spl9_202),
  inference(forward_demodulation,[],[f43901,f221])).
fof(f43901,plain,(
  ( ! [X1] : (~d(lazy_impl(false1,sK7)) | true = prop(sK0(sK7,X1))) ) | ~spl9_202),
  inference(superposition,[],[f10007,f252])).
fof(f43463,plain,(
  spl9_13 | spl9_173),
  inference(avatar_contradiction_clause,[],[f43462])).
fof(f43462,plain,(
  $false | (spl9_13 | spl9_173)),
  inference(subsumption_resolution,[],[f43461,f637])).
fof(f43461,plain,(
  true = prop(sK2(sK7,sK8)) | spl9_173),
  inference(subsumption_resolution,[],[f43448,f221])).
fof(f43448,plain,(
  true != lazy_impl(false1,sK7) | true = prop(sK2(sK7,sK8)) | spl9_173),
  inference(superposition,[],[f8416,f252])).
fof(f8416,plain,(
  true != lazy_impl(prop(sK2(sK7,sK8)),sK7) | spl9_173),
  inference(avatar_component_clause,[],[f8415])).
fof(f43460,plain,(
  spl9_32 | spl9_173 | spl9_182),
  inference(avatar_contradiction_clause,[],[f43459])).
fof(f43459,plain,(
  $false | (spl9_32 | spl9_173 | spl9_182)),
  inference(subsumption_resolution,[],[f43458,f8568])).
fof(f43458,plain,(
  false1 = sK2(sK7,sK8) | (spl9_32 | spl9_173)),
  inference(subsumption_resolution,[],[f43457,f1182])).
fof(f43457,plain,(
  true = sK2(sK7,sK8) | false1 = sK2(sK7,sK8) | spl9_173),
  inference(subsumption_resolution,[],[f43447,f221])).
fof(f43447,plain,(
  true != lazy_impl(false1,sK7) | true = sK2(sK7,sK8) | false1 = sK2(sK7,sK8) | spl9_173),
  inference(superposition,[],[f8416,f282])).
fof(f43093,plain,(
  ~spl9_47 | ~spl9_113 | ~spl9_167 | ~spl9_427),
  inference(avatar_contradiction_clause,[],[f43092])).
fof(f43092,plain,(
  $false | (~spl9_47 | ~spl9_113 | ~spl9_167 | ~spl9_427)),
  inference(subsumption_resolution,[],[f43091,f2970])).
fof(f43091,plain,(
  bool(sK8) | (~spl9_113 | ~spl9_167 | ~spl9_427)),
  inference(forward_demodulation,[],[f43065,f4119])).
fof(f43065,plain,(
  bool(lazy_impl(true,sK8)) | (~spl9_167 | ~spl9_427)),
  inference(superposition,[],[f43042,f8121])).
fof(f43042,plain,(
  ( ! [X33] : (bool(lazy_impl(prop(X33),sK8))) ) | ~spl9_427),
  inference(avatar_component_clause,[],[f43041])).
fof(f39742,plain,(
  spl9_233 | ~spl9_114 | ~spl9_232),
  inference(avatar_split_clause,[],[f39741,f11755,f4966,f11771])).
fof(f39741,plain,(
  prop(sK7) = lazy_impl(true,prop(sK7)) | (~spl9_114 | ~spl9_232)),
  inference(subsumption_resolution,[],[f39705,f6849])).
fof(f39705,plain,(
  bool(err) | prop(sK7) = lazy_impl(true,prop(sK7)) | ~spl9_232),
  inference(superposition,[],[f11757,f212])).
fof(f11757,plain,(
  bool(lazy_impl(true,prop(sK7))) | ~spl9_232),
  inference(avatar_component_clause,[],[f11755])).
fof(f37601,plain,(
  spl9_33 | spl9_57 | spl9_70 | ~spl9_108 | ~spl9_112 | ~spl9_262),
  inference(avatar_contradiction_clause,[],[f37600])).
fof(f37600,plain,(
  $false | (spl9_33 | spl9_57 | spl9_70 | ~spl9_108 | ~spl9_112 | ~spl9_262)),
  inference(subsumption_resolution,[],[f37599,f1889])).
fof(f1889,plain,(
  true != prop(sK7) | spl9_57),
  inference(avatar_component_clause,[],[f1888])).
fof(f37599,plain,(
  true = prop(sK7) | (spl9_33 | spl9_70 | ~spl9_108 | ~spl9_112 | ~spl9_262)),
  inference(subsumption_resolution,[],[f37598,f7094])).
fof(f37598,plain,(
  sK7 != lazy_impl(true,sK7) | true = prop(sK7) | (spl9_33 | spl9_70 | ~spl9_108 | ~spl9_112 | ~spl9_262)),
  inference(superposition,[],[f36894,f346])).
fof(f346,plain,(
  ( ! [X0,X1] : (lazy_impl(true,X0) = or1(X0,X1) | true = prop(X0)) )),
  inference(resolution,[],[f230,f151])).
fof(f36894,plain,(
  sK7 != or1(sK7,sK8) | (spl9_33 | spl9_70 | ~spl9_108 | ~spl9_112 | ~spl9_262)),
  inference(forward_demodulation,[],[f36893,f7094])).
fof(f36893,plain,(
  or1(sK7,sK8) != lazy_impl(true,sK7) | (spl9_33 | spl9_70 | ~spl9_108 | ~spl9_112 | ~spl9_262)),
  inference(forward_demodulation,[],[f36892,f7094])).
fof(f36892,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,sK7)) | (spl9_33 | ~spl9_108 | ~spl9_112 | ~spl9_262)),
  inference(forward_demodulation,[],[f36891,f15327])).
fof(f36891,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(sK7,sK8))) | (spl9_33 | ~spl9_108 | ~spl9_112 | ~spl9_262)),
  inference(forward_demodulation,[],[f36553,f15327])).
fof(f36553,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(impl(sK7,false1),sK8))) | (spl9_33 | ~spl9_108 | ~spl9_112)),
  inference(forward_demodulation,[],[f36552,f15924])).
fof(f36552,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(impl(sK7,false1),impl(sK8,false1)))) | (spl9_33 | ~spl9_108 | ~spl9_112)),
  inference(forward_demodulation,[],[f1187,f15924])).
fof(f1187,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(impl(sK7,false1),impl(impl(sK8,false1),false1)))) | spl9_33),
  inference(avatar_component_clause,[],[f1185])).
fof(f35065,plain,(
  spl9_406),
  inference(avatar_split_clause,[],[f277,f34986])).
fof(f277,plain,(
  false1 = or1(false1,false1)),
  inference(resolution,[],[f232,f243])).
fof(f30272,plain,(
  ~spl9_283 | spl9_289),
  inference(avatar_contradiction_clause,[],[f30271])).
fof(f30271,plain,(
  $false | (~spl9_283 | spl9_289)),
  inference(subsumption_resolution,[],[f30270,f269])).
fof(f269,plain,(
  false1 = lazy_impl(true,false1)),
  inference(resolution,[],[f214,f209])).
fof(f209,plain,(
  d(false1)),
  inference(definition_unfolding,[],[f138,f189])).
fof(f138,plain,(
  d(false)),
  inference(cnf_transformation,[],[f3])).
fof(f30270,plain,(
  false1 != lazy_impl(true,false1) | (~spl9_283 | spl9_289)),
  inference(forward_demodulation,[],[f16429,f16385])).
fof(f16429,plain,(
  false1 != lazy_impl(true,sK6) | spl9_289),
  inference(avatar_component_clause,[],[f16428])).
fof(f16428,plain,(
  spl9_289 <=> false1 = lazy_impl(true,sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_289])])).
fof(f30237,plain,(
  ~spl9_6 | ~spl9_283),
  inference(avatar_contradiction_clause,[],[f30236])).
fof(f30236,plain,(
  $false | (~spl9_6 | ~spl9_283)),
  inference(subsumption_resolution,[],[f30235,f207])).
fof(f30235,plain,(
  err = false1 | (~spl9_6 | ~spl9_283)),
  inference(backward_demodulation,[],[f30227,f269])).
fof(f30227,plain,(
  err = lazy_impl(true,false1) | (~spl9_6 | ~spl9_283)),
  inference(backward_demodulation,[],[f16385,f459])).
fof(f459,plain,(
  err = lazy_impl(true,sK6) | ~spl9_6),
  inference(avatar_component_clause,[],[f457])).
fof(f30225,plain,(
  spl9_283 | ~spl9_1 | spl9_275),
  inference(avatar_split_clause,[],[f30224,f16076,f335,f16383])).
fof(f16076,plain,(
  spl9_275 <=> true = sK6),
  introduced(avatar_definition,[new_symbols(naming,[spl9_275])])).
fof(f30224,plain,(
  false1 = sK6 | (~spl9_1 | spl9_275)),
  inference(subsumption_resolution,[],[f30221,f16077])).
fof(f16077,plain,(
  true != sK6 | spl9_275),
  inference(avatar_component_clause,[],[f16076])).
fof(f30221,plain,(
  false1 = sK6 | true = sK6 | ~spl9_1),
  inference(resolution,[],[f30213,f206])).
fof(f30213,plain,(
  bool(sK6) | ~spl9_1),
  inference(trivial_inequality_removal,[],[f30212])).
fof(f30212,plain,(
  true != true | bool(sK6) | ~spl9_1),
  inference(superposition,[],[f150,f337])).
fof(f21239,plain,(
  spl9_5 | ~spl9_7 | ~spl9_275),
  inference(avatar_split_clause,[],[f21238,f16076,f461,f372])).
fof(f372,plain,(
  spl9_5 <=> ! [X0] : ~forallprefers(lazy_impl(prop(X0),X0),true)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_5])])).
fof(f461,plain,(
  spl9_7 <=> ! [X0] : ~forallprefers(lazy_impl(prop(X0),X0),sK6)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_7])])).
fof(f21238,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),X0),true)) ) | (~spl9_7 | ~spl9_275)),
  inference(forward_demodulation,[],[f462,f16078])).
fof(f16078,plain,(
  true = sK6 | ~spl9_275),
  inference(avatar_component_clause,[],[f16076])).
fof(f462,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),X0),sK6)) ) | ~spl9_7),
  inference(avatar_component_clause,[],[f461])).
fof(f21231,plain,(
  ~spl9_6 | ~spl9_275),
  inference(avatar_contradiction_clause,[],[f21230])).
fof(f21230,plain,(
  $false | (~spl9_6 | ~spl9_275)),
  inference(subsumption_resolution,[],[f21229,f135])).
fof(f21229,plain,(
  true = err | (~spl9_6 | ~spl9_275)),
  inference(forward_demodulation,[],[f21228,f267])).
fof(f21228,plain,(
  err = lazy_impl(true,true) | (~spl9_6 | ~spl9_275)),
  inference(forward_demodulation,[],[f459,f16078])).
fof(f16508,plain,(
  ~spl9_283 | spl9_8 | ~spl9_289),
  inference(avatar_split_clause,[],[f16500,f16428,f465,f16383])).
fof(f16500,plain,(
  false1 != sK6 | (spl9_8 | ~spl9_289)),
  inference(backward_demodulation,[],[f16430,f466])).
fof(f466,plain,(
  sK6 != lazy_impl(true,sK6) | spl9_8),
  inference(avatar_component_clause,[],[f465])).
fof(f16430,plain,(
  false1 = lazy_impl(true,sK6) | ~spl9_289),
  inference(avatar_component_clause,[],[f16428])).
fof(f15956,plain,(
  ~spl9_47 | spl9_92 | ~spl9_108 | ~spl9_112 | ~spl9_113),
  inference(avatar_contradiction_clause,[],[f15955])).
fof(f15955,plain,(
  $false | (~spl9_47 | spl9_92 | ~spl9_108 | ~spl9_112 | ~spl9_113)),
  inference(subsumption_resolution,[],[f15954,f4119])).
fof(f15954,plain,(
  sK8 != lazy_impl(true,sK8) | (~spl9_47 | spl9_92 | ~spl9_108 | ~spl9_112)),
  inference(forward_demodulation,[],[f15953,f15918])).
fof(f15918,plain,(
  sK8 = impl(false1,sK8) | (~spl9_47 | ~spl9_112)),
  inference(backward_demodulation,[],[f4114,f3094])).
fof(f3094,plain,(
  not1(sK8) = impl(false1,sK8) | ~spl9_47),
  inference(forward_demodulation,[],[f3073,f3071])).
fof(f3073,plain,(
  lazy_impl(true,sK8) = impl(false1,sK8) | ~spl9_47),
  inference(resolution,[],[f2970,f525])).
fof(f15953,plain,(
  impl(false1,sK8) != lazy_impl(true,impl(false1,sK8)) | (spl9_92 | ~spl9_108 | ~spl9_112)),
  inference(forward_demodulation,[],[f15952,f15924])).
fof(f15952,plain,(
  impl(false1,impl(sK8,false1)) != lazy_impl(true,impl(false1,impl(sK8,false1))) | (spl9_92 | ~spl9_112)),
  inference(forward_demodulation,[],[f3151,f4114])).
fof(f3151,plain,(
  impl(false1,impl(not1(sK8),false1)) != lazy_impl(true,impl(false1,impl(not1(sK8),false1))) | spl9_92),
  inference(avatar_component_clause,[],[f3150])).
fof(f15948,plain,(
  ~spl9_265 | ~spl9_95 | spl9_98 | ~spl9_112),
  inference(avatar_split_clause,[],[f15932,f4112,f3177,f3163,f15598])).
fof(f15932,plain,(
  sK8 != impl(true,sK8) | (~spl9_95 | spl9_98 | ~spl9_112)),
  inference(backward_demodulation,[],[f4114,f15865])).
fof(f15865,plain,(
  not1(sK8) != impl(true,not1(sK8)) | (~spl9_95 | spl9_98)),
  inference(forward_demodulation,[],[f3178,f3164])).
fof(f3178,plain,(
  impl(false1,impl(not1(sK8),false1)) != impl(true,impl(false1,impl(not1(sK8),false1))) | spl9_98),
  inference(avatar_component_clause,[],[f3177])).
fof(f15939,plain,(
  spl9_107 | ~spl9_47 | ~spl9_112),
  inference(avatar_split_clause,[],[f15922,f4112,f1576,f4089])).
fof(f15922,plain,(
  sK8 = or1(false1,sK8) | (~spl9_47 | ~spl9_112)),
  inference(backward_demodulation,[],[f4114,f3105])).
fof(f3105,plain,(
  or1(false1,sK8) = not1(sK8) | ~spl9_47),
  inference(forward_demodulation,[],[f3077,f3071])).
fof(f3077,plain,(
  or1(false1,sK8) = lazy_impl(true,sK8) | ~spl9_47),
  inference(resolution,[],[f2970,f547])).
fof(f547,plain,(
  ( ! [X1] : (bool(X1) | lazy_impl(true,X1) = or1(false1,X1)) )),
  inference(resolution,[],[f231,f243])).
fof(f15937,plain,(
  spl9_265 | ~spl9_47 | ~spl9_112),
  inference(avatar_split_clause,[],[f15917,f4112,f1576,f15598])).
fof(f15917,plain,(
  sK8 = impl(true,sK8) | (~spl9_47 | ~spl9_112)),
  inference(backward_demodulation,[],[f4114,f3093])).
fof(f3093,plain,(
  not1(sK8) = impl(true,sK8) | ~spl9_47),
  inference(forward_demodulation,[],[f3072,f3071])).
fof(f3072,plain,(
  lazy_impl(true,sK8) = impl(true,sK8) | ~spl9_47),
  inference(resolution,[],[f2970,f524])).
fof(f15793,plain,(
  spl9_266 | spl9_59 | ~spl9_111),
  inference(avatar_split_clause,[],[f15792,f4108,f1897,f15602])).
fof(f15792,plain,(
  ( ! [X0] : (err = impl(sK8,X0)) ) | (spl9_59 | ~spl9_111)),
  inference(subsumption_resolution,[],[f4723,f1898])).
fof(f4723,plain,(
  ( ! [X0] : (err = impl(sK8,X0) | true = prop(sK8)) ) | ~spl9_111),
  inference(superposition,[],[f4110,f291])).
fof(f15787,plain,(
  spl9_266 | spl9_85 | ~spl9_111),
  inference(avatar_split_clause,[],[f15786,f4108,f2853,f15602])).
fof(f15786,plain,(
  ( ! [X3] : (err = impl(sK8,X3)) ) | (spl9_85 | ~spl9_111)),
  inference(subsumption_resolution,[],[f4726,f2855])).
fof(f4726,plain,(
  ( ! [X3] : (err = impl(sK8,X3) | true = or1(true,sK8)) ) | ~spl9_111),
  inference(superposition,[],[f4110,f294])).
fof(f15752,plain,(
  spl9_271 | spl9_266 | spl9_37 | ~spl9_111),
  inference(avatar_split_clause,[],[f15751,f4108,f1207,f15602,f15747])).
fof(f15751,plain,(
  ( ! [X60] : (err = impl(sK8,X60) | false1 = sK8) ) | (spl9_37 | ~spl9_111)),
  inference(subsumption_resolution,[],[f5162,f1208])).
fof(f1208,plain,(
  true != sK8 | spl9_37),
  inference(avatar_component_clause,[],[f1207])).
fof(f5162,plain,(
  ( ! [X60] : (err = impl(sK8,X60) | false1 = sK8 | true = sK8) ) | ~spl9_111),
  inference(superposition,[],[f295,f4110])).
fof(f15499,plain,(
  spl9_1 | spl9_5),
  inference(avatar_split_clause,[],[f410,f372,f335])).
fof(f410,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),X0),true) | true = prop(sK6)) )),
  inference(forward_demodulation,[],[f392,f221])).
fof(f392,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),X0),lazy_impl(false1,sK6)) | true = prop(sK6)) )),
  inference(superposition,[],[f236,f252])).
fof(f236,plain,(
  ( ! [X1] : (~forallprefers(lazy_impl(prop(X1),X1),lazy_impl(prop(sK6),sK6))) )),
  inference(definition_unfolding,[],[f192,f190,f190])).
fof(f192,plain,(
  ( ! [X1] : (~forallprefers(f7(X1),f7(sK6))) )),
  inference(cnf_transformation,[],[f128])).
fof(f13492,plain,(
  spl9_256 | spl9_257 | spl9_57 | spl9_70),
  inference(avatar_split_clause,[],[f13485,f2172,f1888,f13490,f13487])).
fof(f13485,plain,(
  ( ! [X33,X34] : (~d(lazy_impl(prop(X34),sK7)) | ~bool(lazy_impl(prop(sK2(sK7,X33)),sK7)) | bool(lazy_impl(prop(X34),sK7)) | ~d(lazy_impl(prop(sK2(sK7,X33)),sK7))) ) | (spl9_57 | spl9_70)),
  inference(forward_demodulation,[],[f13484,f8300])).
fof(f8300,plain,(
  ( ! [X0] : (sK7 = impl(sK7,X0)) ) | (spl9_57 | spl9_70)),
  inference(subsumption_resolution,[],[f8256,f1889])).
fof(f8256,plain,(
  ( ! [X0] : (sK7 = impl(sK7,X0) | true = prop(sK7)) ) | spl9_70),
  inference(superposition,[],[f7094,f291])).
fof(f13484,plain,(
  ( ! [X33,X34] : (~d(lazy_impl(prop(X34),impl(sK7,impl(impl(X33,X34),X34)))) | ~bool(lazy_impl(prop(sK2(sK7,X33)),sK7)) | bool(lazy_impl(prop(X34),sK7)) | ~d(lazy_impl(prop(sK2(sK7,X33)),sK7))) ) | (spl9_57 | spl9_70)),
  inference(forward_demodulation,[],[f13483,f8300])).
fof(f13483,plain,(
  ( ! [X33,X34] : (~bool(lazy_impl(prop(sK2(sK7,X33)),sK7)) | bool(lazy_impl(prop(X34),sK7)) | ~d(lazy_impl(prop(sK2(sK7,X33)),sK7)) | ~d(lazy_impl(prop(X34),impl(impl(sK7,X34),impl(impl(X33,X34),X34))))) ) | (spl9_57 | spl9_70)),
  inference(forward_demodulation,[],[f13482,f8300])).
fof(f13482,plain,(
  ( ! [X33,X34] : (~bool(lazy_impl(prop(sK2(sK7,X33)),impl(sK7,impl(impl(X33,sK2(sK7,X33)),sK2(sK7,X33))))) | bool(lazy_impl(prop(X34),sK7)) | ~d(lazy_impl(prop(sK2(sK7,X33)),sK7)) | ~d(lazy_impl(prop(X34),impl(impl(sK7,X34),impl(impl(X33,X34),X34))))) ) | (spl9_57 | spl9_70)),
  inference(forward_demodulation,[],[f13481,f8300])).
fof(f13481,plain,(
  ( ! [X33,X34] : (bool(lazy_impl(prop(X34),sK7)) | ~d(lazy_impl(prop(sK2(sK7,X33)),sK7)) | ~bool(lazy_impl(prop(sK2(sK7,X33)),impl(impl(sK7,sK2(sK7,X33)),impl(impl(X33,sK2(sK7,X33)),sK2(sK7,X33))))) | ~d(lazy_impl(prop(X34),impl(impl(sK7,X34),impl(impl(X33,X34),X34))))) ) | (spl9_57 | spl9_70)),
  inference(forward_demodulation,[],[f13480,f8300])).
fof(f13480,plain,(
  ( ! [X33,X34] : (bool(lazy_impl(prop(X34),impl(sK7,impl(impl(X33,X34),X34)))) | ~d(lazy_impl(prop(sK2(sK7,X33)),sK7)) | ~bool(lazy_impl(prop(sK2(sK7,X33)),impl(impl(sK7,sK2(sK7,X33)),impl(impl(X33,sK2(sK7,X33)),sK2(sK7,X33))))) | ~d(lazy_impl(prop(X34),impl(impl(sK7,X34),impl(impl(X33,X34),X34))))) ) | (spl9_57 | spl9_70)),
  inference(forward_demodulation,[],[f11031,f8300])).
fof(f11031,plain,(
  ( ! [X33,X34] : (~d(lazy_impl(prop(sK2(sK7,X33)),sK7)) | bool(lazy_impl(prop(X34),impl(impl(sK7,X34),impl(impl(X33,X34),X34)))) | ~bool(lazy_impl(prop(sK2(sK7,X33)),impl(impl(sK7,sK2(sK7,X33)),impl(impl(X33,sK2(sK7,X33)),sK2(sK7,X33))))) | ~d(lazy_impl(prop(X34),impl(impl(sK7,X34),impl(impl(X33,X34),X34))))) ) | (spl9_57 | spl9_70)),
  inference(forward_demodulation,[],[f10969,f8300])).
fof(f10969,plain,(
  ( ! [X33,X34] : (~d(lazy_impl(prop(sK2(sK7,X33)),impl(sK7,impl(impl(X33,sK2(sK7,X33)),sK2(sK7,X33))))) | bool(lazy_impl(prop(X34),impl(impl(sK7,X34),impl(impl(X33,X34),X34)))) | ~bool(lazy_impl(prop(sK2(sK7,X33)),impl(impl(sK7,sK2(sK7,X33)),impl(impl(X33,sK2(sK7,X33)),sK2(sK7,X33))))) | ~d(lazy_impl(prop(X34),impl(impl(sK7,X34),impl(impl(X33,X34),X34))))) ) | (spl9_57 | spl9_70)),
  inference(superposition,[],[f1469,f8300])).
fof(f13437,plain,(
  ~spl9_32 | ~spl9_33 | spl9_49 | spl9_57 | spl9_58 | spl9_70 | ~spl9_93 | ~spl9_108),
  inference(avatar_contradiction_clause,[],[f13436])).
fof(f13436,plain,(
  $false | (~spl9_32 | ~spl9_33 | spl9_49 | spl9_57 | spl9_58 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(subsumption_resolution,[],[f13435,f7094])).
fof(f13435,plain,(
  sK7 != lazy_impl(true,sK7) | (~spl9_32 | ~spl9_33 | spl9_49 | spl9_57 | spl9_58 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f13434,f7094])).
fof(f13434,plain,(
  sK7 != lazy_impl(true,lazy_impl(true,sK7)) | (~spl9_32 | ~spl9_33 | spl9_49 | spl9_57 | spl9_58 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f13433,f248])).
fof(f13433,plain,(
  sK7 != lazy_impl(true,lazy_impl(prop(true),sK7)) | (~spl9_32 | ~spl9_33 | spl9_49 | spl9_57 | spl9_58 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f8692,f1183])).
fof(f8692,plain,(
  sK7 != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),sK7)) | (~spl9_33 | spl9_49 | spl9_57 | spl9_58 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(backward_demodulation,[],[f8688,f8301])).
fof(f8301,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),sK7)) | (spl9_49 | spl9_57 | spl9_58 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(backward_demodulation,[],[f8300,f7095])).
fof(f7095,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(sK7,err))) | (spl9_49 | spl9_58 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(backward_demodulation,[],[f7094,f6798])).
fof(f6798,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(lazy_impl(true,sK7),err))) | (spl9_49 | spl9_58 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f6797,f1864])).
fof(f6797,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(lazy_impl(true,sK7),impl(err,sK2(sK7,sK8))))) | (spl9_58 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f1894,f4193])).
fof(f4193,plain,(
  ( ! [X6] : (err = impl(sK8,X6)) ) | (~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f4094,f3155])).
fof(f8688,plain,(
  sK7 = or1(sK7,sK8) | (~spl9_33 | spl9_49 | spl9_57 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f8687,f7094])).
fof(f8687,plain,(
  or1(sK7,sK8) = lazy_impl(true,sK7) | (~spl9_33 | spl9_49 | spl9_57 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f8686,f7094])).
fof(f8686,plain,(
  or1(sK7,sK8) = lazy_impl(true,lazy_impl(true,sK7)) | (~spl9_33 | spl9_49 | spl9_57 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f8685,f8300])).
fof(f8685,plain,(
  or1(sK7,sK8) = lazy_impl(true,lazy_impl(true,impl(sK7,err))) | (~spl9_33 | spl9_49 | spl9_57 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f8684,f8300])).
fof(f8684,plain,(
  or1(sK7,sK8) = lazy_impl(true,lazy_impl(true,impl(impl(sK7,false1),err))) | (~spl9_33 | spl9_49 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f8683,f1864])).
fof(f8683,plain,(
  or1(sK7,sK8) = lazy_impl(true,lazy_impl(true,impl(impl(sK7,false1),impl(err,false1)))) | (~spl9_33 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f1186,f4193])).
fof(f13376,plain,(
  ~spl9_13 | spl9_69 | ~spl9_173),
  inference(avatar_contradiction_clause,[],[f13375])).
fof(f13375,plain,(
  $false | (~spl9_13 | spl9_69 | ~spl9_173)),
  inference(subsumption_resolution,[],[f13374,f2169])).
fof(f13374,plain,(
  true = lazy_impl(true,sK7) | (~spl9_13 | ~spl9_173)),
  inference(forward_demodulation,[],[f8417,f638])).
fof(f13265,plain,(
  ~spl9_33 | spl9_49 | spl9_57 | spl9_58 | spl9_70 | ~spl9_93 | ~spl9_108 | ~spl9_182),
  inference(avatar_contradiction_clause,[],[f13264])).
fof(f13264,plain,(
  $false | (~spl9_33 | spl9_49 | spl9_57 | spl9_58 | spl9_70 | ~spl9_93 | ~spl9_108 | ~spl9_182)),
  inference(subsumption_resolution,[],[f13263,f7094])).
fof(f13263,plain,(
  sK7 != lazy_impl(true,sK7) | (~spl9_33 | spl9_49 | spl9_57 | spl9_58 | spl9_70 | ~spl9_93 | ~spl9_108 | ~spl9_182)),
  inference(forward_demodulation,[],[f13262,f7094])).
fof(f13262,plain,(
  sK7 != lazy_impl(true,lazy_impl(true,sK7)) | (~spl9_33 | spl9_49 | spl9_57 | spl9_58 | spl9_70 | ~spl9_93 | ~spl9_108 | ~spl9_182)),
  inference(forward_demodulation,[],[f13239,f249])).
fof(f13239,plain,(
  sK7 != lazy_impl(true,lazy_impl(prop(false1),sK7)) | (~spl9_33 | spl9_49 | spl9_57 | spl9_58 | spl9_70 | ~spl9_93 | ~spl9_108 | ~spl9_182)),
  inference(backward_demodulation,[],[f8569,f8692])).
fof(f12375,plain,(
  spl9_167 | spl9_166),
  inference(avatar_split_clause,[],[f8894,f8115,f8120])).
fof(f8115,plain,(
  spl9_166 <=> ! [X1] : ~forallprefers(lazy_impl(prop(X1),impl(true,X1)),true)),
  introduced(avatar_definition,[new_symbols(naming,[spl9_166])])).
fof(f8894,plain,(
  ( ! [X2,X3] : (~forallprefers(lazy_impl(prop(X3),impl(true,X3)),true) | true = prop(sK1(false1,X2))) )),
  inference(forward_demodulation,[],[f8863,f221])).
fof(f8863,plain,(
  ( ! [X2,X3] : (~forallprefers(lazy_impl(prop(X3),impl(true,X3)),lazy_impl(false1,impl(true,sK1(false1,X2)))) | true = prop(sK1(false1,X2))) )),
  inference(superposition,[],[f1432,f252])).
fof(f1432,plain,(
  ( ! [X0,X1] : (~forallprefers(lazy_impl(prop(X1),impl(true,X1)),lazy_impl(prop(sK1(false1,X0)),impl(true,sK1(false1,X0))))) )),
  inference(forward_demodulation,[],[f1422,f221])).
fof(f1422,plain,(
  ( ! [X0,X1] : (~forallprefers(lazy_impl(prop(X1),impl(true,X1)),lazy_impl(prop(sK1(false1,X0)),impl(lazy_impl(false1,impl(X0,sK1(false1,X0))),sK1(false1,X0))))) )),
  inference(superposition,[],[f229,f221])).
fof(f12287,plain,(
  ~spl9_166),
  inference(avatar_contradiction_clause,[],[f12286])).
fof(f12286,plain,(
  $false | ~spl9_166),
  inference(subsumption_resolution,[],[f12285,f245])).
fof(f12285,plain,(
  ~forallprefers(false1,true) | ~spl9_166),
  inference(forward_demodulation,[],[f12284,f269])).
fof(f12284,plain,(
  ~forallprefers(lazy_impl(true,false1),true) | ~spl9_166),
  inference(forward_demodulation,[],[f12247,f249])).
fof(f12247,plain,(
  ~forallprefers(lazy_impl(prop(false1),false1),true) | ~spl9_166),
  inference(superposition,[],[f8116,f258])).
fof(f8116,plain,(
  ( ! [X1] : (~forallprefers(lazy_impl(prop(X1),impl(true,X1)),true)) ) | ~spl9_166),
  inference(avatar_component_clause,[],[f8115])).
fof(f10036,plain,(
  spl9_204 | spl9_203 | spl9_57 | spl9_70),
  inference(avatar_split_clause,[],[f10032,f2172,f1888,f10009,f10034])).
fof(f10032,plain,(
  ( ! [X41,X40] : (d(lazy_impl(prop(X41),sK7)) | ~d(lazy_impl(prop(sK2(sK7,X40)),sK7))) ) | (spl9_57 | spl9_70)),
  inference(forward_demodulation,[],[f10031,f8300])).
fof(f10031,plain,(
  ( ! [X41,X40] : (d(lazy_impl(prop(X41),impl(sK7,impl(impl(X40,X41),X41)))) | ~d(lazy_impl(prop(sK2(sK7,X40)),sK7))) ) | (spl9_57 | spl9_70)),
  inference(forward_demodulation,[],[f10030,f8300])).
fof(f10030,plain,(
  ( ! [X41,X40] : (~d(lazy_impl(prop(sK2(sK7,X40)),sK7)) | d(lazy_impl(prop(X41),impl(impl(sK7,X41),impl(impl(X40,X41),X41))))) ) | (spl9_57 | spl9_70)),
  inference(forward_demodulation,[],[f9974,f8300])).
fof(f9974,plain,(
  ( ! [X41,X40] : (~d(lazy_impl(prop(sK2(sK7,X40)),impl(sK7,impl(impl(X40,sK2(sK7,X40)),sK2(sK7,X40))))) | d(lazy_impl(prop(X41),impl(impl(sK7,X41),impl(impl(X40,X41),X41))))) ) | (spl9_57 | spl9_70)),
  inference(superposition,[],[f1470,f8300])).
fof(f8744,plain,(
  spl9_39 | spl9_57 | spl9_70),
  inference(avatar_contradiction_clause,[],[f8743])).
fof(f8743,plain,(
  $false | (spl9_39 | spl9_57 | spl9_70)),
  inference(subsumption_resolution,[],[f8742,f1889])).
fof(f8742,plain,(
  true = prop(sK7) | (spl9_39 | spl9_57 | spl9_70)),
  inference(subsumption_resolution,[],[f8734,f7094])).
fof(f8734,plain,(
  sK7 != lazy_impl(true,sK7) | true = prop(sK7) | (spl9_39 | spl9_57 | spl9_70)),
  inference(superposition,[],[f8339,f346])).
fof(f8339,plain,(
  sK7 != or1(sK7,false1) | (spl9_39 | spl9_57 | spl9_70)),
  inference(forward_demodulation,[],[f8338,f7094])).
fof(f8338,plain,(
  or1(sK7,false1) != lazy_impl(true,sK7) | (spl9_39 | spl9_57 | spl9_70)),
  inference(forward_demodulation,[],[f8337,f7094])).
fof(f8337,plain,(
  or1(sK7,false1) != lazy_impl(true,lazy_impl(true,sK7)) | (spl9_39 | spl9_57 | spl9_70)),
  inference(forward_demodulation,[],[f8319,f8300])).
fof(f8319,plain,(
  or1(sK7,false1) != lazy_impl(true,lazy_impl(true,impl(sK7,sK2(sK7,false1)))) | (spl9_39 | spl9_57 | spl9_70)),
  inference(backward_demodulation,[],[f8300,f1220])).
fof(f1220,plain,(
  or1(sK7,false1) != lazy_impl(true,lazy_impl(true,impl(impl(sK7,sK2(sK7,false1)),sK2(sK7,false1)))) | spl9_39),
  inference(avatar_component_clause,[],[f1218])).
fof(f8676,plain,(
  spl9_33 | spl9_49 | spl9_57 | spl9_70 | ~spl9_93 | ~spl9_108),
  inference(avatar_contradiction_clause,[],[f8675])).
fof(f8675,plain,(
  $false | (spl9_33 | spl9_49 | spl9_57 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(subsumption_resolution,[],[f8674,f1889])).
fof(f8674,plain,(
  true = prop(sK7) | (spl9_33 | spl9_49 | spl9_57 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(subsumption_resolution,[],[f8666,f7094])).
fof(f8666,plain,(
  sK7 != lazy_impl(true,sK7) | true = prop(sK7) | (spl9_33 | spl9_49 | spl9_57 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(superposition,[],[f8323,f346])).
fof(f8323,plain,(
  sK7 != or1(sK7,sK8) | (spl9_33 | spl9_49 | spl9_57 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f8322,f7094])).
fof(f8322,plain,(
  or1(sK7,sK8) != lazy_impl(true,sK7) | (spl9_33 | spl9_49 | spl9_57 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f8321,f7094])).
fof(f8321,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,sK7)) | (spl9_33 | spl9_49 | spl9_57 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f8305,f8300])).
fof(f8305,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(sK7,err))) | (spl9_33 | spl9_49 | spl9_57 | spl9_70 | ~spl9_93 | ~spl9_108)),
  inference(backward_demodulation,[],[f8300,f6392])).
fof(f6392,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(impl(sK7,false1),err))) | (spl9_33 | spl9_49 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f6391,f1864])).
fof(f6391,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(true,impl(impl(sK7,false1),impl(err,false1)))) | (spl9_33 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f1187,f4193])).
fof(f6482,plain,(
  ~spl9_47 | spl9_49 | spl9_83 | ~spl9_86 | ~spl9_93 | ~spl9_95 | ~spl9_108),
  inference(avatar_contradiction_clause,[],[f6481])).
fof(f6481,plain,(
  $false | (~spl9_47 | spl9_49 | spl9_83 | ~spl9_86 | ~spl9_93 | ~spl9_95 | ~spl9_108)),
  inference(subsumption_resolution,[],[f6480,f268])).
fof(f6480,plain,(
  err != lazy_impl(true,err) | (~spl9_47 | spl9_49 | spl9_83 | ~spl9_86 | ~spl9_93 | ~spl9_95 | ~spl9_108)),
  inference(forward_demodulation,[],[f6479,f268])).
fof(f6479,plain,(
  err != lazy_impl(true,lazy_impl(true,err)) | (~spl9_47 | spl9_49 | spl9_83 | ~spl9_86 | ~spl9_93 | ~spl9_95 | ~spl9_108)),
  inference(forward_demodulation,[],[f6478,f4138])).
fof(f4138,plain,(
  err = or1(true,sK8) | (~spl9_47 | ~spl9_93)),
  inference(backward_demodulation,[],[f3155,f3097])).
fof(f6478,plain,(
  lazy_impl(true,lazy_impl(true,err)) != or1(true,sK8) | (spl9_49 | spl9_83 | ~spl9_86 | ~spl9_93 | ~spl9_95 | ~spl9_108)),
  inference(forward_demodulation,[],[f6352,f6428])).
fof(f6428,plain,(
  err = impl(false1,err) | (spl9_49 | ~spl9_93 | ~spl9_95)),
  inference(forward_demodulation,[],[f6427,f1864])).
fof(f6427,plain,(
  err = impl(false1,impl(err,false1)) | (~spl9_93 | ~spl9_95)),
  inference(forward_demodulation,[],[f3164,f3155])).
fof(f6352,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(true,impl(false1,err))) | (spl9_49 | spl9_83 | ~spl9_86 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f6351,f1864])).
fof(f6351,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(true,impl(false1,impl(err,false1)))) | (spl9_83 | ~spl9_86 | ~spl9_93 | ~spl9_108)),
  inference(forward_demodulation,[],[f2936,f4193])).
fof(f2936,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(true,impl(false1,impl(impl(sK8,false1),false1)))) | (spl9_83 | ~spl9_86)),
  inference(forward_demodulation,[],[f2892,f249])).
fof(f2892,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(false1),impl(false1,impl(impl(sK8,false1),false1)))) | (spl9_83 | ~spl9_86)),
  inference(backward_demodulation,[],[f2860,f2730])).
fof(f6451,plain,(
  spl9_49 | ~spl9_50 | spl9_89 | ~spl9_93 | ~spl9_95),
  inference(avatar_contradiction_clause,[],[f6450])).
fof(f6450,plain,(
  $false | (spl9_49 | ~spl9_50 | spl9_89 | ~spl9_93 | ~spl9_95)),
  inference(subsumption_resolution,[],[f6449,f268])).
fof(f6449,plain,(
  err != lazy_impl(true,err) | (spl9_49 | ~spl9_50 | spl9_89 | ~spl9_93 | ~spl9_95)),
  inference(forward_demodulation,[],[f5986,f6428])).
fof(f5986,plain,(
  err != lazy_impl(true,impl(false1,err)) | (~spl9_50 | spl9_89 | ~spl9_93)),
  inference(forward_demodulation,[],[f5985,f1719])).
fof(f5985,plain,(
  err != lazy_impl(true,impl(false1,not1(err))) | (spl9_89 | ~spl9_93)),
  inference(forward_demodulation,[],[f3137,f3155])).
fof(f3137,plain,(
  not1(sK8) != lazy_impl(true,impl(false1,not1(not1(sK8)))) | spl9_89),
  inference(avatar_component_clause,[],[f3135])).
fof(f3135,plain,(
  spl9_89 <=> not1(sK8) = lazy_impl(true,impl(false1,not1(not1(sK8))))),
  introduced(avatar_definition,[new_symbols(naming,[spl9_89])])).
fof(f6415,plain,(
  ~spl9_116 | spl9_133),
  inference(avatar_contradiction_clause,[],[f6414])).
fof(f6414,plain,(
  $false | (~spl9_116 | spl9_133)),
  inference(subsumption_resolution,[],[f6407,f267])).
fof(f6407,plain,(
  true != lazy_impl(true,true) | (~spl9_116 | spl9_133)),
  inference(backward_demodulation,[],[f5311,f5879])).
fof(f5311,plain,(
  true = impl(false1,err) | ~spl9_116),
  inference(avatar_component_clause,[],[f5309])).
fof(f6277,plain,(
  spl9_114),
  inference(avatar_contradiction_clause,[],[f6276])).
fof(f6276,plain,(
  $false | spl9_114),
  inference(subsumption_resolution,[],[f6275,f207])).
fof(f6275,plain,(
  err = false1 | spl9_114),
  inference(subsumption_resolution,[],[f6274,f135])).
fof(f6274,plain,(
  true = err | err = false1 | spl9_114),
  inference(trivial_inequality_removal,[],[f6271])).
fof(f6271,plain,(
  false1 != false1 | true = err | err = false1 | spl9_114),
  inference(superposition,[],[f4967,f282])).
fof(f4967,plain,(
  false1 != prop(err) | spl9_114),
  inference(avatar_component_clause,[],[f4966])).
fof(f5951,plain,(
  spl9_49 | ~spl9_50 | ~spl9_89 | ~spl9_93 | spl9_94),
  inference(avatar_contradiction_clause,[],[f5950])).
fof(f5950,plain,(
  $false | (spl9_49 | ~spl9_50 | ~spl9_89 | ~spl9_93 | spl9_94)),
  inference(subsumption_resolution,[],[f5949,f5471])).
fof(f5471,plain,(
  err = lazy_impl(true,impl(false1,err)) | (~spl9_50 | ~spl9_89 | ~spl9_93)),
  inference(forward_demodulation,[],[f5470,f1719])).
fof(f5470,plain,(
  err = lazy_impl(true,impl(false1,not1(err))) | (~spl9_89 | ~spl9_93)),
  inference(forward_demodulation,[],[f3136,f3155])).
fof(f3136,plain,(
  not1(sK8) = lazy_impl(true,impl(false1,not1(not1(sK8)))) | ~spl9_89),
  inference(avatar_component_clause,[],[f3135])).
fof(f5949,plain,(
  err != lazy_impl(true,impl(false1,err)) | (spl9_49 | ~spl9_93 | spl9_94)),
  inference(forward_demodulation,[],[f5948,f1864])).
fof(f5948,plain,(
  err != lazy_impl(true,impl(false1,impl(err,false1))) | (~spl9_93 | spl9_94)),
  inference(forward_demodulation,[],[f3160,f3155])).
fof(f3160,plain,(
  err != lazy_impl(true,impl(false1,impl(not1(sK8),false1))) | spl9_94),
  inference(avatar_component_clause,[],[f3159])).
fof(f5721,plain,(
  ~spl9_47 | ~spl9_59),
  inference(avatar_contradiction_clause,[],[f5720])).
fof(f5720,plain,(
  $false | (~spl9_47 | ~spl9_59)),
  inference(subsumption_resolution,[],[f5719,f208])).
fof(f5719,plain,(
  true = false1 | (~spl9_47 | ~spl9_59)),
  inference(backward_demodulation,[],[f1899,f1578])).
fof(f5268,plain,(
  spl9_49 | ~spl9_93 | spl9_95 | ~spl9_114),
  inference(avatar_contradiction_clause,[],[f5267])).
fof(f5267,plain,(
  $false | (spl9_49 | ~spl9_93 | spl9_95 | ~spl9_114)),
  inference(subsumption_resolution,[],[f5266,f4170])).
fof(f4170,plain,(
  err != impl(false1,err) | (spl9_49 | ~spl9_93 | spl9_95)),
  inference(forward_demodulation,[],[f4148,f1864])).
fof(f4148,plain,(
  err != impl(false1,impl(err,false1)) | (~spl9_93 | spl9_95)),
  inference(backward_demodulation,[],[f3155,f3165])).
fof(f3165,plain,(
  not1(sK8) != impl(false1,impl(not1(sK8),false1)) | spl9_95),
  inference(avatar_component_clause,[],[f3163])).
fof(f5266,plain,(
  err = impl(false1,err) | ~spl9_114),
  inference(forward_demodulation,[],[f5250,f268])).
fof(f5250,plain,(
  lazy_impl(true,err) = impl(false1,err) | ~spl9_114),
  inference(resolution,[],[f5117,f525])).
fof(f5117,plain,(
  ~bool(err) | ~spl9_114),
  inference(trivial_inequality_removal,[],[f5102])).
fof(f5102,plain,(
  false1 != false1 | ~bool(err) | ~spl9_114),
  inference(superposition,[],[f216,f4968])).
fof(f4192,plain,(
  spl9_35 | ~spl9_47 | ~spl9_93),
  inference(avatar_split_clause,[],[f4146,f3154,f1576,f1194])).
fof(f4146,plain,(
  err = or1(false1,sK8) | (~spl9_47 | ~spl9_93)),
  inference(backward_demodulation,[],[f3155,f3105])).
fof(f4159,plain,(
  spl9_18 | ~spl9_34 | ~spl9_47 | ~spl9_93),
  inference(avatar_contradiction_clause,[],[f4158])).
fof(f4158,plain,(
  $false | (spl9_18 | ~spl9_34 | ~spl9_47 | ~spl9_93)),
  inference(subsumption_resolution,[],[f4138,f2635])).
fof(f2635,plain,(
  err != or1(true,sK8) | (spl9_18 | ~spl9_34)),
  inference(backward_demodulation,[],[f1192,f660])).
fof(f660,plain,(
  err != or1(sK7,sK8) | spl9_18),
  inference(avatar_component_clause,[],[f658])).
fof(f4152,plain,(
  spl9_111 | ~spl9_47 | ~spl9_93),
  inference(avatar_split_clause,[],[f4128,f3154,f1576,f4108])).
fof(f4128,plain,(
  err = lazy_impl(true,sK8) | (~spl9_47 | ~spl9_93)),
  inference(backward_demodulation,[],[f3155,f3071])).
fof(f4127,plain,(
  spl9_113 | spl9_93 | ~spl9_47),
  inference(avatar_split_clause,[],[f3967,f1576,f3154,f4117])).
fof(f3967,plain,(
  err = not1(sK8) | sK8 = lazy_impl(true,sK8) | ~spl9_47),
  inference(superposition,[],[f3071,f212])).
fof(f4126,plain,(
  spl9_108 | ~spl9_47),
  inference(avatar_split_clause,[],[f3080,f1576,f4093])).
fof(f3080,plain,(
  ( ! [X0] : (not1(sK8) = impl(sK8,X0)) ) | ~spl9_47),
  inference(backward_demodulation,[],[f3071,f3066])).
fof(f3066,plain,(
  ( ! [X0] : (lazy_impl(true,sK8) = impl(sK8,X0)) ) | ~spl9_47),
  inference(resolution,[],[f2970,f217])).
fof(f4125,plain,(
  spl9_111 | spl9_112 | ~spl9_47),
  inference(avatar_split_clause,[],[f3968,f1576,f4112,f4108])).
fof(f3968,plain,(
  sK8 = not1(sK8) | err = lazy_impl(true,sK8) | ~spl9_47),
  inference(superposition,[],[f3071,f212])).
fof(f4073,plain,(
  ~spl9_47 | spl9_93 | spl9_97),
  inference(avatar_contradiction_clause,[],[f4072])).
fof(f4072,plain,(
  $false | (~spl9_47 | spl9_93 | spl9_97)),
  inference(subsumption_resolution,[],[f4071,f4000])).
fof(f4000,plain,(
  sK8 = not1(sK8) | (~spl9_47 | spl9_93)),
  inference(backward_demodulation,[],[f3999,f3071])).
fof(f3999,plain,(
  sK8 = lazy_impl(true,sK8) | (~spl9_47 | spl9_93)),
  inference(subsumption_resolution,[],[f3967,f3156])).
fof(f3156,plain,(
  err != not1(sK8) | spl9_93),
  inference(avatar_component_clause,[],[f3154])).
fof(f4071,plain,(
  sK8 != not1(sK8) | (~spl9_47 | spl9_93 | spl9_97)),
  inference(forward_demodulation,[],[f4070,f4008])).
fof(f4008,plain,(
  sK8 = impl(false1,sK8) | (~spl9_47 | spl9_93)),
  inference(backward_demodulation,[],[f4000,f3094])).
fof(f4070,plain,(
  sK8 != not1(impl(false1,sK8)) | (~spl9_47 | spl9_93 | spl9_97)),
  inference(forward_demodulation,[],[f4022,f4001])).
fof(f4001,plain,(
  ( ! [X0] : (sK8 = impl(sK8,X0)) ) | (~spl9_47 | spl9_93)),
  inference(backward_demodulation,[],[f4000,f3080])).
fof(f4022,plain,(
  sK8 != not1(impl(false1,impl(sK8,false1))) | (~spl9_47 | spl9_93 | spl9_97)),
  inference(backward_demodulation,[],[f4000,f3174])).
fof(f3174,plain,(
  not1(sK8) != not1(impl(false1,impl(not1(sK8),false1))) | spl9_97),
  inference(avatar_component_clause,[],[f3172])).
fof(f4069,plain,(
  ~spl9_47 | spl9_93 | spl9_95),
  inference(avatar_contradiction_clause,[],[f4068])).
fof(f4068,plain,(
  $false | (~spl9_47 | spl9_93 | spl9_95)),
  inference(subsumption_resolution,[],[f4067,f4008])).
fof(f4067,plain,(
  sK8 != impl(false1,sK8) | (~spl9_47 | spl9_93 | spl9_95)),
  inference(forward_demodulation,[],[f4021,f4001])).
fof(f4021,plain,(
  sK8 != impl(false1,impl(sK8,false1)) | (~spl9_47 | spl9_93 | spl9_95)),
  inference(backward_demodulation,[],[f4000,f3165])).
fof(f4033,plain,(
  ~spl9_34 | ~spl9_47 | spl9_67 | ~spl9_86 | spl9_93),
  inference(avatar_contradiction_clause,[],[f4032])).
fof(f4032,plain,(
  $false | (~spl9_34 | ~spl9_47 | spl9_67 | ~spl9_86 | spl9_93)),
  inference(subsumption_resolution,[],[f4031,f3999])).
fof(f4031,plain,(
  sK8 != lazy_impl(true,sK8) | (~spl9_34 | ~spl9_47 | spl9_67 | ~spl9_86 | spl9_93)),
  inference(forward_demodulation,[],[f4030,f3999])).
fof(f4030,plain,(
  sK8 != lazy_impl(true,lazy_impl(true,sK8)) | (~spl9_34 | ~spl9_47 | spl9_67 | ~spl9_86 | spl9_93)),
  inference(forward_demodulation,[],[f4029,f4008])).
fof(f4029,plain,(
  sK8 != lazy_impl(true,lazy_impl(true,impl(false1,sK8))) | (~spl9_34 | ~spl9_47 | spl9_67 | ~spl9_86 | spl9_93)),
  inference(forward_demodulation,[],[f4012,f4001])).
fof(f4012,plain,(
  sK8 != lazy_impl(true,lazy_impl(true,impl(false1,impl(sK8,false1)))) | (~spl9_34 | ~spl9_47 | spl9_67 | ~spl9_86 | spl9_93)),
  inference(backward_demodulation,[],[f4000,f3098])).
fof(f3098,plain,(
  not1(sK8) != lazy_impl(true,lazy_impl(true,impl(false1,impl(not1(sK8),false1)))) | (~spl9_34 | ~spl9_47 | spl9_67 | ~spl9_86)),
  inference(backward_demodulation,[],[f3097,f2933])).
fof(f2933,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(true,impl(false1,impl(not1(sK8),false1)))) | (~spl9_34 | spl9_67 | ~spl9_86)),
  inference(forward_demodulation,[],[f2932,f249])).
fof(f2932,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(false1),impl(false1,impl(not1(sK8),false1)))) | (~spl9_34 | spl9_67 | ~spl9_86)),
  inference(forward_demodulation,[],[f2889,f258])).
fof(f2889,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(false1),impl(impl(true,false1),impl(not1(sK8),false1)))) | (~spl9_34 | spl9_67 | ~spl9_86)),
  inference(backward_demodulation,[],[f2860,f2653])).
fof(f2653,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(not1(sK8),sK2(true,sK8))))) | (~spl9_34 | spl9_67)),
  inference(backward_demodulation,[],[f1192,f2081])).
fof(f2874,plain,(
  spl9_85 | ~spl9_14 | ~spl9_34),
  inference(avatar_split_clause,[],[f2873,f1190,f640,f2853])).
fof(f2873,plain,(
  true = or1(true,sK8) | (~spl9_14 | ~spl9_34)),
  inference(forward_demodulation,[],[f641,f1192])).
fof(f2847,plain,(
  spl9_47 | spl9_14 | ~spl9_34),
  inference(avatar_split_clause,[],[f2822,f1190,f640,f1576])).
fof(f2822,plain,(
  false1 = prop(sK8) | (spl9_14 | ~spl9_34)),
  inference(trivial_inequality_removal,[],[f2820])).
fof(f2820,plain,(
  true != true | false1 = prop(sK8) | (spl9_14 | ~spl9_34)),
  inference(superposition,[],[f2633,f265])).
fof(f2633,plain,(
  true != or1(true,sK8) | (spl9_14 | ~spl9_34)),
  inference(backward_demodulation,[],[f1192,f642])).
fof(f2731,plain,(
  spl9_82 | ~spl9_83 | ~spl9_34),
  inference(avatar_split_clause,[],[f2681,f1190,f2728,f2724])).
fof(f2681,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(sK2(true,sK8)),impl(sK2(true,sK8),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8))))) | false1 = prop(sK2(true,sK8)) | ~spl9_34),
  inference(superposition,[],[f2632,f259])).
fof(f259,plain,(
  ( ! [X0] : (impl(true,X0) = X0 | prop(X0) = false1) )),
  inference(resolution,[],[f157,f215])).
fof(f2632,plain,(
  or1(true,sK8) != lazy_impl(true,lazy_impl(prop(sK2(true,sK8)),impl(impl(true,sK2(true,sK8)),impl(impl(sK8,sK2(true,sK8)),sK2(true,sK8))))) | ~spl9_34),
  inference(backward_demodulation,[],[f1192,f241])).
fof(f2630,plain,(
  ~spl9_69 | ~spl9_70),
  inference(avatar_contradiction_clause,[],[f2629])).
fof(f2629,plain,(
  $false | (~spl9_69 | ~spl9_70)),
  inference(subsumption_resolution,[],[f2628,f135])).
fof(f2628,plain,(
  true = err | (~spl9_69 | ~spl9_70)),
  inference(forward_demodulation,[],[f2173,f2168])).
fof(f2175,plain,(
  spl9_57 | ~spl9_70 | spl9_18),
  inference(avatar_split_clause,[],[f2158,f658,f2172,f1888])).
fof(f2158,plain,(
  err != lazy_impl(true,sK7) | true = prop(sK7) | spl9_18),
  inference(superposition,[],[f660,f346])).
fof(f2082,plain,(
  spl9_59 | ~spl9_67),
  inference(avatar_split_clause,[],[f2060,f2079,f1897])).
fof(f2060,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),impl(not1(sK8),sK2(sK7,sK8))))) | true = prop(sK8)),
  inference(superposition,[],[f241,f1862])).
fof(f2077,plain,(
  spl9_57 | ~spl9_66),
  inference(avatar_split_clause,[],[f2059,f2074,f1888])).
fof(f2059,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(not1(sK7),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | true = prop(sK7)),
  inference(superposition,[],[f241,f1862])).
fof(f2067,plain,(
  spl9_53 | ~spl9_64),
  inference(avatar_split_clause,[],[f2039,f2064,f1870])).
fof(f2039,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(impl(sK7,sK2(sK7,sK8)),not1(impl(sK8,sK2(sK7,sK8)))))) | true = prop(impl(sK8,sK2(sK7,sK8)))),
  inference(superposition,[],[f241,f1862])).
fof(f1895,plain,(
  spl9_57 | ~spl9_58),
  inference(avatar_split_clause,[],[f1845,f1892,f1888])).
fof(f1845,plain,(
  or1(sK7,sK8) != lazy_impl(true,lazy_impl(prop(sK2(sK7,sK8)),impl(lazy_impl(true,sK7),impl(impl(sK8,sK2(sK7,sK8)),sK2(sK7,sK8))))) | true = prop(sK7)),
  inference(superposition,[],[f241,f291])).
fof(f1757,plain,(
  ~spl9_49),
  inference(avatar_contradiction_clause,[],[f1756])).
fof(f1756,plain,(
  $false | ~spl9_49),
  inference(subsumption_resolution,[],[f1755,f207])).
fof(f1755,plain,(
  err = false1 | ~spl9_49),
  inference(subsumption_resolution,[],[f1754,f135])).
fof(f1754,plain,(
  true = err | err = false1 | ~spl9_49),
  inference(subsumption_resolution,[],[f1740,f208])).
fof(f1740,plain,(
  true = false1 | true = err | err = false1 | ~spl9_49),
  inference(superposition,[],[f282,f1715])).
fof(f1715,plain,(
  true = prop(err) | ~spl9_49),
  inference(avatar_component_clause,[],[f1713])).
fof(f1721,plain,(
  spl9_49 | spl9_50),
  inference(avatar_split_clause,[],[f1704,f1717,f1713])).
fof(f1704,plain,(
  err = not1(err) | true = prop(err)),
  inference(superposition,[],[f268,f283])).
fof(f1381,plain,(
  ~spl9_5),
  inference(avatar_contradiction_clause,[],[f1380])).
fof(f1380,plain,(
  $false | ~spl9_5),
  inference(subsumption_resolution,[],[f1379,f245])).
fof(f1379,plain,(
  ~forallprefers(false1,true) | ~spl9_5),
  inference(forward_demodulation,[],[f1377,f269])).
fof(f1377,plain,(
  ~forallprefers(lazy_impl(true,false1),true) | ~spl9_5),
  inference(superposition,[],[f373,f249])).
fof(f373,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),X0),true)) ) | ~spl9_5),
  inference(avatar_component_clause,[],[f372])).
fof(f497,plain,(
  ~spl9_2),
  inference(avatar_contradiction_clause,[],[f496])).
fof(f496,plain,(
  $false | ~spl9_2),
  inference(subsumption_resolution,[],[f493,f245])).
fof(f493,plain,(
  ~forallprefers(false1,true) | ~spl9_2),
  inference(backward_demodulation,[],[f491,f368])).
fof(f368,plain,(
  ~forallprefers(false1,lazy_impl(prop(sK6),sK6))),
  inference(forward_demodulation,[],[f356,f269])).
fof(f356,plain,(
  ~forallprefers(lazy_impl(true,false1),lazy_impl(prop(sK6),sK6))),
  inference(superposition,[],[f236,f249])).
fof(f491,plain,(
  true = lazy_impl(prop(sK6),sK6) | ~spl9_2),
  inference(subsumption_resolution,[],[f490,f135])).
fof(f490,plain,(
  true = err | true = lazy_impl(prop(sK6),sK6) | ~spl9_2),
  inference(forward_demodulation,[],[f484,f482])).
fof(f482,plain,(
  true = lazy_impl(true,lazy_impl(prop(sK6),sK6)) | ~spl9_2),
  inference(backward_demodulation,[],[f341,f237])).
fof(f341,plain,(
  true = false2 | ~spl9_2),
  inference(avatar_component_clause,[],[f339])).
fof(f484,plain,(
  true = lazy_impl(prop(sK6),sK6) | err = lazy_impl(true,lazy_impl(prop(sK6),sK6)) | ~spl9_2),
  inference(superposition,[],[f482,f212])).
fof(f463,plain,(
  spl9_6 | spl9_7 | ~spl9_1),
  inference(avatar_split_clause,[],[f435,f335,f461,f457])).
fof(f435,plain,(
  ( ! [X0] : (~forallprefers(lazy_impl(prop(X0),X0),sK6) | err = lazy_impl(true,sK6)) ) | ~spl9_1),
  inference(superposition,[],[f416,f212])).
fof(f416,plain,(
  ( ! [X1] : (~forallprefers(lazy_impl(prop(X1),X1),lazy_impl(true,sK6))) ) | ~spl9_1),
  inference(backward_demodulation,[],[f337,f236])).
% SZS output end Proof for SWW098+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 380334
% Time elapsed: 55.124 s
% ------------------------------
% ------------------------------
