% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SYN248-1
% SZS output start Proof for SYN248-1
fof(f2245,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f463,f1486,f1487,f2238])).
fof(f2238,plain,(
  ~spl0_21 | ~spl0_76),
  inference(avatar_contradiction_clause,[],[f2237])).
fof(f2237,plain,(
  $false | (~spl0_21 | ~spl0_76)),
  inference(subsumption_resolution,[],[f2236,f708])).
fof(f708,plain,(
  k1(c) | ~spl0_76),
  inference(avatar_component_clause,[],[f707])).
fof(f707,plain,(
  spl0_76 <=> k1(c)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_76])])).
fof(f2236,plain,(
  ~k1(c) | ~spl0_21),
  inference(subsumption_resolution,[],[f2233,f17])).
fof(f17,axiom,(
  ( ! [X0] : (q0(X0,d)) )),
  file('./p/SYN248-1.tptp',unknown)).
fof(f2233,plain,(
  ~q0(d,d) | ~k1(c) | ~spl0_21),
  inference(resolution,[],[f462,f369])).
fof(f369,axiom,(
  ~m1(d,c,c)),
  file('./p/SYN248-1.tptp',unknown)).
fof(f462,plain,(
  ( ! [X10,X11] : (m1(X10,X11,X11) | ~q0(X10,X10) | ~k1(X11)) ) | ~spl0_21),
  inference(avatar_component_clause,[],[f461])).
fof(f461,plain,(
  spl0_21 <=> ! [X11,X10] : (~q0(X10,X10) | m1(X10,X11,X11) | ~k1(X11))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_21])])).
fof(f1487,plain,(
  spl0_76),
  inference(avatar_split_clause,[],[f1464,f707])).
fof(f1464,plain,(
  k1(c)),
  inference(resolution,[],[f39,f26])).
fof(f26,axiom,(
  n0(d,c)),
  file('./p/SYN248-1.tptp',unknown)).
fof(f39,axiom,(
  ( ! [X2,X3] : (~n0(X3,X2) | k1(X2)) )),
  file('./p/SYN248-1.tptp',unknown)).
fof(f1486,plain,(
  spl0_20),
  inference(avatar_split_clause,[],[f1468,f457])).
fof(f457,plain,(
  spl0_20 <=> k1(a)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_20])])).
fof(f1468,plain,(
  k1(a)),
  inference(resolution,[],[f39,f37])).
fof(f37,axiom,(
  n0(b,a)),
  file('./p/SYN248-1.tptp',unknown)).
fof(f463,plain,(
  ~spl0_20 | spl0_21),
  inference(avatar_split_clause,[],[f72,f461,f457])).
fof(f72,axiom,(
  ( ! [X10,X11] : (~q0(X10,X10) | ~k1(X11) | ~k1(a) | m1(X10,X11,X11)) )),
  file('./p/SYN248-1.tptp',unknown)).
% SZS output end Proof for SYN248-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 6140
% Time elapsed: 0.010 s
% ------------------------------
% ------------------------------
