% Refutation found. Thanks to Tanya!
% SZS status Theorem for SYN332+1
% SZS output start Proof for SYN332+1
fof(f96,plain,(
  $false),
  inference(subsumption_resolution,[],[f95,f80])).
fof(f80,plain,(
  ( ! [X4,X3] : (big_f(X4,X3)) )),
  inference(subsumption_resolution,[],[f79,f63])).
fof(f63,plain,(
  ( ! [X4,X5] : (~big_f(X4,X5) | big_f(X5,X4)) )),
  inference(resolution,[],[f18,f52])).
fof(f52,plain,(
  ( ! [X0] : (big_f(X0,X0)) )),
  inference(duplicate_literal_removal,[],[f45])).
fof(f45,plain,(
  ( ! [X0] : (big_f(X0,X0) | big_f(X0,X0) | big_f(X0,X0)) )),
  inference(resolution,[],[f21,f24])).
fof(f24,plain,(
  ( ! [X2,X3] : (big_f(sK0(X2,X3),X2) | big_f(X2,X3)) )),
  inference(resolution,[],[f12,f9])).
fof(f9,plain,(
  ( ! [X0,X1] : (big_f(X0,sK0(X0,X1)) | big_f(X0,X1)) )),
  inference(cnf_transformation,[],[f8])).
fof(f8,plain,(
  ! [X0,X1] : ((~big_f(sK0(X0,X1),X1) | ((~big_f(X1,X0) | ~big_f(X0,X1)) & (big_f(X1,X0) | big_f(X0,X1)))) & (big_f(sK0(X0,X1),X1) | ((big_f(X0,X1) | ~big_f(X1,X0)) & (big_f(X1,X0) | ~big_f(X0,X1)))) & (big_f(X0,X1) | ~big_f(X1,X0) | ~big_f(sK0(X0,X1),sK0(X0,X1))) & (big_f(sK0(X0,X1),sK0(X0,X1)) | (~big_f(X0,X1) & big_f(X1,X0))) & (big_f(X0,sK0(X0,X1)) | ~big_f(X1,sK0(X0,X1))) & (big_f(X1,sK0(X0,X1)) | ~big_f(X0,sK0(X0,X1))) & (big_f(X0,sK0(X0,X1)) | ~big_f(sK0(X0,X1),X0)) & (big_f(sK0(X0,X1),X0) | ~big_f(X0,sK0(X0,X1))) & (~big_f(X0,sK0(X0,X1)) | ~big_f(X1,X0) | ~big_f(X0,X1)) & (big_f(X0,sK0(X0,X1)) | (big_f(X1,X0) & big_f(X0,X1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f6,f7])).
fof(f7,plain,(
  ! [X1,X0] : (? [X2] : ((~big_f(X2,X1) | ((~big_f(X1,X0) | ~big_f(X0,X1)) & (big_f(X1,X0) | big_f(X0,X1)))) & (big_f(X2,X1) | ((big_f(X0,X1) | ~big_f(X1,X0)) & (big_f(X1,X0) | ~big_f(X0,X1)))) & (big_f(X0,X1) | ~big_f(X1,X0) | ~big_f(X2,X2)) & (big_f(X2,X2) | (~big_f(X0,X1) & big_f(X1,X0))) & (big_f(X0,X2) | ~big_f(X1,X2)) & (big_f(X1,X2) | ~big_f(X0,X2)) & (big_f(X0,X2) | ~big_f(X2,X0)) & (big_f(X2,X0) | ~big_f(X0,X2)) & (~big_f(X0,X2) | ~big_f(X1,X0) | ~big_f(X0,X1)) & (big_f(X0,X2) | (big_f(X1,X0) & big_f(X0,X1)))) => ((~big_f(sK0(X0,X1),X1) | ((~big_f(X1,X0) | ~big_f(X0,X1)) & (big_f(X1,X0) | big_f(X0,X1)))) & (big_f(sK0(X0,X1),X1) | ((big_f(X0,X1) | ~big_f(X1,X0)) & (big_f(X1,X0) | ~big_f(X0,X1)))) & (big_f(X0,X1) | ~big_f(X1,X0) | ~big_f(sK0(X0,X1),sK0(X0,X1))) & (big_f(sK0(X0,X1),sK0(X0,X1)) | (~big_f(X0,X1) & big_f(X1,X0))) & (big_f(X0,sK0(X0,X1)) | ~big_f(X1,sK0(X0,X1))) & (big_f(X1,sK0(X0,X1)) | ~big_f(X0,sK0(X0,X1))) & (big_f(X0,sK0(X0,X1)) | ~big_f(sK0(X0,X1),X0)) & (big_f(sK0(X0,X1),X0) | ~big_f(X0,sK0(X0,X1))) & (~big_f(X0,sK0(X0,X1)) | ~big_f(X1,X0) | ~big_f(X0,X1)) & (big_f(X0,sK0(X0,X1)) | (big_f(X1,X0) & big_f(X0,X1)))))),
  introduced(choice_axiom,[])).
fof(f6,plain,(
  ! [X0,X1] : ? [X2] : ((~big_f(X2,X1) | ((~big_f(X1,X0) | ~big_f(X0,X1)) & (big_f(X1,X0) | big_f(X0,X1)))) & (big_f(X2,X1) | ((big_f(X0,X1) | ~big_f(X1,X0)) & (big_f(X1,X0) | ~big_f(X0,X1)))) & (big_f(X0,X1) | ~big_f(X1,X0) | ~big_f(X2,X2)) & (big_f(X2,X2) | (~big_f(X0,X1) & big_f(X1,X0))) & (big_f(X0,X2) | ~big_f(X1,X2)) & (big_f(X1,X2) | ~big_f(X0,X2)) & (big_f(X0,X2) | ~big_f(X2,X0)) & (big_f(X2,X0) | ~big_f(X0,X2)) & (~big_f(X0,X2) | ~big_f(X1,X0) | ~big_f(X0,X1)) & (big_f(X0,X2) | (big_f(X1,X0) & big_f(X0,X1))))),
  inference(flattening,[],[f5])).
fof(f5,plain,(
  ! [X0,X1] : ? [X2] : (((~big_f(X2,X1) | ((~big_f(X1,X0) | ~big_f(X0,X1)) & (big_f(X1,X0) | big_f(X0,X1)))) & (big_f(X2,X1) | ((big_f(X0,X1) | ~big_f(X1,X0)) & (big_f(X1,X0) | ~big_f(X0,X1))))) & (((big_f(X0,X1) | ~big_f(X1,X0)) | ~big_f(X2,X2)) & (big_f(X2,X2) | (~big_f(X0,X1) & big_f(X1,X0)))) & ((big_f(X0,X2) | ~big_f(X1,X2)) & (big_f(X1,X2) | ~big_f(X0,X2))) & ((big_f(X0,X2) | ~big_f(X2,X0)) & (big_f(X2,X0) | ~big_f(X0,X2))) & ((~big_f(X0,X2) | (~big_f(X1,X0) | ~big_f(X0,X1))) & (big_f(X0,X2) | (big_f(X1,X0) & big_f(X0,X1)))))),
  inference(nnf_transformation,[],[f4])).
fof(f4,plain,(
  ! [X0,X1] : ? [X2] : (((big_f(X0,X1) <=> big_f(X1,X0)) <~> big_f(X2,X1)) & ((big_f(X0,X1) | ~big_f(X1,X0)) <=> big_f(X2,X2)) & (big_f(X0,X2) <=> big_f(X1,X2)) & (big_f(X0,X2) <=> big_f(X2,X0)) & ((big_f(X1,X0) & big_f(X0,X1)) <~> big_f(X0,X2)))),
  inference(flattening,[],[f3])).
fof(f3,plain,(
  ! [X0,X1] : ? [X2] : ((((((big_f(X0,X1) <=> big_f(X1,X0)) <~> big_f(X2,X1)) & ((big_f(X0,X1) | ~big_f(X1,X0)) <=> big_f(X2,X2))) & (big_f(X0,X2) <=> big_f(X1,X2))) & (big_f(X0,X2) <=> big_f(X2,X0))) & ((big_f(X1,X0) & big_f(X0,X1)) <~> big_f(X0,X2)))),
  inference(ennf_transformation,[],[f2])).
fof(f2,negated_conjecture,(
  ~? [X0,X1] : ! [X2] : (((big_f(X1,X0) & big_f(X0,X1)) <~> big_f(X0,X2)) => ((big_f(X0,X2) <=> big_f(X2,X0)) => ((big_f(X0,X2) <=> big_f(X1,X2)) => (((big_f(X1,X0) => big_f(X0,X1)) <=> big_f(X2,X2)) => ((big_f(X0,X1) <=> big_f(X1,X0)) <=> big_f(X2,X1))))))),
  inference(negated_conjecture,[],[f1])).
fof(f1,conjecture,(
  ? [X0,X1] : ! [X2] : (((big_f(X1,X0) & big_f(X0,X1)) <~> big_f(X0,X2)) => ((big_f(X0,X2) <=> big_f(X2,X0)) => ((big_f(X0,X2) <=> big_f(X1,X2)) => (((big_f(X1,X0) => big_f(X0,X1)) <=> big_f(X2,X2)) => ((big_f(X0,X1) <=> big_f(X1,X0)) <=> big_f(X2,X1))))))),
  file('./p/SYN332+1.tptp',unknown)).
fof(f12,plain,(
  ( ! [X0,X1] : (~big_f(X0,sK0(X0,X1)) | big_f(sK0(X0,X1),X0)) )),
  inference(cnf_transformation,[],[f8])).
fof(f21,plain,(
  ( ! [X0,X1] : (~big_f(sK0(X0,X1),X1) | big_f(X1,X0) | big_f(X0,X1)) )),
  inference(cnf_transformation,[],[f8])).
fof(f18,plain,(
  ( ! [X0,X1] : (~big_f(sK0(X0,X1),sK0(X0,X1)) | ~big_f(X1,X0) | big_f(X0,X1)) )),
  inference(cnf_transformation,[],[f8])).
fof(f79,plain,(
  ( ! [X4,X3] : (big_f(X3,X4) | big_f(X4,X3)) )),
  inference(duplicate_literal_removal,[],[f77])).
fof(f77,plain,(
  ( ! [X4,X3] : (big_f(X3,X4) | big_f(X4,X3) | big_f(X3,X4)) )),
  inference(resolution,[],[f66,f21])).
fof(f66,plain,(
  ( ! [X4,X5] : (big_f(sK0(X4,X5),X5) | big_f(X4,X5)) )),
  inference(resolution,[],[f63,f27])).
fof(f27,plain,(
  ( ! [X2,X3] : (big_f(X2,sK0(X3,X2)) | big_f(X3,X2)) )),
  inference(resolution,[],[f14,f9])).
fof(f14,plain,(
  ( ! [X0,X1] : (~big_f(X0,sK0(X0,X1)) | big_f(X1,sK0(X0,X1))) )),
  inference(cnf_transformation,[],[f8])).
fof(f95,plain,(
  ( ! [X0,X1] : (~big_f(X1,X0)) )),
  inference(subsumption_resolution,[],[f86,f80])).
fof(f86,plain,(
  ( ! [X0,X1] : (~big_f(X0,X1) | ~big_f(X1,X0)) )),
  inference(resolution,[],[f80,f11])).
fof(f11,plain,(
  ( ! [X0,X1] : (~big_f(X0,sK0(X0,X1)) | ~big_f(X1,X0) | ~big_f(X0,X1)) )),
  inference(cnf_transformation,[],[f8])).
% SZS output end Proof for SYN332+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 4861
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
