% Refutation found. Thanks to Tanya!
% SZS status Theorem for SYN474+1
% SZS output start Proof for SYN474+1
fof(f5131,plain,(
  $false),
  inference(avatar_sat_refutation,[],[f259,f272,f294,f303,f312,f317,f326,f342,f354,f358,f362,f363,f371,f376,f384,f385,f393,f397,f398,f402,f406,f410,f418,f422,f423,f424,f428,f429,f433,f437,f438,f445,f446,f447,f451,f452,f456,f460,f465,f469,f470,f471,f475,f476,f484,f485,f490,f495,f499,f500,f501,f506,f507,f508,f512,f517,f518,f519,f520,f524,f525,f531,f536,f540,f541,f545,f546,f551,f556,f561,f567,f572,f577,f583,f588,f593,f599,f604,f609,f615,f620,f625,f631,f636,f641,f647,f652,f657,f663,f668,f673,f679,f684,f689,f695,f700,f705,f711,f716,f721,f727,f732,f737,f743,f748,f753,f759,f764,f769,f775,f780,f785,f791,f796,f801,f807,f812,f817,f823,f828,f833,f839,f844,f849,f855,f860,f865,f866,f871,f876,f881,f887,f892,f897,f903,f908,f913,f914,f919,f929,f935,f940,f945,f951,f956,f961,f967,f972,f977,f983,f988,f993,f1015,f1020,f1025,f1031,f1036,f1041,f1047,f1052,f1057,f1058,f1152,f1153,f1212,f1245,f1287,f1302,f1313,f1319,f1329,f1333,f1507,f1614,f1707,f1745,f1809,f1811,f1852,f1856,f1858,f1864,f1882,f1974,f1978,f2004,f2093,f2097,f2105,f2109,f2146,f2164,f2166,f2294,f2320,f2327,f2370,f2490,f2527,f2538,f2611,f2617,f2621,f2648,f2650,f2683,f2695,f2696,f2703,f2732,f2748,f2888,f2906,f2947,f2960,f2961,f3000,f3002,f3008,f3075,f3123,f3139,f3169,f3176,f3191,f3253,f3255,f3370,f3451,f3466,f3477,f3571,f3574,f3592,f3643,f3649,f3668,f3669,f3702,f3734,f3740,f3774,f3777,f3779,f3787,f3794,f3817,f3821,f3852,f3888,f3936,f3940,f3949,f3950,f3968,f3997,f4021,f4190,f4191,f4236,f4241,f4267,f4333,f4361,f4385,f4425,f4457,f4541,f4544,f4599,f4601,f4605,f4644,f4659,f4678,f4679,f4680,f4691,f4732,f4795,f4798,f4819,f4839,f4841,f4843,f4897,f4903,f5018,f5060,f5062,f5064,f5066,f5070,f5074,f5129])).
fof(f5129,plain,(
  ~spl0_27 | ~spl0_31 | ~spl0_43 | spl0_87 | ~spl0_163),
  inference(avatar_contradiction_clause,[],[f5128])).
fof(f5128,plain,(
  $false | (~spl0_27 | ~spl0_31 | ~spl0_43 | spl0_87 | ~spl0_163)),
  inference(subsumption_resolution,[],[f5117,f662])).
fof(f662,plain,(
  ~c2_1(a953) | spl0_87),
  inference(avatar_component_clause,[],[f660])).
fof(f660,plain,(
  spl0_87 <=> c2_1(a953)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_87])])).
fof(f5117,plain,(
  c2_1(a953) | (~spl0_27 | ~spl0_31 | ~spl0_43 | ~spl0_163)),
  inference(resolution,[],[f5079,f3474])).
fof(f3474,plain,(
  c0_1(a953) | ~spl0_163),
  inference(avatar_component_clause,[],[f3473])).
fof(f3473,plain,(
  spl0_163 <=> c0_1(a953)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_163])])).
fof(f5079,plain,(
  ( ! [X9] : (~c0_1(X9) | c2_1(X9)) ) | (~spl0_27 | ~spl0_31 | ~spl0_43)),
  inference(subsumption_resolution,[],[f379,f5075])).
fof(f5075,plain,(
  ( ! [X32] : (~c0_1(X32) | c3_1(X32)) ) | (~spl0_27 | ~spl0_43)),
  inference(subsumption_resolution,[],[f432,f361])).
fof(f361,plain,(
  ( ! [X4] : (~c1_1(X4) | c3_1(X4) | ~c0_1(X4)) ) | ~spl0_27),
  inference(avatar_component_clause,[],[f360])).
fof(f360,plain,(
  spl0_27 <=> ! [X4] : (~c1_1(X4) | c3_1(X4) | ~c0_1(X4))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_27])])).
fof(f432,plain,(
  ( ! [X32] : (~c0_1(X32) | c1_1(X32) | c3_1(X32)) ) | ~spl0_43),
  inference(avatar_component_clause,[],[f431])).
fof(f431,plain,(
  spl0_43 <=> ! [X32] : (~c0_1(X32) | c1_1(X32) | c3_1(X32))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_43])])).
fof(f379,plain,(
  ( ! [X9] : (~c3_1(X9) | c2_1(X9) | ~c0_1(X9)) ) | ~spl0_31),
  inference(avatar_component_clause,[],[f378])).
fof(f378,plain,(
  spl0_31 <=> ! [X9] : (~c3_1(X9) | c2_1(X9) | ~c0_1(X9))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_31])])).
fof(f5074,plain,(
  ~spl0_167 | spl0_108 | ~spl0_27 | ~spl0_110),
  inference(avatar_split_clause,[],[f4965,f782,f360,f772,f3791])).
fof(f3791,plain,(
  spl0_167 <=> c0_1(a928)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_167])])).
fof(f772,plain,(
  spl0_108 <=> c3_1(a928)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_108])])).
fof(f782,plain,(
  spl0_110 <=> c1_1(a928)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_110])])).
fof(f4965,plain,(
  c3_1(a928) | ~c0_1(a928) | (~spl0_27 | ~spl0_110)),
  inference(resolution,[],[f361,f784])).
fof(f784,plain,(
  c1_1(a928) | ~spl0_110),
  inference(avatar_component_clause,[],[f782])).
fof(f5070,plain,(
  ~spl0_45 | ~spl0_65 | spl0_81 | spl0_82),
  inference(avatar_contradiction_clause,[],[f5069])).
fof(f5069,plain,(
  $false | (~spl0_45 | ~spl0_65 | spl0_81 | spl0_82)),
  inference(subsumption_resolution,[],[f5057,f635])).
fof(f635,plain,(
  ~c1_1(a969) | spl0_82),
  inference(avatar_component_clause,[],[f633])).
fof(f633,plain,(
  spl0_82 <=> c1_1(a969)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_82])])).
fof(f5057,plain,(
  c1_1(a969) | (~spl0_45 | ~spl0_65 | spl0_81)),
  inference(resolution,[],[f5038,f630])).
fof(f630,plain,(
  ~c2_1(a969) | spl0_81),
  inference(avatar_component_clause,[],[f628])).
fof(f628,plain,(
  spl0_81 <=> c2_1(a969)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_81])])).
fof(f5038,plain,(
  ( ! [X38] : (c2_1(X38) | c1_1(X38)) ) | (~spl0_45 | ~spl0_65)),
  inference(subsumption_resolution,[],[f441,f544])).
fof(f544,plain,(
  ( ! [X109] : (c2_1(X109) | c0_1(X109) | c1_1(X109)) ) | ~spl0_65),
  inference(avatar_component_clause,[],[f543])).
fof(f543,plain,(
  spl0_65 <=> ! [X109] : (c2_1(X109) | c0_1(X109) | c1_1(X109))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_65])])).
fof(f441,plain,(
  ( ! [X38] : (~c0_1(X38) | c1_1(X38) | c2_1(X38)) ) | ~spl0_45),
  inference(avatar_component_clause,[],[f440])).
fof(f440,plain,(
  spl0_45 <=> ! [X38] : (~c0_1(X38) | c1_1(X38) | c2_1(X38))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_45])])).
fof(f5066,plain,(
  ~spl0_45 | ~spl0_65 | spl0_120 | spl0_121),
  inference(avatar_contradiction_clause,[],[f5065])).
fof(f5065,plain,(
  $false | (~spl0_45 | ~spl0_65 | spl0_120 | spl0_121)),
  inference(subsumption_resolution,[],[f5050,f843])).
fof(f843,plain,(
  ~c1_1(a914) | spl0_121),
  inference(avatar_component_clause,[],[f841])).
fof(f841,plain,(
  spl0_121 <=> c1_1(a914)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_121])])).
fof(f5050,plain,(
  c1_1(a914) | (~spl0_45 | ~spl0_65 | spl0_120)),
  inference(resolution,[],[f5038,f838])).
fof(f838,plain,(
  ~c2_1(a914) | spl0_120),
  inference(avatar_component_clause,[],[f836])).
fof(f836,plain,(
  spl0_120 <=> c2_1(a914)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_120])])).
fof(f5064,plain,(
  ~spl0_45 | ~spl0_65 | spl0_127 | spl0_128),
  inference(avatar_contradiction_clause,[],[f5063])).
fof(f5063,plain,(
  $false | (~spl0_45 | ~spl0_65 | spl0_127 | spl0_128)),
  inference(subsumption_resolution,[],[f5049,f880])).
fof(f880,plain,(
  ~c1_1(a912) | spl0_128),
  inference(avatar_component_clause,[],[f878])).
fof(f878,plain,(
  spl0_128 <=> c1_1(a912)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_128])])).
fof(f5049,plain,(
  c1_1(a912) | (~spl0_45 | ~spl0_65 | spl0_127)),
  inference(resolution,[],[f5038,f875])).
fof(f875,plain,(
  ~c2_1(a912) | spl0_127),
  inference(avatar_component_clause,[],[f873])).
fof(f873,plain,(
  spl0_127 <=> c2_1(a912)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_127])])).
fof(f5062,plain,(
  ~spl0_45 | ~spl0_65 | spl0_132 | spl0_133),
  inference(avatar_contradiction_clause,[],[f5061])).
fof(f5061,plain,(
  $false | (~spl0_45 | ~spl0_65 | spl0_132 | spl0_133)),
  inference(subsumption_resolution,[],[f5047,f907])).
fof(f907,plain,(
  ~c1_1(a909) | spl0_133),
  inference(avatar_component_clause,[],[f905])).
fof(f905,plain,(
  spl0_133 <=> c1_1(a909)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_133])])).
fof(f5047,plain,(
  c1_1(a909) | (~spl0_45 | ~spl0_65 | spl0_132)),
  inference(resolution,[],[f5038,f902])).
fof(f902,plain,(
  ~c2_1(a909) | spl0_132),
  inference(avatar_component_clause,[],[f900])).
fof(f900,plain,(
  spl0_132 <=> c2_1(a909)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_132])])).
fof(f5060,plain,(
  ~spl0_45 | ~spl0_65 | spl0_138 | spl0_164),
  inference(avatar_contradiction_clause,[],[f5059])).
fof(f5059,plain,(
  $false | (~spl0_45 | ~spl0_65 | spl0_138 | spl0_164)),
  inference(subsumption_resolution,[],[f5045,f934])).
fof(f934,plain,(
  ~c1_1(a907) | spl0_138),
  inference(avatar_component_clause,[],[f932])).
fof(f932,plain,(
  spl0_138 <=> c1_1(a907)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_138])])).
fof(f5045,plain,(
  c1_1(a907) | (~spl0_45 | ~spl0_65 | spl0_164)),
  inference(resolution,[],[f5038,f3482])).
fof(f3482,plain,(
  ~c2_1(a907) | spl0_164),
  inference(avatar_component_clause,[],[f3480])).
fof(f3480,plain,(
  spl0_164 <=> c2_1(a907)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_164])])).
fof(f5018,plain,(
  ~spl0_44 | spl0_120 | spl0_121 | ~spl0_122),
  inference(avatar_contradiction_clause,[],[f5017])).
fof(f5017,plain,(
  $false | (~spl0_44 | spl0_120 | spl0_121 | ~spl0_122)),
  inference(subsumption_resolution,[],[f5016,f838])).
fof(f5016,plain,(
  c2_1(a914) | (~spl0_44 | spl0_121 | ~spl0_122)),
  inference(subsumption_resolution,[],[f5001,f843])).
fof(f5001,plain,(
  c1_1(a914) | c2_1(a914) | (~spl0_44 | ~spl0_122)),
  inference(resolution,[],[f436,f848])).
fof(f848,plain,(
  c3_1(a914) | ~spl0_122),
  inference(avatar_component_clause,[],[f846])).
fof(f846,plain,(
  spl0_122 <=> c3_1(a914)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_122])])).
fof(f436,plain,(
  ( ! [X34] : (~c3_1(X34) | c1_1(X34) | c2_1(X34)) ) | ~spl0_44),
  inference(avatar_component_clause,[],[f435])).
fof(f435,plain,(
  spl0_44 <=> ! [X34] : (~c3_1(X34) | c1_1(X34) | c2_1(X34))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_44])])).
fof(f4903,plain,(
  ~spl0_43 | ~spl0_64 | spl0_84 | spl0_85),
  inference(avatar_contradiction_clause,[],[f4902])).
fof(f4902,plain,(
  $false | (~spl0_43 | ~spl0_64 | spl0_84 | spl0_85)),
  inference(subsumption_resolution,[],[f4895,f651])).
fof(f651,plain,(
  ~c1_1(a958) | spl0_85),
  inference(avatar_component_clause,[],[f649])).
fof(f649,plain,(
  spl0_85 <=> c1_1(a958)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_85])])).
fof(f4895,plain,(
  c1_1(a958) | (~spl0_43 | ~spl0_64 | spl0_84)),
  inference(resolution,[],[f4847,f646])).
fof(f646,plain,(
  ~c3_1(a958) | spl0_84),
  inference(avatar_component_clause,[],[f644])).
fof(f644,plain,(
  spl0_84 <=> c3_1(a958)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_84])])).
fof(f4847,plain,(
  ( ! [X32] : (c3_1(X32) | c1_1(X32)) ) | (~spl0_43 | ~spl0_64)),
  inference(subsumption_resolution,[],[f432,f539])).
fof(f539,plain,(
  ( ! [X105] : (c3_1(X105) | c0_1(X105) | c1_1(X105)) ) | ~spl0_64),
  inference(avatar_component_clause,[],[f538])).
fof(f538,plain,(
  spl0_64 <=> ! [X105] : (c3_1(X105) | c0_1(X105) | c1_1(X105))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_64])])).
fof(f4897,plain,(
  ~spl0_43 | ~spl0_64 | spl0_126 | spl0_128),
  inference(avatar_contradiction_clause,[],[f4896])).
fof(f4896,plain,(
  $false | (~spl0_43 | ~spl0_64 | spl0_126 | spl0_128)),
  inference(subsumption_resolution,[],[f4889,f880])).
fof(f4889,plain,(
  c1_1(a912) | (~spl0_43 | ~spl0_64 | spl0_126)),
  inference(resolution,[],[f4847,f870])).
fof(f870,plain,(
  ~c3_1(a912) | spl0_126),
  inference(avatar_component_clause,[],[f868])).
fof(f868,plain,(
  spl0_126 <=> c3_1(a912)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_126])])).
fof(f4843,plain,(
  ~spl0_20 | ~spl0_48 | ~spl0_66 | ~spl0_67),
  inference(avatar_contradiction_clause,[],[f4842])).
fof(f4842,plain,(
  $false | (~spl0_20 | ~spl0_48 | ~spl0_66 | ~spl0_67)),
  inference(subsumption_resolution,[],[f4835,f555])).
fof(f555,plain,(
  c2_1(a957) | ~spl0_67),
  inference(avatar_component_clause,[],[f553])).
fof(f553,plain,(
  spl0_67 <=> c2_1(a957)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_67])])).
fof(f4835,plain,(
  ~c2_1(a957) | (~spl0_20 | ~spl0_48 | ~spl0_66)),
  inference(resolution,[],[f4803,f550])).
fof(f550,plain,(
  c3_1(a957) | ~spl0_66),
  inference(avatar_component_clause,[],[f548])).
fof(f548,plain,(
  spl0_66 <=> c3_1(a957)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_66])])).
fof(f4803,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c2_1(X0)) ) | (~spl0_20 | ~spl0_48)),
  inference(subsumption_resolution,[],[f333,f455])).
fof(f455,plain,(
  ( ! [X49] : (~c3_1(X49) | c0_1(X49) | ~c2_1(X49)) ) | ~spl0_48),
  inference(avatar_component_clause,[],[f454])).
fof(f454,plain,(
  spl0_48 <=> ! [X49] : (~c3_1(X49) | c0_1(X49) | ~c2_1(X49))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_48])])).
fof(f333,plain,(
  ( ! [X0] : (~c3_1(X0) | ~c0_1(X0) | ~c2_1(X0)) ) | ~spl0_20),
  inference(avatar_component_clause,[],[f332])).
fof(f332,plain,(
  spl0_20 <=> ! [X0] : (~c3_1(X0) | ~c0_1(X0) | ~c2_1(X0))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_20])])).
fof(f4841,plain,(
  ~spl0_20 | ~spl0_48 | ~spl0_72 | ~spl0_169),
  inference(avatar_contradiction_clause,[],[f4840])).
fof(f4840,plain,(
  $false | (~spl0_20 | ~spl0_48 | ~spl0_72 | ~spl0_169)),
  inference(subsumption_resolution,[],[f4834,f4234])).
fof(f4234,plain,(
  c2_1(a911) | ~spl0_169),
  inference(avatar_component_clause,[],[f4232])).
fof(f4232,plain,(
  spl0_169 <=> c2_1(a911)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_169])])).
fof(f4834,plain,(
  ~c2_1(a911) | (~spl0_20 | ~spl0_48 | ~spl0_72)),
  inference(resolution,[],[f4803,f582])).
fof(f582,plain,(
  c3_1(a911) | ~spl0_72),
  inference(avatar_component_clause,[],[f580])).
fof(f580,plain,(
  spl0_72 <=> c3_1(a911)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_72])])).
fof(f4839,plain,(
  ~spl0_20 | ~spl0_48 | ~spl0_79 | ~spl0_80),
  inference(avatar_contradiction_clause,[],[f4838])).
fof(f4838,plain,(
  $false | (~spl0_20 | ~spl0_48 | ~spl0_79 | ~spl0_80)),
  inference(subsumption_resolution,[],[f4832,f624])).
fof(f624,plain,(
  c2_1(a978) | ~spl0_80),
  inference(avatar_component_clause,[],[f622])).
fof(f622,plain,(
  spl0_80 <=> c2_1(a978)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_80])])).
fof(f4832,plain,(
  ~c2_1(a978) | (~spl0_20 | ~spl0_48 | ~spl0_79)),
  inference(resolution,[],[f4803,f619])).
fof(f619,plain,(
  c3_1(a978) | ~spl0_79),
  inference(avatar_component_clause,[],[f617])).
fof(f617,plain,(
  spl0_79 <=> c3_1(a978)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_79])])).
fof(f4819,plain,(
  ~spl0_36 | ~spl0_61 | spl0_108 | spl0_109),
  inference(avatar_contradiction_clause,[],[f4818])).
fof(f4818,plain,(
  $false | (~spl0_36 | ~spl0_61 | spl0_108 | spl0_109)),
  inference(subsumption_resolution,[],[f4809,f779])).
fof(f779,plain,(
  ~c2_1(a928) | spl0_109),
  inference(avatar_component_clause,[],[f777])).
fof(f777,plain,(
  spl0_109 <=> c2_1(a928)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_109])])).
fof(f4809,plain,(
  c2_1(a928) | (~spl0_36 | ~spl0_61 | spl0_108)),
  inference(resolution,[],[f4800,f774])).
fof(f774,plain,(
  ~c3_1(a928) | spl0_108),
  inference(avatar_component_clause,[],[f772])).
fof(f4800,plain,(
  ( ! [X16] : (c3_1(X16) | c2_1(X16)) ) | (~spl0_36 | ~spl0_61)),
  inference(subsumption_resolution,[],[f401,f523])).
fof(f523,plain,(
  ( ! [X94] : (c3_1(X94) | c0_1(X94) | c2_1(X94)) ) | ~spl0_61),
  inference(avatar_component_clause,[],[f522])).
fof(f522,plain,(
  spl0_61 <=> ! [X94] : (c3_1(X94) | c0_1(X94) | c2_1(X94))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_61])])).
fof(f401,plain,(
  ( ! [X16] : (~c0_1(X16) | c2_1(X16) | c3_1(X16)) ) | ~spl0_36),
  inference(avatar_component_clause,[],[f400])).
fof(f400,plain,(
  spl0_36 <=> ! [X16] : (~c0_1(X16) | c2_1(X16) | c3_1(X16))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_36])])).
fof(f4798,plain,(
  ~spl0_31 | spl0_90 | ~spl0_91 | ~spl0_92),
  inference(avatar_contradiction_clause,[],[f4797])).
fof(f4797,plain,(
  $false | (~spl0_31 | spl0_90 | ~spl0_91 | ~spl0_92)),
  inference(subsumption_resolution,[],[f4796,f688])).
fof(f688,plain,(
  c0_1(a950) | ~spl0_92),
  inference(avatar_component_clause,[],[f686])).
fof(f686,plain,(
  spl0_92 <=> c0_1(a950)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_92])])).
fof(f4796,plain,(
  ~c0_1(a950) | (~spl0_31 | spl0_90 | ~spl0_91)),
  inference(subsumption_resolution,[],[f4739,f678])).
fof(f678,plain,(
  ~c2_1(a950) | spl0_90),
  inference(avatar_component_clause,[],[f676])).
fof(f676,plain,(
  spl0_90 <=> c2_1(a950)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_90])])).
fof(f4739,plain,(
  c2_1(a950) | ~c0_1(a950) | (~spl0_31 | ~spl0_91)),
  inference(resolution,[],[f683,f379])).
fof(f683,plain,(
  c3_1(a950) | ~spl0_91),
  inference(avatar_component_clause,[],[f681])).
fof(f681,plain,(
  spl0_91 <=> c3_1(a950)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_91])])).
fof(f4795,plain,(
  ~spl0_33 | ~spl0_45 | spl0_90 | ~spl0_92),
  inference(avatar_contradiction_clause,[],[f4794])).
fof(f4794,plain,(
  $false | (~spl0_33 | ~spl0_45 | spl0_90 | ~spl0_92)),
  inference(subsumption_resolution,[],[f4787,f678])).
fof(f4787,plain,(
  c2_1(a950) | (~spl0_33 | ~spl0_45 | ~spl0_92)),
  inference(resolution,[],[f4738,f688])).
fof(f4738,plain,(
  ( ! [X38] : (~c0_1(X38) | c2_1(X38)) ) | (~spl0_33 | ~spl0_45)),
  inference(subsumption_resolution,[],[f441,f388])).
fof(f388,plain,(
  ( ! [X12] : (~c1_1(X12) | c2_1(X12) | ~c0_1(X12)) ) | ~spl0_33),
  inference(avatar_component_clause,[],[f387])).
fof(f387,plain,(
  spl0_33 <=> ! [X12] : (~c1_1(X12) | c2_1(X12) | ~c0_1(X12))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_33])])).
fof(f4732,plain,(
  ~spl0_46 | ~spl0_57 | ~spl0_69 | ~spl0_70),
  inference(avatar_contradiction_clause,[],[f4731])).
fof(f4731,plain,(
  $false | (~spl0_46 | ~spl0_57 | ~spl0_69 | ~spl0_70)),
  inference(subsumption_resolution,[],[f4716,f571])).
fof(f571,plain,(
  c1_1(a919) | ~spl0_70),
  inference(avatar_component_clause,[],[f569])).
fof(f569,plain,(
  spl0_70 <=> c1_1(a919)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_70])])).
fof(f4716,plain,(
  ~c1_1(a919) | (~spl0_46 | ~spl0_57 | ~spl0_69)),
  inference(resolution,[],[f4693,f566])).
fof(f566,plain,(
  c2_1(a919) | ~spl0_69),
  inference(avatar_component_clause,[],[f564])).
fof(f564,plain,(
  spl0_69 <=> c2_1(a919)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_69])])).
fof(f4693,plain,(
  ( ! [X68] : (~c2_1(X68) | ~c1_1(X68)) ) | (~spl0_46 | ~spl0_57)),
  inference(subsumption_resolution,[],[f498,f444])).
fof(f444,plain,(
  ( ! [X37] : (~c3_1(X37) | ~c1_1(X37) | ~c2_1(X37)) ) | ~spl0_46),
  inference(avatar_component_clause,[],[f443])).
fof(f443,plain,(
  spl0_46 <=> ! [X37] : (~c3_1(X37) | ~c1_1(X37) | ~c2_1(X37))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_46])])).
fof(f498,plain,(
  ( ! [X68] : (~c2_1(X68) | c3_1(X68) | ~c1_1(X68)) ) | ~spl0_57),
  inference(avatar_component_clause,[],[f497])).
fof(f497,plain,(
  spl0_57 <=> ! [X68] : (~c2_1(X68) | c3_1(X68) | ~c1_1(X68))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_57])])).
fof(f4691,plain,(
  ~spl0_31 | ~spl0_39 | ~spl0_46 | ~spl0_48 | ~spl0_55 | ~spl0_75),
  inference(avatar_contradiction_clause,[],[f4690])).
fof(f4690,plain,(
  $false | (~spl0_31 | ~spl0_39 | ~spl0_46 | ~spl0_48 | ~spl0_55 | ~spl0_75)),
  inference(subsumption_resolution,[],[f598,f4689])).
fof(f4689,plain,(
  ( ! [X37] : (~c3_1(X37)) ) | (~spl0_31 | ~spl0_39 | ~spl0_46 | ~spl0_48 | ~spl0_55)),
  inference(subsumption_resolution,[],[f4666,f4688])).
fof(f4688,plain,(
  ( ! [X21] : (~c3_1(X21) | c1_1(X21)) ) | (~spl0_31 | ~spl0_39 | ~spl0_48 | ~spl0_55)),
  inference(subsumption_resolution,[],[f413,f4665])).
fof(f4665,plain,(
  ( ! [X49] : (~c3_1(X49) | c0_1(X49)) ) | (~spl0_31 | ~spl0_48 | ~spl0_55)),
  inference(subsumption_resolution,[],[f455,f4431])).
fof(f4431,plain,(
  ( ! [X65] : (~c3_1(X65) | c2_1(X65)) ) | (~spl0_31 | ~spl0_55)),
  inference(subsumption_resolution,[],[f488,f379])).
fof(f488,plain,(
  ( ! [X65] : (~c3_1(X65) | c0_1(X65) | c2_1(X65)) ) | ~spl0_55),
  inference(avatar_component_clause,[],[f487])).
fof(f487,plain,(
  spl0_55 <=> ! [X65] : (~c3_1(X65) | c0_1(X65) | c2_1(X65))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_55])])).
fof(f413,plain,(
  ( ! [X21] : (~c3_1(X21) | c1_1(X21) | ~c0_1(X21)) ) | ~spl0_39),
  inference(avatar_component_clause,[],[f412])).
fof(f412,plain,(
  spl0_39 <=> ! [X21] : (~c3_1(X21) | c1_1(X21) | ~c0_1(X21))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_39])])).
fof(f4666,plain,(
  ( ! [X37] : (~c3_1(X37) | ~c1_1(X37)) ) | (~spl0_31 | ~spl0_46 | ~spl0_55)),
  inference(subsumption_resolution,[],[f444,f4431])).
fof(f598,plain,(
  c3_1(a900) | ~spl0_75),
  inference(avatar_component_clause,[],[f596])).
fof(f596,plain,(
  spl0_75 <=> c3_1(a900)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_75])])).
fof(f4680,plain,(
  ~spl0_31 | ~spl0_37 | ~spl0_46 | ~spl0_55 | ~spl0_79),
  inference(avatar_contradiction_clause,[],[f4675])).
fof(f4675,plain,(
  $false | (~spl0_31 | ~spl0_37 | ~spl0_46 | ~spl0_55 | ~spl0_79)),
  inference(resolution,[],[f4667,f619])).
fof(f4667,plain,(
  ( ! [X37] : (~c3_1(X37)) ) | (~spl0_31 | ~spl0_37 | ~spl0_46 | ~spl0_55)),
  inference(subsumption_resolution,[],[f4666,f4664])).
fof(f4664,plain,(
  ( ! [X17] : (~c3_1(X17) | c1_1(X17)) ) | (~spl0_31 | ~spl0_37 | ~spl0_55)),
  inference(subsumption_resolution,[],[f405,f4431])).
fof(f405,plain,(
  ( ! [X17] : (~c3_1(X17) | c1_1(X17) | ~c2_1(X17)) ) | ~spl0_37),
  inference(avatar_component_clause,[],[f404])).
fof(f404,plain,(
  spl0_37 <=> ! [X17] : (~c3_1(X17) | c1_1(X17) | ~c2_1(X17))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_37])])).
fof(f4679,plain,(
  ~spl0_31 | ~spl0_37 | ~spl0_46 | ~spl0_55 | ~spl0_72),
  inference(avatar_contradiction_clause,[],[f4676])).
fof(f4676,plain,(
  $false | (~spl0_31 | ~spl0_37 | ~spl0_46 | ~spl0_55 | ~spl0_72)),
  inference(resolution,[],[f4667,f582])).
fof(f4678,plain,(
  ~spl0_31 | ~spl0_37 | ~spl0_46 | ~spl0_55 | ~spl0_66),
  inference(avatar_contradiction_clause,[],[f4677])).
fof(f4677,plain,(
  $false | (~spl0_31 | ~spl0_37 | ~spl0_46 | ~spl0_55 | ~spl0_66)),
  inference(resolution,[],[f4667,f550])).
fof(f4659,plain,(
  spl0_164 | ~spl0_31 | ~spl0_55 | ~spl0_140),
  inference(avatar_split_clause,[],[f4435,f942,f487,f378,f3480])).
fof(f942,plain,(
  spl0_140 <=> c3_1(a907)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_140])])).
fof(f4435,plain,(
  c2_1(a907) | (~spl0_31 | ~spl0_55 | ~spl0_140)),
  inference(resolution,[],[f4431,f944])).
fof(f944,plain,(
  c3_1(a907) | ~spl0_140),
  inference(avatar_component_clause,[],[f942])).
fof(f4644,plain,(
  ~spl0_37 | ~spl0_42 | spl0_138 | ~spl0_164),
  inference(avatar_contradiction_clause,[],[f4643])).
fof(f4643,plain,(
  $false | (~spl0_37 | ~spl0_42 | spl0_138 | ~spl0_164)),
  inference(subsumption_resolution,[],[f4629,f934])).
fof(f4629,plain,(
  c1_1(a907) | (~spl0_37 | ~spl0_42 | ~spl0_164)),
  inference(resolution,[],[f4571,f3481])).
fof(f3481,plain,(
  c2_1(a907) | ~spl0_164),
  inference(avatar_component_clause,[],[f3480])).
fof(f4571,plain,(
  ( ! [X17] : (~c2_1(X17) | c1_1(X17)) ) | (~spl0_37 | ~spl0_42)),
  inference(subsumption_resolution,[],[f405,f427])).
fof(f427,plain,(
  ( ! [X29] : (~c2_1(X29) | c1_1(X29) | c3_1(X29)) ) | ~spl0_42),
  inference(avatar_component_clause,[],[f426])).
fof(f426,plain,(
  spl0_42 <=> ! [X29] : (~c2_1(X29) | c1_1(X29) | c3_1(X29))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_42])])).
fof(f4605,plain,(
  ~spl0_26 | ~spl0_31 | ~spl0_36 | ~spl0_41 | ~spl0_161),
  inference(avatar_contradiction_clause,[],[f4586])).
fof(f4586,plain,(
  $false | (~spl0_26 | ~spl0_31 | ~spl0_36 | ~spl0_41 | ~spl0_161)),
  inference(resolution,[],[f4575,f1056])).
fof(f1056,plain,(
  c0_1(a898) | ~spl0_161),
  inference(avatar_component_clause,[],[f1054])).
fof(f1054,plain,(
  spl0_161 <=> c0_1(a898)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_161])])).
fof(f4575,plain,(
  ( ! [X2] : (~c0_1(X2)) ) | (~spl0_26 | ~spl0_31 | ~spl0_36 | ~spl0_41)),
  inference(subsumption_resolution,[],[f4574,f4572])).
fof(f4572,plain,(
  ( ! [X16] : (~c0_1(X16) | c2_1(X16)) ) | (~spl0_31 | ~spl0_36)),
  inference(subsumption_resolution,[],[f401,f379])).
fof(f4574,plain,(
  ( ! [X2] : (~c2_1(X2) | ~c0_1(X2)) ) | (~spl0_26 | ~spl0_41)),
  inference(subsumption_resolution,[],[f357,f421])).
fof(f421,plain,(
  ( ! [X23] : (~c2_1(X23) | c1_1(X23) | ~c0_1(X23)) ) | ~spl0_41),
  inference(avatar_component_clause,[],[f420])).
fof(f420,plain,(
  spl0_41 <=> ! [X23] : (~c2_1(X23) | c1_1(X23) | ~c0_1(X23))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_41])])).
fof(f357,plain,(
  ( ! [X2] : (~c2_1(X2) | ~c0_1(X2) | ~c1_1(X2)) ) | ~spl0_26),
  inference(avatar_component_clause,[],[f356])).
fof(f356,plain,(
  spl0_26 <=> ! [X2] : (~c2_1(X2) | ~c0_1(X2) | ~c1_1(X2))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_26])])).
fof(f4601,plain,(
  ~spl0_26 | ~spl0_31 | ~spl0_36 | ~spl0_41 | ~spl0_116),
  inference(avatar_contradiction_clause,[],[f4590])).
fof(f4590,plain,(
  $false | (~spl0_26 | ~spl0_31 | ~spl0_36 | ~spl0_41 | ~spl0_116)),
  inference(resolution,[],[f4575,f816])).
fof(f816,plain,(
  c0_1(a921) | ~spl0_116),
  inference(avatar_component_clause,[],[f814])).
fof(f814,plain,(
  spl0_116 <=> c0_1(a921)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_116])])).
fof(f4599,plain,(
  ~spl0_26 | ~spl0_31 | ~spl0_36 | ~spl0_41 | ~spl0_163),
  inference(avatar_contradiction_clause,[],[f4592])).
fof(f4592,plain,(
  $false | (~spl0_26 | ~spl0_31 | ~spl0_36 | ~spl0_41 | ~spl0_163)),
  inference(resolution,[],[f4575,f3474])).
fof(f4544,plain,(
  ~spl0_39 | ~spl0_91 | ~spl0_92 | spl0_170),
  inference(avatar_contradiction_clause,[],[f4543])).
fof(f4543,plain,(
  $false | (~spl0_39 | ~spl0_91 | ~spl0_92 | spl0_170)),
  inference(subsumption_resolution,[],[f4542,f688])).
fof(f4542,plain,(
  ~c0_1(a950) | (~spl0_39 | ~spl0_91 | spl0_170)),
  inference(subsumption_resolution,[],[f4532,f4240])).
fof(f4240,plain,(
  ~c1_1(a950) | spl0_170),
  inference(avatar_component_clause,[],[f4238])).
fof(f4238,plain,(
  spl0_170 <=> c1_1(a950)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_170])])).
fof(f4532,plain,(
  c1_1(a950) | ~c0_1(a950) | (~spl0_39 | ~spl0_91)),
  inference(resolution,[],[f413,f683])).
fof(f4541,plain,(
  ~spl0_39 | spl0_117 | ~spl0_118 | ~spl0_119),
  inference(avatar_contradiction_clause,[],[f4540])).
fof(f4540,plain,(
  $false | (~spl0_39 | spl0_117 | ~spl0_118 | ~spl0_119)),
  inference(subsumption_resolution,[],[f4539,f832])).
fof(f832,plain,(
  c0_1(a918) | ~spl0_119),
  inference(avatar_component_clause,[],[f830])).
fof(f830,plain,(
  spl0_119 <=> c0_1(a918)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_119])])).
fof(f4539,plain,(
  ~c0_1(a918) | (~spl0_39 | spl0_117 | ~spl0_118)),
  inference(subsumption_resolution,[],[f4530,f822])).
fof(f822,plain,(
  ~c1_1(a918) | spl0_117),
  inference(avatar_component_clause,[],[f820])).
fof(f820,plain,(
  spl0_117 <=> c1_1(a918)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_117])])).
fof(f4530,plain,(
  c1_1(a918) | ~c0_1(a918) | (~spl0_39 | ~spl0_118)),
  inference(resolution,[],[f413,f827])).
fof(f827,plain,(
  c3_1(a918) | ~spl0_118),
  inference(avatar_component_clause,[],[f825])).
fof(f825,plain,(
  spl0_118 <=> c3_1(a918)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_118])])).
fof(f4457,plain,(
  ~spl0_31 | ~spl0_55 | spl0_87 | ~spl0_88),
  inference(avatar_contradiction_clause,[],[f4456])).
fof(f4456,plain,(
  $false | (~spl0_31 | ~spl0_55 | spl0_87 | ~spl0_88)),
  inference(subsumption_resolution,[],[f4443,f662])).
fof(f4443,plain,(
  c2_1(a953) | (~spl0_31 | ~spl0_55 | ~spl0_88)),
  inference(resolution,[],[f4431,f667])).
fof(f667,plain,(
  c3_1(a953) | ~spl0_88),
  inference(avatar_component_clause,[],[f665])).
fof(f665,plain,(
  spl0_88 <=> c3_1(a953)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_88])])).
fof(f4425,plain,(
  ~spl0_51 | ~spl0_59 | ~spl0_63 | ~spl0_65 | spl0_86),
  inference(avatar_contradiction_clause,[],[f4422])).
fof(f4422,plain,(
  $false | (~spl0_51 | ~spl0_59 | ~spl0_63 | ~spl0_65 | spl0_86)),
  inference(resolution,[],[f4412,f656])).
fof(f656,plain,(
  ~c0_1(a958) | spl0_86),
  inference(avatar_component_clause,[],[f654])).
fof(f654,plain,(
  spl0_86 <=> c0_1(a958)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_86])])).
fof(f4412,plain,(
  ( ! [X54] : (c0_1(X54)) ) | (~spl0_51 | ~spl0_59 | ~spl0_63 | ~spl0_65)),
  inference(subsumption_resolution,[],[f4411,f4410])).
fof(f4410,plain,(
  ( ! [X109] : (c0_1(X109) | c1_1(X109)) ) | (~spl0_63 | ~spl0_65)),
  inference(subsumption_resolution,[],[f544,f535])).
fof(f535,plain,(
  ( ! [X104] : (~c2_1(X104) | c0_1(X104) | c1_1(X104)) ) | ~spl0_63),
  inference(avatar_component_clause,[],[f534])).
fof(f534,plain,(
  spl0_63 <=> ! [X104] : (~c2_1(X104) | c0_1(X104) | c1_1(X104))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_63])])).
fof(f4411,plain,(
  ( ! [X54] : (c0_1(X54) | ~c1_1(X54)) ) | (~spl0_51 | ~spl0_59)),
  inference(subsumption_resolution,[],[f468,f511])).
fof(f511,plain,(
  ( ! [X82] : (~c1_1(X82) | c0_1(X82) | c2_1(X82)) ) | ~spl0_59),
  inference(avatar_component_clause,[],[f510])).
fof(f510,plain,(
  spl0_59 <=> ! [X82] : (~c1_1(X82) | c0_1(X82) | c2_1(X82))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_59])])).
fof(f468,plain,(
  ( ! [X54] : (~c2_1(X54) | c0_1(X54) | ~c1_1(X54)) ) | ~spl0_51),
  inference(avatar_component_clause,[],[f467])).
fof(f467,plain,(
  spl0_51 <=> ! [X54] : (~c2_1(X54) | c0_1(X54) | ~c1_1(X54))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_51])])).
fof(f4385,plain,(
  spl0_169 | ~spl0_28 | ~spl0_35 | ~spl0_73),
  inference(avatar_split_clause,[],[f4382,f585,f395,f365,f4232])).
fof(f365,plain,(
  spl0_28 <=> ! [X7] : (~c3_1(X7) | c2_1(X7) | ~c1_1(X7))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_28])])).
fof(f395,plain,(
  spl0_35 <=> ! [X13] : (~c1_1(X13) | c2_1(X13) | c3_1(X13))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_35])])).
fof(f585,plain,(
  spl0_73 <=> c1_1(a911)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_73])])).
fof(f4382,plain,(
  c2_1(a911) | (~spl0_28 | ~spl0_35 | ~spl0_73)),
  inference(resolution,[],[f4291,f587])).
fof(f587,plain,(
  c1_1(a911) | ~spl0_73),
  inference(avatar_component_clause,[],[f585])).
fof(f4291,plain,(
  ( ! [X13] : (~c1_1(X13) | c2_1(X13)) ) | (~spl0_28 | ~spl0_35)),
  inference(subsumption_resolution,[],[f396,f366])).
fof(f366,plain,(
  ( ! [X7] : (~c3_1(X7) | c2_1(X7) | ~c1_1(X7)) ) | ~spl0_28),
  inference(avatar_component_clause,[],[f365])).
fof(f396,plain,(
  ( ! [X13] : (~c1_1(X13) | c2_1(X13) | c3_1(X13)) ) | ~spl0_35),
  inference(avatar_component_clause,[],[f395])).
fof(f4361,plain,(
  ~spl0_23 | ~spl0_52 | spl0_114 | ~spl0_168),
  inference(avatar_contradiction_clause,[],[f4360])).
fof(f4360,plain,(
  $false | (~spl0_23 | ~spl0_52 | spl0_114 | ~spl0_168)),
  inference(subsumption_resolution,[],[f4344,f806])).
fof(f806,plain,(
  ~c3_1(a921) | spl0_114),
  inference(avatar_component_clause,[],[f804])).
fof(f804,plain,(
  spl0_114 <=> c3_1(a921)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_114])])).
fof(f4344,plain,(
  c3_1(a921) | (~spl0_23 | ~spl0_52 | ~spl0_168)),
  inference(resolution,[],[f4259,f3948])).
fof(f3948,plain,(
  c2_1(a921) | ~spl0_168),
  inference(avatar_component_clause,[],[f3946])).
fof(f3946,plain,(
  spl0_168 <=> c2_1(a921)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_168])])).
fof(f4259,plain,(
  ( ! [X1] : (~c2_1(X1) | c3_1(X1)) ) | (~spl0_23 | ~spl0_52)),
  inference(subsumption_resolution,[],[f345,f474])).
fof(f474,plain,(
  ( ! [X59] : (~c2_1(X59) | c0_1(X59) | c3_1(X59)) ) | ~spl0_52),
  inference(avatar_component_clause,[],[f473])).
fof(f473,plain,(
  spl0_52 <=> ! [X59] : (~c2_1(X59) | c0_1(X59) | c3_1(X59))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_52])])).
fof(f345,plain,(
  ( ! [X1] : (~c2_1(X1) | c3_1(X1) | ~c0_1(X1)) ) | ~spl0_23),
  inference(avatar_component_clause,[],[f344])).
fof(f344,plain,(
  spl0_23 <=> ! [X1] : (~c2_1(X1) | c3_1(X1) | ~c0_1(X1))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_23])])).
fof(f4333,plain,(
  ~spl0_31 | ~spl0_55 | spl0_120 | ~spl0_122),
  inference(avatar_contradiction_clause,[],[f4332])).
fof(f4332,plain,(
  $false | (~spl0_31 | ~spl0_55 | spl0_120 | ~spl0_122)),
  inference(subsumption_resolution,[],[f4325,f838])).
fof(f4325,plain,(
  c2_1(a914) | (~spl0_31 | ~spl0_55 | ~spl0_122)),
  inference(resolution,[],[f4258,f848])).
fof(f4258,plain,(
  ( ! [X9] : (~c3_1(X9) | c2_1(X9)) ) | (~spl0_31 | ~spl0_55)),
  inference(subsumption_resolution,[],[f379,f488])).
fof(f4267,plain,(
  spl0_129 | ~spl0_55 | spl0_130 | ~spl0_165),
  inference(avatar_split_clause,[],[f4266,f3646,f889,f487,f884])).
fof(f884,plain,(
  spl0_129 <=> c2_1(a910)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_129])])).
fof(f889,plain,(
  spl0_130 <=> c0_1(a910)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_130])])).
fof(f3646,plain,(
  spl0_165 <=> c3_1(a910)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_165])])).
fof(f4266,plain,(
  c2_1(a910) | (~spl0_55 | spl0_130 | ~spl0_165)),
  inference(subsumption_resolution,[],[f4161,f891])).
fof(f891,plain,(
  ~c0_1(a910) | spl0_130),
  inference(avatar_component_clause,[],[f889])).
fof(f4161,plain,(
  c0_1(a910) | c2_1(a910) | (~spl0_55 | ~spl0_165)),
  inference(resolution,[],[f488,f3648])).
fof(f3648,plain,(
  c3_1(a910) | ~spl0_165),
  inference(avatar_component_clause,[],[f3646])).
fof(f4241,plain,(
  ~spl0_170 | ~spl0_92 | ~spl0_38 | ~spl0_91),
  inference(avatar_split_clause,[],[f4129,f681,f408,f686,f4238])).
fof(f408,plain,(
  spl0_38 <=> ! [X18] : (~c3_1(X18) | ~c0_1(X18) | ~c1_1(X18))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_38])])).
fof(f4129,plain,(
  ~c0_1(a950) | ~c1_1(a950) | (~spl0_38 | ~spl0_91)),
  inference(resolution,[],[f409,f683])).
fof(f409,plain,(
  ( ! [X18] : (~c3_1(X18) | ~c0_1(X18) | ~c1_1(X18)) ) | ~spl0_38),
  inference(avatar_component_clause,[],[f408])).
fof(f4236,plain,(
  ~spl0_73 | ~spl0_74 | ~spl0_38 | ~spl0_72),
  inference(avatar_split_clause,[],[f4132,f580,f408,f590,f585])).
fof(f590,plain,(
  spl0_74 <=> c0_1(a911)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_74])])).
fof(f4132,plain,(
  ~c0_1(a911) | ~c1_1(a911) | (~spl0_38 | ~spl0_72)),
  inference(resolution,[],[f409,f582])).
fof(f4191,plain,(
  spl0_109 | ~spl0_59 | ~spl0_110 | spl0_167),
  inference(avatar_split_clause,[],[f4181,f3791,f782,f510,f777])).
fof(f4181,plain,(
  c2_1(a928) | (~spl0_59 | ~spl0_110 | spl0_167)),
  inference(subsumption_resolution,[],[f4177,f3793])).
fof(f3793,plain,(
  ~c0_1(a928) | spl0_167),
  inference(avatar_component_clause,[],[f3791])).
fof(f4177,plain,(
  c0_1(a928) | c2_1(a928) | (~spl0_59 | ~spl0_110)),
  inference(resolution,[],[f511,f784])).
fof(f4190,plain,(
  spl0_87 | ~spl0_59 | ~spl0_89 | spl0_163),
  inference(avatar_split_clause,[],[f4185,f3473,f670,f510,f660])).
fof(f670,plain,(
  spl0_89 <=> c1_1(a953)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_89])])).
fof(f4185,plain,(
  c2_1(a953) | (~spl0_59 | ~spl0_89 | spl0_163)),
  inference(subsumption_resolution,[],[f4179,f3475])).
fof(f3475,plain,(
  ~c0_1(a953) | spl0_163),
  inference(avatar_component_clause,[],[f3473])).
fof(f4179,plain,(
  c0_1(a953) | c2_1(a953) | (~spl0_59 | ~spl0_89)),
  inference(resolution,[],[f511,f672])).
fof(f672,plain,(
  c1_1(a953) | ~spl0_89),
  inference(avatar_component_clause,[],[f670])).
fof(f4021,plain,(
  ~spl0_27 | ~spl0_31 | ~spl0_36 | ~spl0_38 | ~spl0_41 | ~spl0_125),
  inference(avatar_contradiction_clause,[],[f4010])).
fof(f4010,plain,(
  $false | (~spl0_27 | ~spl0_31 | ~spl0_36 | ~spl0_38 | ~spl0_41 | ~spl0_125)),
  inference(resolution,[],[f4006,f864])).
fof(f864,plain,(
  c0_1(a913) | ~spl0_125),
  inference(avatar_component_clause,[],[f862])).
fof(f862,plain,(
  spl0_125 <=> c0_1(a913)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_125])])).
fof(f4006,plain,(
  ( ! [X23] : (~c0_1(X23)) ) | (~spl0_27 | ~spl0_31 | ~spl0_36 | ~spl0_38 | ~spl0_41)),
  inference(subsumption_resolution,[],[f4005,f3979])).
fof(f3979,plain,(
  ( ! [X9] : (~c0_1(X9) | c2_1(X9)) ) | (~spl0_31 | ~spl0_36)),
  inference(subsumption_resolution,[],[f379,f401])).
fof(f4005,plain,(
  ( ! [X23] : (~c2_1(X23) | ~c0_1(X23)) ) | (~spl0_27 | ~spl0_38 | ~spl0_41)),
  inference(subsumption_resolution,[],[f421,f3951])).
fof(f3951,plain,(
  ( ! [X4] : (~c1_1(X4) | ~c0_1(X4)) ) | (~spl0_27 | ~spl0_38)),
  inference(subsumption_resolution,[],[f361,f409])).
fof(f3997,plain,(
  spl0_123 | ~spl0_23 | ~spl0_36 | ~spl0_125),
  inference(avatar_split_clause,[],[f3984,f862,f400,f344,f852])).
fof(f852,plain,(
  spl0_123 <=> c3_1(a913)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_123])])).
fof(f3984,plain,(
  c3_1(a913) | (~spl0_23 | ~spl0_36 | ~spl0_125)),
  inference(resolution,[],[f3952,f864])).
fof(f3952,plain,(
  ( ! [X1] : (~c0_1(X1) | c3_1(X1)) ) | (~spl0_23 | ~spl0_36)),
  inference(subsumption_resolution,[],[f345,f401])).
fof(f3968,plain,(
  ~spl0_27 | ~spl0_38 | ~spl0_160 | ~spl0_161),
  inference(avatar_contradiction_clause,[],[f3967])).
fof(f3967,plain,(
  $false | (~spl0_27 | ~spl0_38 | ~spl0_160 | ~spl0_161)),
  inference(subsumption_resolution,[],[f3955,f1056])).
fof(f3955,plain,(
  ~c0_1(a898) | (~spl0_27 | ~spl0_38 | ~spl0_160)),
  inference(resolution,[],[f3951,f1051])).
fof(f1051,plain,(
  c1_1(a898) | ~spl0_160),
  inference(avatar_component_clause,[],[f1049])).
fof(f1049,plain,(
  spl0_160 <=> c1_1(a898)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_160])])).
fof(f3950,plain,(
  spl0_114 | spl0_168 | ~spl0_35 | ~spl0_115),
  inference(avatar_split_clause,[],[f3893,f809,f395,f3946,f804])).
fof(f809,plain,(
  spl0_115 <=> c1_1(a921)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_115])])).
fof(f3893,plain,(
  c2_1(a921) | c3_1(a921) | (~spl0_35 | ~spl0_115)),
  inference(resolution,[],[f396,f811])).
fof(f811,plain,(
  c1_1(a921) | ~spl0_115),
  inference(avatar_component_clause,[],[f809])).
fof(f3949,plain,(
  spl0_114 | spl0_168 | ~spl0_36 | ~spl0_116),
  inference(avatar_split_clause,[],[f3911,f814,f400,f3946,f804])).
fof(f3911,plain,(
  c2_1(a921) | c3_1(a921) | (~spl0_36 | ~spl0_116)),
  inference(resolution,[],[f401,f816])).
fof(f3940,plain,(
  ~spl0_52 | spl0_99 | spl0_100 | ~spl0_101),
  inference(avatar_contradiction_clause,[],[f3939])).
fof(f3939,plain,(
  $false | (~spl0_52 | spl0_99 | spl0_100 | ~spl0_101)),
  inference(subsumption_resolution,[],[f3938,f726])).
fof(f726,plain,(
  ~c3_1(a937) | spl0_99),
  inference(avatar_component_clause,[],[f724])).
fof(f724,plain,(
  spl0_99 <=> c3_1(a937)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_99])])).
fof(f3938,plain,(
  c3_1(a937) | (~spl0_52 | spl0_100 | ~spl0_101)),
  inference(subsumption_resolution,[],[f3926,f731])).
fof(f731,plain,(
  ~c0_1(a937) | spl0_100),
  inference(avatar_component_clause,[],[f729])).
fof(f729,plain,(
  spl0_100 <=> c0_1(a937)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_100])])).
fof(f3926,plain,(
  c0_1(a937) | c3_1(a937) | (~spl0_52 | ~spl0_101)),
  inference(resolution,[],[f474,f736])).
fof(f736,plain,(
  c2_1(a937) | ~spl0_101),
  inference(avatar_component_clause,[],[f734])).
fof(f734,plain,(
  spl0_101 <=> c2_1(a937)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_101])])).
fof(f3936,plain,(
  ~spl0_26 | ~spl0_37 | ~spl0_42 | ~spl0_52 | spl0_111 | ~spl0_112),
  inference(avatar_contradiction_clause,[],[f3935])).
fof(f3935,plain,(
  $false | (~spl0_26 | ~spl0_37 | ~spl0_42 | ~spl0_52 | spl0_111 | ~spl0_112)),
  inference(subsumption_resolution,[],[f3934,f790])).
fof(f790,plain,(
  ~c3_1(a923) | spl0_111),
  inference(avatar_component_clause,[],[f788])).
fof(f788,plain,(
  spl0_111 <=> c3_1(a923)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_111])])).
fof(f3934,plain,(
  c3_1(a923) | (~spl0_26 | ~spl0_37 | ~spl0_42 | ~spl0_52 | ~spl0_112)),
  inference(subsumption_resolution,[],[f3924,f3860])).
fof(f3860,plain,(
  ~c0_1(a923) | (~spl0_26 | ~spl0_37 | ~spl0_42 | ~spl0_112)),
  inference(resolution,[],[f3853,f795])).
fof(f795,plain,(
  c2_1(a923) | ~spl0_112),
  inference(avatar_component_clause,[],[f793])).
fof(f793,plain,(
  spl0_112 <=> c2_1(a923)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_112])])).
fof(f3853,plain,(
  ( ! [X2] : (~c2_1(X2) | ~c0_1(X2)) ) | (~spl0_26 | ~spl0_37 | ~spl0_42)),
  inference(subsumption_resolution,[],[f357,f3796])).
fof(f3796,plain,(
  ( ! [X29] : (~c2_1(X29) | c1_1(X29)) ) | (~spl0_37 | ~spl0_42)),
  inference(subsumption_resolution,[],[f427,f405])).
fof(f3924,plain,(
  c0_1(a923) | c3_1(a923) | (~spl0_52 | ~spl0_112)),
  inference(resolution,[],[f474,f795])).
fof(f3888,plain,(
  ~spl0_27 | spl0_114 | ~spl0_115 | ~spl0_116),
  inference(avatar_contradiction_clause,[],[f3887])).
fof(f3887,plain,(
  $false | (~spl0_27 | spl0_114 | ~spl0_115 | ~spl0_116)),
  inference(subsumption_resolution,[],[f3886,f816])).
fof(f3886,plain,(
  ~c0_1(a921) | (~spl0_27 | spl0_114 | ~spl0_115)),
  inference(subsumption_resolution,[],[f3875,f806])).
fof(f3875,plain,(
  c3_1(a921) | ~c0_1(a921) | (~spl0_27 | ~spl0_115)),
  inference(resolution,[],[f361,f811])).
fof(f3852,plain,(
  spl0_105 | ~spl0_37 | ~spl0_42 | ~spl0_106),
  inference(avatar_split_clause,[],[f3841,f761,f426,f404,f756])).
fof(f756,plain,(
  spl0_105 <=> c1_1(a929)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_105])])).
fof(f761,plain,(
  spl0_106 <=> c2_1(a929)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_106])])).
fof(f3841,plain,(
  c1_1(a929) | (~spl0_37 | ~spl0_42 | ~spl0_106)),
  inference(resolution,[],[f3796,f763])).
fof(f763,plain,(
  c2_1(a929) | ~spl0_106),
  inference(avatar_component_clause,[],[f761])).
fof(f3821,plain,(
  ~spl0_62 | ~spl0_64 | spl0_85 | spl0_86),
  inference(avatar_contradiction_clause,[],[f3820])).
fof(f3820,plain,(
  $false | (~spl0_62 | ~spl0_64 | spl0_85 | spl0_86)),
  inference(subsumption_resolution,[],[f3812,f656])).
fof(f3812,plain,(
  c0_1(a958) | (~spl0_62 | ~spl0_64 | spl0_85)),
  inference(resolution,[],[f3789,f651])).
fof(f3789,plain,(
  ( ! [X97] : (c1_1(X97) | c0_1(X97)) ) | (~spl0_62 | ~spl0_64)),
  inference(subsumption_resolution,[],[f528,f539])).
fof(f528,plain,(
  ( ! [X97] : (~c3_1(X97) | c0_1(X97) | c1_1(X97)) ) | ~spl0_62),
  inference(avatar_component_clause,[],[f527])).
fof(f527,plain,(
  spl0_62 <=> ! [X97] : (~c3_1(X97) | c0_1(X97) | c1_1(X97))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_62])])).
fof(f3817,plain,(
  ~spl0_62 | ~spl0_64 | spl0_138 | spl0_139),
  inference(avatar_contradiction_clause,[],[f3816])).
fof(f3816,plain,(
  $false | (~spl0_62 | ~spl0_64 | spl0_138 | spl0_139)),
  inference(subsumption_resolution,[],[f3804,f939])).
fof(f939,plain,(
  ~c0_1(a907) | spl0_139),
  inference(avatar_component_clause,[],[f937])).
fof(f937,plain,(
  spl0_139 <=> c0_1(a907)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_139])])).
fof(f3804,plain,(
  c0_1(a907) | (~spl0_62 | ~spl0_64 | spl0_138)),
  inference(resolution,[],[f3789,f934])).
fof(f3794,plain,(
  ~spl0_167 | spl0_109 | ~spl0_33 | ~spl0_110),
  inference(avatar_split_clause,[],[f3752,f782,f387,f777,f3791])).
fof(f3752,plain,(
  c2_1(a928) | ~c0_1(a928) | (~spl0_33 | ~spl0_110)),
  inference(resolution,[],[f388,f784])).
fof(f3787,plain,(
  ~spl0_166 | ~spl0_20 | ~spl0_66 | ~spl0_67),
  inference(avatar_split_clause,[],[f3786,f553,f548,f332,f3697])).
fof(f3697,plain,(
  spl0_166 <=> c0_1(a957)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_166])])).
fof(f3786,plain,(
  ~c0_1(a957) | (~spl0_20 | ~spl0_66 | ~spl0_67)),
  inference(subsumption_resolution,[],[f3783,f555])).
fof(f3783,plain,(
  ~c0_1(a957) | ~c2_1(a957) | (~spl0_20 | ~spl0_66)),
  inference(resolution,[],[f550,f333])).
fof(f3779,plain,(
  spl0_159 | ~spl0_33 | ~spl0_160 | ~spl0_161),
  inference(avatar_split_clause,[],[f3778,f1054,f1049,f387,f1044])).
fof(f1044,plain,(
  spl0_159 <=> c2_1(a898)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_159])])).
fof(f3778,plain,(
  c2_1(a898) | (~spl0_33 | ~spl0_160 | ~spl0_161)),
  inference(subsumption_resolution,[],[f3744,f1056])).
fof(f3744,plain,(
  c2_1(a898) | ~c0_1(a898) | (~spl0_33 | ~spl0_160)),
  inference(resolution,[],[f388,f1051])).
fof(f3777,plain,(
  ~spl0_20 | ~spl0_33 | ~spl0_72 | ~spl0_73 | ~spl0_74),
  inference(avatar_contradiction_clause,[],[f3776])).
fof(f3776,plain,(
  $false | (~spl0_20 | ~spl0_33 | ~spl0_72 | ~spl0_73 | ~spl0_74)),
  inference(subsumption_resolution,[],[f3775,f592])).
fof(f592,plain,(
  c0_1(a911) | ~spl0_74),
  inference(avatar_component_clause,[],[f590])).
fof(f3775,plain,(
  ~c0_1(a911) | (~spl0_20 | ~spl0_33 | ~spl0_72 | ~spl0_73 | ~spl0_74)),
  inference(subsumption_resolution,[],[f3755,f3491])).
fof(f3491,plain,(
  ~c2_1(a911) | (~spl0_20 | ~spl0_72 | ~spl0_74)),
  inference(subsumption_resolution,[],[f3490,f592])).
fof(f3490,plain,(
  ~c0_1(a911) | ~c2_1(a911) | (~spl0_20 | ~spl0_72)),
  inference(resolution,[],[f582,f333])).
fof(f3755,plain,(
  c2_1(a911) | ~c0_1(a911) | (~spl0_33 | ~spl0_73)),
  inference(resolution,[],[f388,f587])).
fof(f3774,plain,(
  ~spl0_33 | spl0_87 | ~spl0_89 | ~spl0_163),
  inference(avatar_contradiction_clause,[],[f3773])).
fof(f3773,plain,(
  $false | (~spl0_33 | spl0_87 | ~spl0_89 | ~spl0_163)),
  inference(subsumption_resolution,[],[f3772,f3474])).
fof(f3772,plain,(
  ~c0_1(a953) | (~spl0_33 | spl0_87 | ~spl0_89)),
  inference(subsumption_resolution,[],[f3754,f662])).
fof(f3754,plain,(
  c2_1(a953) | ~c0_1(a953) | (~spl0_33 | ~spl0_89)),
  inference(resolution,[],[f388,f672])).
fof(f3740,plain,(
  ~spl0_68 | ~spl0_51 | ~spl0_67 | spl0_166),
  inference(avatar_split_clause,[],[f3739,f3697,f553,f467,f558])).
fof(f558,plain,(
  spl0_68 <=> c1_1(a957)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_68])])).
fof(f3739,plain,(
  ~c1_1(a957) | (~spl0_51 | ~spl0_67 | spl0_166)),
  inference(subsumption_resolution,[],[f3731,f3699])).
fof(f3699,plain,(
  ~c0_1(a957) | spl0_166),
  inference(avatar_component_clause,[],[f3697])).
fof(f3731,plain,(
  c0_1(a957) | ~c1_1(a957) | (~spl0_51 | ~spl0_67)),
  inference(resolution,[],[f468,f555])).
fof(f3734,plain,(
  ~spl0_51 | spl0_156 | ~spl0_157 | ~spl0_158),
  inference(avatar_contradiction_clause,[],[f3733])).
fof(f3733,plain,(
  $false | (~spl0_51 | spl0_156 | ~spl0_157 | ~spl0_158)),
  inference(subsumption_resolution,[],[f3732,f1040])).
fof(f1040,plain,(
  c1_1(a899) | ~spl0_158),
  inference(avatar_component_clause,[],[f1038])).
fof(f1038,plain,(
  spl0_158 <=> c1_1(a899)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_158])])).
fof(f3732,plain,(
  ~c1_1(a899) | (~spl0_51 | spl0_156 | ~spl0_157)),
  inference(subsumption_resolution,[],[f3721,f1030])).
fof(f1030,plain,(
  ~c0_1(a899) | spl0_156),
  inference(avatar_component_clause,[],[f1028])).
fof(f1028,plain,(
  spl0_156 <=> c0_1(a899)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_156])])).
fof(f3721,plain,(
  c0_1(a899) | ~c1_1(a899) | (~spl0_51 | ~spl0_157)),
  inference(resolution,[],[f468,f1035])).
fof(f1035,plain,(
  c2_1(a899) | ~spl0_157),
  inference(avatar_component_clause,[],[f1033])).
fof(f1033,plain,(
  spl0_157 <=> c2_1(a899)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_157])])).
fof(f3702,plain,(
  spl0_105 | ~spl0_41 | ~spl0_106 | ~spl0_107),
  inference(avatar_split_clause,[],[f3701,f766,f761,f420,f756])).
fof(f766,plain,(
  spl0_107 <=> c0_1(a929)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_107])])).
fof(f3701,plain,(
  c1_1(a929) | (~spl0_41 | ~spl0_106 | ~spl0_107)),
  inference(subsumption_resolution,[],[f3686,f768])).
fof(f768,plain,(
  c0_1(a929) | ~spl0_107),
  inference(avatar_component_clause,[],[f766])).
fof(f3686,plain,(
  c1_1(a929) | ~c0_1(a929) | (~spl0_41 | ~spl0_106)),
  inference(resolution,[],[f421,f763])).
fof(f3669,plain,(
  ~spl0_68 | ~spl0_26 | ~spl0_51 | ~spl0_67),
  inference(avatar_split_clause,[],[f3664,f553,f467,f356,f558])).
fof(f3664,plain,(
  ~c1_1(a957) | (~spl0_26 | ~spl0_51 | ~spl0_67)),
  inference(resolution,[],[f3650,f555])).
fof(f3650,plain,(
  ( ! [X54] : (~c2_1(X54) | ~c1_1(X54)) ) | (~spl0_26 | ~spl0_51)),
  inference(subsumption_resolution,[],[f468,f357])).
fof(f3668,plain,(
  ~spl0_26 | ~spl0_51 | ~spl0_112 | ~spl0_113),
  inference(avatar_contradiction_clause,[],[f3667])).
fof(f3667,plain,(
  $false | (~spl0_26 | ~spl0_51 | ~spl0_112 | ~spl0_113)),
  inference(subsumption_resolution,[],[f3658,f800])).
fof(f800,plain,(
  c1_1(a923) | ~spl0_113),
  inference(avatar_component_clause,[],[f798])).
fof(f798,plain,(
  spl0_113 <=> c1_1(a923)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_113])])).
fof(f3658,plain,(
  ~c1_1(a923) | (~spl0_26 | ~spl0_51 | ~spl0_112)),
  inference(resolution,[],[f3650,f795])).
fof(f3649,plain,(
  spl0_165 | spl0_129 | ~spl0_35 | ~spl0_131),
  inference(avatar_split_clause,[],[f3612,f894,f395,f884,f3646])).
fof(f894,plain,(
  spl0_131 <=> c1_1(a910)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_131])])).
fof(f3612,plain,(
  c2_1(a910) | c3_1(a910) | (~spl0_35 | ~spl0_131)),
  inference(resolution,[],[f396,f896])).
fof(f896,plain,(
  c1_1(a910) | ~spl0_131),
  inference(avatar_component_clause,[],[f894])).
fof(f3643,plain,(
  ~spl0_36 | spl0_96 | spl0_97 | ~spl0_98),
  inference(avatar_contradiction_clause,[],[f3642])).
fof(f3642,plain,(
  $false | (~spl0_36 | spl0_96 | spl0_97 | ~spl0_98)),
  inference(subsumption_resolution,[],[f3641,f710])).
fof(f710,plain,(
  ~c3_1(a938) | spl0_96),
  inference(avatar_component_clause,[],[f708])).
fof(f708,plain,(
  spl0_96 <=> c3_1(a938)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_96])])).
fof(f3641,plain,(
  c3_1(a938) | (~spl0_36 | spl0_97 | ~spl0_98)),
  inference(subsumption_resolution,[],[f3631,f715])).
fof(f715,plain,(
  ~c2_1(a938) | spl0_97),
  inference(avatar_component_clause,[],[f713])).
fof(f713,plain,(
  spl0_97 <=> c2_1(a938)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_97])])).
fof(f3631,plain,(
  c2_1(a938) | c3_1(a938) | (~spl0_36 | ~spl0_98)),
  inference(resolution,[],[f401,f720])).
fof(f720,plain,(
  c0_1(a938) | ~spl0_98),
  inference(avatar_component_clause,[],[f718])).
fof(f718,plain,(
  spl0_98 <=> c0_1(a938)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_98])])).
fof(f3592,plain,(
  ~spl0_26 | ~spl0_41 | ~spl0_52 | spl0_111 | ~spl0_112),
  inference(avatar_contradiction_clause,[],[f3591])).
fof(f3591,plain,(
  $false | (~spl0_26 | ~spl0_41 | ~spl0_52 | spl0_111 | ~spl0_112)),
  inference(subsumption_resolution,[],[f3580,f790])).
fof(f3580,plain,(
  c3_1(a923) | (~spl0_26 | ~spl0_41 | ~spl0_52 | ~spl0_112)),
  inference(resolution,[],[f3553,f795])).
fof(f3553,plain,(
  ( ! [X59] : (~c2_1(X59) | c3_1(X59)) ) | (~spl0_26 | ~spl0_41 | ~spl0_52)),
  inference(subsumption_resolution,[],[f474,f3552])).
fof(f3552,plain,(
  ( ! [X23] : (~c2_1(X23) | ~c0_1(X23)) ) | (~spl0_26 | ~spl0_41)),
  inference(subsumption_resolution,[],[f421,f357])).
fof(f3574,plain,(
  ~spl0_71 | ~spl0_26 | ~spl0_41 | ~spl0_69),
  inference(avatar_split_clause,[],[f3565,f564,f420,f356,f574])).
fof(f574,plain,(
  spl0_71 <=> c0_1(a919)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_71])])).
fof(f3565,plain,(
  ~c0_1(a919) | (~spl0_26 | ~spl0_41 | ~spl0_69)),
  inference(resolution,[],[f3552,f566])).
fof(f3571,plain,(
  ~spl0_77 | ~spl0_26 | ~spl0_41 | ~spl0_76),
  inference(avatar_split_clause,[],[f3564,f601,f420,f356,f606])).
fof(f606,plain,(
  spl0_77 <=> c0_1(a900)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_77])])).
fof(f601,plain,(
  spl0_76 <=> c2_1(a900)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_76])])).
fof(f3564,plain,(
  ~c0_1(a900) | (~spl0_26 | ~spl0_41 | ~spl0_76)),
  inference(resolution,[],[f3552,f603])).
fof(f603,plain,(
  c2_1(a900) | ~spl0_76),
  inference(avatar_component_clause,[],[f601])).
fof(f3477,plain,(
  ~spl0_89 | ~spl0_38 | ~spl0_49 | ~spl0_88),
  inference(avatar_split_clause,[],[f3335,f665,f458,f408,f670])).
fof(f458,plain,(
  spl0_49 <=> ! [X50] : (~c3_1(X50) | c0_1(X50) | ~c1_1(X50))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_49])])).
fof(f3335,plain,(
  ~c1_1(a953) | (~spl0_38 | ~spl0_49 | ~spl0_88)),
  inference(resolution,[],[f3209,f667])).
fof(f3209,plain,(
  ( ! [X50] : (~c3_1(X50) | ~c1_1(X50)) ) | (~spl0_38 | ~spl0_49)),
  inference(subsumption_resolution,[],[f459,f409])).
fof(f459,plain,(
  ( ! [X50] : (~c3_1(X50) | c0_1(X50) | ~c1_1(X50)) ) | ~spl0_49),
  inference(avatar_component_clause,[],[f458])).
fof(f3466,plain,(
  ~spl0_26 | ~spl0_33 | ~spl0_115 | ~spl0_116),
  inference(avatar_contradiction_clause,[],[f3465])).
fof(f3465,plain,(
  $false | (~spl0_26 | ~spl0_33 | ~spl0_115 | ~spl0_116)),
  inference(subsumption_resolution,[],[f3457,f816])).
fof(f3457,plain,(
  ~c0_1(a921) | (~spl0_26 | ~spl0_33 | ~spl0_115)),
  inference(resolution,[],[f3401,f811])).
fof(f3401,plain,(
  ( ! [X12] : (~c1_1(X12) | ~c0_1(X12)) ) | (~spl0_26 | ~spl0_33)),
  inference(subsumption_resolution,[],[f388,f357])).
fof(f3451,plain,(
  ~spl0_45 | ~spl0_55 | ~spl0_61 | spl0_81 | spl0_82),
  inference(avatar_contradiction_clause,[],[f3450])).
fof(f3450,plain,(
  $false | (~spl0_45 | ~spl0_55 | ~spl0_61 | spl0_81 | spl0_82)),
  inference(subsumption_resolution,[],[f3445,f635])).
fof(f3445,plain,(
  c1_1(a969) | (~spl0_45 | ~spl0_55 | ~spl0_61 | spl0_81)),
  inference(resolution,[],[f3400,f630])).
fof(f3400,plain,(
  ( ! [X38] : (c2_1(X38) | c1_1(X38)) ) | (~spl0_45 | ~spl0_55 | ~spl0_61)),
  inference(subsumption_resolution,[],[f441,f3374])).
fof(f3374,plain,(
  ( ! [X65] : (c2_1(X65) | c0_1(X65)) ) | (~spl0_55 | ~spl0_61)),
  inference(subsumption_resolution,[],[f488,f523])).
fof(f3370,plain,(
  ~spl0_68 | ~spl0_38 | ~spl0_49 | ~spl0_66),
  inference(avatar_split_clause,[],[f3337,f548,f458,f408,f558])).
fof(f3337,plain,(
  ~c1_1(a957) | (~spl0_38 | ~spl0_49 | ~spl0_66)),
  inference(resolution,[],[f3209,f550])).
fof(f3255,plain,(
  ~spl0_42 | ~spl0_44 | ~spl0_47 | spl0_84 | spl0_85),
  inference(avatar_contradiction_clause,[],[f3254])).
fof(f3254,plain,(
  $false | (~spl0_42 | ~spl0_44 | ~spl0_47 | spl0_84 | spl0_85)),
  inference(subsumption_resolution,[],[f3248,f651])).
fof(f3248,plain,(
  c1_1(a958) | (~spl0_42 | ~spl0_44 | ~spl0_47 | spl0_84)),
  inference(resolution,[],[f3179,f646])).
fof(f3179,plain,(
  ( ! [X29] : (c3_1(X29) | c1_1(X29)) ) | (~spl0_42 | ~spl0_44 | ~spl0_47)),
  inference(subsumption_resolution,[],[f427,f3105])).
fof(f3105,plain,(
  ( ! [X44] : (c2_1(X44) | c1_1(X44)) ) | (~spl0_44 | ~spl0_47)),
  inference(subsumption_resolution,[],[f450,f436])).
fof(f450,plain,(
  ( ! [X44] : (c3_1(X44) | c1_1(X44) | c2_1(X44)) ) | ~spl0_47),
  inference(avatar_component_clause,[],[f449])).
fof(f449,plain,(
  spl0_47 <=> ! [X44] : (c3_1(X44) | c1_1(X44) | c2_1(X44))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_47])])).
fof(f3253,plain,(
  ~spl0_42 | ~spl0_44 | ~spl0_47 | spl0_123 | spl0_124),
  inference(avatar_contradiction_clause,[],[f3252])).
fof(f3252,plain,(
  $false | (~spl0_42 | ~spl0_44 | ~spl0_47 | spl0_123 | spl0_124)),
  inference(subsumption_resolution,[],[f3241,f859])).
fof(f859,plain,(
  ~c1_1(a913) | spl0_124),
  inference(avatar_component_clause,[],[f857])).
fof(f857,plain,(
  spl0_124 <=> c1_1(a913)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_124])])).
fof(f3241,plain,(
  c1_1(a913) | (~spl0_42 | ~spl0_44 | ~spl0_47 | spl0_123)),
  inference(resolution,[],[f3179,f854])).
fof(f854,plain,(
  ~c3_1(a913) | spl0_123),
  inference(avatar_component_clause,[],[f852])).
fof(f3191,plain,(
  ~spl0_41 | ~spl0_42 | ~spl0_44 | ~spl0_47 | ~spl0_62 | spl0_85),
  inference(avatar_contradiction_clause,[],[f3190])).
fof(f3190,plain,(
  $false | (~spl0_41 | ~spl0_42 | ~spl0_44 | ~spl0_47 | ~spl0_62 | spl0_85)),
  inference(resolution,[],[f3180,f651])).
fof(f3180,plain,(
  ( ! [X29] : (c1_1(X29)) ) | (~spl0_41 | ~spl0_42 | ~spl0_44 | ~spl0_47 | ~spl0_62)),
  inference(subsumption_resolution,[],[f3179,f3178])).
fof(f3178,plain,(
  ( ! [X97] : (~c3_1(X97) | c1_1(X97)) ) | (~spl0_41 | ~spl0_44 | ~spl0_47 | ~spl0_62)),
  inference(subsumption_resolution,[],[f528,f3177])).
fof(f3177,plain,(
  ( ! [X23] : (c1_1(X23) | ~c0_1(X23)) ) | (~spl0_41 | ~spl0_44 | ~spl0_47)),
  inference(subsumption_resolution,[],[f421,f3105])).
fof(f3176,plain,(
  ~spl0_20 | ~spl0_75 | ~spl0_76 | ~spl0_77),
  inference(avatar_contradiction_clause,[],[f3175])).
fof(f3175,plain,(
  $false | (~spl0_20 | ~spl0_75 | ~spl0_76 | ~spl0_77)),
  inference(subsumption_resolution,[],[f3174,f603])).
fof(f3174,plain,(
  ~c2_1(a900) | (~spl0_20 | ~spl0_75 | ~spl0_77)),
  inference(subsumption_resolution,[],[f2659,f608])).
fof(f608,plain,(
  c0_1(a900) | ~spl0_77),
  inference(avatar_component_clause,[],[f606])).
fof(f2659,plain,(
  ~c0_1(a900) | ~c2_1(a900) | (~spl0_20 | ~spl0_75)),
  inference(resolution,[],[f598,f333])).
fof(f3169,plain,(
  ~spl0_41 | ~spl0_45 | spl0_124 | ~spl0_125),
  inference(avatar_contradiction_clause,[],[f3168])).
fof(f3168,plain,(
  $false | (~spl0_41 | ~spl0_45 | spl0_124 | ~spl0_125)),
  inference(subsumption_resolution,[],[f3157,f859])).
fof(f3157,plain,(
  c1_1(a913) | (~spl0_41 | ~spl0_45 | ~spl0_125)),
  inference(resolution,[],[f3107,f864])).
fof(f3107,plain,(
  ( ! [X23] : (~c0_1(X23) | c1_1(X23)) ) | (~spl0_41 | ~spl0_45)),
  inference(subsumption_resolution,[],[f421,f441])).
fof(f3139,plain,(
  ~spl0_55 | ~spl0_61 | spl0_154 | spl0_155),
  inference(avatar_contradiction_clause,[],[f3138])).
fof(f3138,plain,(
  $false | (~spl0_55 | ~spl0_61 | spl0_154 | spl0_155)),
  inference(subsumption_resolution,[],[f3131,f1024])).
fof(f1024,plain,(
  ~c0_1(a901) | spl0_155),
  inference(avatar_component_clause,[],[f1022])).
fof(f1022,plain,(
  spl0_155 <=> c0_1(a901)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_155])])).
fof(f3131,plain,(
  c0_1(a901) | (~spl0_55 | ~spl0_61 | spl0_154)),
  inference(resolution,[],[f3095,f1019])).
fof(f1019,plain,(
  ~c2_1(a901) | spl0_154),
  inference(avatar_component_clause,[],[f1017])).
fof(f1017,plain,(
  spl0_154 <=> c2_1(a901)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_154])])).
fof(f3095,plain,(
  ( ! [X94] : (c2_1(X94) | c0_1(X94)) ) | (~spl0_55 | ~spl0_61)),
  inference(subsumption_resolution,[],[f523,f488])).
fof(f3123,plain,(
  ~spl0_49 | ~spl0_53 | spl0_156 | ~spl0_158),
  inference(avatar_contradiction_clause,[],[f3122])).
fof(f3122,plain,(
  $false | (~spl0_49 | ~spl0_53 | spl0_156 | ~spl0_158)),
  inference(subsumption_resolution,[],[f3112,f1030])).
fof(f3112,plain,(
  c0_1(a899) | (~spl0_49 | ~spl0_53 | ~spl0_158)),
  inference(resolution,[],[f3094,f1040])).
fof(f3094,plain,(
  ( ! [X61] : (~c1_1(X61) | c0_1(X61)) ) | (~spl0_49 | ~spl0_53)),
  inference(subsumption_resolution,[],[f479,f459])).
fof(f479,plain,(
  ( ! [X61] : (~c1_1(X61) | c0_1(X61) | c3_1(X61)) ) | ~spl0_53),
  inference(avatar_component_clause,[],[f478])).
fof(f478,plain,(
  spl0_53 <=> ! [X61] : (~c1_1(X61) | c0_1(X61) | c3_1(X61))),
  introduced(avatar_definition,[new_symbols(naming,[spl0_53])])).
fof(f3075,plain,(
  ~spl0_119 | ~spl0_20 | ~spl0_28 | ~spl0_44 | ~spl0_118),
  inference(avatar_split_clause,[],[f3074,f825,f435,f365,f332,f830])).
fof(f3074,plain,(
  ~c0_1(a918) | (~spl0_20 | ~spl0_28 | ~spl0_44 | ~spl0_118)),
  inference(subsumption_resolution,[],[f2563,f3039])).
fof(f3039,plain,(
  c2_1(a918) | (~spl0_28 | ~spl0_44 | ~spl0_118)),
  inference(resolution,[],[f3009,f827])).
fof(f3009,plain,(
  ( ! [X7] : (~c3_1(X7) | c2_1(X7)) ) | (~spl0_28 | ~spl0_44)),
  inference(subsumption_resolution,[],[f366,f436])).
fof(f2563,plain,(
  ~c0_1(a918) | ~c2_1(a918) | (~spl0_20 | ~spl0_118)),
  inference(resolution,[],[f827,f333])).
fof(f3008,plain,(
  ~spl0_37 | ~spl0_44 | spl0_117 | ~spl0_118),
  inference(avatar_contradiction_clause,[],[f3007])).
fof(f3007,plain,(
  $false | (~spl0_37 | ~spl0_44 | spl0_117 | ~spl0_118)),
  inference(subsumption_resolution,[],[f2994,f822])).
fof(f2994,plain,(
  c1_1(a918) | (~spl0_37 | ~spl0_44 | ~spl0_118)),
  inference(resolution,[],[f2927,f827])).
fof(f2927,plain,(
  ( ! [X17] : (~c3_1(X17) | c1_1(X17)) ) | (~spl0_37 | ~spl0_44)),
  inference(subsumption_resolution,[],[f405,f436])).
fof(f3002,plain,(
  ~spl0_37 | ~spl0_44 | spl0_138 | ~spl0_140),
  inference(avatar_contradiction_clause,[],[f3001])).
fof(f3001,plain,(
  $false | (~spl0_37 | ~spl0_44 | spl0_138 | ~spl0_140)),
  inference(subsumption_resolution,[],[f2990,f934])).
fof(f2990,plain,(
  c1_1(a907) | (~spl0_37 | ~spl0_44 | ~spl0_140)),
  inference(resolution,[],[f2927,f944])).
fof(f3000,plain,(
  ~spl0_37 | ~spl0_44 | spl0_141 | ~spl0_142),
  inference(avatar_contradiction_clause,[],[f2999])).
fof(f2999,plain,(
  $false | (~spl0_37 | ~spl0_44 | spl0_141 | ~spl0_142)),
  inference(subsumption_resolution,[],[f2989,f950])).
fof(f950,plain,(
  ~c1_1(a906) | spl0_141),
  inference(avatar_component_clause,[],[f948])).
fof(f948,plain,(
  spl0_141 <=> c1_1(a906)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_141])])).
fof(f2989,plain,(
  c1_1(a906) | (~spl0_37 | ~spl0_44 | ~spl0_142)),
  inference(resolution,[],[f2927,f955])).
fof(f955,plain,(
  c3_1(a906) | ~spl0_142),
  inference(avatar_component_clause,[],[f953])).
fof(f953,plain,(
  spl0_142 <=> c3_1(a906)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_142])])).
fof(f2961,plain,(
  spl0_130 | ~spl0_55 | ~spl0_61 | spl0_129),
  inference(avatar_split_clause,[],[f2953,f884,f522,f487,f889])).
fof(f2953,plain,(
  c0_1(a910) | (~spl0_55 | ~spl0_61 | spl0_129)),
  inference(resolution,[],[f2925,f886])).
fof(f886,plain,(
  ~c2_1(a910) | spl0_129),
  inference(avatar_component_clause,[],[f884])).
fof(f2925,plain,(
  ( ! [X65] : (c2_1(X65) | c0_1(X65)) ) | (~spl0_55 | ~spl0_61)),
  inference(subsumption_resolution,[],[f488,f523])).
fof(f2960,plain,(
  ~spl0_55 | ~spl0_61 | spl0_132 | spl0_134),
  inference(avatar_contradiction_clause,[],[f2959])).
fof(f2959,plain,(
  $false | (~spl0_55 | ~spl0_61 | spl0_132 | spl0_134)),
  inference(subsumption_resolution,[],[f2952,f912])).
fof(f912,plain,(
  ~c0_1(a909) | spl0_134),
  inference(avatar_component_clause,[],[f910])).
fof(f910,plain,(
  spl0_134 <=> c0_1(a909)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_134])])).
fof(f2952,plain,(
  c0_1(a909) | (~spl0_55 | ~spl0_61 | spl0_132)),
  inference(resolution,[],[f2925,f902])).
fof(f2947,plain,(
  ~spl0_39 | ~spl0_43 | spl0_82 | ~spl0_83),
  inference(avatar_contradiction_clause,[],[f2946])).
fof(f2946,plain,(
  $false | (~spl0_39 | ~spl0_43 | spl0_82 | ~spl0_83)),
  inference(subsumption_resolution,[],[f2941,f635])).
fof(f2941,plain,(
  c1_1(a969) | (~spl0_39 | ~spl0_43 | ~spl0_83)),
  inference(resolution,[],[f2924,f640])).
fof(f640,plain,(
  c0_1(a969) | ~spl0_83),
  inference(avatar_component_clause,[],[f638])).
fof(f638,plain,(
  spl0_83 <=> c0_1(a969)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_83])])).
fof(f2924,plain,(
  ( ! [X32] : (~c0_1(X32) | c1_1(X32)) ) | (~spl0_39 | ~spl0_43)),
  inference(subsumption_resolution,[],[f432,f413])).
fof(f2906,plain,(
  ~spl0_20 | ~spl0_31 | ~spl0_49 | ~spl0_145 | ~spl0_146),
  inference(avatar_contradiction_clause,[],[f2905])).
fof(f2905,plain,(
  $false | (~spl0_20 | ~spl0_31 | ~spl0_49 | ~spl0_145 | ~spl0_146)),
  inference(subsumption_resolution,[],[f2773,f976])).
fof(f976,plain,(
  c1_1(a905) | ~spl0_146),
  inference(avatar_component_clause,[],[f974])).
fof(f974,plain,(
  spl0_146 <=> c1_1(a905)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_146])])).
fof(f2773,plain,(
  ~c1_1(a905) | (~spl0_20 | ~spl0_31 | ~spl0_49 | ~spl0_145)),
  inference(resolution,[],[f2735,f971])).
fof(f971,plain,(
  c3_1(a905) | ~spl0_145),
  inference(avatar_component_clause,[],[f969])).
fof(f969,plain,(
  spl0_145 <=> c3_1(a905)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_145])])).
fof(f2735,plain,(
  ( ! [X50] : (~c3_1(X50) | ~c1_1(X50)) ) | (~spl0_20 | ~spl0_31 | ~spl0_49)),
  inference(subsumption_resolution,[],[f459,f2629])).
fof(f2629,plain,(
  ( ! [X9] : (~c3_1(X9) | ~c0_1(X9)) ) | (~spl0_20 | ~spl0_31)),
  inference(subsumption_resolution,[],[f379,f333])).
fof(f2888,plain,(
  ~spl0_46 | ~spl0_57 | ~spl0_67 | ~spl0_68),
  inference(avatar_contradiction_clause,[],[f2887])).
fof(f2887,plain,(
  $false | (~spl0_46 | ~spl0_57 | ~spl0_67 | ~spl0_68)),
  inference(subsumption_resolution,[],[f2801,f560])).
fof(f560,plain,(
  c1_1(a957) | ~spl0_68),
  inference(avatar_component_clause,[],[f558])).
fof(f2801,plain,(
  ~c1_1(a957) | (~spl0_46 | ~spl0_57 | ~spl0_67)),
  inference(resolution,[],[f2736,f555])).
fof(f2736,plain,(
  ( ! [X68] : (~c2_1(X68) | ~c1_1(X68)) ) | (~spl0_46 | ~spl0_57)),
  inference(subsumption_resolution,[],[f498,f444])).
fof(f2748,plain,(
  ~spl0_35 | ~spl0_42 | ~spl0_47 | spl0_84 | spl0_85),
  inference(avatar_contradiction_clause,[],[f2747])).
fof(f2747,plain,(
  $false | (~spl0_35 | ~spl0_42 | ~spl0_47 | spl0_84 | spl0_85)),
  inference(subsumption_resolution,[],[f2744,f651])).
fof(f2744,plain,(
  c1_1(a958) | (~spl0_35 | ~spl0_42 | ~spl0_47 | spl0_84)),
  inference(resolution,[],[f2627,f646])).
fof(f2627,plain,(
  ( ! [X29] : (c3_1(X29) | c1_1(X29)) ) | (~spl0_35 | ~spl0_42 | ~spl0_47)),
  inference(subsumption_resolution,[],[f427,f2624])).
fof(f2624,plain,(
  ( ! [X13] : (c3_1(X13) | c2_1(X13)) ) | (~spl0_35 | ~spl0_47)),
  inference(subsumption_resolution,[],[f396,f450])).
fof(f2732,plain,(
  ~spl0_35 | ~spl0_47 | ~spl0_51 | ~spl0_55 | spl0_94 | ~spl0_95),
  inference(avatar_contradiction_clause,[],[f2731])).
fof(f2731,plain,(
  $false | (~spl0_35 | ~spl0_47 | ~spl0_51 | ~spl0_55 | spl0_94 | ~spl0_95)),
  inference(subsumption_resolution,[],[f2718,f699])).
fof(f699,plain,(
  ~c0_1(a939) | spl0_94),
  inference(avatar_component_clause,[],[f697])).
fof(f697,plain,(
  spl0_94 <=> c0_1(a939)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_94])])).
fof(f2718,plain,(
  c0_1(a939) | (~spl0_35 | ~spl0_47 | ~spl0_51 | ~spl0_55 | ~spl0_95)),
  inference(resolution,[],[f2626,f704])).
fof(f704,plain,(
  c1_1(a939) | ~spl0_95),
  inference(avatar_component_clause,[],[f702])).
fof(f702,plain,(
  spl0_95 <=> c1_1(a939)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_95])])).
fof(f2626,plain,(
  ( ! [X54] : (~c1_1(X54) | c0_1(X54)) ) | (~spl0_35 | ~spl0_47 | ~spl0_51 | ~spl0_55)),
  inference(subsumption_resolution,[],[f468,f2625])).
fof(f2625,plain,(
  ( ! [X65] : (c2_1(X65) | c0_1(X65)) ) | (~spl0_35 | ~spl0_47 | ~spl0_55)),
  inference(subsumption_resolution,[],[f488,f2624])).
fof(f2703,plain,(
  ~spl0_20 | ~spl0_23 | ~spl0_35 | ~spl0_47 | spl0_123 | ~spl0_125),
  inference(avatar_contradiction_clause,[],[f2702])).
fof(f2702,plain,(
  $false | (~spl0_20 | ~spl0_23 | ~spl0_35 | ~spl0_47 | spl0_123 | ~spl0_125)),
  inference(subsumption_resolution,[],[f2701,f864])).
fof(f2701,plain,(
  ~c0_1(a913) | (~spl0_20 | ~spl0_23 | ~spl0_35 | ~spl0_47 | spl0_123)),
  inference(resolution,[],[f2636,f1252])).
fof(f1252,plain,(
  ( ! [X1] : (~c2_1(X1) | ~c0_1(X1)) ) | (~spl0_20 | ~spl0_23)),
  inference(subsumption_resolution,[],[f345,f333])).
fof(f2636,plain,(
  c2_1(a913) | (~spl0_35 | ~spl0_47 | spl0_123)),
  inference(resolution,[],[f2624,f854])).
fof(f2696,plain,(
  spl0_134 | ~spl0_35 | ~spl0_47 | ~spl0_55 | spl0_132),
  inference(avatar_split_clause,[],[f2672,f900,f487,f449,f395,f910])).
fof(f2672,plain,(
  c0_1(a909) | (~spl0_35 | ~spl0_47 | ~spl0_55 | spl0_132)),
  inference(resolution,[],[f2625,f902])).
fof(f2695,plain,(
  spl0_130 | ~spl0_35 | ~spl0_47 | ~spl0_55 | spl0_129),
  inference(avatar_split_clause,[],[f2673,f884,f487,f449,f395,f889])).
fof(f2673,plain,(
  c0_1(a910) | (~spl0_35 | ~spl0_47 | ~spl0_55 | spl0_129)),
  inference(resolution,[],[f2625,f886])).
fof(f2683,plain,(
  ~spl0_35 | ~spl0_47 | ~spl0_55 | spl0_154 | spl0_155),
  inference(avatar_contradiction_clause,[],[f2682])).
fof(f2682,plain,(
  $false | (~spl0_35 | ~spl0_47 | ~spl0_55 | spl0_154 | spl0_155)),
  inference(subsumption_resolution,[],[f2669,f1024])).
fof(f2669,plain,(
  c0_1(a901) | (~spl0_35 | ~spl0_47 | ~spl0_55 | spl0_154)),
  inference(resolution,[],[f2625,f1019])).
fof(f2650,plain,(
  ~spl0_35 | ~spl0_47 | spl0_126 | spl0_127),
  inference(avatar_contradiction_clause,[],[f2649])).
fof(f2649,plain,(
  $false | (~spl0_35 | ~spl0_47 | spl0_126 | spl0_127)),
  inference(subsumption_resolution,[],[f2635,f875])).
fof(f2635,plain,(
  c2_1(a912) | (~spl0_35 | ~spl0_47 | spl0_126)),
  inference(resolution,[],[f2624,f870])).
fof(f2648,plain,(
  ~spl0_35 | ~spl0_47 | spl0_153 | spl0_154),
  inference(avatar_contradiction_clause,[],[f2647])).
fof(f2647,plain,(
  $false | (~spl0_35 | ~spl0_47 | spl0_153 | spl0_154)),
  inference(subsumption_resolution,[],[f2633,f1019])).
fof(f2633,plain,(
  c2_1(a901) | (~spl0_35 | ~spl0_47 | spl0_153)),
  inference(resolution,[],[f2624,f1014])).
fof(f1014,plain,(
  ~c3_1(a901) | spl0_153),
  inference(avatar_component_clause,[],[f1012])).
fof(f1012,plain,(
  spl0_153 <=> c3_1(a901)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_153])])).
fof(f2621,plain,(
  ~spl0_28 | ~spl0_35 | spl0_87 | ~spl0_89),
  inference(avatar_contradiction_clause,[],[f2620])).
fof(f2620,plain,(
  $false | (~spl0_28 | ~spl0_35 | spl0_87 | ~spl0_89)),
  inference(subsumption_resolution,[],[f2606,f662])).
fof(f2606,plain,(
  c2_1(a953) | (~spl0_28 | ~spl0_35 | ~spl0_89)),
  inference(resolution,[],[f2582,f672])).
fof(f2582,plain,(
  ( ! [X13] : (~c1_1(X13) | c2_1(X13)) ) | (~spl0_28 | ~spl0_35)),
  inference(subsumption_resolution,[],[f396,f366])).
fof(f2617,plain,(
  ~spl0_28 | ~spl0_35 | spl0_109 | ~spl0_110),
  inference(avatar_contradiction_clause,[],[f2616])).
fof(f2616,plain,(
  $false | (~spl0_28 | ~spl0_35 | spl0_109 | ~spl0_110)),
  inference(subsumption_resolution,[],[f2603,f779])).
fof(f2603,plain,(
  c2_1(a928) | (~spl0_28 | ~spl0_35 | ~spl0_110)),
  inference(resolution,[],[f2582,f784])).
fof(f2611,plain,(
  ~spl0_28 | ~spl0_35 | spl0_159 | ~spl0_160),
  inference(avatar_contradiction_clause,[],[f2610])).
fof(f2610,plain,(
  $false | (~spl0_28 | ~spl0_35 | spl0_159 | ~spl0_160)),
  inference(subsumption_resolution,[],[f2596,f1046])).
fof(f1046,plain,(
  ~c2_1(a898) | spl0_159),
  inference(avatar_component_clause,[],[f1044])).
fof(f2596,plain,(
  c2_1(a898) | (~spl0_28 | ~spl0_35 | ~spl0_160)),
  inference(resolution,[],[f2582,f1051])).
fof(f2538,plain,(
  ~spl0_20 | ~spl0_23 | ~spl0_36 | ~spl0_38 | ~spl0_59 | spl0_109 | ~spl0_110),
  inference(avatar_contradiction_clause,[],[f2537])).
fof(f2537,plain,(
  $false | (~spl0_20 | ~spl0_23 | ~spl0_36 | ~spl0_38 | ~spl0_59 | spl0_109 | ~spl0_110)),
  inference(subsumption_resolution,[],[f2536,f779])).
fof(f2536,plain,(
  c2_1(a928) | (~spl0_20 | ~spl0_23 | ~spl0_36 | ~spl0_38 | ~spl0_59 | ~spl0_110)),
  inference(subsumption_resolution,[],[f2518,f2431])).
fof(f2431,plain,(
  ~c0_1(a928) | (~spl0_20 | ~spl0_23 | ~spl0_36 | ~spl0_38 | ~spl0_110)),
  inference(resolution,[],[f2420,f784])).
fof(f2420,plain,(
  ( ! [X18] : (~c1_1(X18) | ~c0_1(X18)) ) | (~spl0_20 | ~spl0_23 | ~spl0_36 | ~spl0_38)),
  inference(subsumption_resolution,[],[f409,f2385])).
fof(f2385,plain,(
  ( ! [X16] : (~c0_1(X16) | c3_1(X16)) ) | (~spl0_20 | ~spl0_23 | ~spl0_36)),
  inference(subsumption_resolution,[],[f401,f1252])).
fof(f2518,plain,(
  c0_1(a928) | c2_1(a928) | (~spl0_59 | ~spl0_110)),
  inference(resolution,[],[f511,f784])).
fof(f2527,plain,(
  ~spl0_59 | spl0_129 | spl0_130 | ~spl0_131),
  inference(avatar_contradiction_clause,[],[f2526])).
fof(f2526,plain,(
  $false | (~spl0_59 | spl0_129 | spl0_130 | ~spl0_131)),
  inference(subsumption_resolution,[],[f2525,f886])).
fof(f2525,plain,(
  c2_1(a910) | (~spl0_59 | spl0_130 | ~spl0_131)),
  inference(subsumption_resolution,[],[f2513,f891])).
fof(f2513,plain,(
  c0_1(a910) | c2_1(a910) | (~spl0_59 | ~spl0_131)),
  inference(resolution,[],[f511,f896])).
fof(f2490,plain,(
  ~spl0_53 | spl0_93 | spl0_94 | ~spl0_95),
  inference(avatar_contradiction_clause,[],[f2489])).
fof(f2489,plain,(
  $false | (~spl0_53 | spl0_93 | spl0_94 | ~spl0_95)),
  inference(subsumption_resolution,[],[f2488,f694])).
fof(f694,plain,(
  ~c3_1(a939) | spl0_93),
  inference(avatar_component_clause,[],[f692])).
fof(f692,plain,(
  spl0_93 <=> c3_1(a939)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_93])])).
fof(f2488,plain,(
  c3_1(a939) | (~spl0_53 | spl0_94 | ~spl0_95)),
  inference(subsumption_resolution,[],[f2462,f699])).
fof(f2462,plain,(
  c0_1(a939) | c3_1(a939) | (~spl0_53 | ~spl0_95)),
  inference(resolution,[],[f479,f704])).
fof(f2370,plain,(
  ~spl0_49 | ~spl0_53 | ~spl0_62 | spl0_94 | ~spl0_95),
  inference(avatar_contradiction_clause,[],[f2369])).
fof(f2369,plain,(
  $false | (~spl0_49 | ~spl0_53 | ~spl0_62 | spl0_94 | ~spl0_95)),
  inference(subsumption_resolution,[],[f2356,f699])).
fof(f2356,plain,(
  c0_1(a939) | (~spl0_49 | ~spl0_53 | ~spl0_62 | ~spl0_95)),
  inference(resolution,[],[f2300,f704])).
fof(f2300,plain,(
  ( ! [X61] : (~c1_1(X61) | c0_1(X61)) ) | (~spl0_49 | ~spl0_53 | ~spl0_62)),
  inference(subsumption_resolution,[],[f479,f2299])).
fof(f2299,plain,(
  ( ! [X50] : (~c3_1(X50) | c0_1(X50)) ) | (~spl0_49 | ~spl0_62)),
  inference(subsumption_resolution,[],[f459,f528])).
fof(f2327,plain,(
  ~spl0_27 | ~spl0_43 | spl0_96 | ~spl0_98),
  inference(avatar_contradiction_clause,[],[f2326])).
fof(f2326,plain,(
  $false | (~spl0_27 | ~spl0_43 | spl0_96 | ~spl0_98)),
  inference(subsumption_resolution,[],[f2325,f710])).
fof(f2325,plain,(
  c3_1(a938) | (~spl0_27 | ~spl0_43 | ~spl0_98)),
  inference(resolution,[],[f720,f2240])).
fof(f2240,plain,(
  ( ! [X32] : (~c0_1(X32) | c3_1(X32)) ) | (~spl0_27 | ~spl0_43)),
  inference(subsumption_resolution,[],[f432,f361])).
fof(f2320,plain,(
  ~spl0_20 | ~spl0_23 | ~spl0_49 | ~spl0_62 | ~spl0_66 | ~spl0_67),
  inference(avatar_contradiction_clause,[],[f2319])).
fof(f2319,plain,(
  $false | (~spl0_20 | ~spl0_23 | ~spl0_49 | ~spl0_62 | ~spl0_66 | ~spl0_67)),
  inference(subsumption_resolution,[],[f2316,f2304])).
fof(f2304,plain,(
  ~c0_1(a957) | (~spl0_20 | ~spl0_23 | ~spl0_67)),
  inference(resolution,[],[f555,f1252])).
fof(f2316,plain,(
  c0_1(a957) | (~spl0_49 | ~spl0_62 | ~spl0_66)),
  inference(resolution,[],[f2299,f550])).
fof(f2294,plain,(
  ~spl0_20 | ~spl0_23 | ~spl0_35 | ~spl0_52 | spl0_93 | ~spl0_95),
  inference(avatar_contradiction_clause,[],[f2293])).
fof(f2293,plain,(
  $false | (~spl0_20 | ~spl0_23 | ~spl0_35 | ~spl0_52 | spl0_93 | ~spl0_95)),
  inference(subsumption_resolution,[],[f2279,f694])).
fof(f2279,plain,(
  c3_1(a939) | (~spl0_20 | ~spl0_23 | ~spl0_35 | ~spl0_52 | ~spl0_95)),
  inference(resolution,[],[f2238,f704])).
fof(f2238,plain,(
  ( ! [X13] : (~c1_1(X13) | c3_1(X13)) ) | (~spl0_20 | ~spl0_23 | ~spl0_35 | ~spl0_52)),
  inference(subsumption_resolution,[],[f396,f2035])).
fof(f2035,plain,(
  ( ! [X59] : (~c2_1(X59) | c3_1(X59)) ) | (~spl0_20 | ~spl0_23 | ~spl0_52)),
  inference(subsumption_resolution,[],[f474,f1252])).
fof(f2166,plain,(
  ~spl0_39 | ~spl0_62 | spl0_121 | ~spl0_122),
  inference(avatar_contradiction_clause,[],[f2165])).
fof(f2165,plain,(
  $false | (~spl0_39 | ~spl0_62 | spl0_121 | ~spl0_122)),
  inference(subsumption_resolution,[],[f2154,f843])).
fof(f2154,plain,(
  c1_1(a914) | (~spl0_39 | ~spl0_62 | ~spl0_122)),
  inference(resolution,[],[f2130,f848])).
fof(f2130,plain,(
  ( ! [X97] : (~c3_1(X97) | c1_1(X97)) ) | (~spl0_39 | ~spl0_62)),
  inference(subsumption_resolution,[],[f528,f413])).
fof(f2164,plain,(
  ~spl0_39 | ~spl0_62 | spl0_138 | ~spl0_140),
  inference(avatar_contradiction_clause,[],[f2163])).
fof(f2163,plain,(
  $false | (~spl0_39 | ~spl0_62 | spl0_138 | ~spl0_140)),
  inference(subsumption_resolution,[],[f2152,f934])).
fof(f2152,plain,(
  c1_1(a907) | (~spl0_39 | ~spl0_62 | ~spl0_140)),
  inference(resolution,[],[f2130,f944])).
fof(f2146,plain,(
  ~spl0_28 | ~spl0_44 | spl0_135 | ~spl0_137),
  inference(avatar_contradiction_clause,[],[f2145])).
fof(f2145,plain,(
  $false | (~spl0_28 | ~spl0_44 | spl0_135 | ~spl0_137)),
  inference(subsumption_resolution,[],[f2135,f918])).
fof(f918,plain,(
  ~c2_1(a908) | spl0_135),
  inference(avatar_component_clause,[],[f916])).
fof(f916,plain,(
  spl0_135 <=> c2_1(a908)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_135])])).
fof(f2135,plain,(
  c2_1(a908) | (~spl0_28 | ~spl0_44 | ~spl0_137)),
  inference(resolution,[],[f2111,f928])).
fof(f928,plain,(
  c3_1(a908) | ~spl0_137),
  inference(avatar_component_clause,[],[f926])).
fof(f926,plain,(
  spl0_137 <=> c3_1(a908)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_137])])).
fof(f2111,plain,(
  ( ! [X7] : (~c3_1(X7) | c2_1(X7)) ) | (~spl0_28 | ~spl0_44)),
  inference(subsumption_resolution,[],[f366,f436])).
fof(f2109,plain,(
  ~spl0_20 | ~spl0_23 | ~spl0_51 | ~spl0_59 | ~spl0_67 | ~spl0_68),
  inference(avatar_contradiction_clause,[],[f2108])).
fof(f2108,plain,(
  $false | (~spl0_20 | ~spl0_23 | ~spl0_51 | ~spl0_59 | ~spl0_67 | ~spl0_68)),
  inference(subsumption_resolution,[],[f2091,f1722])).
fof(f1722,plain,(
  ~c0_1(a957) | (~spl0_20 | ~spl0_23 | ~spl0_67)),
  inference(resolution,[],[f555,f1252])).
fof(f2091,plain,(
  c0_1(a957) | (~spl0_51 | ~spl0_59 | ~spl0_68)),
  inference(resolution,[],[f2037,f560])).
fof(f2037,plain,(
  ( ! [X54] : (~c1_1(X54) | c0_1(X54)) ) | (~spl0_51 | ~spl0_59)),
  inference(subsumption_resolution,[],[f468,f511])).
fof(f2105,plain,(
  ~spl0_51 | ~spl0_59 | spl0_94 | ~spl0_95),
  inference(avatar_contradiction_clause,[],[f2104])).
fof(f2104,plain,(
  $false | (~spl0_51 | ~spl0_59 | spl0_94 | ~spl0_95)),
  inference(subsumption_resolution,[],[f2089,f699])).
fof(f2089,plain,(
  c0_1(a939) | (~spl0_51 | ~spl0_59 | ~spl0_95)),
  inference(resolution,[],[f2037,f704])).
fof(f2097,plain,(
  ~spl0_51 | ~spl0_59 | spl0_144 | ~spl0_146),
  inference(avatar_contradiction_clause,[],[f2096])).
fof(f2096,plain,(
  $false | (~spl0_51 | ~spl0_59 | spl0_144 | ~spl0_146)),
  inference(subsumption_resolution,[],[f2085,f966])).
fof(f966,plain,(
  ~c0_1(a905) | spl0_144),
  inference(avatar_component_clause,[],[f964])).
fof(f964,plain,(
  spl0_144 <=> c0_1(a905)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_144])])).
fof(f2085,plain,(
  c0_1(a905) | (~spl0_51 | ~spl0_59 | ~spl0_146)),
  inference(resolution,[],[f2037,f976])).
fof(f2093,plain,(
  ~spl0_51 | ~spl0_59 | spl0_156 | ~spl0_158),
  inference(avatar_contradiction_clause,[],[f2092])).
fof(f2092,plain,(
  $false | (~spl0_51 | ~spl0_59 | spl0_156 | ~spl0_158)),
  inference(subsumption_resolution,[],[f2083,f1030])).
fof(f2083,plain,(
  c0_1(a899) | (~spl0_51 | ~spl0_59 | ~spl0_158)),
  inference(resolution,[],[f2037,f1040])).
fof(f2004,plain,(
  spl0_147 | ~spl0_42 | spl0_148 | ~spl0_149),
  inference(avatar_split_clause,[],[f2001,f990,f985,f426,f980])).
fof(f980,plain,(
  spl0_147 <=> c3_1(a904)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_147])])).
fof(f985,plain,(
  spl0_148 <=> c1_1(a904)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_148])])).
fof(f990,plain,(
  spl0_149 <=> c2_1(a904)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_149])])).
fof(f2001,plain,(
  c3_1(a904) | (~spl0_42 | spl0_148 | ~spl0_149)),
  inference(subsumption_resolution,[],[f1997,f987])).
fof(f987,plain,(
  ~c1_1(a904) | spl0_148),
  inference(avatar_component_clause,[],[f985])).
fof(f1997,plain,(
  c1_1(a904) | c3_1(a904) | (~spl0_42 | ~spl0_149)),
  inference(resolution,[],[f992,f427])).
fof(f992,plain,(
  c2_1(a904) | ~spl0_149),
  inference(avatar_component_clause,[],[f990])).
fof(f1978,plain,(
  ~spl0_49 | ~spl0_62 | spl0_139 | ~spl0_140),
  inference(avatar_contradiction_clause,[],[f1977])).
fof(f1977,plain,(
  $false | (~spl0_49 | ~spl0_62 | spl0_139 | ~spl0_140)),
  inference(subsumption_resolution,[],[f1962,f939])).
fof(f1962,plain,(
  c0_1(a907) | (~spl0_49 | ~spl0_62 | ~spl0_140)),
  inference(resolution,[],[f1925,f944])).
fof(f1925,plain,(
  ( ! [X50] : (~c3_1(X50) | c0_1(X50)) ) | (~spl0_49 | ~spl0_62)),
  inference(subsumption_resolution,[],[f459,f528])).
fof(f1974,plain,(
  ~spl0_49 | ~spl0_62 | spl0_144 | ~spl0_145),
  inference(avatar_contradiction_clause,[],[f1973])).
fof(f1973,plain,(
  $false | (~spl0_49 | ~spl0_62 | spl0_144 | ~spl0_145)),
  inference(subsumption_resolution,[],[f1960,f966])).
fof(f1960,plain,(
  c0_1(a905) | (~spl0_49 | ~spl0_62 | ~spl0_145)),
  inference(resolution,[],[f1925,f971])).
fof(f1882,plain,(
  ~spl0_20 | ~spl0_23 | ~spl0_42 | ~spl0_48 | ~spl0_57 | ~spl0_157),
  inference(avatar_contradiction_clause,[],[f1873])).
fof(f1873,plain,(
  $false | (~spl0_20 | ~spl0_23 | ~spl0_42 | ~spl0_48 | ~spl0_57 | ~spl0_157)),
  inference(resolution,[],[f1869,f1035])).
fof(f1869,plain,(
  ( ! [X68] : (~c2_1(X68)) ) | (~spl0_20 | ~spl0_23 | ~spl0_42 | ~spl0_48 | ~spl0_57)),
  inference(subsumption_resolution,[],[f1868,f1867])).
fof(f1867,plain,(
  ( ! [X49] : (~c3_1(X49) | ~c2_1(X49)) ) | (~spl0_20 | ~spl0_23 | ~spl0_48)),
  inference(subsumption_resolution,[],[f455,f1252])).
fof(f1868,plain,(
  ( ! [X68] : (~c2_1(X68) | c3_1(X68)) ) | (~spl0_42 | ~spl0_57)),
  inference(subsumption_resolution,[],[f498,f427])).
fof(f1864,plain,(
  ~spl0_48 | ~spl0_55 | spl0_78 | ~spl0_79),
  inference(avatar_contradiction_clause,[],[f1863])).
fof(f1863,plain,(
  $false | (~spl0_48 | ~spl0_55 | spl0_78 | ~spl0_79)),
  inference(subsumption_resolution,[],[f1847,f614])).
fof(f614,plain,(
  ~c0_1(a978) | spl0_78),
  inference(avatar_component_clause,[],[f612])).
fof(f612,plain,(
  spl0_78 <=> c0_1(a978)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_78])])).
fof(f1847,plain,(
  c0_1(a978) | (~spl0_48 | ~spl0_55 | ~spl0_79)),
  inference(resolution,[],[f1836,f619])).
fof(f1836,plain,(
  ( ! [X49] : (~c3_1(X49) | c0_1(X49)) ) | (~spl0_48 | ~spl0_55)),
  inference(subsumption_resolution,[],[f455,f488])).
fof(f1858,plain,(
  ~spl0_42 | ~spl0_48 | ~spl0_55 | spl0_102 | spl0_103 | ~spl0_104),
  inference(avatar_contradiction_clause,[],[f1857])).
fof(f1857,plain,(
  $false | (~spl0_42 | ~spl0_48 | ~spl0_55 | spl0_102 | spl0_103 | ~spl0_104)),
  inference(subsumption_resolution,[],[f1843,f747])).
fof(f747,plain,(
  ~c0_1(a930) | spl0_103),
  inference(avatar_component_clause,[],[f745])).
fof(f745,plain,(
  spl0_103 <=> c0_1(a930)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_103])])).
fof(f1843,plain,(
  c0_1(a930) | (~spl0_42 | ~spl0_48 | ~spl0_55 | spl0_102 | ~spl0_104)),
  inference(resolution,[],[f1836,f1683])).
fof(f1683,plain,(
  c3_1(a930) | (~spl0_42 | spl0_102 | ~spl0_104)),
  inference(subsumption_resolution,[],[f1681,f742])).
fof(f742,plain,(
  ~c1_1(a930) | spl0_102),
  inference(avatar_component_clause,[],[f740])).
fof(f740,plain,(
  spl0_102 <=> c1_1(a930)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_102])])).
fof(f1681,plain,(
  c1_1(a930) | c3_1(a930) | (~spl0_42 | ~spl0_104)),
  inference(resolution,[],[f752,f427])).
fof(f752,plain,(
  c2_1(a930) | ~spl0_104),
  inference(avatar_component_clause,[],[f750])).
fof(f750,plain,(
  spl0_104 <=> c2_1(a930)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_104])])).
fof(f1856,plain,(
  ~spl0_48 | ~spl0_55 | spl0_139 | ~spl0_140),
  inference(avatar_contradiction_clause,[],[f1855])).
fof(f1855,plain,(
  $false | (~spl0_48 | ~spl0_55 | spl0_139 | ~spl0_140)),
  inference(subsumption_resolution,[],[f1840,f939])).
fof(f1840,plain,(
  c0_1(a907) | (~spl0_48 | ~spl0_55 | ~spl0_140)),
  inference(resolution,[],[f1836,f944])).
fof(f1852,plain,(
  ~spl0_48 | ~spl0_55 | spl0_144 | ~spl0_145),
  inference(avatar_contradiction_clause,[],[f1851])).
fof(f1851,plain,(
  $false | (~spl0_48 | ~spl0_55 | spl0_144 | ~spl0_145)),
  inference(subsumption_resolution,[],[f1838,f966])).
fof(f1838,plain,(
  c0_1(a905) | (~spl0_48 | ~spl0_55 | ~spl0_145)),
  inference(resolution,[],[f1836,f971])).
fof(f1811,plain,(
  ~spl0_37 | ~spl0_42 | spl0_102 | ~spl0_104),
  inference(avatar_contradiction_clause,[],[f1810])).
fof(f1810,plain,(
  $false | (~spl0_37 | ~spl0_42 | spl0_102 | ~spl0_104)),
  inference(subsumption_resolution,[],[f1803,f742])).
fof(f1803,plain,(
  c1_1(a930) | (~spl0_37 | ~spl0_42 | ~spl0_104)),
  inference(resolution,[],[f1761,f752])).
fof(f1761,plain,(
  ( ! [X17] : (~c2_1(X17) | c1_1(X17)) ) | (~spl0_37 | ~spl0_42)),
  inference(subsumption_resolution,[],[f405,f427])).
fof(f1809,plain,(
  ~spl0_37 | ~spl0_42 | spl0_141 | ~spl0_143),
  inference(avatar_contradiction_clause,[],[f1808])).
fof(f1808,plain,(
  $false | (~spl0_37 | ~spl0_42 | spl0_141 | ~spl0_143)),
  inference(subsumption_resolution,[],[f1802,f950])).
fof(f1802,plain,(
  c1_1(a906) | (~spl0_37 | ~spl0_42 | ~spl0_143)),
  inference(resolution,[],[f1761,f960])).
fof(f960,plain,(
  c2_1(a906) | ~spl0_143),
  inference(avatar_component_clause,[],[f958])).
fof(f958,plain,(
  spl0_143 <=> c2_1(a906)),
  introduced(avatar_definition,[new_symbols(naming,[spl0_143])])).
fof(f1745,plain,(
  spl0_90 | ~spl0_28 | ~spl0_44 | ~spl0_91),
  inference(avatar_split_clause,[],[f1739,f681,f435,f365,f676])).
fof(f1739,plain,(
  c2_1(a950) | (~spl0_28 | ~spl0_44 | ~spl0_91)),
  inference(resolution,[],[f1716,f683])).
fof(f1716,plain,(
  ( ! [X7] : (~c3_1(X7) | c2_1(X7)) ) | (~spl0_28 | ~spl0_44)),
  inference(subsumption_resolution,[],[f366,f436])).
fof(f1707,plain,(
  ~spl0_44 | ~spl0_47 | spl0_132 | spl0_133),
  inference(avatar_contradiction_clause,[],[f1706])).
fof(f1706,plain,(
  $false | (~spl0_44 | ~spl0_47 | spl0_132 | spl0_133)),
  inference(subsumption_resolution,[],[f1693,f907])).
fof(f1693,plain,(
  c1_1(a909) | (~spl0_44 | ~spl0_47 | spl0_132)),
  inference(resolution,[],[f1678,f902])).
fof(f1678,plain,(
  ( ! [X44] : (c2_1(X44) | c1_1(X44)) ) | (~spl0_44 | ~spl0_47)),
  inference(subsumption_resolution,[],[f450,f436])).
fof(f1614,plain,(
  ~spl0_20 | ~spl0_23 | ~spl0_48 | ~spl0_52 | ~spl0_157),
  inference(avatar_contradiction_clause,[],[f1599])).
fof(f1599,plain,(
  $false | (~spl0_20 | ~spl0_23 | ~spl0_48 | ~spl0_52 | ~spl0_157)),
  inference(resolution,[],[f1598,f1035])).
fof(f1598,plain,(
  ( ! [X49] : (~c2_1(X49)) ) | (~spl0_20 | ~spl0_23 | ~spl0_48 | ~spl0_52)),
  inference(subsumption_resolution,[],[f1597,f1252])).
fof(f1597,plain,(
  ( ! [X49] : (c0_1(X49) | ~c2_1(X49)) ) | (~spl0_20 | ~spl0_23 | ~spl0_48 | ~spl0_52)),
  inference(subsumption_resolution,[],[f455,f1556])).
fof(f1556,plain,(
  ( ! [X59] : (~c2_1(X59) | c3_1(X59)) ) | (~spl0_20 | ~spl0_23 | ~spl0_52)),
  inference(subsumption_resolution,[],[f474,f1252])).
fof(f1507,plain,(
  ~spl0_35 | spl0_108 | spl0_109 | ~spl0_110),
  inference(avatar_contradiction_clause,[],[f1506])).
fof(f1506,plain,(
  $false | (~spl0_35 | spl0_108 | spl0_109 | ~spl0_110)),
  inference(subsumption_resolution,[],[f1505,f774])).
fof(f1505,plain,(
  c3_1(a928) | (~spl0_35 | spl0_109 | ~spl0_110)),
  inference(subsumption_resolution,[],[f1501,f779])).
fof(f1501,plain,(
  c2_1(a928) | c3_1(a928) | (~spl0_35 | ~spl0_110)),
  inference(resolution,[],[f396,f784])).
fof(f1333,plain,(
  ~spl0_20 | ~spl0_23 | ~spl0_76 | ~spl0_77),
  inference(avatar_contradiction_clause,[],[f1332])).
fof(f1332,plain,(
  $false | (~spl0_20 | ~spl0_23 | ~spl0_76 | ~spl0_77)),
  inference(subsumption_resolution,[],[f1326,f608])).
fof(f1326,plain,(
  ~c0_1(a900) | (~spl0_20 | ~spl0_23 | ~spl0_76)),
  inference(resolution,[],[f1252,f603])).
fof(f1329,plain,(
  ~spl0_20 | ~spl0_23 | ~spl0_106 | ~spl0_107),
  inference(avatar_contradiction_clause,[],[f1328])).
fof(f1328,plain,(
  $false | (~spl0_20 | ~spl0_23 | ~spl0_106 | ~spl0_107)),
  inference(subsumption_resolution,[],[f1322,f768])).
fof(f1322,plain,(
  ~c0_1(a929) | (~spl0_20 | ~spl0_23 | ~spl0_106)),
  inference(resolution,[],[f1252,f763])).
fof(f1319,plain,(
  ~spl0_20 | ~spl0_31 | ~spl0_72 | ~spl0_74),
  inference(avatar_contradiction_clause,[],[f1318])).
fof(f1318,plain,(
  $false | (~spl0_20 | ~spl0_31 | ~spl0_72 | ~spl0_74)),
  inference(subsumption_resolution,[],[f1310,f592])).
fof(f1310,plain,(
  ~c0_1(a911) | (~spl0_20 | ~spl0_31 | ~spl0_72)),
  inference(resolution,[],[f1185,f582])).
fof(f1185,plain,(
  ( ! [X9] : (~c3_1(X9) | ~c0_1(X9)) ) | (~spl0_20 | ~spl0_31)),
  inference(subsumption_resolution,[],[f379,f333])).
fof(f1313,plain,(
  ~spl0_20 | ~spl0_31 | ~spl0_118 | ~spl0_119),
  inference(avatar_contradiction_clause,[],[f1312])).
fof(f1312,plain,(
  $false | (~spl0_20 | ~spl0_31 | ~spl0_118 | ~spl0_119)),
  inference(subsumption_resolution,[],[f1307,f832])).
fof(f1307,plain,(
  ~c0_1(a918) | (~spl0_20 | ~spl0_31 | ~spl0_118)),
  inference(resolution,[],[f1185,f827])).
fof(f1302,plain,(
  ~spl0_20 | ~spl0_31 | ~spl0_48 | ~spl0_55 | ~spl0_66),
  inference(avatar_contradiction_clause,[],[f1301])).
fof(f1301,plain,(
  $false | (~spl0_20 | ~spl0_31 | ~spl0_48 | ~spl0_55 | ~spl0_66)),
  inference(resolution,[],[f1296,f550])).
fof(f1296,plain,(
  ( ! [X49] : (~c3_1(X49)) ) | (~spl0_20 | ~spl0_31 | ~spl0_48 | ~spl0_55)),
  inference(subsumption_resolution,[],[f1295,f1249])).
fof(f1249,plain,(
  ( ! [X65] : (~c3_1(X65) | c2_1(X65)) ) | (~spl0_20 | ~spl0_31 | ~spl0_55)),
  inference(subsumption_resolution,[],[f488,f1185])).
fof(f1295,plain,(
  ( ! [X49] : (~c3_1(X49) | ~c2_1(X49)) ) | (~spl0_20 | ~spl0_31 | ~spl0_48)),
  inference(subsumption_resolution,[],[f455,f1185])).
fof(f1287,plain,(
  ~spl0_20 | ~spl0_23 | ~spl0_33 | ~spl0_45 | ~spl0_77),
  inference(avatar_contradiction_clause,[],[f1284])).
fof(f1284,plain,(
  $false | (~spl0_20 | ~spl0_23 | ~spl0_33 | ~spl0_45 | ~spl0_77)),
  inference(resolution,[],[f1274,f608])).
fof(f1274,plain,(
  ( ! [X12] : (~c0_1(X12)) ) | (~spl0_20 | ~spl0_23 | ~spl0_33 | ~spl0_45)),
  inference(subsumption_resolution,[],[f1273,f1272])).
fof(f1272,plain,(
  ( ! [X38] : (~c0_1(X38) | c1_1(X38)) ) | (~spl0_20 | ~spl0_23 | ~spl0_45)),
  inference(subsumption_resolution,[],[f441,f1252])).
fof(f1273,plain,(
  ( ! [X12] : (~c1_1(X12) | ~c0_1(X12)) ) | (~spl0_20 | ~spl0_23 | ~spl0_33)),
  inference(subsumption_resolution,[],[f388,f1252])).
fof(f1245,plain,(
  ~spl0_59 | ~spl0_65 | spl0_132 | spl0_134),
  inference(avatar_contradiction_clause,[],[f1244])).
fof(f1244,plain,(
  $false | (~spl0_59 | ~spl0_65 | spl0_132 | spl0_134)),
  inference(subsumption_resolution,[],[f1237,f912])).
fof(f1237,plain,(
  c0_1(a909) | (~spl0_59 | ~spl0_65 | spl0_132)),
  inference(resolution,[],[f1059,f902])).
fof(f1059,plain,(
  ( ! [X82] : (c2_1(X82) | c0_1(X82)) ) | (~spl0_59 | ~spl0_65)),
  inference(subsumption_resolution,[],[f511,f544])).
fof(f1212,plain,(
  ~spl0_26 | ~spl0_37 | ~spl0_42 | ~spl0_49 | ~spl0_53 | ~spl0_67),
  inference(avatar_contradiction_clause,[],[f1211])).
fof(f1211,plain,(
  $false | (~spl0_26 | ~spl0_37 | ~spl0_42 | ~spl0_49 | ~spl0_53 | ~spl0_67)),
  inference(resolution,[],[f1198,f555])).
fof(f1198,plain,(
  ( ! [X17] : (~c2_1(X17)) ) | (~spl0_26 | ~spl0_37 | ~spl0_42 | ~spl0_49 | ~spl0_53)),
  inference(subsumption_resolution,[],[f1197,f1195])).
fof(f1195,plain,(
  ( ! [X2] : (~c2_1(X2) | ~c1_1(X2)) ) | (~spl0_26 | ~spl0_49 | ~spl0_53)),
  inference(subsumption_resolution,[],[f357,f1061])).
fof(f1061,plain,(
  ( ! [X50] : (c0_1(X50) | ~c1_1(X50)) ) | (~spl0_49 | ~spl0_53)),
  inference(subsumption_resolution,[],[f459,f479])).
fof(f1197,plain,(
  ( ! [X17] : (c1_1(X17) | ~c2_1(X17)) ) | (~spl0_26 | ~spl0_37 | ~spl0_42 | ~spl0_49 | ~spl0_53)),
  inference(subsumption_resolution,[],[f405,f1196])).
fof(f1196,plain,(
  ( ! [X29] : (~c2_1(X29) | c3_1(X29)) ) | (~spl0_26 | ~spl0_42 | ~spl0_49 | ~spl0_53)),
  inference(subsumption_resolution,[],[f427,f1195])).
fof(f1153,plain,(
  ~spl0_31 | ~spl0_36 | ~spl0_59 | ~spl0_65 | spl0_109),
  inference(avatar_contradiction_clause,[],[f1150])).
fof(f1150,plain,(
  $false | (~spl0_31 | ~spl0_36 | ~spl0_59 | ~spl0_65 | spl0_109)),
  inference(resolution,[],[f1141,f779])).
fof(f1141,plain,(
  ( ! [X9] : (c2_1(X9)) ) | (~spl0_31 | ~spl0_36 | ~spl0_59 | ~spl0_65)),
  inference(subsumption_resolution,[],[f1140,f1059])).
fof(f1140,plain,(
  ( ! [X9] : (c2_1(X9) | ~c0_1(X9)) ) | (~spl0_31 | ~spl0_36 | ~spl0_59 | ~spl0_65)),
  inference(subsumption_resolution,[],[f379,f1138])).
fof(f1138,plain,(
  ( ! [X16] : (c2_1(X16) | c3_1(X16)) ) | (~spl0_36 | ~spl0_59 | ~spl0_65)),
  inference(subsumption_resolution,[],[f401,f1059])).
fof(f1152,plain,(
  ~spl0_31 | ~spl0_36 | ~spl0_59 | ~spl0_65 | spl0_81),
  inference(avatar_contradiction_clause,[],[f1151])).
fof(f1151,plain,(
  $false | (~spl0_31 | ~spl0_36 | ~spl0_59 | ~spl0_65 | spl0_81)),
  inference(resolution,[],[f1141,f630])).
fof(f1058,plain,(
  ~spl0_14 | spl0_19),
  inference(avatar_split_clause,[],[f7,f328,f305])).
fof(f305,plain,(
  spl0_14 <=> hskp0),
  introduced(avatar_definition,[new_symbols(naming,[spl0_14])])).
fof(f328,plain,(
  spl0_19 <=> ndr1_0),
  introduced(avatar_definition,[new_symbols(naming,[spl0_19])])).
fof(f7,plain,(
  ndr1_0 | ~hskp0),
  inference(cnf_transformation,[],[f6])).
fof(f6,plain,(
  (hskp9 | hskp7 | hskp27) & (hskp22 | hskp26 | hskp23) & (hskp25 | hskp24 | hskp3) & (hskp17 | hskp22 | hskp28) & (hskp19 | hskp22 | hskp15) & (hskp25 | hskp21 | hskp0) & (hskp9 | hskp12 | hskp0) & (hskp27 | hskp31 | hskp29) & (hskp2 | hskp13 | ! [X0] : (~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0) | ~ndr1_0)) & (hskp10 | hskp16 | ! [X1] : (~c2_1(X1) | ~c0_1(X1) | c3_1(X1) | ~ndr1_0)) & (hskp9 | ! [X2] : (~c2_1(X2) | ~c1_1(X2) | ~c0_1(X2) | ~ndr1_0) | ! [X3] : (~c2_1(X3) | ~c0_1(X3) | c3_1(X3) | ~ndr1_0)) & (hskp17 | hskp26 | ! [X4] : (~c1_1(X4) | ~c0_1(X4) | c3_1(X4) | ~ndr1_0)) & (hskp21 | ! [X5] : (~c3_1(X5) | ~c2_1(X5) | ~c0_1(X5) | ~ndr1_0) | ! [X6] : (~c1_1(X6) | ~c0_1(X6) | c3_1(X6) | ~ndr1_0)) & (hskp13 | hskp1 | ! [X7] : (~c3_1(X7) | ~c1_1(X7) | c2_1(X7) | ~ndr1_0)) & (hskp6 | hskp12 | ! [X8] : (~c3_1(X8) | ~c1_1(X8) | c2_1(X8) | ~ndr1_0)) & (hskp5 | hskp1 | ! [X9] : (~c3_1(X9) | ~c0_1(X9) | c2_1(X9) | ~ndr1_0)) & (hskp0 | ! [X10] : (~c3_1(X10) | ~c2_1(X10) | ~c0_1(X10) | ~ndr1_0) | ! [X11] : (~c3_1(X11) | ~c0_1(X11) | c2_1(X11) | ~ndr1_0)) & (hskp20 | hskp1 | ! [X12] : (~c1_1(X12) | ~c0_1(X12) | c2_1(X12) | ~ndr1_0)) & (hskp25 | ! [X13] : (~c1_1(X13) | c3_1(X13) | c2_1(X13) | ~ndr1_0)) & (hskp31 | ! [X14] : (~c1_1(X14) | ~c0_1(X14) | c2_1(X14) | ~ndr1_0) | ! [X15] : (~c1_1(X15) | c3_1(X15) | c2_1(X15) | ~ndr1_0)) & (hskp7 | hskp10 | ! [X16] : (~c0_1(X16) | c3_1(X16) | c2_1(X16) | ~ndr1_0)) & (hskp17 | hskp24 | ! [X17] : (~c3_1(X17) | ~c2_1(X17) | c1_1(X17) | ~ndr1_0)) & (hskp28 | ! [X18] : (~c3_1(X18) | ~c1_1(X18) | ~c0_1(X18) | ~ndr1_0) | ! [X19] : (~c3_1(X19) | ~c2_1(X19) | c1_1(X19) | ~ndr1_0)) & (hskp18 | ! [X20] : (~c0_1(X20) | c3_1(X20) | c2_1(X20) | ~ndr1_0) | ! [X21] : (~c3_1(X21) | ~c0_1(X21) | c1_1(X21) | ~ndr1_0)) & (hskp23 | ! [X22] : (~c3_1(X22) | ~c2_1(X22) | ~c0_1(X22) | ~ndr1_0) | ! [X23] : (~c2_1(X23) | ~c0_1(X23) | c1_1(X23) | ~ndr1_0)) & (hskp7 | ! [X24] : (~c1_1(X24) | ~c0_1(X24) | c2_1(X24) | ~ndr1_0) | ! [X25] : (~c2_1(X25) | ~c0_1(X25) | c1_1(X25) | ~ndr1_0)) & (! [X26] : (~c3_1(X26) | ~c1_1(X26) | c2_1(X26) | ~ndr1_0) | ! [X27] : (~c0_1(X27) | c3_1(X27) | c2_1(X27) | ~ndr1_0) | ! [X28] : (~c2_1(X28) | ~c0_1(X28) | c1_1(X28) | ~ndr1_0)) & (hskp15 | hskp29 | ! [X29] : (~c2_1(X29) | c3_1(X29) | c1_1(X29) | ~ndr1_0)) & (hskp21 | ! [X30] : (~c1_1(X30) | ~c0_1(X30) | c2_1(X30) | ~ndr1_0) | ! [X31] : (~c2_1(X31) | c3_1(X31) | c1_1(X31) | ~ndr1_0)) & (hskp12 | hskp18 | ! [X32] : (~c0_1(X32) | c3_1(X32) | c1_1(X32) | ~ndr1_0)) & (hskp18 | ! [X33] : (~c1_1(X33) | c3_1(X33) | c2_1(X33) | ~ndr1_0) | ! [X34] : (~c3_1(X34) | c2_1(X34) | c1_1(X34) | ~ndr1_0)) & (hskp12 | ! [X35] : (~c2_1(X35) | c3_1(X35) | c1_1(X35) | ~ndr1_0) | ! [X36] : (~c3_1(X36) | c2_1(X36) | c1_1(X36) | ~ndr1_0)) & (hskp1 | ! [X37] : (~c3_1(X37) | ~c2_1(X37) | ~c1_1(X37) | ~ndr1_0) | ! [X38] : (~c0_1(X38) | c2_1(X38) | c1_1(X38) | ~ndr1_0)) & (hskp22 | ! [X39] : (~c3_1(X39) | ~c0_1(X39) | c2_1(X39) | ~ndr1_0) | ! [X40] : (~c0_1(X40) | c2_1(X40) | c1_1(X40) | ~ndr1_0)) & (hskp22 | ! [X41] : (~c3_1(X41) | ~c2_1(X41) | c1_1(X41) | ~ndr1_0) | ! [X42] : (~c0_1(X42) | c2_1(X42) | c1_1(X42) | ~ndr1_0)) & (hskp21 | ! [X43] : (~c2_1(X43) | ~c1_1(X43) | ~c0_1(X43) | ~ndr1_0) | ! [X44] : (c3_1(X44) | c2_1(X44) | c1_1(X44) | ~ndr1_0)) & (hskp20 | ! [X45] : (~c2_1(X45) | ~c0_1(X45) | c3_1(X45) | ~ndr1_0) | ! [X46] : (c3_1(X46) | c2_1(X46) | c1_1(X46) | ~ndr1_0)) & (! [X47] : (~c1_1(X47) | c3_1(X47) | c2_1(X47) | ~ndr1_0) | ! [X48] : (~c3_1(X48) | c2_1(X48) | c1_1(X48) | ~ndr1_0) | ! [X49] : (~c3_1(X49) | ~c2_1(X49) | c0_1(X49) | ~ndr1_0)) & (hskp10 | hskp12 | ! [X50] : (~c3_1(X50) | ~c1_1(X50) | c0_1(X50) | ~ndr1_0)) & (hskp14 | ! [X51] : (~c2_1(X51) | ~c0_1(X51) | c3_1(X51) | ~ndr1_0) | ! [X52] : (~c3_1(X52) | ~c1_1(X52) | c0_1(X52) | ~ndr1_0)) & (hskp9 | ! [X53] : (~c3_1(X53) | ~c1_1(X53) | ~c0_1(X53) | ~ndr1_0) | ! [X54] : (~c2_1(X54) | ~c1_1(X54) | c0_1(X54) | ~ndr1_0)) & (hskp15 | ! [X55] : (~c2_1(X55) | ~c0_1(X55) | c3_1(X55) | ~ndr1_0) | ! [X56] : (~c2_1(X56) | ~c1_1(X56) | c0_1(X56) | ~ndr1_0)) & (hskp14 | ! [X57] : (~c0_1(X57) | c3_1(X57) | c2_1(X57) | ~ndr1_0) | ! [X58] : (~c2_1(X58) | ~c1_1(X58) | c0_1(X58) | ~ndr1_0)) & (hskp19 | hskp18 | ! [X59] : (~c2_1(X59) | c3_1(X59) | c0_1(X59) | ~ndr1_0)) & (hskp17 | hskp0 | ! [X60] : (~c2_1(X60) | c3_1(X60) | c0_1(X60) | ~ndr1_0)) & (hskp7 | hskp4 | ! [X61] : (~c1_1(X61) | c3_1(X61) | c0_1(X61) | ~ndr1_0)) & (! [X62] : (~c2_1(X62) | c3_1(X62) | c1_1(X62) | ~ndr1_0) | ! [X63] : (~c3_1(X63) | ~c2_1(X63) | c0_1(X63) | ~ndr1_0) | ! [X64] : (~c1_1(X64) | c3_1(X64) | c0_1(X64) | ~ndr1_0)) & (hskp9 | hskp16 | ! [X65] : (~c3_1(X65) | c2_1(X65) | c0_1(X65) | ~ndr1_0)) & (hskp5 | hskp15 | ! [X66] : (~c3_1(X66) | c2_1(X66) | c0_1(X66) | ~ndr1_0)) & (hskp1 | hskp30 | ! [X67] : (~c3_1(X67) | c2_1(X67) | c0_1(X67) | ~ndr1_0)) & (hskp14 | ! [X68] : (~c2_1(X68) | ~c1_1(X68) | c3_1(X68) | ~ndr1_0) | ! [X69] : (~c3_1(X69) | c2_1(X69) | c0_1(X69) | ~ndr1_0)) & (hskp12 | ! [X70] : (~c2_1(X70) | ~c0_1(X70) | c3_1(X70) | ~ndr1_0) | ! [X71] : (~c3_1(X71) | c2_1(X71) | c0_1(X71) | ~ndr1_0)) & (hskp4 | ! [X72] : (~c1_1(X72) | ~c0_1(X72) | c3_1(X72) | ~ndr1_0) | ! [X73] : (~c3_1(X73) | c2_1(X73) | c0_1(X73) | ~ndr1_0)) & (hskp8 | ! [X74] : (~c1_1(X74) | c3_1(X74) | c2_1(X74) | ~ndr1_0) | ! [X75] : (~c3_1(X75) | c2_1(X75) | c0_1(X75) | ~ndr1_0)) & (! [X76] : (~c1_1(X76) | ~c0_1(X76) | c2_1(X76) | ~ndr1_0) | ! [X77] : (~c3_1(X77) | ~c0_1(X77) | c1_1(X77) | ~ndr1_0) | ! [X78] : (~c3_1(X78) | c2_1(X78) | c0_1(X78) | ~ndr1_0)) & (! [X79] : (~c3_1(X79) | ~c1_1(X79) | c2_1(X79) | ~ndr1_0) | ! [X80] : (~c2_1(X80) | c3_1(X80) | c1_1(X80) | ~ndr1_0) | ! [X81] : (~c3_1(X81) | c2_1(X81) | c0_1(X81) | ~ndr1_0)) & (hskp13 | hskp12 | ! [X82] : (~c1_1(X82) | c2_1(X82) | c0_1(X82) | ~ndr1_0)) & (hskp11 | ! [X83] : (~c3_1(X83) | ~c1_1(X83) | ~c0_1(X83) | ~ndr1_0) | ! [X84] : (~c1_1(X84) | c2_1(X84) | c0_1(X84) | ~ndr1_0)) & (hskp29 | ! [X85] : (~c1_1(X85) | ~c0_1(X85) | c3_1(X85) | ~ndr1_0) | ! [X86] : (~c1_1(X86) | c2_1(X86) | c0_1(X86) | ~ndr1_0)) & (hskp10 | ! [X87] : (c3_1(X87) | c2_1(X87) | c1_1(X87) | ~ndr1_0) | ! [X88] : (~c1_1(X88) | c2_1(X88) | c0_1(X88) | ~ndr1_0)) & (! [X89] : (c3_1(X89) | c2_1(X89) | c1_1(X89) | ~ndr1_0) | ! [X90] : (~c3_1(X90) | ~c2_1(X90) | c0_1(X90) | ~ndr1_0) | ! [X91] : (~c1_1(X91) | c2_1(X91) | c0_1(X91) | ~ndr1_0)) & (! [X92] : (~c2_1(X92) | ~c0_1(X92) | c3_1(X92) | ~ndr1_0) | ! [X93] : (~c3_1(X93) | ~c2_1(X93) | c1_1(X93) | ~ndr1_0) | ! [X94] : (c3_1(X94) | c2_1(X94) | c0_1(X94) | ~ndr1_0)) & (hskp9 | ! [X95] : (c3_1(X95) | c2_1(X95) | c1_1(X95) | ~ndr1_0) | ! [X96] : (c3_1(X96) | c2_1(X96) | c0_1(X96) | ~ndr1_0)) & (hskp8 | hskp7 | ! [X97] : (~c3_1(X97) | c1_1(X97) | c0_1(X97) | ~ndr1_0)) & (hskp6 | hskp5 | ! [X98] : (~c3_1(X98) | c1_1(X98) | c0_1(X98) | ~ndr1_0)) & (hskp4 | ! [X99] : (~c3_1(X99) | ~c0_1(X99) | c2_1(X99) | ~ndr1_0) | ! [X100] : (~c3_1(X100) | c1_1(X100) | c0_1(X100) | ~ndr1_0)) & (hskp3 | ! [X101] : (~c2_1(X101) | ~c0_1(X101) | c1_1(X101) | ~ndr1_0) | ! [X102] : (~c3_1(X102) | c1_1(X102) | c0_1(X102) | ~ndr1_0)) & (hskp1 | ! [X103] : (~c1_1(X103) | ~c0_1(X103) | c3_1(X103) | ~ndr1_0) | ! [X104] : (~c2_1(X104) | c1_1(X104) | c0_1(X104) | ~ndr1_0)) & (hskp2 | hskp28 | ! [X105] : (c3_1(X105) | c1_1(X105) | c0_1(X105) | ~ndr1_0)) & (! [X106] : (~c3_1(X106) | c2_1(X106) | c1_1(X106) | ~ndr1_0) | ! [X107] : (~c3_1(X107) | c2_1(X107) | c0_1(X107) | ~ndr1_0) | ! [X108] : (c3_1(X108) | c1_1(X108) | c0_1(X108) | ~ndr1_0)) & (hskp1 | hskp0 | ! [X109] : (c2_1(X109) | c1_1(X109) | c0_1(X109) | ~ndr1_0)) & (! [X110] : (~c3_1(X110) | ~c2_1(X110) | ~c0_1(X110) | ~ndr1_0) | ! [X111] : (~c2_1(X111) | c3_1(X111) | c1_1(X111) | ~ndr1_0) | ! [X112] : (c2_1(X112) | c1_1(X112) | c0_1(X112) | ~ndr1_0)) & ((c3_1(a957) & c2_1(a957) & c1_1(a957) & ndr1_0) | ~hskp31) & ((c2_1(a919) & c1_1(a919) & c0_1(a919) & ndr1_0) | ~hskp30) & ((c3_1(a911) & c1_1(a911) & c0_1(a911) & ndr1_0) | ~hskp29) & ((c3_1(a900) & c2_1(a900) & c0_1(a900) & ndr1_0) | ~hskp28) & ((~c0_1(a978) & c3_1(a978) & c2_1(a978) & ndr1_0) | ~hskp27) & ((~c2_1(a969) & ~c1_1(a969) & c0_1(a969) & ndr1_0) | ~hskp26) & ((~c3_1(a958) & ~c1_1(a958) & ~c0_1(a958) & ndr1_0) | ~hskp25) & ((~c2_1(a953) & c3_1(a953) & c1_1(a953) & ndr1_0) | ~hskp24) & ((~c2_1(a950) & c3_1(a950) & c0_1(a950) & ndr1_0) | ~hskp23) & ((~c3_1(a939) & ~c0_1(a939) & c1_1(a939) & ndr1_0) | ~hskp22) & ((~c3_1(a938) & ~c2_1(a938) & c0_1(a938) & ndr1_0) | ~hskp21) & ((~c3_1(a937) & ~c0_1(a937) & c2_1(a937) & ndr1_0) | ~hskp20) & ((~c1_1(a930) & ~c0_1(a930) & c2_1(a930) & ndr1_0) | ~hskp19) & ((~c1_1(a929) & c2_1(a929) & c0_1(a929) & ndr1_0) | ~hskp18) & ((~c3_1(a928) & ~c2_1(a928) & c1_1(a928) & ndr1_0) | ~hskp17) & ((~c3_1(a923) & c2_1(a923) & c1_1(a923) & ndr1_0) | ~hskp16) & ((~c3_1(a921) & c1_1(a921) & c0_1(a921) & ndr1_0) | ~hskp15) & ((~c1_1(a918) & c3_1(a918) & c0_1(a918) & ndr1_0) | ~hskp14) & ((~c2_1(a914) & ~c1_1(a914) & c3_1(a914) & ndr1_0) | ~hskp13) & ((~c3_1(a913) & ~c1_1(a913) & c0_1(a913) & ndr1_0) | ~hskp12) & ((~c3_1(a912) & ~c2_1(a912) & ~c1_1(a912) & ndr1_0) | ~hskp11) & ((~c2_1(a910) & ~c0_1(a910) & c1_1(a910) & ndr1_0) | ~hskp10) & ((~c2_1(a909) & ~c1_1(a909) & ~c0_1(a909) & ndr1_0) | ~hskp9) & ((~c2_1(a908) & ~c0_1(a908) & c3_1(a908) & ndr1_0) | ~hskp8) & ((~c1_1(a907) & ~c0_1(a907) & c3_1(a907) & ndr1_0) | ~hskp7) & ((~c1_1(a906) & c3_1(a906) & c2_1(a906) & ndr1_0) | ~hskp6) & ((~c0_1(a905) & c3_1(a905) & c1_1(a905) & ndr1_0) | ~hskp5) & ((~c3_1(a904) & ~c1_1(a904) & c2_1(a904) & ndr1_0) | ~hskp4) & ((~c3_1(a903) & c2_1(a903) & c0_1(a903) & ndr1_0) | ~hskp3) & ((~c3_1(a901) & ~c2_1(a901) & ~c0_1(a901) & ndr1_0) | ~hskp2) & ((~c0_1(a899) & c2_1(a899) & c1_1(a899) & ndr1_0) | ~hskp1) & ((~c2_1(a898) & c1_1(a898) & c0_1(a898) & ndr1_0) | ~hskp0)),
  inference(flattening,[],[f5])).
fof(f5,plain,(
  (hskp9 | hskp7 | hskp27) & (hskp22 | hskp26 | hskp23) & (hskp25 | hskp24 | hskp3) & (hskp17 | hskp22 | hskp28) & (hskp19 | hskp22 | hskp15) & (hskp25 | hskp21 | hskp0) & (hskp9 | hskp12 | hskp0) & (hskp27 | hskp31 | hskp29) & (hskp2 | hskp13 | ! [X0] : ((~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0)) | ~ndr1_0)) & (hskp10 | hskp16 | ! [X1] : ((~c2_1(X1) | ~c0_1(X1) | c3_1(X1)) | ~ndr1_0)) & (hskp9 | ! [X2] : ((~c2_1(X2) | ~c1_1(X2) | ~c0_1(X2)) | ~ndr1_0) | ! [X3] : ((~c2_1(X3) | ~c0_1(X3) | c3_1(X3)) | ~ndr1_0)) & (hskp17 | hskp26 | ! [X4] : ((~c1_1(X4) | ~c0_1(X4) | c3_1(X4)) | ~ndr1_0)) & (hskp21 | ! [X5] : ((~c3_1(X5) | ~c2_1(X5) | ~c0_1(X5)) | ~ndr1_0) | ! [X6] : ((~c1_1(X6) | ~c0_1(X6) | c3_1(X6)) | ~ndr1_0)) & (hskp13 | hskp1 | ! [X7] : ((~c3_1(X7) | ~c1_1(X7) | c2_1(X7)) | ~ndr1_0)) & (hskp6 | hskp12 | ! [X8] : ((~c3_1(X8) | ~c1_1(X8) | c2_1(X8)) | ~ndr1_0)) & (hskp5 | hskp1 | ! [X9] : ((~c3_1(X9) | ~c0_1(X9) | c2_1(X9)) | ~ndr1_0)) & (hskp0 | ! [X10] : ((~c3_1(X10) | ~c2_1(X10) | ~c0_1(X10)) | ~ndr1_0) | ! [X11] : ((~c3_1(X11) | ~c0_1(X11) | c2_1(X11)) | ~ndr1_0)) & (hskp20 | hskp1 | ! [X12] : ((~c1_1(X12) | ~c0_1(X12) | c2_1(X12)) | ~ndr1_0)) & (hskp25 | ! [X13] : ((~c1_1(X13) | c3_1(X13) | c2_1(X13)) | ~ndr1_0)) & (hskp31 | ! [X14] : ((~c1_1(X14) | ~c0_1(X14) | c2_1(X14)) | ~ndr1_0) | ! [X15] : ((~c1_1(X15) | c3_1(X15) | c2_1(X15)) | ~ndr1_0)) & (hskp7 | hskp10 | ! [X16] : ((~c0_1(X16) | c3_1(X16) | c2_1(X16)) | ~ndr1_0)) & (hskp17 | hskp24 | ! [X17] : ((~c3_1(X17) | ~c2_1(X17) | c1_1(X17)) | ~ndr1_0)) & (hskp28 | ! [X18] : ((~c3_1(X18) | ~c1_1(X18) | ~c0_1(X18)) | ~ndr1_0) | ! [X19] : ((~c3_1(X19) | ~c2_1(X19) | c1_1(X19)) | ~ndr1_0)) & (hskp18 | ! [X20] : ((~c0_1(X20) | c3_1(X20) | c2_1(X20)) | ~ndr1_0) | ! [X21] : ((~c3_1(X21) | ~c0_1(X21) | c1_1(X21)) | ~ndr1_0)) & (hskp23 | ! [X22] : ((~c3_1(X22) | ~c2_1(X22) | ~c0_1(X22)) | ~ndr1_0) | ! [X23] : ((~c2_1(X23) | ~c0_1(X23) | c1_1(X23)) | ~ndr1_0)) & (hskp7 | ! [X24] : ((~c1_1(X24) | ~c0_1(X24) | c2_1(X24)) | ~ndr1_0) | ! [X25] : ((~c2_1(X25) | ~c0_1(X25) | c1_1(X25)) | ~ndr1_0)) & (! [X26] : ((~c3_1(X26) | ~c1_1(X26) | c2_1(X26)) | ~ndr1_0) | ! [X27] : ((~c0_1(X27) | c3_1(X27) | c2_1(X27)) | ~ndr1_0) | ! [X28] : ((~c2_1(X28) | ~c0_1(X28) | c1_1(X28)) | ~ndr1_0)) & (hskp15 | hskp29 | ! [X29] : ((~c2_1(X29) | c3_1(X29) | c1_1(X29)) | ~ndr1_0)) & (hskp21 | ! [X30] : ((~c1_1(X30) | ~c0_1(X30) | c2_1(X30)) | ~ndr1_0) | ! [X31] : ((~c2_1(X31) | c3_1(X31) | c1_1(X31)) | ~ndr1_0)) & (hskp12 | hskp18 | ! [X32] : ((~c0_1(X32) | c3_1(X32) | c1_1(X32)) | ~ndr1_0)) & (hskp18 | ! [X33] : ((~c1_1(X33) | c3_1(X33) | c2_1(X33)) | ~ndr1_0) | ! [X34] : ((~c3_1(X34) | c2_1(X34) | c1_1(X34)) | ~ndr1_0)) & (hskp12 | ! [X35] : ((~c2_1(X35) | c3_1(X35) | c1_1(X35)) | ~ndr1_0) | ! [X36] : ((~c3_1(X36) | c2_1(X36) | c1_1(X36)) | ~ndr1_0)) & (hskp1 | ! [X37] : ((~c3_1(X37) | ~c2_1(X37) | ~c1_1(X37)) | ~ndr1_0) | ! [X38] : ((~c0_1(X38) | c2_1(X38) | c1_1(X38)) | ~ndr1_0)) & (hskp22 | ! [X39] : ((~c3_1(X39) | ~c0_1(X39) | c2_1(X39)) | ~ndr1_0) | ! [X40] : ((~c0_1(X40) | c2_1(X40) | c1_1(X40)) | ~ndr1_0)) & (hskp22 | ! [X41] : ((~c3_1(X41) | ~c2_1(X41) | c1_1(X41)) | ~ndr1_0) | ! [X42] : ((~c0_1(X42) | c2_1(X42) | c1_1(X42)) | ~ndr1_0)) & (hskp21 | ! [X43] : ((~c2_1(X43) | ~c1_1(X43) | ~c0_1(X43)) | ~ndr1_0) | ! [X44] : ((c3_1(X44) | c2_1(X44) | c1_1(X44)) | ~ndr1_0)) & (hskp20 | ! [X45] : ((~c2_1(X45) | ~c0_1(X45) | c3_1(X45)) | ~ndr1_0) | ! [X46] : ((c3_1(X46) | c2_1(X46) | c1_1(X46)) | ~ndr1_0)) & (! [X47] : ((~c1_1(X47) | c3_1(X47) | c2_1(X47)) | ~ndr1_0) | ! [X48] : ((~c3_1(X48) | c2_1(X48) | c1_1(X48)) | ~ndr1_0) | ! [X49] : ((~c3_1(X49) | ~c2_1(X49) | c0_1(X49)) | ~ndr1_0)) & (hskp10 | hskp12 | ! [X50] : ((~c3_1(X50) | ~c1_1(X50) | c0_1(X50)) | ~ndr1_0)) & (hskp14 | ! [X51] : ((~c2_1(X51) | ~c0_1(X51) | c3_1(X51)) | ~ndr1_0) | ! [X52] : ((~c3_1(X52) | ~c1_1(X52) | c0_1(X52)) | ~ndr1_0)) & (hskp9 | ! [X53] : ((~c3_1(X53) | ~c1_1(X53) | ~c0_1(X53)) | ~ndr1_0) | ! [X54] : ((~c2_1(X54) | ~c1_1(X54) | c0_1(X54)) | ~ndr1_0)) & (hskp15 | ! [X55] : ((~c2_1(X55) | ~c0_1(X55) | c3_1(X55)) | ~ndr1_0) | ! [X56] : ((~c2_1(X56) | ~c1_1(X56) | c0_1(X56)) | ~ndr1_0)) & (hskp14 | ! [X57] : ((~c0_1(X57) | c3_1(X57) | c2_1(X57)) | ~ndr1_0) | ! [X58] : ((~c2_1(X58) | ~c1_1(X58) | c0_1(X58)) | ~ndr1_0)) & (hskp19 | hskp18 | ! [X59] : ((~c2_1(X59) | c3_1(X59) | c0_1(X59)) | ~ndr1_0)) & (hskp17 | hskp0 | ! [X60] : ((~c2_1(X60) | c3_1(X60) | c0_1(X60)) | ~ndr1_0)) & (hskp7 | hskp4 | ! [X61] : ((~c1_1(X61) | c3_1(X61) | c0_1(X61)) | ~ndr1_0)) & (! [X62] : ((~c2_1(X62) | c3_1(X62) | c1_1(X62)) | ~ndr1_0) | ! [X63] : ((~c3_1(X63) | ~c2_1(X63) | c0_1(X63)) | ~ndr1_0) | ! [X64] : ((~c1_1(X64) | c3_1(X64) | c0_1(X64)) | ~ndr1_0)) & (hskp9 | hskp16 | ! [X65] : ((~c3_1(X65) | c2_1(X65) | c0_1(X65)) | ~ndr1_0)) & (hskp5 | hskp15 | ! [X66] : ((~c3_1(X66) | c2_1(X66) | c0_1(X66)) | ~ndr1_0)) & (hskp1 | hskp30 | ! [X67] : ((~c3_1(X67) | c2_1(X67) | c0_1(X67)) | ~ndr1_0)) & (hskp14 | ! [X68] : ((~c2_1(X68) | ~c1_1(X68) | c3_1(X68)) | ~ndr1_0) | ! [X69] : ((~c3_1(X69) | c2_1(X69) | c0_1(X69)) | ~ndr1_0)) & (hskp12 | ! [X70] : ((~c2_1(X70) | ~c0_1(X70) | c3_1(X70)) | ~ndr1_0) | ! [X71] : ((~c3_1(X71) | c2_1(X71) | c0_1(X71)) | ~ndr1_0)) & (hskp4 | ! [X72] : ((~c1_1(X72) | ~c0_1(X72) | c3_1(X72)) | ~ndr1_0) | ! [X73] : ((~c3_1(X73) | c2_1(X73) | c0_1(X73)) | ~ndr1_0)) & (hskp8 | ! [X74] : ((~c1_1(X74) | c3_1(X74) | c2_1(X74)) | ~ndr1_0) | ! [X75] : ((~c3_1(X75) | c2_1(X75) | c0_1(X75)) | ~ndr1_0)) & (! [X76] : ((~c1_1(X76) | ~c0_1(X76) | c2_1(X76)) | ~ndr1_0) | ! [X77] : ((~c3_1(X77) | ~c0_1(X77) | c1_1(X77)) | ~ndr1_0) | ! [X78] : ((~c3_1(X78) | c2_1(X78) | c0_1(X78)) | ~ndr1_0)) & (! [X79] : ((~c3_1(X79) | ~c1_1(X79) | c2_1(X79)) | ~ndr1_0) | ! [X80] : ((~c2_1(X80) | c3_1(X80) | c1_1(X80)) | ~ndr1_0) | ! [X81] : ((~c3_1(X81) | c2_1(X81) | c0_1(X81)) | ~ndr1_0)) & (hskp13 | hskp12 | ! [X82] : ((~c1_1(X82) | c2_1(X82) | c0_1(X82)) | ~ndr1_0)) & (hskp11 | ! [X83] : ((~c3_1(X83) | ~c1_1(X83) | ~c0_1(X83)) | ~ndr1_0) | ! [X84] : ((~c1_1(X84) | c2_1(X84) | c0_1(X84)) | ~ndr1_0)) & (hskp29 | ! [X85] : ((~c1_1(X85) | ~c0_1(X85) | c3_1(X85)) | ~ndr1_0) | ! [X86] : ((~c1_1(X86) | c2_1(X86) | c0_1(X86)) | ~ndr1_0)) & (hskp10 | ! [X87] : ((c3_1(X87) | c2_1(X87) | c1_1(X87)) | ~ndr1_0) | ! [X88] : ((~c1_1(X88) | c2_1(X88) | c0_1(X88)) | ~ndr1_0)) & (! [X89] : ((c3_1(X89) | c2_1(X89) | c1_1(X89)) | ~ndr1_0) | ! [X90] : ((~c3_1(X90) | ~c2_1(X90) | c0_1(X90)) | ~ndr1_0) | ! [X91] : ((~c1_1(X91) | c2_1(X91) | c0_1(X91)) | ~ndr1_0)) & (! [X92] : ((~c2_1(X92) | ~c0_1(X92) | c3_1(X92)) | ~ndr1_0) | ! [X93] : ((~c3_1(X93) | ~c2_1(X93) | c1_1(X93)) | ~ndr1_0) | ! [X94] : ((c3_1(X94) | c2_1(X94) | c0_1(X94)) | ~ndr1_0)) & (hskp9 | ! [X95] : ((c3_1(X95) | c2_1(X95) | c1_1(X95)) | ~ndr1_0) | ! [X96] : ((c3_1(X96) | c2_1(X96) | c0_1(X96)) | ~ndr1_0)) & (hskp8 | hskp7 | ! [X97] : ((~c3_1(X97) | c1_1(X97) | c0_1(X97)) | ~ndr1_0)) & (hskp6 | hskp5 | ! [X98] : ((~c3_1(X98) | c1_1(X98) | c0_1(X98)) | ~ndr1_0)) & (hskp4 | ! [X99] : ((~c3_1(X99) | ~c0_1(X99) | c2_1(X99)) | ~ndr1_0) | ! [X100] : ((~c3_1(X100) | c1_1(X100) | c0_1(X100)) | ~ndr1_0)) & (hskp3 | ! [X101] : ((~c2_1(X101) | ~c0_1(X101) | c1_1(X101)) | ~ndr1_0) | ! [X102] : ((~c3_1(X102) | c1_1(X102) | c0_1(X102)) | ~ndr1_0)) & (hskp1 | ! [X103] : ((~c1_1(X103) | ~c0_1(X103) | c3_1(X103)) | ~ndr1_0) | ! [X104] : ((~c2_1(X104) | c1_1(X104) | c0_1(X104)) | ~ndr1_0)) & (hskp2 | hskp28 | ! [X105] : ((c3_1(X105) | c1_1(X105) | c0_1(X105)) | ~ndr1_0)) & (! [X106] : ((~c3_1(X106) | c2_1(X106) | c1_1(X106)) | ~ndr1_0) | ! [X107] : ((~c3_1(X107) | c2_1(X107) | c0_1(X107)) | ~ndr1_0) | ! [X108] : ((c3_1(X108) | c1_1(X108) | c0_1(X108)) | ~ndr1_0)) & (hskp1 | hskp0 | ! [X109] : ((c2_1(X109) | c1_1(X109) | c0_1(X109)) | ~ndr1_0)) & (! [X110] : ((~c3_1(X110) | ~c2_1(X110) | ~c0_1(X110)) | ~ndr1_0) | ! [X111] : ((~c2_1(X111) | c3_1(X111) | c1_1(X111)) | ~ndr1_0) | ! [X112] : ((c2_1(X112) | c1_1(X112) | c0_1(X112)) | ~ndr1_0)) & ((c3_1(a957) & c2_1(a957) & c1_1(a957) & ndr1_0) | ~hskp31) & ((c2_1(a919) & c1_1(a919) & c0_1(a919) & ndr1_0) | ~hskp30) & ((c3_1(a911) & c1_1(a911) & c0_1(a911) & ndr1_0) | ~hskp29) & ((c3_1(a900) & c2_1(a900) & c0_1(a900) & ndr1_0) | ~hskp28) & ((~c0_1(a978) & c3_1(a978) & c2_1(a978) & ndr1_0) | ~hskp27) & ((~c2_1(a969) & ~c1_1(a969) & c0_1(a969) & ndr1_0) | ~hskp26) & ((~c3_1(a958) & ~c1_1(a958) & ~c0_1(a958) & ndr1_0) | ~hskp25) & ((~c2_1(a953) & c3_1(a953) & c1_1(a953) & ndr1_0) | ~hskp24) & ((~c2_1(a950) & c3_1(a950) & c0_1(a950) & ndr1_0) | ~hskp23) & ((~c3_1(a939) & ~c0_1(a939) & c1_1(a939) & ndr1_0) | ~hskp22) & ((~c3_1(a938) & ~c2_1(a938) & c0_1(a938) & ndr1_0) | ~hskp21) & ((~c3_1(a937) & ~c0_1(a937) & c2_1(a937) & ndr1_0) | ~hskp20) & ((~c1_1(a930) & ~c0_1(a930) & c2_1(a930) & ndr1_0) | ~hskp19) & ((~c1_1(a929) & c2_1(a929) & c0_1(a929) & ndr1_0) | ~hskp18) & ((~c3_1(a928) & ~c2_1(a928) & c1_1(a928) & ndr1_0) | ~hskp17) & ((~c3_1(a923) & c2_1(a923) & c1_1(a923) & ndr1_0) | ~hskp16) & ((~c3_1(a921) & c1_1(a921) & c0_1(a921) & ndr1_0) | ~hskp15) & ((~c1_1(a918) & c3_1(a918) & c0_1(a918) & ndr1_0) | ~hskp14) & ((~c2_1(a914) & ~c1_1(a914) & c3_1(a914) & ndr1_0) | ~hskp13) & ((~c3_1(a913) & ~c1_1(a913) & c0_1(a913) & ndr1_0) | ~hskp12) & ((~c3_1(a912) & ~c2_1(a912) & ~c1_1(a912) & ndr1_0) | ~hskp11) & ((~c2_1(a910) & ~c0_1(a910) & c1_1(a910) & ndr1_0) | ~hskp10) & ((~c2_1(a909) & ~c1_1(a909) & ~c0_1(a909) & ndr1_0) | ~hskp9) & ((~c2_1(a908) & ~c0_1(a908) & c3_1(a908) & ndr1_0) | ~hskp8) & ((~c1_1(a907) & ~c0_1(a907) & c3_1(a907) & ndr1_0) | ~hskp7) & ((~c1_1(a906) & c3_1(a906) & c2_1(a906) & ndr1_0) | ~hskp6) & ((~c0_1(a905) & c3_1(a905) & c1_1(a905) & ndr1_0) | ~hskp5) & ((~c3_1(a904) & ~c1_1(a904) & c2_1(a904) & ndr1_0) | ~hskp4) & ((~c3_1(a903) & c2_1(a903) & c0_1(a903) & ndr1_0) | ~hskp3) & ((~c3_1(a901) & ~c2_1(a901) & ~c0_1(a901) & ndr1_0) | ~hskp2) & ((~c0_1(a899) & c2_1(a899) & c1_1(a899) & ndr1_0) | ~hskp1) & ((~c2_1(a898) & c1_1(a898) & c0_1(a898) & ndr1_0) | ~hskp0)),
  inference(ennf_transformation,[],[f4])).
fof(f4,plain,(
  (hskp9 | hskp7 | hskp27) & (hskp22 | hskp26 | hskp23) & (hskp25 | hskp24 | hskp3) & (hskp17 | hskp22 | hskp28) & (hskp19 | hskp22 | hskp15) & (hskp25 | hskp21 | hskp0) & (hskp9 | hskp12 | hskp0) & (hskp27 | hskp31 | hskp29) & (hskp2 | hskp13 | ! [X0] : (ndr1_0 => (~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0)))) & (hskp10 | hskp16 | ! [X1] : (ndr1_0 => (~c2_1(X1) | ~c0_1(X1) | c3_1(X1)))) & (hskp9 | ! [X2] : (ndr1_0 => (~c2_1(X2) | ~c1_1(X2) | ~c0_1(X2))) | ! [X3] : (ndr1_0 => (~c2_1(X3) | ~c0_1(X3) | c3_1(X3)))) & (hskp17 | hskp26 | ! [X4] : (ndr1_0 => (~c1_1(X4) | ~c0_1(X4) | c3_1(X4)))) & (hskp21 | ! [X5] : (ndr1_0 => (~c3_1(X5) | ~c2_1(X5) | ~c0_1(X5))) | ! [X6] : (ndr1_0 => (~c1_1(X6) | ~c0_1(X6) | c3_1(X6)))) & (hskp13 | hskp1 | ! [X7] : (ndr1_0 => (~c3_1(X7) | ~c1_1(X7) | c2_1(X7)))) & (hskp6 | hskp12 | ! [X8] : (ndr1_0 => (~c3_1(X8) | ~c1_1(X8) | c2_1(X8)))) & (hskp5 | hskp1 | ! [X9] : (ndr1_0 => (~c3_1(X9) | ~c0_1(X9) | c2_1(X9)))) & (hskp0 | ! [X10] : (ndr1_0 => (~c3_1(X10) | ~c2_1(X10) | ~c0_1(X10))) | ! [X11] : (ndr1_0 => (~c3_1(X11) | ~c0_1(X11) | c2_1(X11)))) & (hskp20 | hskp1 | ! [X12] : (ndr1_0 => (~c1_1(X12) | ~c0_1(X12) | c2_1(X12)))) & (hskp25 | ! [X13] : (ndr1_0 => (~c1_1(X13) | c3_1(X13) | c2_1(X13)))) & (hskp31 | ! [X14] : (ndr1_0 => (~c1_1(X14) | ~c0_1(X14) | c2_1(X14))) | ! [X15] : (ndr1_0 => (~c1_1(X15) | c3_1(X15) | c2_1(X15)))) & (hskp7 | hskp10 | ! [X16] : (ndr1_0 => (~c0_1(X16) | c3_1(X16) | c2_1(X16)))) & (hskp17 | hskp24 | ! [X17] : (ndr1_0 => (~c3_1(X17) | ~c2_1(X17) | c1_1(X17)))) & (hskp28 | ! [X18] : (ndr1_0 => (~c3_1(X18) | ~c1_1(X18) | ~c0_1(X18))) | ! [X19] : (ndr1_0 => (~c3_1(X19) | ~c2_1(X19) | c1_1(X19)))) & (hskp18 | ! [X20] : (ndr1_0 => (~c0_1(X20) | c3_1(X20) | c2_1(X20))) | ! [X21] : (ndr1_0 => (~c3_1(X21) | ~c0_1(X21) | c1_1(X21)))) & (hskp23 | ! [X22] : (ndr1_0 => (~c3_1(X22) | ~c2_1(X22) | ~c0_1(X22))) | ! [X23] : (ndr1_0 => (~c2_1(X23) | ~c0_1(X23) | c1_1(X23)))) & (hskp7 | ! [X24] : (ndr1_0 => (~c1_1(X24) | ~c0_1(X24) | c2_1(X24))) | ! [X25] : (ndr1_0 => (~c2_1(X25) | ~c0_1(X25) | c1_1(X25)))) & (! [X26] : (ndr1_0 => (~c3_1(X26) | ~c1_1(X26) | c2_1(X26))) | ! [X27] : (ndr1_0 => (~c0_1(X27) | c3_1(X27) | c2_1(X27))) | ! [X28] : (ndr1_0 => (~c2_1(X28) | ~c0_1(X28) | c1_1(X28)))) & (hskp15 | hskp29 | ! [X29] : (ndr1_0 => (~c2_1(X29) | c3_1(X29) | c1_1(X29)))) & (hskp21 | ! [X30] : (ndr1_0 => (~c1_1(X30) | ~c0_1(X30) | c2_1(X30))) | ! [X31] : (ndr1_0 => (~c2_1(X31) | c3_1(X31) | c1_1(X31)))) & (hskp12 | hskp18 | ! [X32] : (ndr1_0 => (~c0_1(X32) | c3_1(X32) | c1_1(X32)))) & (hskp18 | ! [X33] : (ndr1_0 => (~c1_1(X33) | c3_1(X33) | c2_1(X33))) | ! [X34] : (ndr1_0 => (~c3_1(X34) | c2_1(X34) | c1_1(X34)))) & (hskp12 | ! [X35] : (ndr1_0 => (~c2_1(X35) | c3_1(X35) | c1_1(X35))) | ! [X36] : (ndr1_0 => (~c3_1(X36) | c2_1(X36) | c1_1(X36)))) & (hskp1 | ! [X37] : (ndr1_0 => (~c3_1(X37) | ~c2_1(X37) | ~c1_1(X37))) | ! [X38] : (ndr1_0 => (~c0_1(X38) | c2_1(X38) | c1_1(X38)))) & (hskp22 | ! [X39] : (ndr1_0 => (~c3_1(X39) | ~c0_1(X39) | c2_1(X39))) | ! [X40] : (ndr1_0 => (~c0_1(X40) | c2_1(X40) | c1_1(X40)))) & (hskp22 | ! [X41] : (ndr1_0 => (~c3_1(X41) | ~c2_1(X41) | c1_1(X41))) | ! [X42] : (ndr1_0 => (~c0_1(X42) | c2_1(X42) | c1_1(X42)))) & (hskp21 | ! [X43] : (ndr1_0 => (~c2_1(X43) | ~c1_1(X43) | ~c0_1(X43))) | ! [X44] : (ndr1_0 => (c3_1(X44) | c2_1(X44) | c1_1(X44)))) & (hskp20 | ! [X45] : (ndr1_0 => (~c2_1(X45) | ~c0_1(X45) | c3_1(X45))) | ! [X46] : (ndr1_0 => (c3_1(X46) | c2_1(X46) | c1_1(X46)))) & (! [X47] : (ndr1_0 => (~c1_1(X47) | c3_1(X47) | c2_1(X47))) | ! [X48] : (ndr1_0 => (~c3_1(X48) | c2_1(X48) | c1_1(X48))) | ! [X49] : (ndr1_0 => (~c3_1(X49) | ~c2_1(X49) | c0_1(X49)))) & (hskp10 | hskp12 | ! [X50] : (ndr1_0 => (~c3_1(X50) | ~c1_1(X50) | c0_1(X50)))) & (hskp14 | ! [X51] : (ndr1_0 => (~c2_1(X51) | ~c0_1(X51) | c3_1(X51))) | ! [X52] : (ndr1_0 => (~c3_1(X52) | ~c1_1(X52) | c0_1(X52)))) & (hskp9 | ! [X53] : (ndr1_0 => (~c3_1(X53) | ~c1_1(X53) | ~c0_1(X53))) | ! [X54] : (ndr1_0 => (~c2_1(X54) | ~c1_1(X54) | c0_1(X54)))) & (hskp15 | ! [X55] : (ndr1_0 => (~c2_1(X55) | ~c0_1(X55) | c3_1(X55))) | ! [X56] : (ndr1_0 => (~c2_1(X56) | ~c1_1(X56) | c0_1(X56)))) & (hskp14 | ! [X57] : (ndr1_0 => (~c0_1(X57) | c3_1(X57) | c2_1(X57))) | ! [X58] : (ndr1_0 => (~c2_1(X58) | ~c1_1(X58) | c0_1(X58)))) & (hskp19 | hskp18 | ! [X59] : (ndr1_0 => (~c2_1(X59) | c3_1(X59) | c0_1(X59)))) & (hskp17 | hskp0 | ! [X60] : (ndr1_0 => (~c2_1(X60) | c3_1(X60) | c0_1(X60)))) & (hskp7 | hskp4 | ! [X61] : (ndr1_0 => (~c1_1(X61) | c3_1(X61) | c0_1(X61)))) & (! [X62] : (ndr1_0 => (~c2_1(X62) | c3_1(X62) | c1_1(X62))) | ! [X63] : (ndr1_0 => (~c3_1(X63) | ~c2_1(X63) | c0_1(X63))) | ! [X64] : (ndr1_0 => (~c1_1(X64) | c3_1(X64) | c0_1(X64)))) & (hskp9 | hskp16 | ! [X65] : (ndr1_0 => (~c3_1(X65) | c2_1(X65) | c0_1(X65)))) & (hskp5 | hskp15 | ! [X66] : (ndr1_0 => (~c3_1(X66) | c2_1(X66) | c0_1(X66)))) & (hskp1 | hskp30 | ! [X67] : (ndr1_0 => (~c3_1(X67) | c2_1(X67) | c0_1(X67)))) & (hskp14 | ! [X68] : (ndr1_0 => (~c2_1(X68) | ~c1_1(X68) | c3_1(X68))) | ! [X69] : (ndr1_0 => (~c3_1(X69) | c2_1(X69) | c0_1(X69)))) & (hskp12 | ! [X70] : (ndr1_0 => (~c2_1(X70) | ~c0_1(X70) | c3_1(X70))) | ! [X71] : (ndr1_0 => (~c3_1(X71) | c2_1(X71) | c0_1(X71)))) & (hskp4 | ! [X72] : (ndr1_0 => (~c1_1(X72) | ~c0_1(X72) | c3_1(X72))) | ! [X73] : (ndr1_0 => (~c3_1(X73) | c2_1(X73) | c0_1(X73)))) & (hskp8 | ! [X74] : (ndr1_0 => (~c1_1(X74) | c3_1(X74) | c2_1(X74))) | ! [X75] : (ndr1_0 => (~c3_1(X75) | c2_1(X75) | c0_1(X75)))) & (! [X76] : (ndr1_0 => (~c1_1(X76) | ~c0_1(X76) | c2_1(X76))) | ! [X77] : (ndr1_0 => (~c3_1(X77) | ~c0_1(X77) | c1_1(X77))) | ! [X78] : (ndr1_0 => (~c3_1(X78) | c2_1(X78) | c0_1(X78)))) & (! [X79] : (ndr1_0 => (~c3_1(X79) | ~c1_1(X79) | c2_1(X79))) | ! [X80] : (ndr1_0 => (~c2_1(X80) | c3_1(X80) | c1_1(X80))) | ! [X81] : (ndr1_0 => (~c3_1(X81) | c2_1(X81) | c0_1(X81)))) & (hskp13 | hskp12 | ! [X82] : (ndr1_0 => (~c1_1(X82) | c2_1(X82) | c0_1(X82)))) & (hskp11 | ! [X83] : (ndr1_0 => (~c3_1(X83) | ~c1_1(X83) | ~c0_1(X83))) | ! [X84] : (ndr1_0 => (~c1_1(X84) | c2_1(X84) | c0_1(X84)))) & (hskp29 | ! [X85] : (ndr1_0 => (~c1_1(X85) | ~c0_1(X85) | c3_1(X85))) | ! [X86] : (ndr1_0 => (~c1_1(X86) | c2_1(X86) | c0_1(X86)))) & (hskp10 | ! [X87] : (ndr1_0 => (c3_1(X87) | c2_1(X87) | c1_1(X87))) | ! [X88] : (ndr1_0 => (~c1_1(X88) | c2_1(X88) | c0_1(X88)))) & (! [X89] : (ndr1_0 => (c3_1(X89) | c2_1(X89) | c1_1(X89))) | ! [X90] : (ndr1_0 => (~c3_1(X90) | ~c2_1(X90) | c0_1(X90))) | ! [X91] : (ndr1_0 => (~c1_1(X91) | c2_1(X91) | c0_1(X91)))) & (! [X92] : (ndr1_0 => (~c2_1(X92) | ~c0_1(X92) | c3_1(X92))) | ! [X93] : (ndr1_0 => (~c3_1(X93) | ~c2_1(X93) | c1_1(X93))) | ! [X94] : (ndr1_0 => (c3_1(X94) | c2_1(X94) | c0_1(X94)))) & (hskp9 | ! [X95] : (ndr1_0 => (c3_1(X95) | c2_1(X95) | c1_1(X95))) | ! [X96] : (ndr1_0 => (c3_1(X96) | c2_1(X96) | c0_1(X96)))) & (hskp8 | hskp7 | ! [X97] : (ndr1_0 => (~c3_1(X97) | c1_1(X97) | c0_1(X97)))) & (hskp6 | hskp5 | ! [X98] : (ndr1_0 => (~c3_1(X98) | c1_1(X98) | c0_1(X98)))) & (hskp4 | ! [X99] : (ndr1_0 => (~c3_1(X99) | ~c0_1(X99) | c2_1(X99))) | ! [X100] : (ndr1_0 => (~c3_1(X100) | c1_1(X100) | c0_1(X100)))) & (hskp3 | ! [X101] : (ndr1_0 => (~c2_1(X101) | ~c0_1(X101) | c1_1(X101))) | ! [X102] : (ndr1_0 => (~c3_1(X102) | c1_1(X102) | c0_1(X102)))) & (hskp1 | ! [X103] : (ndr1_0 => (~c1_1(X103) | ~c0_1(X103) | c3_1(X103))) | ! [X104] : (ndr1_0 => (~c2_1(X104) | c1_1(X104) | c0_1(X104)))) & (hskp2 | hskp28 | ! [X105] : (ndr1_0 => (c3_1(X105) | c1_1(X105) | c0_1(X105)))) & (! [X106] : (ndr1_0 => (~c3_1(X106) | c2_1(X106) | c1_1(X106))) | ! [X107] : (ndr1_0 => (~c3_1(X107) | c2_1(X107) | c0_1(X107))) | ! [X108] : (ndr1_0 => (c3_1(X108) | c1_1(X108) | c0_1(X108)))) & (hskp1 | hskp0 | ! [X109] : (ndr1_0 => (c2_1(X109) | c1_1(X109) | c0_1(X109)))) & (! [X110] : (ndr1_0 => (~c3_1(X110) | ~c2_1(X110) | ~c0_1(X110))) | ! [X111] : (ndr1_0 => (~c2_1(X111) | c3_1(X111) | c1_1(X111))) | ! [X112] : (ndr1_0 => (c2_1(X112) | c1_1(X112) | c0_1(X112)))) & ((c3_1(a957) & c2_1(a957) & c1_1(a957) & ndr1_0) | ~hskp31) & ((c2_1(a919) & c1_1(a919) & c0_1(a919) & ndr1_0) | ~hskp30) & ((c3_1(a911) & c1_1(a911) & c0_1(a911) & ndr1_0) | ~hskp29) & ((c3_1(a900) & c2_1(a900) & c0_1(a900) & ndr1_0) | ~hskp28) & ((~c0_1(a978) & c3_1(a978) & c2_1(a978) & ndr1_0) | ~hskp27) & ((~c2_1(a969) & ~c1_1(a969) & c0_1(a969) & ndr1_0) | ~hskp26) & ((~c3_1(a958) & ~c1_1(a958) & ~c0_1(a958) & ndr1_0) | ~hskp25) & ((~c2_1(a953) & c3_1(a953) & c1_1(a953) & ndr1_0) | ~hskp24) & ((~c2_1(a950) & c3_1(a950) & c0_1(a950) & ndr1_0) | ~hskp23) & ((~c3_1(a939) & ~c0_1(a939) & c1_1(a939) & ndr1_0) | ~hskp22) & ((~c3_1(a938) & ~c2_1(a938) & c0_1(a938) & ndr1_0) | ~hskp21) & ((~c3_1(a937) & ~c0_1(a937) & c2_1(a937) & ndr1_0) | ~hskp20) & ((~c1_1(a930) & ~c0_1(a930) & c2_1(a930) & ndr1_0) | ~hskp19) & ((~c1_1(a929) & c2_1(a929) & c0_1(a929) & ndr1_0) | ~hskp18) & ((~c3_1(a928) & ~c2_1(a928) & c1_1(a928) & ndr1_0) | ~hskp17) & ((~c3_1(a923) & c2_1(a923) & c1_1(a923) & ndr1_0) | ~hskp16) & ((~c3_1(a921) & c1_1(a921) & c0_1(a921) & ndr1_0) | ~hskp15) & ((~c1_1(a918) & c3_1(a918) & c0_1(a918) & ndr1_0) | ~hskp14) & ((~c2_1(a914) & ~c1_1(a914) & c3_1(a914) & ndr1_0) | ~hskp13) & ((~c3_1(a913) & ~c1_1(a913) & c0_1(a913) & ndr1_0) | ~hskp12) & ((~c3_1(a912) & ~c2_1(a912) & ~c1_1(a912) & ndr1_0) | ~hskp11) & ((~c2_1(a910) & ~c0_1(a910) & c1_1(a910) & ndr1_0) | ~hskp10) & ((~c2_1(a909) & ~c1_1(a909) & ~c0_1(a909) & ndr1_0) | ~hskp9) & ((~c2_1(a908) & ~c0_1(a908) & c3_1(a908) & ndr1_0) | ~hskp8) & ((~c1_1(a907) & ~c0_1(a907) & c3_1(a907) & ndr1_0) | ~hskp7) & ((~c1_1(a906) & c3_1(a906) & c2_1(a906) & ndr1_0) | ~hskp6) & ((~c0_1(a905) & c3_1(a905) & c1_1(a905) & ndr1_0) | ~hskp5) & ((~c3_1(a904) & ~c1_1(a904) & c2_1(a904) & ndr1_0) | ~hskp4) & ((~c3_1(a903) & c2_1(a903) & c0_1(a903) & ndr1_0) | ~hskp3) & ((~c3_1(a901) & ~c2_1(a901) & ~c0_1(a901) & ndr1_0) | ~hskp2) & ((~c0_1(a899) & c2_1(a899) & c1_1(a899) & ndr1_0) | ~hskp1) & ((~c2_1(a898) & c1_1(a898) & c0_1(a898) & ndr1_0) | ~hskp0)),
  inference(flattening,[],[f3])).
fof(f3,plain,(
  ~~((hskp9 | hskp7 | hskp27) & (hskp22 | hskp26 | hskp23) & (hskp25 | hskp24 | hskp3) & (hskp17 | hskp22 | hskp28) & (hskp19 | hskp22 | hskp15) & (hskp25 | hskp21 | hskp0) & (hskp9 | hskp12 | hskp0) & (hskp27 | hskp31 | hskp29) & (hskp2 | hskp13 | ! [X0] : (ndr1_0 => (~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0)))) & (hskp10 | hskp16 | ! [X1] : (ndr1_0 => (~c2_1(X1) | ~c0_1(X1) | c3_1(X1)))) & (hskp9 | ! [X2] : (ndr1_0 => (~c2_1(X2) | ~c1_1(X2) | ~c0_1(X2))) | ! [X3] : (ndr1_0 => (~c2_1(X3) | ~c0_1(X3) | c3_1(X3)))) & (hskp17 | hskp26 | ! [X4] : (ndr1_0 => (~c1_1(X4) | ~c0_1(X4) | c3_1(X4)))) & (hskp21 | ! [X5] : (ndr1_0 => (~c3_1(X5) | ~c2_1(X5) | ~c0_1(X5))) | ! [X6] : (ndr1_0 => (~c1_1(X6) | ~c0_1(X6) | c3_1(X6)))) & (hskp13 | hskp1 | ! [X7] : (ndr1_0 => (~c3_1(X7) | ~c1_1(X7) | c2_1(X7)))) & (hskp6 | hskp12 | ! [X8] : (ndr1_0 => (~c3_1(X8) | ~c1_1(X8) | c2_1(X8)))) & (hskp5 | hskp1 | ! [X9] : (ndr1_0 => (~c3_1(X9) | ~c0_1(X9) | c2_1(X9)))) & (hskp0 | ! [X10] : (ndr1_0 => (~c3_1(X10) | ~c2_1(X10) | ~c0_1(X10))) | ! [X11] : (ndr1_0 => (~c3_1(X11) | ~c0_1(X11) | c2_1(X11)))) & (hskp20 | hskp1 | ! [X12] : (ndr1_0 => (~c1_1(X12) | ~c0_1(X12) | c2_1(X12)))) & (hskp25 | ! [X13] : (ndr1_0 => (~c1_1(X13) | c3_1(X13) | c2_1(X13)))) & (hskp31 | ! [X14] : (ndr1_0 => (~c1_1(X14) | ~c0_1(X14) | c2_1(X14))) | ! [X15] : (ndr1_0 => (~c1_1(X15) | c3_1(X15) | c2_1(X15)))) & (hskp7 | hskp10 | ! [X16] : (ndr1_0 => (~c0_1(X16) | c3_1(X16) | c2_1(X16)))) & (hskp17 | hskp24 | ! [X17] : (ndr1_0 => (~c3_1(X17) | ~c2_1(X17) | c1_1(X17)))) & (hskp28 | ! [X18] : (ndr1_0 => (~c3_1(X18) | ~c1_1(X18) | ~c0_1(X18))) | ! [X19] : (ndr1_0 => (~c3_1(X19) | ~c2_1(X19) | c1_1(X19)))) & (hskp18 | ! [X20] : (ndr1_0 => (~c0_1(X20) | c3_1(X20) | c2_1(X20))) | ! [X21] : (ndr1_0 => (~c3_1(X21) | ~c0_1(X21) | c1_1(X21)))) & (hskp23 | ! [X22] : (ndr1_0 => (~c3_1(X22) | ~c2_1(X22) | ~c0_1(X22))) | ! [X23] : (ndr1_0 => (~c2_1(X23) | ~c0_1(X23) | c1_1(X23)))) & (hskp7 | ! [X24] : (ndr1_0 => (~c1_1(X24) | ~c0_1(X24) | c2_1(X24))) | ! [X25] : (ndr1_0 => (~c2_1(X25) | ~c0_1(X25) | c1_1(X25)))) & (! [X26] : (ndr1_0 => (~c3_1(X26) | ~c1_1(X26) | c2_1(X26))) | ! [X27] : (ndr1_0 => (~c0_1(X27) | c3_1(X27) | c2_1(X27))) | ! [X28] : (ndr1_0 => (~c2_1(X28) | ~c0_1(X28) | c1_1(X28)))) & (hskp15 | hskp29 | ! [X29] : (ndr1_0 => (~c2_1(X29) | c3_1(X29) | c1_1(X29)))) & (hskp21 | ! [X30] : (ndr1_0 => (~c1_1(X30) | ~c0_1(X30) | c2_1(X30))) | ! [X31] : (ndr1_0 => (~c2_1(X31) | c3_1(X31) | c1_1(X31)))) & (hskp12 | hskp18 | ! [X32] : (ndr1_0 => (~c0_1(X32) | c3_1(X32) | c1_1(X32)))) & (hskp18 | ! [X33] : (ndr1_0 => (~c1_1(X33) | c3_1(X33) | c2_1(X33))) | ! [X34] : (ndr1_0 => (~c3_1(X34) | c2_1(X34) | c1_1(X34)))) & (hskp12 | ! [X35] : (ndr1_0 => (~c2_1(X35) | c3_1(X35) | c1_1(X35))) | ! [X36] : (ndr1_0 => (~c3_1(X36) | c2_1(X36) | c1_1(X36)))) & (hskp1 | ! [X37] : (ndr1_0 => (~c3_1(X37) | ~c2_1(X37) | ~c1_1(X37))) | ! [X38] : (ndr1_0 => (~c0_1(X38) | c2_1(X38) | c1_1(X38)))) & (hskp22 | ! [X39] : (ndr1_0 => (~c3_1(X39) | ~c0_1(X39) | c2_1(X39))) | ! [X40] : (ndr1_0 => (~c0_1(X40) | c2_1(X40) | c1_1(X40)))) & (hskp22 | ! [X41] : (ndr1_0 => (~c3_1(X41) | ~c2_1(X41) | c1_1(X41))) | ! [X42] : (ndr1_0 => (~c0_1(X42) | c2_1(X42) | c1_1(X42)))) & (hskp21 | ! [X43] : (ndr1_0 => (~c2_1(X43) | ~c1_1(X43) | ~c0_1(X43))) | ! [X44] : (ndr1_0 => (c3_1(X44) | c2_1(X44) | c1_1(X44)))) & (hskp20 | ! [X45] : (ndr1_0 => (~c2_1(X45) | ~c0_1(X45) | c3_1(X45))) | ! [X46] : (ndr1_0 => (c3_1(X46) | c2_1(X46) | c1_1(X46)))) & (! [X47] : (ndr1_0 => (~c1_1(X47) | c3_1(X47) | c2_1(X47))) | ! [X48] : (ndr1_0 => (~c3_1(X48) | c2_1(X48) | c1_1(X48))) | ! [X49] : (ndr1_0 => (~c3_1(X49) | ~c2_1(X49) | c0_1(X49)))) & (hskp10 | hskp12 | ! [X50] : (ndr1_0 => (~c3_1(X50) | ~c1_1(X50) | c0_1(X50)))) & (hskp14 | ! [X51] : (ndr1_0 => (~c2_1(X51) | ~c0_1(X51) | c3_1(X51))) | ! [X52] : (ndr1_0 => (~c3_1(X52) | ~c1_1(X52) | c0_1(X52)))) & (hskp9 | ! [X53] : (ndr1_0 => (~c3_1(X53) | ~c1_1(X53) | ~c0_1(X53))) | ! [X54] : (ndr1_0 => (~c2_1(X54) | ~c1_1(X54) | c0_1(X54)))) & (hskp15 | ! [X55] : (ndr1_0 => (~c2_1(X55) | ~c0_1(X55) | c3_1(X55))) | ! [X56] : (ndr1_0 => (~c2_1(X56) | ~c1_1(X56) | c0_1(X56)))) & (hskp14 | ! [X57] : (ndr1_0 => (~c0_1(X57) | c3_1(X57) | c2_1(X57))) | ! [X58] : (ndr1_0 => (~c2_1(X58) | ~c1_1(X58) | c0_1(X58)))) & (hskp19 | hskp18 | ! [X59] : (ndr1_0 => (~c2_1(X59) | c3_1(X59) | c0_1(X59)))) & (hskp17 | hskp0 | ! [X60] : (ndr1_0 => (~c2_1(X60) | c3_1(X60) | c0_1(X60)))) & (hskp7 | hskp4 | ! [X61] : (ndr1_0 => (~c1_1(X61) | c3_1(X61) | c0_1(X61)))) & (! [X62] : (ndr1_0 => (~c2_1(X62) | c3_1(X62) | c1_1(X62))) | ! [X63] : (ndr1_0 => (~c3_1(X63) | ~c2_1(X63) | c0_1(X63))) | ! [X64] : (ndr1_0 => (~c1_1(X64) | c3_1(X64) | c0_1(X64)))) & (hskp9 | hskp16 | ! [X65] : (ndr1_0 => (~c3_1(X65) | c2_1(X65) | c0_1(X65)))) & (hskp5 | hskp15 | ! [X66] : (ndr1_0 => (~c3_1(X66) | c2_1(X66) | c0_1(X66)))) & (hskp1 | hskp30 | ! [X67] : (ndr1_0 => (~c3_1(X67) | c2_1(X67) | c0_1(X67)))) & (hskp14 | ! [X68] : (ndr1_0 => (~c2_1(X68) | ~c1_1(X68) | c3_1(X68))) | ! [X69] : (ndr1_0 => (~c3_1(X69) | c2_1(X69) | c0_1(X69)))) & (hskp12 | ! [X70] : (ndr1_0 => (~c2_1(X70) | ~c0_1(X70) | c3_1(X70))) | ! [X71] : (ndr1_0 => (~c3_1(X71) | c2_1(X71) | c0_1(X71)))) & (hskp4 | ! [X72] : (ndr1_0 => (~c1_1(X72) | ~c0_1(X72) | c3_1(X72))) | ! [X73] : (ndr1_0 => (~c3_1(X73) | c2_1(X73) | c0_1(X73)))) & (hskp8 | ! [X74] : (ndr1_0 => (~c1_1(X74) | c3_1(X74) | c2_1(X74))) | ! [X75] : (ndr1_0 => (~c3_1(X75) | c2_1(X75) | c0_1(X75)))) & (! [X76] : (ndr1_0 => (~c1_1(X76) | ~c0_1(X76) | c2_1(X76))) | ! [X77] : (ndr1_0 => (~c3_1(X77) | ~c0_1(X77) | c1_1(X77))) | ! [X78] : (ndr1_0 => (~c3_1(X78) | c2_1(X78) | c0_1(X78)))) & (! [X79] : (ndr1_0 => (~c3_1(X79) | ~c1_1(X79) | c2_1(X79))) | ! [X80] : (ndr1_0 => (~c2_1(X80) | c3_1(X80) | c1_1(X80))) | ! [X81] : (ndr1_0 => (~c3_1(X81) | c2_1(X81) | c0_1(X81)))) & (hskp13 | hskp12 | ! [X82] : (ndr1_0 => (~c1_1(X82) | c2_1(X82) | c0_1(X82)))) & (hskp11 | ! [X83] : (ndr1_0 => (~c3_1(X83) | ~c1_1(X83) | ~c0_1(X83))) | ! [X84] : (ndr1_0 => (~c1_1(X84) | c2_1(X84) | c0_1(X84)))) & (hskp29 | ! [X85] : (ndr1_0 => (~c1_1(X85) | ~c0_1(X85) | c3_1(X85))) | ! [X86] : (ndr1_0 => (~c1_1(X86) | c2_1(X86) | c0_1(X86)))) & (hskp10 | ! [X87] : (ndr1_0 => (c3_1(X87) | c2_1(X87) | c1_1(X87))) | ! [X88] : (ndr1_0 => (~c1_1(X88) | c2_1(X88) | c0_1(X88)))) & (! [X89] : (ndr1_0 => (c3_1(X89) | c2_1(X89) | c1_1(X89))) | ! [X90] : (ndr1_0 => (~c3_1(X90) | ~c2_1(X90) | c0_1(X90))) | ! [X91] : (ndr1_0 => (~c1_1(X91) | c2_1(X91) | c0_1(X91)))) & (! [X92] : (ndr1_0 => (~c2_1(X92) | ~c0_1(X92) | c3_1(X92))) | ! [X93] : (ndr1_0 => (~c3_1(X93) | ~c2_1(X93) | c1_1(X93))) | ! [X94] : (ndr1_0 => (c3_1(X94) | c2_1(X94) | c0_1(X94)))) & (hskp9 | ! [X95] : (ndr1_0 => (c3_1(X95) | c2_1(X95) | c1_1(X95))) | ! [X96] : (ndr1_0 => (c3_1(X96) | c2_1(X96) | c0_1(X96)))) & (hskp8 | hskp7 | ! [X97] : (ndr1_0 => (~c3_1(X97) | c1_1(X97) | c0_1(X97)))) & (hskp6 | hskp5 | ! [X98] : (ndr1_0 => (~c3_1(X98) | c1_1(X98) | c0_1(X98)))) & (hskp4 | ! [X99] : (ndr1_0 => (~c3_1(X99) | ~c0_1(X99) | c2_1(X99))) | ! [X100] : (ndr1_0 => (~c3_1(X100) | c1_1(X100) | c0_1(X100)))) & (hskp3 | ! [X101] : (ndr1_0 => (~c2_1(X101) | ~c0_1(X101) | c1_1(X101))) | ! [X102] : (ndr1_0 => (~c3_1(X102) | c1_1(X102) | c0_1(X102)))) & (hskp1 | ! [X103] : (ndr1_0 => (~c1_1(X103) | ~c0_1(X103) | c3_1(X103))) | ! [X104] : (ndr1_0 => (~c2_1(X104) | c1_1(X104) | c0_1(X104)))) & (hskp2 | hskp28 | ! [X105] : (ndr1_0 => (c3_1(X105) | c1_1(X105) | c0_1(X105)))) & (! [X106] : (ndr1_0 => (~c3_1(X106) | c2_1(X106) | c1_1(X106))) | ! [X107] : (ndr1_0 => (~c3_1(X107) | c2_1(X107) | c0_1(X107))) | ! [X108] : (ndr1_0 => (c3_1(X108) | c1_1(X108) | c0_1(X108)))) & (hskp1 | hskp0 | ! [X109] : (ndr1_0 => (c2_1(X109) | c1_1(X109) | c0_1(X109)))) & (! [X110] : (ndr1_0 => (~c3_1(X110) | ~c2_1(X110) | ~c0_1(X110))) | ! [X111] : (ndr1_0 => (~c2_1(X111) | c3_1(X111) | c1_1(X111))) | ! [X112] : (ndr1_0 => (c2_1(X112) | c1_1(X112) | c0_1(X112)))) & ((c3_1(a957) & c2_1(a957) & c1_1(a957) & ndr1_0) | ~hskp31) & ((c2_1(a919) & c1_1(a919) & c0_1(a919) & ndr1_0) | ~hskp30) & ((c3_1(a911) & c1_1(a911) & c0_1(a911) & ndr1_0) | ~hskp29) & ((c3_1(a900) & c2_1(a900) & c0_1(a900) & ndr1_0) | ~hskp28) & ((~c0_1(a978) & c3_1(a978) & c2_1(a978) & ndr1_0) | ~hskp27) & ((~c2_1(a969) & ~c1_1(a969) & c0_1(a969) & ndr1_0) | ~hskp26) & ((~c3_1(a958) & ~c1_1(a958) & ~c0_1(a958) & ndr1_0) | ~hskp25) & ((~c2_1(a953) & c3_1(a953) & c1_1(a953) & ndr1_0) | ~hskp24) & ((~c2_1(a950) & c3_1(a950) & c0_1(a950) & ndr1_0) | ~hskp23) & ((~c3_1(a939) & ~c0_1(a939) & c1_1(a939) & ndr1_0) | ~hskp22) & ((~c3_1(a938) & ~c2_1(a938) & c0_1(a938) & ndr1_0) | ~hskp21) & ((~c3_1(a937) & ~c0_1(a937) & c2_1(a937) & ndr1_0) | ~hskp20) & ((~c1_1(a930) & ~c0_1(a930) & c2_1(a930) & ndr1_0) | ~hskp19) & ((~c1_1(a929) & c2_1(a929) & c0_1(a929) & ndr1_0) | ~hskp18) & ((~c3_1(a928) & ~c2_1(a928) & c1_1(a928) & ndr1_0) | ~hskp17) & ((~c3_1(a923) & c2_1(a923) & c1_1(a923) & ndr1_0) | ~hskp16) & ((~c3_1(a921) & c1_1(a921) & c0_1(a921) & ndr1_0) | ~hskp15) & ((~c1_1(a918) & c3_1(a918) & c0_1(a918) & ndr1_0) | ~hskp14) & ((~c2_1(a914) & ~c1_1(a914) & c3_1(a914) & ndr1_0) | ~hskp13) & ((~c3_1(a913) & ~c1_1(a913) & c0_1(a913) & ndr1_0) | ~hskp12) & ((~c3_1(a912) & ~c2_1(a912) & ~c1_1(a912) & ndr1_0) | ~hskp11) & ((~c2_1(a910) & ~c0_1(a910) & c1_1(a910) & ndr1_0) | ~hskp10) & ((~c2_1(a909) & ~c1_1(a909) & ~c0_1(a909) & ndr1_0) | ~hskp9) & ((~c2_1(a908) & ~c0_1(a908) & c3_1(a908) & ndr1_0) | ~hskp8) & ((~c1_1(a907) & ~c0_1(a907) & c3_1(a907) & ndr1_0) | ~hskp7) & ((~c1_1(a906) & c3_1(a906) & c2_1(a906) & ndr1_0) | ~hskp6) & ((~c0_1(a905) & c3_1(a905) & c1_1(a905) & ndr1_0) | ~hskp5) & ((~c3_1(a904) & ~c1_1(a904) & c2_1(a904) & ndr1_0) | ~hskp4) & ((~c3_1(a903) & c2_1(a903) & c0_1(a903) & ndr1_0) | ~hskp3) & ((~c3_1(a901) & ~c2_1(a901) & ~c0_1(a901) & ndr1_0) | ~hskp2) & ((~c0_1(a899) & c2_1(a899) & c1_1(a899) & ndr1_0) | ~hskp1) & ((~c2_1(a898) & c1_1(a898) & c0_1(a898) & ndr1_0) | ~hskp0))),
  inference(rectify,[],[f2])).
fof(f2,negated_conjecture,(
  ~~((hskp9 | hskp7 | hskp27) & (hskp22 | hskp26 | hskp23) & (hskp25 | hskp24 | hskp3) & (hskp17 | hskp22 | hskp28) & (hskp19 | hskp22 | hskp15) & (hskp25 | hskp21 | hskp0) & (hskp9 | hskp12 | hskp0) & (hskp27 | hskp31 | hskp29) & (hskp2 | hskp13 | ! [X112] : (ndr1_0 => (~c3_1(X112) | ~c2_1(X112) | ~c0_1(X112)))) & (hskp10 | hskp16 | ! [X111] : (ndr1_0 => (~c2_1(X111) | ~c0_1(X111) | c3_1(X111)))) & (hskp9 | ! [X110] : (ndr1_0 => (~c2_1(X110) | ~c1_1(X110) | ~c0_1(X110))) | ! [X109] : (ndr1_0 => (~c2_1(X109) | ~c0_1(X109) | c3_1(X109)))) & (hskp17 | hskp26 | ! [X108] : (ndr1_0 => (~c1_1(X108) | ~c0_1(X108) | c3_1(X108)))) & (hskp21 | ! [X107] : (ndr1_0 => (~c3_1(X107) | ~c2_1(X107) | ~c0_1(X107))) | ! [X106] : (ndr1_0 => (~c1_1(X106) | ~c0_1(X106) | c3_1(X106)))) & (hskp13 | hskp1 | ! [X105] : (ndr1_0 => (~c3_1(X105) | ~c1_1(X105) | c2_1(X105)))) & (hskp6 | hskp12 | ! [X104] : (ndr1_0 => (~c3_1(X104) | ~c1_1(X104) | c2_1(X104)))) & (hskp5 | hskp1 | ! [X103] : (ndr1_0 => (~c3_1(X103) | ~c0_1(X103) | c2_1(X103)))) & (hskp0 | ! [X102] : (ndr1_0 => (~c3_1(X102) | ~c2_1(X102) | ~c0_1(X102))) | ! [X101] : (ndr1_0 => (~c3_1(X101) | ~c0_1(X101) | c2_1(X101)))) & (hskp20 | hskp1 | ! [X100] : (ndr1_0 => (~c1_1(X100) | ~c0_1(X100) | c2_1(X100)))) & (hskp25 | ! [X99] : (ndr1_0 => (~c1_1(X99) | c3_1(X99) | c2_1(X99)))) & (hskp31 | ! [X98] : (ndr1_0 => (~c1_1(X98) | ~c0_1(X98) | c2_1(X98))) | ! [X97] : (ndr1_0 => (~c1_1(X97) | c3_1(X97) | c2_1(X97)))) & (hskp7 | hskp10 | ! [X96] : (ndr1_0 => (~c0_1(X96) | c3_1(X96) | c2_1(X96)))) & (hskp17 | hskp24 | ! [X95] : (ndr1_0 => (~c3_1(X95) | ~c2_1(X95) | c1_1(X95)))) & (hskp28 | ! [X94] : (ndr1_0 => (~c3_1(X94) | ~c1_1(X94) | ~c0_1(X94))) | ! [X93] : (ndr1_0 => (~c3_1(X93) | ~c2_1(X93) | c1_1(X93)))) & (hskp18 | ! [X92] : (ndr1_0 => (~c0_1(X92) | c3_1(X92) | c2_1(X92))) | ! [X91] : (ndr1_0 => (~c3_1(X91) | ~c0_1(X91) | c1_1(X91)))) & (hskp23 | ! [X90] : (ndr1_0 => (~c3_1(X90) | ~c2_1(X90) | ~c0_1(X90))) | ! [X89] : (ndr1_0 => (~c2_1(X89) | ~c0_1(X89) | c1_1(X89)))) & (hskp7 | ! [X88] : (ndr1_0 => (~c1_1(X88) | ~c0_1(X88) | c2_1(X88))) | ! [X87] : (ndr1_0 => (~c2_1(X87) | ~c0_1(X87) | c1_1(X87)))) & (! [X86] : (ndr1_0 => (~c3_1(X86) | ~c1_1(X86) | c2_1(X86))) | ! [X85] : (ndr1_0 => (~c0_1(X85) | c3_1(X85) | c2_1(X85))) | ! [X84] : (ndr1_0 => (~c2_1(X84) | ~c0_1(X84) | c1_1(X84)))) & (hskp15 | hskp29 | ! [X83] : (ndr1_0 => (~c2_1(X83) | c3_1(X83) | c1_1(X83)))) & (hskp21 | ! [X82] : (ndr1_0 => (~c1_1(X82) | ~c0_1(X82) | c2_1(X82))) | ! [X81] : (ndr1_0 => (~c2_1(X81) | c3_1(X81) | c1_1(X81)))) & (hskp12 | hskp18 | ! [X80] : (ndr1_0 => (~c0_1(X80) | c3_1(X80) | c1_1(X80)))) & (hskp18 | ! [X79] : (ndr1_0 => (~c1_1(X79) | c3_1(X79) | c2_1(X79))) | ! [X78] : (ndr1_0 => (~c3_1(X78) | c2_1(X78) | c1_1(X78)))) & (hskp12 | ! [X77] : (ndr1_0 => (~c2_1(X77) | c3_1(X77) | c1_1(X77))) | ! [X76] : (ndr1_0 => (~c3_1(X76) | c2_1(X76) | c1_1(X76)))) & (hskp1 | ! [X75] : (ndr1_0 => (~c3_1(X75) | ~c2_1(X75) | ~c1_1(X75))) | ! [X74] : (ndr1_0 => (~c0_1(X74) | c2_1(X74) | c1_1(X74)))) & (hskp22 | ! [X73] : (ndr1_0 => (~c3_1(X73) | ~c0_1(X73) | c2_1(X73))) | ! [X72] : (ndr1_0 => (~c0_1(X72) | c2_1(X72) | c1_1(X72)))) & (hskp22 | ! [X71] : (ndr1_0 => (~c3_1(X71) | ~c2_1(X71) | c1_1(X71))) | ! [X70] : (ndr1_0 => (~c0_1(X70) | c2_1(X70) | c1_1(X70)))) & (hskp21 | ! [X69] : (ndr1_0 => (~c2_1(X69) | ~c1_1(X69) | ~c0_1(X69))) | ! [X68] : (ndr1_0 => (c3_1(X68) | c2_1(X68) | c1_1(X68)))) & (hskp20 | ! [X67] : (ndr1_0 => (~c2_1(X67) | ~c0_1(X67) | c3_1(X67))) | ! [X66] : (ndr1_0 => (c3_1(X66) | c2_1(X66) | c1_1(X66)))) & (! [X65] : (ndr1_0 => (~c1_1(X65) | c3_1(X65) | c2_1(X65))) | ! [X64] : (ndr1_0 => (~c3_1(X64) | c2_1(X64) | c1_1(X64))) | ! [X63] : (ndr1_0 => (~c3_1(X63) | ~c2_1(X63) | c0_1(X63)))) & (hskp10 | hskp12 | ! [X62] : (ndr1_0 => (~c3_1(X62) | ~c1_1(X62) | c0_1(X62)))) & (hskp14 | ! [X61] : (ndr1_0 => (~c2_1(X61) | ~c0_1(X61) | c3_1(X61))) | ! [X60] : (ndr1_0 => (~c3_1(X60) | ~c1_1(X60) | c0_1(X60)))) & (hskp9 | ! [X59] : (ndr1_0 => (~c3_1(X59) | ~c1_1(X59) | ~c0_1(X59))) | ! [X58] : (ndr1_0 => (~c2_1(X58) | ~c1_1(X58) | c0_1(X58)))) & (hskp15 | ! [X57] : (ndr1_0 => (~c2_1(X57) | ~c0_1(X57) | c3_1(X57))) | ! [X56] : (ndr1_0 => (~c2_1(X56) | ~c1_1(X56) | c0_1(X56)))) & (hskp14 | ! [X55] : (ndr1_0 => (~c0_1(X55) | c3_1(X55) | c2_1(X55))) | ! [X54] : (ndr1_0 => (~c2_1(X54) | ~c1_1(X54) | c0_1(X54)))) & (hskp19 | hskp18 | ! [X53] : (ndr1_0 => (~c2_1(X53) | c3_1(X53) | c0_1(X53)))) & (hskp17 | hskp0 | ! [X52] : (ndr1_0 => (~c2_1(X52) | c3_1(X52) | c0_1(X52)))) & (hskp7 | hskp4 | ! [X51] : (ndr1_0 => (~c1_1(X51) | c3_1(X51) | c0_1(X51)))) & (! [X50] : (ndr1_0 => (~c2_1(X50) | c3_1(X50) | c1_1(X50))) | ! [X49] : (ndr1_0 => (~c3_1(X49) | ~c2_1(X49) | c0_1(X49))) | ! [X48] : (ndr1_0 => (~c1_1(X48) | c3_1(X48) | c0_1(X48)))) & (hskp9 | hskp16 | ! [X47] : (ndr1_0 => (~c3_1(X47) | c2_1(X47) | c0_1(X47)))) & (hskp5 | hskp15 | ! [X46] : (ndr1_0 => (~c3_1(X46) | c2_1(X46) | c0_1(X46)))) & (hskp1 | hskp30 | ! [X45] : (ndr1_0 => (~c3_1(X45) | c2_1(X45) | c0_1(X45)))) & (hskp14 | ! [X44] : (ndr1_0 => (~c2_1(X44) | ~c1_1(X44) | c3_1(X44))) | ! [X43] : (ndr1_0 => (~c3_1(X43) | c2_1(X43) | c0_1(X43)))) & (hskp12 | ! [X42] : (ndr1_0 => (~c2_1(X42) | ~c0_1(X42) | c3_1(X42))) | ! [X41] : (ndr1_0 => (~c3_1(X41) | c2_1(X41) | c0_1(X41)))) & (hskp4 | ! [X40] : (ndr1_0 => (~c1_1(X40) | ~c0_1(X40) | c3_1(X40))) | ! [X39] : (ndr1_0 => (~c3_1(X39) | c2_1(X39) | c0_1(X39)))) & (hskp8 | ! [X38] : (ndr1_0 => (~c1_1(X38) | c3_1(X38) | c2_1(X38))) | ! [X37] : (ndr1_0 => (~c3_1(X37) | c2_1(X37) | c0_1(X37)))) & (! [X36] : (ndr1_0 => (~c1_1(X36) | ~c0_1(X36) | c2_1(X36))) | ! [X35] : (ndr1_0 => (~c3_1(X35) | ~c0_1(X35) | c1_1(X35))) | ! [X34] : (ndr1_0 => (~c3_1(X34) | c2_1(X34) | c0_1(X34)))) & (! [X33] : (ndr1_0 => (~c3_1(X33) | ~c1_1(X33) | c2_1(X33))) | ! [X32] : (ndr1_0 => (~c2_1(X32) | c3_1(X32) | c1_1(X32))) | ! [X31] : (ndr1_0 => (~c3_1(X31) | c2_1(X31) | c0_1(X31)))) & (hskp13 | hskp12 | ! [X30] : (ndr1_0 => (~c1_1(X30) | c2_1(X30) | c0_1(X30)))) & (hskp11 | ! [X29] : (ndr1_0 => (~c3_1(X29) | ~c1_1(X29) | ~c0_1(X29))) | ! [X28] : (ndr1_0 => (~c1_1(X28) | c2_1(X28) | c0_1(X28)))) & (hskp29 | ! [X27] : (ndr1_0 => (~c1_1(X27) | ~c0_1(X27) | c3_1(X27))) | ! [X26] : (ndr1_0 => (~c1_1(X26) | c2_1(X26) | c0_1(X26)))) & (hskp10 | ! [X25] : (ndr1_0 => (c3_1(X25) | c2_1(X25) | c1_1(X25))) | ! [X24] : (ndr1_0 => (~c1_1(X24) | c2_1(X24) | c0_1(X24)))) & (! [X23] : (ndr1_0 => (c3_1(X23) | c2_1(X23) | c1_1(X23))) | ! [X22] : (ndr1_0 => (~c3_1(X22) | ~c2_1(X22) | c0_1(X22))) | ! [X21] : (ndr1_0 => (~c1_1(X21) | c2_1(X21) | c0_1(X21)))) & (! [X20] : (ndr1_0 => (~c2_1(X20) | ~c0_1(X20) | c3_1(X20))) | ! [X19] : (ndr1_0 => (~c3_1(X19) | ~c2_1(X19) | c1_1(X19))) | ! [X18] : (ndr1_0 => (c3_1(X18) | c2_1(X18) | c0_1(X18)))) & (hskp9 | ! [X17] : (ndr1_0 => (c3_1(X17) | c2_1(X17) | c1_1(X17))) | ! [X16] : (ndr1_0 => (c3_1(X16) | c2_1(X16) | c0_1(X16)))) & (hskp8 | hskp7 | ! [X15] : (ndr1_0 => (~c3_1(X15) | c1_1(X15) | c0_1(X15)))) & (hskp6 | hskp5 | ! [X14] : (ndr1_0 => (~c3_1(X14) | c1_1(X14) | c0_1(X14)))) & (hskp4 | ! [X13] : (ndr1_0 => (~c3_1(X13) | ~c0_1(X13) | c2_1(X13))) | ! [X12] : (ndr1_0 => (~c3_1(X12) | c1_1(X12) | c0_1(X12)))) & (hskp3 | ! [X11] : (ndr1_0 => (~c2_1(X11) | ~c0_1(X11) | c1_1(X11))) | ! [X10] : (ndr1_0 => (~c3_1(X10) | c1_1(X10) | c0_1(X10)))) & (hskp1 | ! [X9] : (ndr1_0 => (~c1_1(X9) | ~c0_1(X9) | c3_1(X9))) | ! [X8] : (ndr1_0 => (~c2_1(X8) | c1_1(X8) | c0_1(X8)))) & (hskp2 | hskp28 | ! [X7] : (ndr1_0 => (c3_1(X7) | c1_1(X7) | c0_1(X7)))) & (! [X6] : (ndr1_0 => (~c3_1(X6) | c2_1(X6) | c1_1(X6))) | ! [X5] : (ndr1_0 => (~c3_1(X5) | c2_1(X5) | c0_1(X5))) | ! [X4] : (ndr1_0 => (c3_1(X4) | c1_1(X4) | c0_1(X4)))) & (hskp1 | hskp0 | ! [X3] : (ndr1_0 => (c2_1(X3) | c1_1(X3) | c0_1(X3)))) & (! [X2] : (ndr1_0 => (~c3_1(X2) | ~c2_1(X2) | ~c0_1(X2))) | ! [X1] : (ndr1_0 => (~c2_1(X1) | c3_1(X1) | c1_1(X1))) | ! [X0] : (ndr1_0 => (c2_1(X0) | c1_1(X0) | c0_1(X0)))) & ((c3_1(a957) & c2_1(a957) & c1_1(a957) & ndr1_0) | ~hskp31) & ((c2_1(a919) & c1_1(a919) & c0_1(a919) & ndr1_0) | ~hskp30) & ((c3_1(a911) & c1_1(a911) & c0_1(a911) & ndr1_0) | ~hskp29) & ((c3_1(a900) & c2_1(a900) & c0_1(a900) & ndr1_0) | ~hskp28) & ((~c0_1(a978) & c3_1(a978) & c2_1(a978) & ndr1_0) | ~hskp27) & ((~c2_1(a969) & ~c1_1(a969) & c0_1(a969) & ndr1_0) | ~hskp26) & ((~c3_1(a958) & ~c1_1(a958) & ~c0_1(a958) & ndr1_0) | ~hskp25) & ((~c2_1(a953) & c3_1(a953) & c1_1(a953) & ndr1_0) | ~hskp24) & ((~c2_1(a950) & c3_1(a950) & c0_1(a950) & ndr1_0) | ~hskp23) & ((~c3_1(a939) & ~c0_1(a939) & c1_1(a939) & ndr1_0) | ~hskp22) & ((~c3_1(a938) & ~c2_1(a938) & c0_1(a938) & ndr1_0) | ~hskp21) & ((~c3_1(a937) & ~c0_1(a937) & c2_1(a937) & ndr1_0) | ~hskp20) & ((~c1_1(a930) & ~c0_1(a930) & c2_1(a930) & ndr1_0) | ~hskp19) & ((~c1_1(a929) & c2_1(a929) & c0_1(a929) & ndr1_0) | ~hskp18) & ((~c3_1(a928) & ~c2_1(a928) & c1_1(a928) & ndr1_0) | ~hskp17) & ((~c3_1(a923) & c2_1(a923) & c1_1(a923) & ndr1_0) | ~hskp16) & ((~c3_1(a921) & c1_1(a921) & c0_1(a921) & ndr1_0) | ~hskp15) & ((~c1_1(a918) & c3_1(a918) & c0_1(a918) & ndr1_0) | ~hskp14) & ((~c2_1(a914) & ~c1_1(a914) & c3_1(a914) & ndr1_0) | ~hskp13) & ((~c3_1(a913) & ~c1_1(a913) & c0_1(a913) & ndr1_0) | ~hskp12) & ((~c3_1(a912) & ~c2_1(a912) & ~c1_1(a912) & ndr1_0) | ~hskp11) & ((~c2_1(a910) & ~c0_1(a910) & c1_1(a910) & ndr1_0) | ~hskp10) & ((~c2_1(a909) & ~c1_1(a909) & ~c0_1(a909) & ndr1_0) | ~hskp9) & ((~c2_1(a908) & ~c0_1(a908) & c3_1(a908) & ndr1_0) | ~hskp8) & ((~c1_1(a907) & ~c0_1(a907) & c3_1(a907) & ndr1_0) | ~hskp7) & ((~c1_1(a906) & c3_1(a906) & c2_1(a906) & ndr1_0) | ~hskp6) & ((~c0_1(a905) & c3_1(a905) & c1_1(a905) & ndr1_0) | ~hskp5) & ((~c3_1(a904) & ~c1_1(a904) & c2_1(a904) & ndr1_0) | ~hskp4) & ((~c3_1(a903) & c2_1(a903) & c0_1(a903) & ndr1_0) | ~hskp3) & ((~c3_1(a901) & ~c2_1(a901) & ~c0_1(a901) & ndr1_0) | ~hskp2) & ((~c0_1(a899) & c2_1(a899) & c1_1(a899) & ndr1_0) | ~hskp1) & ((~c2_1(a898) & c1_1(a898) & c0_1(a898) & ndr1_0) | ~hskp0))),
  inference(negated_conjecture,[],[f1])).
fof(f1,conjecture,(
  ~((hskp9 | hskp7 | hskp27) & (hskp22 | hskp26 | hskp23) & (hskp25 | hskp24 | hskp3) & (hskp17 | hskp22 | hskp28) & (hskp19 | hskp22 | hskp15) & (hskp25 | hskp21 | hskp0) & (hskp9 | hskp12 | hskp0) & (hskp27 | hskp31 | hskp29) & (hskp2 | hskp13 | ! [X112] : (ndr1_0 => (~c3_1(X112) | ~c2_1(X112) | ~c0_1(X112)))) & (hskp10 | hskp16 | ! [X111] : (ndr1_0 => (~c2_1(X111) | ~c0_1(X111) | c3_1(X111)))) & (hskp9 | ! [X110] : (ndr1_0 => (~c2_1(X110) | ~c1_1(X110) | ~c0_1(X110))) | ! [X109] : (ndr1_0 => (~c2_1(X109) | ~c0_1(X109) | c3_1(X109)))) & (hskp17 | hskp26 | ! [X108] : (ndr1_0 => (~c1_1(X108) | ~c0_1(X108) | c3_1(X108)))) & (hskp21 | ! [X107] : (ndr1_0 => (~c3_1(X107) | ~c2_1(X107) | ~c0_1(X107))) | ! [X106] : (ndr1_0 => (~c1_1(X106) | ~c0_1(X106) | c3_1(X106)))) & (hskp13 | hskp1 | ! [X105] : (ndr1_0 => (~c3_1(X105) | ~c1_1(X105) | c2_1(X105)))) & (hskp6 | hskp12 | ! [X104] : (ndr1_0 => (~c3_1(X104) | ~c1_1(X104) | c2_1(X104)))) & (hskp5 | hskp1 | ! [X103] : (ndr1_0 => (~c3_1(X103) | ~c0_1(X103) | c2_1(X103)))) & (hskp0 | ! [X102] : (ndr1_0 => (~c3_1(X102) | ~c2_1(X102) | ~c0_1(X102))) | ! [X101] : (ndr1_0 => (~c3_1(X101) | ~c0_1(X101) | c2_1(X101)))) & (hskp20 | hskp1 | ! [X100] : (ndr1_0 => (~c1_1(X100) | ~c0_1(X100) | c2_1(X100)))) & (hskp25 | ! [X99] : (ndr1_0 => (~c1_1(X99) | c3_1(X99) | c2_1(X99)))) & (hskp31 | ! [X98] : (ndr1_0 => (~c1_1(X98) | ~c0_1(X98) | c2_1(X98))) | ! [X97] : (ndr1_0 => (~c1_1(X97) | c3_1(X97) | c2_1(X97)))) & (hskp7 | hskp10 | ! [X96] : (ndr1_0 => (~c0_1(X96) | c3_1(X96) | c2_1(X96)))) & (hskp17 | hskp24 | ! [X95] : (ndr1_0 => (~c3_1(X95) | ~c2_1(X95) | c1_1(X95)))) & (hskp28 | ! [X94] : (ndr1_0 => (~c3_1(X94) | ~c1_1(X94) | ~c0_1(X94))) | ! [X93] : (ndr1_0 => (~c3_1(X93) | ~c2_1(X93) | c1_1(X93)))) & (hskp18 | ! [X92] : (ndr1_0 => (~c0_1(X92) | c3_1(X92) | c2_1(X92))) | ! [X91] : (ndr1_0 => (~c3_1(X91) | ~c0_1(X91) | c1_1(X91)))) & (hskp23 | ! [X90] : (ndr1_0 => (~c3_1(X90) | ~c2_1(X90) | ~c0_1(X90))) | ! [X89] : (ndr1_0 => (~c2_1(X89) | ~c0_1(X89) | c1_1(X89)))) & (hskp7 | ! [X88] : (ndr1_0 => (~c1_1(X88) | ~c0_1(X88) | c2_1(X88))) | ! [X87] : (ndr1_0 => (~c2_1(X87) | ~c0_1(X87) | c1_1(X87)))) & (! [X86] : (ndr1_0 => (~c3_1(X86) | ~c1_1(X86) | c2_1(X86))) | ! [X85] : (ndr1_0 => (~c0_1(X85) | c3_1(X85) | c2_1(X85))) | ! [X84] : (ndr1_0 => (~c2_1(X84) | ~c0_1(X84) | c1_1(X84)))) & (hskp15 | hskp29 | ! [X83] : (ndr1_0 => (~c2_1(X83) | c3_1(X83) | c1_1(X83)))) & (hskp21 | ! [X82] : (ndr1_0 => (~c1_1(X82) | ~c0_1(X82) | c2_1(X82))) | ! [X81] : (ndr1_0 => (~c2_1(X81) | c3_1(X81) | c1_1(X81)))) & (hskp12 | hskp18 | ! [X80] : (ndr1_0 => (~c0_1(X80) | c3_1(X80) | c1_1(X80)))) & (hskp18 | ! [X79] : (ndr1_0 => (~c1_1(X79) | c3_1(X79) | c2_1(X79))) | ! [X78] : (ndr1_0 => (~c3_1(X78) | c2_1(X78) | c1_1(X78)))) & (hskp12 | ! [X77] : (ndr1_0 => (~c2_1(X77) | c3_1(X77) | c1_1(X77))) | ! [X76] : (ndr1_0 => (~c3_1(X76) | c2_1(X76) | c1_1(X76)))) & (hskp1 | ! [X75] : (ndr1_0 => (~c3_1(X75) | ~c2_1(X75) | ~c1_1(X75))) | ! [X74] : (ndr1_0 => (~c0_1(X74) | c2_1(X74) | c1_1(X74)))) & (hskp22 | ! [X73] : (ndr1_0 => (~c3_1(X73) | ~c0_1(X73) | c2_1(X73))) | ! [X72] : (ndr1_0 => (~c0_1(X72) | c2_1(X72) | c1_1(X72)))) & (hskp22 | ! [X71] : (ndr1_0 => (~c3_1(X71) | ~c2_1(X71) | c1_1(X71))) | ! [X70] : (ndr1_0 => (~c0_1(X70) | c2_1(X70) | c1_1(X70)))) & (hskp21 | ! [X69] : (ndr1_0 => (~c2_1(X69) | ~c1_1(X69) | ~c0_1(X69))) | ! [X68] : (ndr1_0 => (c3_1(X68) | c2_1(X68) | c1_1(X68)))) & (hskp20 | ! [X67] : (ndr1_0 => (~c2_1(X67) | ~c0_1(X67) | c3_1(X67))) | ! [X66] : (ndr1_0 => (c3_1(X66) | c2_1(X66) | c1_1(X66)))) & (! [X65] : (ndr1_0 => (~c1_1(X65) | c3_1(X65) | c2_1(X65))) | ! [X64] : (ndr1_0 => (~c3_1(X64) | c2_1(X64) | c1_1(X64))) | ! [X63] : (ndr1_0 => (~c3_1(X63) | ~c2_1(X63) | c0_1(X63)))) & (hskp10 | hskp12 | ! [X62] : (ndr1_0 => (~c3_1(X62) | ~c1_1(X62) | c0_1(X62)))) & (hskp14 | ! [X61] : (ndr1_0 => (~c2_1(X61) | ~c0_1(X61) | c3_1(X61))) | ! [X60] : (ndr1_0 => (~c3_1(X60) | ~c1_1(X60) | c0_1(X60)))) & (hskp9 | ! [X59] : (ndr1_0 => (~c3_1(X59) | ~c1_1(X59) | ~c0_1(X59))) | ! [X58] : (ndr1_0 => (~c2_1(X58) | ~c1_1(X58) | c0_1(X58)))) & (hskp15 | ! [X57] : (ndr1_0 => (~c2_1(X57) | ~c0_1(X57) | c3_1(X57))) | ! [X56] : (ndr1_0 => (~c2_1(X56) | ~c1_1(X56) | c0_1(X56)))) & (hskp14 | ! [X55] : (ndr1_0 => (~c0_1(X55) | c3_1(X55) | c2_1(X55))) | ! [X54] : (ndr1_0 => (~c2_1(X54) | ~c1_1(X54) | c0_1(X54)))) & (hskp19 | hskp18 | ! [X53] : (ndr1_0 => (~c2_1(X53) | c3_1(X53) | c0_1(X53)))) & (hskp17 | hskp0 | ! [X52] : (ndr1_0 => (~c2_1(X52) | c3_1(X52) | c0_1(X52)))) & (hskp7 | hskp4 | ! [X51] : (ndr1_0 => (~c1_1(X51) | c3_1(X51) | c0_1(X51)))) & (! [X50] : (ndr1_0 => (~c2_1(X50) | c3_1(X50) | c1_1(X50))) | ! [X49] : (ndr1_0 => (~c3_1(X49) | ~c2_1(X49) | c0_1(X49))) | ! [X48] : (ndr1_0 => (~c1_1(X48) | c3_1(X48) | c0_1(X48)))) & (hskp9 | hskp16 | ! [X47] : (ndr1_0 => (~c3_1(X47) | c2_1(X47) | c0_1(X47)))) & (hskp5 | hskp15 | ! [X46] : (ndr1_0 => (~c3_1(X46) | c2_1(X46) | c0_1(X46)))) & (hskp1 | hskp30 | ! [X45] : (ndr1_0 => (~c3_1(X45) | c2_1(X45) | c0_1(X45)))) & (hskp14 | ! [X44] : (ndr1_0 => (~c2_1(X44) | ~c1_1(X44) | c3_1(X44))) | ! [X43] : (ndr1_0 => (~c3_1(X43) | c2_1(X43) | c0_1(X43)))) & (hskp12 | ! [X42] : (ndr1_0 => (~c2_1(X42) | ~c0_1(X42) | c3_1(X42))) | ! [X41] : (ndr1_0 => (~c3_1(X41) | c2_1(X41) | c0_1(X41)))) & (hskp4 | ! [X40] : (ndr1_0 => (~c1_1(X40) | ~c0_1(X40) | c3_1(X40))) | ! [X39] : (ndr1_0 => (~c3_1(X39) | c2_1(X39) | c0_1(X39)))) & (hskp8 | ! [X38] : (ndr1_0 => (~c1_1(X38) | c3_1(X38) | c2_1(X38))) | ! [X37] : (ndr1_0 => (~c3_1(X37) | c2_1(X37) | c0_1(X37)))) & (! [X36] : (ndr1_0 => (~c1_1(X36) | ~c0_1(X36) | c2_1(X36))) | ! [X35] : (ndr1_0 => (~c3_1(X35) | ~c0_1(X35) | c1_1(X35))) | ! [X34] : (ndr1_0 => (~c3_1(X34) | c2_1(X34) | c0_1(X34)))) & (! [X33] : (ndr1_0 => (~c3_1(X33) | ~c1_1(X33) | c2_1(X33))) | ! [X32] : (ndr1_0 => (~c2_1(X32) | c3_1(X32) | c1_1(X32))) | ! [X31] : (ndr1_0 => (~c3_1(X31) | c2_1(X31) | c0_1(X31)))) & (hskp13 | hskp12 | ! [X30] : (ndr1_0 => (~c1_1(X30) | c2_1(X30) | c0_1(X30)))) & (hskp11 | ! [X29] : (ndr1_0 => (~c3_1(X29) | ~c1_1(X29) | ~c0_1(X29))) | ! [X28] : (ndr1_0 => (~c1_1(X28) | c2_1(X28) | c0_1(X28)))) & (hskp29 | ! [X27] : (ndr1_0 => (~c1_1(X27) | ~c0_1(X27) | c3_1(X27))) | ! [X26] : (ndr1_0 => (~c1_1(X26) | c2_1(X26) | c0_1(X26)))) & (hskp10 | ! [X25] : (ndr1_0 => (c3_1(X25) | c2_1(X25) | c1_1(X25))) | ! [X24] : (ndr1_0 => (~c1_1(X24) | c2_1(X24) | c0_1(X24)))) & (! [X23] : (ndr1_0 => (c3_1(X23) | c2_1(X23) | c1_1(X23))) | ! [X22] : (ndr1_0 => (~c3_1(X22) | ~c2_1(X22) | c0_1(X22))) | ! [X21] : (ndr1_0 => (~c1_1(X21) | c2_1(X21) | c0_1(X21)))) & (! [X20] : (ndr1_0 => (~c2_1(X20) | ~c0_1(X20) | c3_1(X20))) | ! [X19] : (ndr1_0 => (~c3_1(X19) | ~c2_1(X19) | c1_1(X19))) | ! [X18] : (ndr1_0 => (c3_1(X18) | c2_1(X18) | c0_1(X18)))) & (hskp9 | ! [X17] : (ndr1_0 => (c3_1(X17) | c2_1(X17) | c1_1(X17))) | ! [X16] : (ndr1_0 => (c3_1(X16) | c2_1(X16) | c0_1(X16)))) & (hskp8 | hskp7 | ! [X15] : (ndr1_0 => (~c3_1(X15) | c1_1(X15) | c0_1(X15)))) & (hskp6 | hskp5 | ! [X14] : (ndr1_0 => (~c3_1(X14) | c1_1(X14) | c0_1(X14)))) & (hskp4 | ! [X13] : (ndr1_0 => (~c3_1(X13) | ~c0_1(X13) | c2_1(X13))) | ! [X12] : (ndr1_0 => (~c3_1(X12) | c1_1(X12) | c0_1(X12)))) & (hskp3 | ! [X11] : (ndr1_0 => (~c2_1(X11) | ~c0_1(X11) | c1_1(X11))) | ! [X10] : (ndr1_0 => (~c3_1(X10) | c1_1(X10) | c0_1(X10)))) & (hskp1 | ! [X9] : (ndr1_0 => (~c1_1(X9) | ~c0_1(X9) | c3_1(X9))) | ! [X8] : (ndr1_0 => (~c2_1(X8) | c1_1(X8) | c0_1(X8)))) & (hskp2 | hskp28 | ! [X7] : (ndr1_0 => (c3_1(X7) | c1_1(X7) | c0_1(X7)))) & (! [X6] : (ndr1_0 => (~c3_1(X6) | c2_1(X6) | c1_1(X6))) | ! [X5] : (ndr1_0 => (~c3_1(X5) | c2_1(X5) | c0_1(X5))) | ! [X4] : (ndr1_0 => (c3_1(X4) | c1_1(X4) | c0_1(X4)))) & (hskp1 | hskp0 | ! [X3] : (ndr1_0 => (c2_1(X3) | c1_1(X3) | c0_1(X3)))) & (! [X2] : (ndr1_0 => (~c3_1(X2) | ~c2_1(X2) | ~c0_1(X2))) | ! [X1] : (ndr1_0 => (~c2_1(X1) | c3_1(X1) | c1_1(X1))) | ! [X0] : (ndr1_0 => (c2_1(X0) | c1_1(X0) | c0_1(X0)))) & ((c3_1(a957) & c2_1(a957) & c1_1(a957) & ndr1_0) | ~hskp31) & ((c2_1(a919) & c1_1(a919) & c0_1(a919) & ndr1_0) | ~hskp30) & ((c3_1(a911) & c1_1(a911) & c0_1(a911) & ndr1_0) | ~hskp29) & ((c3_1(a900) & c2_1(a900) & c0_1(a900) & ndr1_0) | ~hskp28) & ((~c0_1(a978) & c3_1(a978) & c2_1(a978) & ndr1_0) | ~hskp27) & ((~c2_1(a969) & ~c1_1(a969) & c0_1(a969) & ndr1_0) | ~hskp26) & ((~c3_1(a958) & ~c1_1(a958) & ~c0_1(a958) & ndr1_0) | ~hskp25) & ((~c2_1(a953) & c3_1(a953) & c1_1(a953) & ndr1_0) | ~hskp24) & ((~c2_1(a950) & c3_1(a950) & c0_1(a950) & ndr1_0) | ~hskp23) & ((~c3_1(a939) & ~c0_1(a939) & c1_1(a939) & ndr1_0) | ~hskp22) & ((~c3_1(a938) & ~c2_1(a938) & c0_1(a938) & ndr1_0) | ~hskp21) & ((~c3_1(a937) & ~c0_1(a937) & c2_1(a937) & ndr1_0) | ~hskp20) & ((~c1_1(a930) & ~c0_1(a930) & c2_1(a930) & ndr1_0) | ~hskp19) & ((~c1_1(a929) & c2_1(a929) & c0_1(a929) & ndr1_0) | ~hskp18) & ((~c3_1(a928) & ~c2_1(a928) & c1_1(a928) & ndr1_0) | ~hskp17) & ((~c3_1(a923) & c2_1(a923) & c1_1(a923) & ndr1_0) | ~hskp16) & ((~c3_1(a921) & c1_1(a921) & c0_1(a921) & ndr1_0) | ~hskp15) & ((~c1_1(a918) & c3_1(a918) & c0_1(a918) & ndr1_0) | ~hskp14) & ((~c2_1(a914) & ~c1_1(a914) & c3_1(a914) & ndr1_0) | ~hskp13) & ((~c3_1(a913) & ~c1_1(a913) & c0_1(a913) & ndr1_0) | ~hskp12) & ((~c3_1(a912) & ~c2_1(a912) & ~c1_1(a912) & ndr1_0) | ~hskp11) & ((~c2_1(a910) & ~c0_1(a910) & c1_1(a910) & ndr1_0) | ~hskp10) & ((~c2_1(a909) & ~c1_1(a909) & ~c0_1(a909) & ndr1_0) | ~hskp9) & ((~c2_1(a908) & ~c0_1(a908) & c3_1(a908) & ndr1_0) | ~hskp8) & ((~c1_1(a907) & ~c0_1(a907) & c3_1(a907) & ndr1_0) | ~hskp7) & ((~c1_1(a906) & c3_1(a906) & c2_1(a906) & ndr1_0) | ~hskp6) & ((~c0_1(a905) & c3_1(a905) & c1_1(a905) & ndr1_0) | ~hskp5) & ((~c3_1(a904) & ~c1_1(a904) & c2_1(a904) & ndr1_0) | ~hskp4) & ((~c3_1(a903) & c2_1(a903) & c0_1(a903) & ndr1_0) | ~hskp3) & ((~c3_1(a901) & ~c2_1(a901) & ~c0_1(a901) & ndr1_0) | ~hskp2) & ((~c0_1(a899) & c2_1(a899) & c1_1(a899) & ndr1_0) | ~hskp1) & ((~c2_1(a898) & c1_1(a898) & c0_1(a898) & ndr1_0) | ~hskp0))),
  file('./p/SYN474+1.tptp',unknown)).
fof(f1057,plain,(
  ~spl0_14 | spl0_161),
  inference(avatar_split_clause,[],[f8,f1054,f305])).
fof(f8,plain,(
  c0_1(a898) | ~hskp0),
  inference(cnf_transformation,[],[f6])).
fof(f1052,plain,(
  ~spl0_14 | spl0_160),
  inference(avatar_split_clause,[],[f9,f1049,f305])).
fof(f9,plain,(
  c1_1(a898) | ~hskp0),
  inference(cnf_transformation,[],[f6])).
fof(f1047,plain,(
  ~spl0_14 | ~spl0_159),
  inference(avatar_split_clause,[],[f10,f1044,f305])).
fof(f10,plain,(
  ~c2_1(a898) | ~hskp0),
  inference(cnf_transformation,[],[f6])).
fof(f1041,plain,(
  ~spl0_29 | spl0_158),
  inference(avatar_split_clause,[],[f12,f1038,f368])).
fof(f368,plain,(
  spl0_29 <=> hskp1),
  introduced(avatar_definition,[new_symbols(naming,[spl0_29])])).
fof(f12,plain,(
  c1_1(a899) | ~hskp1),
  inference(cnf_transformation,[],[f6])).
fof(f1036,plain,(
  ~spl0_29 | spl0_157),
  inference(avatar_split_clause,[],[f13,f1033,f368])).
fof(f13,plain,(
  c2_1(a899) | ~hskp1),
  inference(cnf_transformation,[],[f6])).
fof(f1031,plain,(
  ~spl0_29 | ~spl0_156),
  inference(avatar_split_clause,[],[f14,f1028,f368])).
fof(f14,plain,(
  ~c0_1(a899) | ~hskp1),
  inference(cnf_transformation,[],[f6])).
fof(f1025,plain,(
  ~spl0_22 | ~spl0_155),
  inference(avatar_split_clause,[],[f16,f1022,f339])).
fof(f339,plain,(
  spl0_22 <=> hskp2),
  introduced(avatar_definition,[new_symbols(naming,[spl0_22])])).
fof(f16,plain,(
  ~c0_1(a901) | ~hskp2),
  inference(cnf_transformation,[],[f6])).
fof(f1020,plain,(
  ~spl0_22 | ~spl0_154),
  inference(avatar_split_clause,[],[f17,f1017,f339])).
fof(f17,plain,(
  ~c2_1(a901) | ~hskp2),
  inference(cnf_transformation,[],[f6])).
fof(f1015,plain,(
  ~spl0_22 | ~spl0_153),
  inference(avatar_split_clause,[],[f18,f1012,f339])).
fof(f18,plain,(
  ~c3_1(a901) | ~hskp2),
  inference(cnf_transformation,[],[f6])).
fof(f993,plain,(
  ~spl0_54 | spl0_149),
  inference(avatar_split_clause,[],[f24,f990,f481])).
fof(f481,plain,(
  spl0_54 <=> hskp4),
  introduced(avatar_definition,[new_symbols(naming,[spl0_54])])).
fof(f24,plain,(
  c2_1(a904) | ~hskp4),
  inference(cnf_transformation,[],[f6])).
fof(f988,plain,(
  ~spl0_54 | ~spl0_148),
  inference(avatar_split_clause,[],[f25,f985,f481])).
fof(f25,plain,(
  ~c1_1(a904) | ~hskp4),
  inference(cnf_transformation,[],[f6])).
fof(f983,plain,(
  ~spl0_54 | ~spl0_147),
  inference(avatar_split_clause,[],[f26,f980,f481])).
fof(f26,plain,(
  ~c3_1(a904) | ~hskp4),
  inference(cnf_transformation,[],[f6])).
fof(f977,plain,(
  ~spl0_32 | spl0_146),
  inference(avatar_split_clause,[],[f28,f974,f381])).
fof(f381,plain,(
  spl0_32 <=> hskp5),
  introduced(avatar_definition,[new_symbols(naming,[spl0_32])])).
fof(f28,plain,(
  c1_1(a905) | ~hskp5),
  inference(cnf_transformation,[],[f6])).
fof(f972,plain,(
  ~spl0_32 | spl0_145),
  inference(avatar_split_clause,[],[f29,f969,f381])).
fof(f29,plain,(
  c3_1(a905) | ~hskp5),
  inference(cnf_transformation,[],[f6])).
fof(f967,plain,(
  ~spl0_32 | ~spl0_144),
  inference(avatar_split_clause,[],[f30,f964,f381])).
fof(f30,plain,(
  ~c0_1(a905) | ~hskp5),
  inference(cnf_transformation,[],[f6])).
fof(f961,plain,(
  ~spl0_30 | spl0_143),
  inference(avatar_split_clause,[],[f32,f958,f373])).
fof(f373,plain,(
  spl0_30 <=> hskp6),
  introduced(avatar_definition,[new_symbols(naming,[spl0_30])])).
fof(f32,plain,(
  c2_1(a906) | ~hskp6),
  inference(cnf_transformation,[],[f6])).
fof(f956,plain,(
  ~spl0_30 | spl0_142),
  inference(avatar_split_clause,[],[f33,f953,f373])).
fof(f33,plain,(
  c3_1(a906) | ~hskp6),
  inference(cnf_transformation,[],[f6])).
fof(f951,plain,(
  ~spl0_30 | ~spl0_141),
  inference(avatar_split_clause,[],[f34,f948,f373])).
fof(f34,plain,(
  ~c1_1(a906) | ~hskp6),
  inference(cnf_transformation,[],[f6])).
fof(f945,plain,(
  ~spl0_2 | spl0_140),
  inference(avatar_split_clause,[],[f36,f942,f252])).
fof(f252,plain,(
  spl0_2 <=> hskp7),
  introduced(avatar_definition,[new_symbols(naming,[spl0_2])])).
fof(f36,plain,(
  c3_1(a907) | ~hskp7),
  inference(cnf_transformation,[],[f6])).
fof(f940,plain,(
  ~spl0_2 | ~spl0_139),
  inference(avatar_split_clause,[],[f37,f937,f252])).
fof(f37,plain,(
  ~c0_1(a907) | ~hskp7),
  inference(cnf_transformation,[],[f6])).
fof(f935,plain,(
  ~spl0_2 | ~spl0_138),
  inference(avatar_split_clause,[],[f38,f932,f252])).
fof(f38,plain,(
  ~c1_1(a907) | ~hskp7),
  inference(cnf_transformation,[],[f6])).
fof(f929,plain,(
  ~spl0_58 | spl0_137),
  inference(avatar_split_clause,[],[f40,f926,f503])).
fof(f503,plain,(
  spl0_58 <=> hskp8),
  introduced(avatar_definition,[new_symbols(naming,[spl0_58])])).
fof(f40,plain,(
  c3_1(a908) | ~hskp8),
  inference(cnf_transformation,[],[f6])).
fof(f919,plain,(
  ~spl0_58 | ~spl0_135),
  inference(avatar_split_clause,[],[f42,f916,f503])).
fof(f42,plain,(
  ~c2_1(a908) | ~hskp8),
  inference(cnf_transformation,[],[f6])).
fof(f914,plain,(
  ~spl0_3 | spl0_19),
  inference(avatar_split_clause,[],[f43,f328,f256])).
fof(f256,plain,(
  spl0_3 <=> hskp9),
  introduced(avatar_definition,[new_symbols(naming,[spl0_3])])).
fof(f43,plain,(
  ndr1_0 | ~hskp9),
  inference(cnf_transformation,[],[f6])).
fof(f913,plain,(
  ~spl0_3 | ~spl0_134),
  inference(avatar_split_clause,[],[f44,f910,f256])).
fof(f44,plain,(
  ~c0_1(a909) | ~hskp9),
  inference(cnf_transformation,[],[f6])).
fof(f908,plain,(
  ~spl0_3 | ~spl0_133),
  inference(avatar_split_clause,[],[f45,f905,f256])).
fof(f45,plain,(
  ~c1_1(a909) | ~hskp9),
  inference(cnf_transformation,[],[f6])).
fof(f903,plain,(
  ~spl0_3 | ~spl0_132),
  inference(avatar_split_clause,[],[f46,f900,f256])).
fof(f46,plain,(
  ~c2_1(a909) | ~hskp9),
  inference(cnf_transformation,[],[f6])).
fof(f897,plain,(
  ~spl0_25 | spl0_131),
  inference(avatar_split_clause,[],[f48,f894,f351])).
fof(f351,plain,(
  spl0_25 <=> hskp10),
  introduced(avatar_definition,[new_symbols(naming,[spl0_25])])).
fof(f48,plain,(
  c1_1(a910) | ~hskp10),
  inference(cnf_transformation,[],[f6])).
fof(f892,plain,(
  ~spl0_25 | ~spl0_130),
  inference(avatar_split_clause,[],[f49,f889,f351])).
fof(f49,plain,(
  ~c0_1(a910) | ~hskp10),
  inference(cnf_transformation,[],[f6])).
fof(f887,plain,(
  ~spl0_25 | ~spl0_129),
  inference(avatar_split_clause,[],[f50,f884,f351])).
fof(f50,plain,(
  ~c2_1(a910) | ~hskp10),
  inference(cnf_transformation,[],[f6])).
fof(f881,plain,(
  ~spl0_60 | ~spl0_128),
  inference(avatar_split_clause,[],[f52,f878,f514])).
fof(f514,plain,(
  spl0_60 <=> hskp11),
  introduced(avatar_definition,[new_symbols(naming,[spl0_60])])).
fof(f52,plain,(
  ~c1_1(a912) | ~hskp11),
  inference(cnf_transformation,[],[f6])).
fof(f876,plain,(
  ~spl0_60 | ~spl0_127),
  inference(avatar_split_clause,[],[f53,f873,f514])).
fof(f53,plain,(
  ~c2_1(a912) | ~hskp11),
  inference(cnf_transformation,[],[f6])).
fof(f871,plain,(
  ~spl0_60 | ~spl0_126),
  inference(avatar_split_clause,[],[f54,f868,f514])).
fof(f54,plain,(
  ~c3_1(a912) | ~hskp11),
  inference(cnf_transformation,[],[f6])).
fof(f866,plain,(
  ~spl0_16 | spl0_19),
  inference(avatar_split_clause,[],[f55,f328,f314])).
fof(f314,plain,(
  spl0_16 <=> hskp12),
  introduced(avatar_definition,[new_symbols(naming,[spl0_16])])).
fof(f55,plain,(
  ndr1_0 | ~hskp12),
  inference(cnf_transformation,[],[f6])).
fof(f865,plain,(
  ~spl0_16 | spl0_125),
  inference(avatar_split_clause,[],[f56,f862,f314])).
fof(f56,plain,(
  c0_1(a913) | ~hskp12),
  inference(cnf_transformation,[],[f6])).
fof(f860,plain,(
  ~spl0_16 | ~spl0_124),
  inference(avatar_split_clause,[],[f57,f857,f314])).
fof(f57,plain,(
  ~c1_1(a913) | ~hskp12),
  inference(cnf_transformation,[],[f6])).
fof(f855,plain,(
  ~spl0_16 | ~spl0_123),
  inference(avatar_split_clause,[],[f58,f852,f314])).
fof(f58,plain,(
  ~c3_1(a913) | ~hskp12),
  inference(cnf_transformation,[],[f6])).
fof(f849,plain,(
  ~spl0_21 | spl0_122),
  inference(avatar_split_clause,[],[f60,f846,f335])).
fof(f335,plain,(
  spl0_21 <=> hskp13),
  introduced(avatar_definition,[new_symbols(naming,[spl0_21])])).
fof(f60,plain,(
  c3_1(a914) | ~hskp13),
  inference(cnf_transformation,[],[f6])).
fof(f844,plain,(
  ~spl0_21 | ~spl0_121),
  inference(avatar_split_clause,[],[f61,f841,f335])).
fof(f61,plain,(
  ~c1_1(a914) | ~hskp13),
  inference(cnf_transformation,[],[f6])).
fof(f839,plain,(
  ~spl0_21 | ~spl0_120),
  inference(avatar_split_clause,[],[f62,f836,f335])).
fof(f62,plain,(
  ~c2_1(a914) | ~hskp13),
  inference(cnf_transformation,[],[f6])).
fof(f833,plain,(
  ~spl0_50 | spl0_119),
  inference(avatar_split_clause,[],[f64,f830,f462])).
fof(f462,plain,(
  spl0_50 <=> hskp14),
  introduced(avatar_definition,[new_symbols(naming,[spl0_50])])).
fof(f64,plain,(
  c0_1(a918) | ~hskp14),
  inference(cnf_transformation,[],[f6])).
fof(f828,plain,(
  ~spl0_50 | spl0_118),
  inference(avatar_split_clause,[],[f65,f825,f462])).
fof(f65,plain,(
  c3_1(a918) | ~hskp14),
  inference(cnf_transformation,[],[f6])).
fof(f823,plain,(
  ~spl0_50 | ~spl0_117),
  inference(avatar_split_clause,[],[f66,f820,f462])).
fof(f66,plain,(
  ~c1_1(a918) | ~hskp14),
  inference(cnf_transformation,[],[f6])).
fof(f817,plain,(
  ~spl0_12 | spl0_116),
  inference(avatar_split_clause,[],[f68,f814,f296])).
fof(f296,plain,(
  spl0_12 <=> hskp15),
  introduced(avatar_definition,[new_symbols(naming,[spl0_12])])).
fof(f68,plain,(
  c0_1(a921) | ~hskp15),
  inference(cnf_transformation,[],[f6])).
fof(f812,plain,(
  ~spl0_12 | spl0_115),
  inference(avatar_split_clause,[],[f69,f809,f296])).
fof(f69,plain,(
  c1_1(a921) | ~hskp15),
  inference(cnf_transformation,[],[f6])).
fof(f807,plain,(
  ~spl0_12 | ~spl0_114),
  inference(avatar_split_clause,[],[f70,f804,f296])).
fof(f70,plain,(
  ~c3_1(a921) | ~hskp15),
  inference(cnf_transformation,[],[f6])).
fof(f801,plain,(
  ~spl0_24 | spl0_113),
  inference(avatar_split_clause,[],[f72,f798,f347])).
fof(f347,plain,(
  spl0_24 <=> hskp16),
  introduced(avatar_definition,[new_symbols(naming,[spl0_24])])).
fof(f72,plain,(
  c1_1(a923) | ~hskp16),
  inference(cnf_transformation,[],[f6])).
fof(f796,plain,(
  ~spl0_24 | spl0_112),
  inference(avatar_split_clause,[],[f73,f793,f347])).
fof(f73,plain,(
  c2_1(a923) | ~hskp16),
  inference(cnf_transformation,[],[f6])).
fof(f791,plain,(
  ~spl0_24 | ~spl0_111),
  inference(avatar_split_clause,[],[f74,f788,f347])).
fof(f74,plain,(
  ~c3_1(a923) | ~hskp16),
  inference(cnf_transformation,[],[f6])).
fof(f785,plain,(
  ~spl0_11 | spl0_110),
  inference(avatar_split_clause,[],[f76,f782,f291])).
fof(f291,plain,(
  spl0_11 <=> hskp17),
  introduced(avatar_definition,[new_symbols(naming,[spl0_11])])).
fof(f76,plain,(
  c1_1(a928) | ~hskp17),
  inference(cnf_transformation,[],[f6])).
fof(f780,plain,(
  ~spl0_11 | ~spl0_109),
  inference(avatar_split_clause,[],[f77,f777,f291])).
fof(f77,plain,(
  ~c2_1(a928) | ~hskp17),
  inference(cnf_transformation,[],[f6])).
fof(f775,plain,(
  ~spl0_11 | ~spl0_108),
  inference(avatar_split_clause,[],[f78,f772,f291])).
fof(f78,plain,(
  ~c3_1(a928) | ~hskp17),
  inference(cnf_transformation,[],[f6])).
fof(f769,plain,(
  ~spl0_40 | spl0_107),
  inference(avatar_split_clause,[],[f80,f766,f415])).
fof(f415,plain,(
  spl0_40 <=> hskp18),
  introduced(avatar_definition,[new_symbols(naming,[spl0_40])])).
fof(f80,plain,(
  c0_1(a929) | ~hskp18),
  inference(cnf_transformation,[],[f6])).
fof(f764,plain,(
  ~spl0_40 | spl0_106),
  inference(avatar_split_clause,[],[f81,f761,f415])).
fof(f81,plain,(
  c2_1(a929) | ~hskp18),
  inference(cnf_transformation,[],[f6])).
fof(f759,plain,(
  ~spl0_40 | ~spl0_105),
  inference(avatar_split_clause,[],[f82,f756,f415])).
fof(f82,plain,(
  ~c1_1(a929) | ~hskp18),
  inference(cnf_transformation,[],[f6])).
fof(f753,plain,(
  ~spl0_13 | spl0_104),
  inference(avatar_split_clause,[],[f84,f750,f300])).
fof(f300,plain,(
  spl0_13 <=> hskp19),
  introduced(avatar_definition,[new_symbols(naming,[spl0_13])])).
fof(f84,plain,(
  c2_1(a930) | ~hskp19),
  inference(cnf_transformation,[],[f6])).
fof(f748,plain,(
  ~spl0_13 | ~spl0_103),
  inference(avatar_split_clause,[],[f85,f745,f300])).
fof(f85,plain,(
  ~c0_1(a930) | ~hskp19),
  inference(cnf_transformation,[],[f6])).
fof(f743,plain,(
  ~spl0_13 | ~spl0_102),
  inference(avatar_split_clause,[],[f86,f740,f300])).
fof(f86,plain,(
  ~c1_1(a930) | ~hskp19),
  inference(cnf_transformation,[],[f6])).
fof(f737,plain,(
  ~spl0_34 | spl0_101),
  inference(avatar_split_clause,[],[f88,f734,f390])).
fof(f390,plain,(
  spl0_34 <=> hskp20),
  introduced(avatar_definition,[new_symbols(naming,[spl0_34])])).
fof(f88,plain,(
  c2_1(a937) | ~hskp20),
  inference(cnf_transformation,[],[f6])).
fof(f732,plain,(
  ~spl0_34 | ~spl0_100),
  inference(avatar_split_clause,[],[f89,f729,f390])).
fof(f89,plain,(
  ~c0_1(a937) | ~hskp20),
  inference(cnf_transformation,[],[f6])).
fof(f727,plain,(
  ~spl0_34 | ~spl0_99),
  inference(avatar_split_clause,[],[f90,f724,f390])).
fof(f90,plain,(
  ~c3_1(a937) | ~hskp20),
  inference(cnf_transformation,[],[f6])).
fof(f721,plain,(
  ~spl0_15 | spl0_98),
  inference(avatar_split_clause,[],[f92,f718,f309])).
fof(f309,plain,(
  spl0_15 <=> hskp21),
  introduced(avatar_definition,[new_symbols(naming,[spl0_15])])).
fof(f92,plain,(
  c0_1(a938) | ~hskp21),
  inference(cnf_transformation,[],[f6])).
fof(f716,plain,(
  ~spl0_15 | ~spl0_97),
  inference(avatar_split_clause,[],[f93,f713,f309])).
fof(f93,plain,(
  ~c2_1(a938) | ~hskp21),
  inference(cnf_transformation,[],[f6])).
fof(f711,plain,(
  ~spl0_15 | ~spl0_96),
  inference(avatar_split_clause,[],[f94,f708,f309])).
fof(f94,plain,(
  ~c3_1(a938) | ~hskp21),
  inference(cnf_transformation,[],[f6])).
fof(f705,plain,(
  ~spl0_6 | spl0_95),
  inference(avatar_split_clause,[],[f96,f702,f269])).
fof(f269,plain,(
  spl0_6 <=> hskp22),
  introduced(avatar_definition,[new_symbols(naming,[spl0_6])])).
fof(f96,plain,(
  c1_1(a939) | ~hskp22),
  inference(cnf_transformation,[],[f6])).
fof(f700,plain,(
  ~spl0_6 | ~spl0_94),
  inference(avatar_split_clause,[],[f97,f697,f269])).
fof(f97,plain,(
  ~c0_1(a939) | ~hskp22),
  inference(cnf_transformation,[],[f6])).
fof(f695,plain,(
  ~spl0_6 | ~spl0_93),
  inference(avatar_split_clause,[],[f98,f692,f269])).
fof(f98,plain,(
  ~c3_1(a939) | ~hskp22),
  inference(cnf_transformation,[],[f6])).
fof(f689,plain,(
  ~spl0_4 | spl0_92),
  inference(avatar_split_clause,[],[f100,f686,f261])).
fof(f261,plain,(
  spl0_4 <=> hskp23),
  introduced(avatar_definition,[new_symbols(naming,[spl0_4])])).
fof(f100,plain,(
  c0_1(a950) | ~hskp23),
  inference(cnf_transformation,[],[f6])).
fof(f684,plain,(
  ~spl0_4 | spl0_91),
  inference(avatar_split_clause,[],[f101,f681,f261])).
fof(f101,plain,(
  c3_1(a950) | ~hskp23),
  inference(cnf_transformation,[],[f6])).
fof(f679,plain,(
  ~spl0_4 | ~spl0_90),
  inference(avatar_split_clause,[],[f102,f676,f261])).
fof(f102,plain,(
  ~c2_1(a950) | ~hskp23),
  inference(cnf_transformation,[],[f6])).
fof(f673,plain,(
  ~spl0_8 | spl0_89),
  inference(avatar_split_clause,[],[f104,f670,f278])).
fof(f278,plain,(
  spl0_8 <=> hskp24),
  introduced(avatar_definition,[new_symbols(naming,[spl0_8])])).
fof(f104,plain,(
  c1_1(a953) | ~hskp24),
  inference(cnf_transformation,[],[f6])).
fof(f668,plain,(
  ~spl0_8 | spl0_88),
  inference(avatar_split_clause,[],[f105,f665,f278])).
fof(f105,plain,(
  c3_1(a953) | ~hskp24),
  inference(cnf_transformation,[],[f6])).
fof(f663,plain,(
  ~spl0_8 | ~spl0_87),
  inference(avatar_split_clause,[],[f106,f660,f278])).
fof(f106,plain,(
  ~c2_1(a953) | ~hskp24),
  inference(cnf_transformation,[],[f6])).
fof(f657,plain,(
  ~spl0_9 | ~spl0_86),
  inference(avatar_split_clause,[],[f108,f654,f282])).
fof(f282,plain,(
  spl0_9 <=> hskp25),
  introduced(avatar_definition,[new_symbols(naming,[spl0_9])])).
fof(f108,plain,(
  ~c0_1(a958) | ~hskp25),
  inference(cnf_transformation,[],[f6])).
fof(f652,plain,(
  ~spl0_9 | ~spl0_85),
  inference(avatar_split_clause,[],[f109,f649,f282])).
fof(f109,plain,(
  ~c1_1(a958) | ~hskp25),
  inference(cnf_transformation,[],[f6])).
fof(f647,plain,(
  ~spl0_9 | ~spl0_84),
  inference(avatar_split_clause,[],[f110,f644,f282])).
fof(f110,plain,(
  ~c3_1(a958) | ~hskp25),
  inference(cnf_transformation,[],[f6])).
fof(f641,plain,(
  ~spl0_5 | spl0_83),
  inference(avatar_split_clause,[],[f112,f638,f265])).
fof(f265,plain,(
  spl0_5 <=> hskp26),
  introduced(avatar_definition,[new_symbols(naming,[spl0_5])])).
fof(f112,plain,(
  c0_1(a969) | ~hskp26),
  inference(cnf_transformation,[],[f6])).
fof(f636,plain,(
  ~spl0_5 | ~spl0_82),
  inference(avatar_split_clause,[],[f113,f633,f265])).
fof(f113,plain,(
  ~c1_1(a969) | ~hskp26),
  inference(cnf_transformation,[],[f6])).
fof(f631,plain,(
  ~spl0_5 | ~spl0_81),
  inference(avatar_split_clause,[],[f114,f628,f265])).
fof(f114,plain,(
  ~c2_1(a969) | ~hskp26),
  inference(cnf_transformation,[],[f6])).
fof(f625,plain,(
  ~spl0_1 | spl0_80),
  inference(avatar_split_clause,[],[f116,f622,f248])).
fof(f248,plain,(
  spl0_1 <=> hskp27),
  introduced(avatar_definition,[new_symbols(naming,[spl0_1])])).
fof(f116,plain,(
  c2_1(a978) | ~hskp27),
  inference(cnf_transformation,[],[f6])).
fof(f620,plain,(
  ~spl0_1 | spl0_79),
  inference(avatar_split_clause,[],[f117,f617,f248])).
fof(f117,plain,(
  c3_1(a978) | ~hskp27),
  inference(cnf_transformation,[],[f6])).
fof(f615,plain,(
  ~spl0_1 | ~spl0_78),
  inference(avatar_split_clause,[],[f118,f612,f248])).
fof(f118,plain,(
  ~c0_1(a978) | ~hskp27),
  inference(cnf_transformation,[],[f6])).
fof(f609,plain,(
  ~spl0_10 | spl0_77),
  inference(avatar_split_clause,[],[f120,f606,f287])).
fof(f287,plain,(
  spl0_10 <=> hskp28),
  introduced(avatar_definition,[new_symbols(naming,[spl0_10])])).
fof(f120,plain,(
  c0_1(a900) | ~hskp28),
  inference(cnf_transformation,[],[f6])).
fof(f604,plain,(
  ~spl0_10 | spl0_76),
  inference(avatar_split_clause,[],[f121,f601,f287])).
fof(f121,plain,(
  c2_1(a900) | ~hskp28),
  inference(cnf_transformation,[],[f6])).
fof(f599,plain,(
  ~spl0_10 | spl0_75),
  inference(avatar_split_clause,[],[f122,f596,f287])).
fof(f122,plain,(
  c3_1(a900) | ~hskp28),
  inference(cnf_transformation,[],[f6])).
fof(f593,plain,(
  ~spl0_17 | spl0_74),
  inference(avatar_split_clause,[],[f124,f590,f319])).
fof(f319,plain,(
  spl0_17 <=> hskp29),
  introduced(avatar_definition,[new_symbols(naming,[spl0_17])])).
fof(f124,plain,(
  c0_1(a911) | ~hskp29),
  inference(cnf_transformation,[],[f6])).
fof(f588,plain,(
  ~spl0_17 | spl0_73),
  inference(avatar_split_clause,[],[f125,f585,f319])).
fof(f125,plain,(
  c1_1(a911) | ~hskp29),
  inference(cnf_transformation,[],[f6])).
fof(f583,plain,(
  ~spl0_17 | spl0_72),
  inference(avatar_split_clause,[],[f126,f580,f319])).
fof(f126,plain,(
  c3_1(a911) | ~hskp29),
  inference(cnf_transformation,[],[f6])).
fof(f577,plain,(
  ~spl0_56 | spl0_71),
  inference(avatar_split_clause,[],[f128,f574,f492])).
fof(f492,plain,(
  spl0_56 <=> hskp30),
  introduced(avatar_definition,[new_symbols(naming,[spl0_56])])).
fof(f128,plain,(
  c0_1(a919) | ~hskp30),
  inference(cnf_transformation,[],[f6])).
fof(f572,plain,(
  ~spl0_56 | spl0_70),
  inference(avatar_split_clause,[],[f129,f569,f492])).
fof(f129,plain,(
  c1_1(a919) | ~hskp30),
  inference(cnf_transformation,[],[f6])).
fof(f567,plain,(
  ~spl0_56 | spl0_69),
  inference(avatar_split_clause,[],[f130,f564,f492])).
fof(f130,plain,(
  c2_1(a919) | ~hskp30),
  inference(cnf_transformation,[],[f6])).
fof(f561,plain,(
  ~spl0_18 | spl0_68),
  inference(avatar_split_clause,[],[f132,f558,f323])).
fof(f323,plain,(
  spl0_18 <=> hskp31),
  introduced(avatar_definition,[new_symbols(naming,[spl0_18])])).
fof(f132,plain,(
  c1_1(a957) | ~hskp31),
  inference(cnf_transformation,[],[f6])).
fof(f556,plain,(
  ~spl0_18 | spl0_67),
  inference(avatar_split_clause,[],[f133,f553,f323])).
fof(f133,plain,(
  c2_1(a957) | ~hskp31),
  inference(cnf_transformation,[],[f6])).
fof(f551,plain,(
  ~spl0_18 | spl0_66),
  inference(avatar_split_clause,[],[f134,f548,f323])).
fof(f134,plain,(
  c3_1(a957) | ~hskp31),
  inference(cnf_transformation,[],[f6])).
fof(f546,plain,(
  spl0_65 | spl0_42 | ~spl0_19 | spl0_20),
  inference(avatar_split_clause,[],[f207,f332,f328,f426,f543])).
fof(f207,plain,(
  ( ! [X111,X112,X110] : (~c3_1(X110) | ~c2_1(X110) | ~c0_1(X110) | ~ndr1_0 | ~c2_1(X111) | c3_1(X111) | c1_1(X111) | c2_1(X112) | c1_1(X112) | c0_1(X112)) )),
  inference(duplicate_literal_removal,[],[f135])).
fof(f135,plain,(
  ( ! [X111,X112,X110] : (~c3_1(X110) | ~c2_1(X110) | ~c0_1(X110) | ~ndr1_0 | ~c2_1(X111) | c3_1(X111) | c1_1(X111) | ~ndr1_0 | c2_1(X112) | c1_1(X112) | c0_1(X112) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f545,plain,(
  ~spl0_19 | spl0_65 | spl0_14 | spl0_29),
  inference(avatar_split_clause,[],[f136,f368,f305,f543,f328])).
fof(f136,plain,(
  ( ! [X109] : (hskp1 | hskp0 | c2_1(X109) | c1_1(X109) | c0_1(X109) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f541,plain,(
  spl0_64 | spl0_55 | ~spl0_19 | spl0_44),
  inference(avatar_split_clause,[],[f208,f435,f328,f487,f538])).
fof(f208,plain,(
  ( ! [X107,X108,X106] : (~c3_1(X106) | c2_1(X106) | c1_1(X106) | ~ndr1_0 | ~c3_1(X107) | c2_1(X107) | c0_1(X107) | c3_1(X108) | c1_1(X108) | c0_1(X108)) )),
  inference(duplicate_literal_removal,[],[f137])).
fof(f137,plain,(
  ( ! [X107,X108,X106] : (~c3_1(X106) | c2_1(X106) | c1_1(X106) | ~ndr1_0 | ~c3_1(X107) | c2_1(X107) | c0_1(X107) | ~ndr1_0 | c3_1(X108) | c1_1(X108) | c0_1(X108) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f540,plain,(
  ~spl0_19 | spl0_64 | spl0_10 | spl0_22),
  inference(avatar_split_clause,[],[f138,f339,f287,f538,f328])).
fof(f138,plain,(
  ( ! [X105] : (hskp2 | hskp28 | c3_1(X105) | c1_1(X105) | c0_1(X105) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f536,plain,(
  spl0_63 | ~spl0_19 | spl0_27 | spl0_29),
  inference(avatar_split_clause,[],[f209,f368,f360,f328,f534])).
fof(f209,plain,(
  ( ! [X103,X104] : (hskp1 | ~c1_1(X103) | ~c0_1(X103) | c3_1(X103) | ~ndr1_0 | ~c2_1(X104) | c1_1(X104) | c0_1(X104)) )),
  inference(duplicate_literal_removal,[],[f139])).
fof(f139,plain,(
  ( ! [X103,X104] : (hskp1 | ~c1_1(X103) | ~c0_1(X103) | c3_1(X103) | ~ndr1_0 | ~c2_1(X104) | c1_1(X104) | c0_1(X104) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f531,plain,(
  spl0_62 | ~spl0_19 | spl0_31 | spl0_54),
  inference(avatar_split_clause,[],[f211,f481,f378,f328,f527])).
fof(f211,plain,(
  ( ! [X99,X100] : (hskp4 | ~c3_1(X99) | ~c0_1(X99) | c2_1(X99) | ~ndr1_0 | ~c3_1(X100) | c1_1(X100) | c0_1(X100)) )),
  inference(duplicate_literal_removal,[],[f141])).
fof(f141,plain,(
  ( ! [X99,X100] : (hskp4 | ~c3_1(X99) | ~c0_1(X99) | c2_1(X99) | ~ndr1_0 | ~c3_1(X100) | c1_1(X100) | c0_1(X100) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f525,plain,(
  spl0_61 | ~spl0_19 | spl0_47 | spl0_3),
  inference(avatar_split_clause,[],[f212,f256,f449,f328,f522])).
fof(f212,plain,(
  ( ! [X95,X96] : (hskp9 | c3_1(X95) | c2_1(X95) | c1_1(X95) | ~ndr1_0 | c3_1(X96) | c2_1(X96) | c0_1(X96)) )),
  inference(duplicate_literal_removal,[],[f144])).
fof(f144,plain,(
  ( ! [X95,X96] : (hskp9 | c3_1(X95) | c2_1(X95) | c1_1(X95) | ~ndr1_0 | c3_1(X96) | c2_1(X96) | c0_1(X96) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f524,plain,(
  spl0_61 | spl0_37 | ~spl0_19 | spl0_23),
  inference(avatar_split_clause,[],[f213,f344,f328,f404,f522])).
fof(f213,plain,(
  ( ! [X94,X92,X93] : (~c2_1(X92) | ~c0_1(X92) | c3_1(X92) | ~ndr1_0 | ~c3_1(X93) | ~c2_1(X93) | c1_1(X93) | c3_1(X94) | c2_1(X94) | c0_1(X94)) )),
  inference(duplicate_literal_removal,[],[f145])).
fof(f145,plain,(
  ( ! [X94,X92,X93] : (~c2_1(X92) | ~c0_1(X92) | c3_1(X92) | ~ndr1_0 | ~c3_1(X93) | ~c2_1(X93) | c1_1(X93) | ~ndr1_0 | c3_1(X94) | c2_1(X94) | c0_1(X94) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f520,plain,(
  spl0_59 | spl0_48 | ~spl0_19 | spl0_47),
  inference(avatar_split_clause,[],[f214,f449,f328,f454,f510])).
fof(f214,plain,(
  ( ! [X90,X91,X89] : (c3_1(X89) | c2_1(X89) | c1_1(X89) | ~ndr1_0 | ~c3_1(X90) | ~c2_1(X90) | c0_1(X90) | ~c1_1(X91) | c2_1(X91) | c0_1(X91)) )),
  inference(duplicate_literal_removal,[],[f146])).
fof(f146,plain,(
  ( ! [X90,X91,X89] : (c3_1(X89) | c2_1(X89) | c1_1(X89) | ~ndr1_0 | ~c3_1(X90) | ~c2_1(X90) | c0_1(X90) | ~ndr1_0 | ~c1_1(X91) | c2_1(X91) | c0_1(X91) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f519,plain,(
  spl0_59 | ~spl0_19 | spl0_47 | spl0_25),
  inference(avatar_split_clause,[],[f215,f351,f449,f328,f510])).
fof(f215,plain,(
  ( ! [X88,X87] : (hskp10 | c3_1(X87) | c2_1(X87) | c1_1(X87) | ~ndr1_0 | ~c1_1(X88) | c2_1(X88) | c0_1(X88)) )),
  inference(duplicate_literal_removal,[],[f147])).
fof(f147,plain,(
  ( ! [X88,X87] : (hskp10 | c3_1(X87) | c2_1(X87) | c1_1(X87) | ~ndr1_0 | ~c1_1(X88) | c2_1(X88) | c0_1(X88) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f518,plain,(
  spl0_59 | ~spl0_19 | spl0_27 | spl0_17),
  inference(avatar_split_clause,[],[f216,f319,f360,f328,f510])).
fof(f216,plain,(
  ( ! [X85,X86] : (hskp29 | ~c1_1(X85) | ~c0_1(X85) | c3_1(X85) | ~ndr1_0 | ~c1_1(X86) | c2_1(X86) | c0_1(X86)) )),
  inference(duplicate_literal_removal,[],[f148])).
fof(f148,plain,(
  ( ! [X85,X86] : (hskp29 | ~c1_1(X85) | ~c0_1(X85) | c3_1(X85) | ~ndr1_0 | ~c1_1(X86) | c2_1(X86) | c0_1(X86) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f517,plain,(
  spl0_59 | ~spl0_19 | spl0_38 | spl0_60),
  inference(avatar_split_clause,[],[f217,f514,f408,f328,f510])).
fof(f217,plain,(
  ( ! [X83,X84] : (hskp11 | ~c3_1(X83) | ~c1_1(X83) | ~c0_1(X83) | ~ndr1_0 | ~c1_1(X84) | c2_1(X84) | c0_1(X84)) )),
  inference(duplicate_literal_removal,[],[f149])).
fof(f149,plain,(
  ( ! [X83,X84] : (hskp11 | ~c3_1(X83) | ~c1_1(X83) | ~c0_1(X83) | ~ndr1_0 | ~c1_1(X84) | c2_1(X84) | c0_1(X84) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f512,plain,(
  ~spl0_19 | spl0_59 | spl0_16 | spl0_21),
  inference(avatar_split_clause,[],[f150,f335,f314,f510,f328])).
fof(f150,plain,(
  ( ! [X82] : (hskp13 | hskp12 | ~c1_1(X82) | c2_1(X82) | c0_1(X82) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f508,plain,(
  spl0_55 | spl0_42 | ~spl0_19 | spl0_28),
  inference(avatar_split_clause,[],[f218,f365,f328,f426,f487])).
fof(f218,plain,(
  ( ! [X80,X81,X79] : (~c3_1(X79) | ~c1_1(X79) | c2_1(X79) | ~ndr1_0 | ~c2_1(X80) | c3_1(X80) | c1_1(X80) | ~c3_1(X81) | c2_1(X81) | c0_1(X81)) )),
  inference(duplicate_literal_removal,[],[f151])).
fof(f151,plain,(
  ( ! [X80,X81,X79] : (~c3_1(X79) | ~c1_1(X79) | c2_1(X79) | ~ndr1_0 | ~c2_1(X80) | c3_1(X80) | c1_1(X80) | ~ndr1_0 | ~c3_1(X81) | c2_1(X81) | c0_1(X81) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f507,plain,(
  spl0_55 | spl0_39 | ~spl0_19 | spl0_33),
  inference(avatar_split_clause,[],[f219,f387,f328,f412,f487])).
fof(f219,plain,(
  ( ! [X78,X76,X77] : (~c1_1(X76) | ~c0_1(X76) | c2_1(X76) | ~ndr1_0 | ~c3_1(X77) | ~c0_1(X77) | c1_1(X77) | ~c3_1(X78) | c2_1(X78) | c0_1(X78)) )),
  inference(duplicate_literal_removal,[],[f152])).
fof(f152,plain,(
  ( ! [X78,X76,X77] : (~c1_1(X76) | ~c0_1(X76) | c2_1(X76) | ~ndr1_0 | ~c3_1(X77) | ~c0_1(X77) | c1_1(X77) | ~ndr1_0 | ~c3_1(X78) | c2_1(X78) | c0_1(X78) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f506,plain,(
  spl0_55 | ~spl0_19 | spl0_35 | spl0_58),
  inference(avatar_split_clause,[],[f220,f503,f395,f328,f487])).
fof(f220,plain,(
  ( ! [X74,X75] : (hskp8 | ~c1_1(X74) | c3_1(X74) | c2_1(X74) | ~ndr1_0 | ~c3_1(X75) | c2_1(X75) | c0_1(X75)) )),
  inference(duplicate_literal_removal,[],[f153])).
fof(f153,plain,(
  ( ! [X74,X75] : (hskp8 | ~c1_1(X74) | c3_1(X74) | c2_1(X74) | ~ndr1_0 | ~c3_1(X75) | c2_1(X75) | c0_1(X75) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f501,plain,(
  spl0_55 | ~spl0_19 | spl0_27 | spl0_54),
  inference(avatar_split_clause,[],[f221,f481,f360,f328,f487])).
fof(f221,plain,(
  ( ! [X72,X73] : (hskp4 | ~c1_1(X72) | ~c0_1(X72) | c3_1(X72) | ~ndr1_0 | ~c3_1(X73) | c2_1(X73) | c0_1(X73)) )),
  inference(duplicate_literal_removal,[],[f154])).
fof(f154,plain,(
  ( ! [X72,X73] : (hskp4 | ~c1_1(X72) | ~c0_1(X72) | c3_1(X72) | ~ndr1_0 | ~c3_1(X73) | c2_1(X73) | c0_1(X73) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f500,plain,(
  spl0_55 | ~spl0_19 | spl0_23 | spl0_16),
  inference(avatar_split_clause,[],[f222,f314,f344,f328,f487])).
fof(f222,plain,(
  ( ! [X70,X71] : (hskp12 | ~c2_1(X70) | ~c0_1(X70) | c3_1(X70) | ~ndr1_0 | ~c3_1(X71) | c2_1(X71) | c0_1(X71)) )),
  inference(duplicate_literal_removal,[],[f155])).
fof(f155,plain,(
  ( ! [X70,X71] : (hskp12 | ~c2_1(X70) | ~c0_1(X70) | c3_1(X70) | ~ndr1_0 | ~c3_1(X71) | c2_1(X71) | c0_1(X71) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f499,plain,(
  spl0_55 | ~spl0_19 | spl0_57 | spl0_50),
  inference(avatar_split_clause,[],[f223,f462,f497,f328,f487])).
fof(f223,plain,(
  ( ! [X68,X69] : (hskp14 | ~c2_1(X68) | ~c1_1(X68) | c3_1(X68) | ~ndr1_0 | ~c3_1(X69) | c2_1(X69) | c0_1(X69)) )),
  inference(duplicate_literal_removal,[],[f156])).
fof(f156,plain,(
  ( ! [X68,X69] : (hskp14 | ~c2_1(X68) | ~c1_1(X68) | c3_1(X68) | ~ndr1_0 | ~c3_1(X69) | c2_1(X69) | c0_1(X69) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f495,plain,(
  ~spl0_19 | spl0_55 | spl0_56 | spl0_29),
  inference(avatar_split_clause,[],[f157,f368,f492,f487,f328])).
fof(f157,plain,(
  ( ! [X67] : (hskp1 | hskp30 | ~c3_1(X67) | c2_1(X67) | c0_1(X67) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f490,plain,(
  ~spl0_19 | spl0_55 | spl0_12 | spl0_32),
  inference(avatar_split_clause,[],[f158,f381,f296,f487,f328])).
fof(f158,plain,(
  ( ! [X66] : (hskp5 | hskp15 | ~c3_1(X66) | c2_1(X66) | c0_1(X66) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f485,plain,(
  spl0_53 | spl0_48 | ~spl0_19 | spl0_42),
  inference(avatar_split_clause,[],[f224,f426,f328,f454,f478])).
fof(f224,plain,(
  ( ! [X64,X62,X63] : (~c2_1(X62) | c3_1(X62) | c1_1(X62) | ~ndr1_0 | ~c3_1(X63) | ~c2_1(X63) | c0_1(X63) | ~c1_1(X64) | c3_1(X64) | c0_1(X64)) )),
  inference(duplicate_literal_removal,[],[f160])).
fof(f160,plain,(
  ( ! [X64,X62,X63] : (~c2_1(X62) | c3_1(X62) | c1_1(X62) | ~ndr1_0 | ~c3_1(X63) | ~c2_1(X63) | c0_1(X63) | ~ndr1_0 | ~c1_1(X64) | c3_1(X64) | c0_1(X64) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f484,plain,(
  ~spl0_19 | spl0_53 | spl0_54 | spl0_2),
  inference(avatar_split_clause,[],[f161,f252,f481,f478,f328])).
fof(f161,plain,(
  ( ! [X61] : (hskp7 | hskp4 | ~c1_1(X61) | c3_1(X61) | c0_1(X61) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f476,plain,(
  ~spl0_19 | spl0_52 | spl0_14 | spl0_11),
  inference(avatar_split_clause,[],[f162,f291,f305,f473,f328])).
fof(f162,plain,(
  ( ! [X60] : (hskp17 | hskp0 | ~c2_1(X60) | c3_1(X60) | c0_1(X60) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f475,plain,(
  ~spl0_19 | spl0_52 | spl0_40 | spl0_13),
  inference(avatar_split_clause,[],[f163,f300,f415,f473,f328])).
fof(f163,plain,(
  ( ! [X59] : (hskp19 | hskp18 | ~c2_1(X59) | c3_1(X59) | c0_1(X59) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f471,plain,(
  spl0_51 | ~spl0_19 | spl0_36 | spl0_50),
  inference(avatar_split_clause,[],[f225,f462,f400,f328,f467])).
fof(f225,plain,(
  ( ! [X57,X58] : (hskp14 | ~c0_1(X57) | c3_1(X57) | c2_1(X57) | ~ndr1_0 | ~c2_1(X58) | ~c1_1(X58) | c0_1(X58)) )),
  inference(duplicate_literal_removal,[],[f164])).
fof(f164,plain,(
  ( ! [X57,X58] : (hskp14 | ~c0_1(X57) | c3_1(X57) | c2_1(X57) | ~ndr1_0 | ~c2_1(X58) | ~c1_1(X58) | c0_1(X58) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f470,plain,(
  spl0_51 | ~spl0_19 | spl0_23 | spl0_12),
  inference(avatar_split_clause,[],[f226,f296,f344,f328,f467])).
fof(f226,plain,(
  ( ! [X56,X55] : (hskp15 | ~c2_1(X55) | ~c0_1(X55) | c3_1(X55) | ~ndr1_0 | ~c2_1(X56) | ~c1_1(X56) | c0_1(X56)) )),
  inference(duplicate_literal_removal,[],[f165])).
fof(f165,plain,(
  ( ! [X56,X55] : (hskp15 | ~c2_1(X55) | ~c0_1(X55) | c3_1(X55) | ~ndr1_0 | ~c2_1(X56) | ~c1_1(X56) | c0_1(X56) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f469,plain,(
  spl0_51 | ~spl0_19 | spl0_38 | spl0_3),
  inference(avatar_split_clause,[],[f227,f256,f408,f328,f467])).
fof(f227,plain,(
  ( ! [X54,X53] : (hskp9 | ~c3_1(X53) | ~c1_1(X53) | ~c0_1(X53) | ~ndr1_0 | ~c2_1(X54) | ~c1_1(X54) | c0_1(X54)) )),
  inference(duplicate_literal_removal,[],[f166])).
fof(f166,plain,(
  ( ! [X54,X53] : (hskp9 | ~c3_1(X53) | ~c1_1(X53) | ~c0_1(X53) | ~ndr1_0 | ~c2_1(X54) | ~c1_1(X54) | c0_1(X54) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f465,plain,(
  spl0_49 | ~spl0_19 | spl0_23 | spl0_50),
  inference(avatar_split_clause,[],[f228,f462,f344,f328,f458])).
fof(f228,plain,(
  ( ! [X52,X51] : (hskp14 | ~c2_1(X51) | ~c0_1(X51) | c3_1(X51) | ~ndr1_0 | ~c3_1(X52) | ~c1_1(X52) | c0_1(X52)) )),
  inference(duplicate_literal_removal,[],[f167])).
fof(f167,plain,(
  ( ! [X52,X51] : (hskp14 | ~c2_1(X51) | ~c0_1(X51) | c3_1(X51) | ~ndr1_0 | ~c3_1(X52) | ~c1_1(X52) | c0_1(X52) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f460,plain,(
  ~spl0_19 | spl0_49 | spl0_16 | spl0_25),
  inference(avatar_split_clause,[],[f168,f351,f314,f458,f328])).
fof(f168,plain,(
  ( ! [X50] : (hskp10 | hskp12 | ~c3_1(X50) | ~c1_1(X50) | c0_1(X50) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f456,plain,(
  spl0_48 | spl0_44 | ~spl0_19 | spl0_35),
  inference(avatar_split_clause,[],[f229,f395,f328,f435,f454])).
fof(f229,plain,(
  ( ! [X47,X48,X49] : (~c1_1(X47) | c3_1(X47) | c2_1(X47) | ~ndr1_0 | ~c3_1(X48) | c2_1(X48) | c1_1(X48) | ~c3_1(X49) | ~c2_1(X49) | c0_1(X49)) )),
  inference(duplicate_literal_removal,[],[f169])).
fof(f169,plain,(
  ( ! [X47,X48,X49] : (~c1_1(X47) | c3_1(X47) | c2_1(X47) | ~ndr1_0 | ~c3_1(X48) | c2_1(X48) | c1_1(X48) | ~ndr1_0 | ~c3_1(X49) | ~c2_1(X49) | c0_1(X49) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f452,plain,(
  spl0_47 | ~spl0_19 | spl0_23 | spl0_34),
  inference(avatar_split_clause,[],[f230,f390,f344,f328,f449])).
fof(f230,plain,(
  ( ! [X45,X46] : (hskp20 | ~c2_1(X45) | ~c0_1(X45) | c3_1(X45) | ~ndr1_0 | c3_1(X46) | c2_1(X46) | c1_1(X46)) )),
  inference(duplicate_literal_removal,[],[f170])).
fof(f170,plain,(
  ( ! [X45,X46] : (hskp20 | ~c2_1(X45) | ~c0_1(X45) | c3_1(X45) | ~ndr1_0 | c3_1(X46) | c2_1(X46) | c1_1(X46) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f451,plain,(
  spl0_47 | ~spl0_19 | spl0_26 | spl0_15),
  inference(avatar_split_clause,[],[f231,f309,f356,f328,f449])).
fof(f231,plain,(
  ( ! [X43,X44] : (hskp21 | ~c2_1(X43) | ~c1_1(X43) | ~c0_1(X43) | ~ndr1_0 | c3_1(X44) | c2_1(X44) | c1_1(X44)) )),
  inference(duplicate_literal_removal,[],[f171])).
fof(f171,plain,(
  ( ! [X43,X44] : (hskp21 | ~c2_1(X43) | ~c1_1(X43) | ~c0_1(X43) | ~ndr1_0 | c3_1(X44) | c2_1(X44) | c1_1(X44) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f447,plain,(
  spl0_45 | ~spl0_19 | spl0_37 | spl0_6),
  inference(avatar_split_clause,[],[f232,f269,f404,f328,f440])).
fof(f232,plain,(
  ( ! [X41,X42] : (hskp22 | ~c3_1(X41) | ~c2_1(X41) | c1_1(X41) | ~ndr1_0 | ~c0_1(X42) | c2_1(X42) | c1_1(X42)) )),
  inference(duplicate_literal_removal,[],[f172])).
fof(f172,plain,(
  ( ! [X41,X42] : (hskp22 | ~c3_1(X41) | ~c2_1(X41) | c1_1(X41) | ~ndr1_0 | ~c0_1(X42) | c2_1(X42) | c1_1(X42) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f446,plain,(
  spl0_45 | ~spl0_19 | spl0_31 | spl0_6),
  inference(avatar_split_clause,[],[f233,f269,f378,f328,f440])).
fof(f233,plain,(
  ( ! [X39,X40] : (hskp22 | ~c3_1(X39) | ~c0_1(X39) | c2_1(X39) | ~ndr1_0 | ~c0_1(X40) | c2_1(X40) | c1_1(X40)) )),
  inference(duplicate_literal_removal,[],[f173])).
fof(f173,plain,(
  ( ! [X39,X40] : (hskp22 | ~c3_1(X39) | ~c0_1(X39) | c2_1(X39) | ~ndr1_0 | ~c0_1(X40) | c2_1(X40) | c1_1(X40) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f445,plain,(
  spl0_45 | ~spl0_19 | spl0_46 | spl0_29),
  inference(avatar_split_clause,[],[f234,f368,f443,f328,f440])).
fof(f234,plain,(
  ( ! [X37,X38] : (hskp1 | ~c3_1(X37) | ~c2_1(X37) | ~c1_1(X37) | ~ndr1_0 | ~c0_1(X38) | c2_1(X38) | c1_1(X38)) )),
  inference(duplicate_literal_removal,[],[f174])).
fof(f174,plain,(
  ( ! [X37,X38] : (hskp1 | ~c3_1(X37) | ~c2_1(X37) | ~c1_1(X37) | ~ndr1_0 | ~c0_1(X38) | c2_1(X38) | c1_1(X38) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f438,plain,(
  spl0_44 | ~spl0_19 | spl0_42 | spl0_16),
  inference(avatar_split_clause,[],[f235,f314,f426,f328,f435])).
fof(f235,plain,(
  ( ! [X35,X36] : (hskp12 | ~c2_1(X35) | c3_1(X35) | c1_1(X35) | ~ndr1_0 | ~c3_1(X36) | c2_1(X36) | c1_1(X36)) )),
  inference(duplicate_literal_removal,[],[f175])).
fof(f175,plain,(
  ( ! [X35,X36] : (hskp12 | ~c2_1(X35) | c3_1(X35) | c1_1(X35) | ~ndr1_0 | ~c3_1(X36) | c2_1(X36) | c1_1(X36) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f437,plain,(
  spl0_44 | ~spl0_19 | spl0_35 | spl0_40),
  inference(avatar_split_clause,[],[f236,f415,f395,f328,f435])).
fof(f236,plain,(
  ( ! [X33,X34] : (hskp18 | ~c1_1(X33) | c3_1(X33) | c2_1(X33) | ~ndr1_0 | ~c3_1(X34) | c2_1(X34) | c1_1(X34)) )),
  inference(duplicate_literal_removal,[],[f176])).
fof(f176,plain,(
  ( ! [X33,X34] : (hskp18 | ~c1_1(X33) | c3_1(X33) | c2_1(X33) | ~ndr1_0 | ~c3_1(X34) | c2_1(X34) | c1_1(X34) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f433,plain,(
  ~spl0_19 | spl0_43 | spl0_40 | spl0_16),
  inference(avatar_split_clause,[],[f177,f314,f415,f431,f328])).
fof(f177,plain,(
  ( ! [X32] : (hskp12 | hskp18 | ~c0_1(X32) | c3_1(X32) | c1_1(X32) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f429,plain,(
  spl0_42 | ~spl0_19 | spl0_33 | spl0_15),
  inference(avatar_split_clause,[],[f237,f309,f387,f328,f426])).
fof(f237,plain,(
  ( ! [X30,X31] : (hskp21 | ~c1_1(X30) | ~c0_1(X30) | c2_1(X30) | ~ndr1_0 | ~c2_1(X31) | c3_1(X31) | c1_1(X31)) )),
  inference(duplicate_literal_removal,[],[f178])).
fof(f178,plain,(
  ( ! [X30,X31] : (hskp21 | ~c1_1(X30) | ~c0_1(X30) | c2_1(X30) | ~ndr1_0 | ~c2_1(X31) | c3_1(X31) | c1_1(X31) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f428,plain,(
  ~spl0_19 | spl0_42 | spl0_17 | spl0_12),
  inference(avatar_split_clause,[],[f179,f296,f319,f426,f328])).
fof(f179,plain,(
  ( ! [X29] : (hskp15 | hskp29 | ~c2_1(X29) | c3_1(X29) | c1_1(X29) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f424,plain,(
  spl0_41 | spl0_36 | ~spl0_19 | spl0_28),
  inference(avatar_split_clause,[],[f238,f365,f328,f400,f420])).
fof(f238,plain,(
  ( ! [X28,X26,X27] : (~c3_1(X26) | ~c1_1(X26) | c2_1(X26) | ~ndr1_0 | ~c0_1(X27) | c3_1(X27) | c2_1(X27) | ~c2_1(X28) | ~c0_1(X28) | c1_1(X28)) )),
  inference(duplicate_literal_removal,[],[f180])).
fof(f180,plain,(
  ( ! [X28,X26,X27] : (~c3_1(X26) | ~c1_1(X26) | c2_1(X26) | ~ndr1_0 | ~c0_1(X27) | c3_1(X27) | c2_1(X27) | ~ndr1_0 | ~c2_1(X28) | ~c0_1(X28) | c1_1(X28) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f423,plain,(
  spl0_41 | ~spl0_19 | spl0_33 | spl0_2),
  inference(avatar_split_clause,[],[f239,f252,f387,f328,f420])).
fof(f239,plain,(
  ( ! [X24,X25] : (hskp7 | ~c1_1(X24) | ~c0_1(X24) | c2_1(X24) | ~ndr1_0 | ~c2_1(X25) | ~c0_1(X25) | c1_1(X25)) )),
  inference(duplicate_literal_removal,[],[f181])).
fof(f181,plain,(
  ( ! [X24,X25] : (hskp7 | ~c1_1(X24) | ~c0_1(X24) | c2_1(X24) | ~ndr1_0 | ~c2_1(X25) | ~c0_1(X25) | c1_1(X25) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f422,plain,(
  spl0_41 | ~spl0_19 | spl0_20 | spl0_4),
  inference(avatar_split_clause,[],[f240,f261,f332,f328,f420])).
fof(f240,plain,(
  ( ! [X23,X22] : (hskp23 | ~c3_1(X22) | ~c2_1(X22) | ~c0_1(X22) | ~ndr1_0 | ~c2_1(X23) | ~c0_1(X23) | c1_1(X23)) )),
  inference(duplicate_literal_removal,[],[f182])).
fof(f182,plain,(
  ( ! [X23,X22] : (hskp23 | ~c3_1(X22) | ~c2_1(X22) | ~c0_1(X22) | ~ndr1_0 | ~c2_1(X23) | ~c0_1(X23) | c1_1(X23) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f418,plain,(
  spl0_39 | ~spl0_19 | spl0_36 | spl0_40),
  inference(avatar_split_clause,[],[f241,f415,f400,f328,f412])).
fof(f241,plain,(
  ( ! [X21,X20] : (hskp18 | ~c0_1(X20) | c3_1(X20) | c2_1(X20) | ~ndr1_0 | ~c3_1(X21) | ~c0_1(X21) | c1_1(X21)) )),
  inference(duplicate_literal_removal,[],[f183])).
fof(f183,plain,(
  ( ! [X21,X20] : (hskp18 | ~c0_1(X20) | c3_1(X20) | c2_1(X20) | ~ndr1_0 | ~c3_1(X21) | ~c0_1(X21) | c1_1(X21) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f410,plain,(
  spl0_37 | ~spl0_19 | spl0_38 | spl0_10),
  inference(avatar_split_clause,[],[f242,f287,f408,f328,f404])).
fof(f242,plain,(
  ( ! [X19,X18] : (hskp28 | ~c3_1(X18) | ~c1_1(X18) | ~c0_1(X18) | ~ndr1_0 | ~c3_1(X19) | ~c2_1(X19) | c1_1(X19)) )),
  inference(duplicate_literal_removal,[],[f184])).
fof(f184,plain,(
  ( ! [X19,X18] : (hskp28 | ~c3_1(X18) | ~c1_1(X18) | ~c0_1(X18) | ~ndr1_0 | ~c3_1(X19) | ~c2_1(X19) | c1_1(X19) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f406,plain,(
  ~spl0_19 | spl0_37 | spl0_8 | spl0_11),
  inference(avatar_split_clause,[],[f185,f291,f278,f404,f328])).
fof(f185,plain,(
  ( ! [X17] : (hskp17 | hskp24 | ~c3_1(X17) | ~c2_1(X17) | c1_1(X17) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f402,plain,(
  ~spl0_19 | spl0_36 | spl0_25 | spl0_2),
  inference(avatar_split_clause,[],[f186,f252,f351,f400,f328])).
fof(f186,plain,(
  ( ! [X16] : (hskp7 | hskp10 | ~c0_1(X16) | c3_1(X16) | c2_1(X16) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f398,plain,(
  spl0_35 | ~spl0_19 | spl0_33 | spl0_18),
  inference(avatar_split_clause,[],[f243,f323,f387,f328,f395])).
fof(f243,plain,(
  ( ! [X14,X15] : (hskp31 | ~c1_1(X14) | ~c0_1(X14) | c2_1(X14) | ~ndr1_0 | ~c1_1(X15) | c3_1(X15) | c2_1(X15)) )),
  inference(duplicate_literal_removal,[],[f187])).
fof(f187,plain,(
  ( ! [X14,X15] : (hskp31 | ~c1_1(X14) | ~c0_1(X14) | c2_1(X14) | ~ndr1_0 | ~c1_1(X15) | c3_1(X15) | c2_1(X15) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f397,plain,(
  ~spl0_19 | spl0_35 | spl0_9),
  inference(avatar_split_clause,[],[f188,f282,f395,f328])).
fof(f188,plain,(
  ( ! [X13] : (hskp25 | ~c1_1(X13) | c3_1(X13) | c2_1(X13) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f393,plain,(
  ~spl0_19 | spl0_33 | spl0_29 | spl0_34),
  inference(avatar_split_clause,[],[f189,f390,f368,f387,f328])).
fof(f189,plain,(
  ( ! [X12] : (hskp20 | hskp1 | ~c1_1(X12) | ~c0_1(X12) | c2_1(X12) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f385,plain,(
  spl0_31 | ~spl0_19 | spl0_20 | spl0_14),
  inference(avatar_split_clause,[],[f244,f305,f332,f328,f378])).
fof(f244,plain,(
  ( ! [X10,X11] : (hskp0 | ~c3_1(X10) | ~c2_1(X10) | ~c0_1(X10) | ~ndr1_0 | ~c3_1(X11) | ~c0_1(X11) | c2_1(X11)) )),
  inference(duplicate_literal_removal,[],[f190])).
fof(f190,plain,(
  ( ! [X10,X11] : (hskp0 | ~c3_1(X10) | ~c2_1(X10) | ~c0_1(X10) | ~ndr1_0 | ~c3_1(X11) | ~c0_1(X11) | c2_1(X11) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f384,plain,(
  ~spl0_19 | spl0_31 | spl0_29 | spl0_32),
  inference(avatar_split_clause,[],[f191,f381,f368,f378,f328])).
fof(f191,plain,(
  ( ! [X9] : (hskp5 | hskp1 | ~c3_1(X9) | ~c0_1(X9) | c2_1(X9) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f376,plain,(
  ~spl0_19 | spl0_28 | spl0_16 | spl0_30),
  inference(avatar_split_clause,[],[f192,f373,f314,f365,f328])).
fof(f192,plain,(
  ( ! [X8] : (hskp6 | hskp12 | ~c3_1(X8) | ~c1_1(X8) | c2_1(X8) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f371,plain,(
  ~spl0_19 | spl0_28 | spl0_29 | spl0_21),
  inference(avatar_split_clause,[],[f193,f335,f368,f365,f328])).
fof(f193,plain,(
  ( ! [X7] : (hskp13 | hskp1 | ~c3_1(X7) | ~c1_1(X7) | c2_1(X7) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f363,plain,(
  spl0_27 | ~spl0_19 | spl0_20 | spl0_15),
  inference(avatar_split_clause,[],[f245,f309,f332,f328,f360])).
fof(f245,plain,(
  ( ! [X6,X5] : (hskp21 | ~c3_1(X5) | ~c2_1(X5) | ~c0_1(X5) | ~ndr1_0 | ~c1_1(X6) | ~c0_1(X6) | c3_1(X6)) )),
  inference(duplicate_literal_removal,[],[f194])).
fof(f194,plain,(
  ( ! [X6,X5] : (hskp21 | ~c3_1(X5) | ~c2_1(X5) | ~c0_1(X5) | ~ndr1_0 | ~c1_1(X6) | ~c0_1(X6) | c3_1(X6) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f362,plain,(
  ~spl0_19 | spl0_27 | spl0_5 | spl0_11),
  inference(avatar_split_clause,[],[f195,f291,f265,f360,f328])).
fof(f195,plain,(
  ( ! [X4] : (hskp17 | hskp26 | ~c1_1(X4) | ~c0_1(X4) | c3_1(X4) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f358,plain,(
  spl0_23 | ~spl0_19 | spl0_26 | spl0_3),
  inference(avatar_split_clause,[],[f246,f256,f356,f328,f344])).
fof(f246,plain,(
  ( ! [X2,X3] : (hskp9 | ~c2_1(X2) | ~c1_1(X2) | ~c0_1(X2) | ~ndr1_0 | ~c2_1(X3) | ~c0_1(X3) | c3_1(X3)) )),
  inference(duplicate_literal_removal,[],[f196])).
fof(f196,plain,(
  ( ! [X2,X3] : (hskp9 | ~c2_1(X2) | ~c1_1(X2) | ~c0_1(X2) | ~ndr1_0 | ~c2_1(X3) | ~c0_1(X3) | c3_1(X3) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f354,plain,(
  ~spl0_19 | spl0_23 | spl0_24 | spl0_25),
  inference(avatar_split_clause,[],[f197,f351,f347,f344,f328])).
fof(f197,plain,(
  ( ! [X1] : (hskp10 | hskp16 | ~c2_1(X1) | ~c0_1(X1) | c3_1(X1) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f342,plain,(
  ~spl0_19 | spl0_20 | spl0_21 | spl0_22),
  inference(avatar_split_clause,[],[f198,f339,f335,f332,f328])).
fof(f198,plain,(
  ( ! [X0] : (hskp2 | hskp13 | ~c3_1(X0) | ~c2_1(X0) | ~c0_1(X0) | ~ndr1_0) )),
  inference(cnf_transformation,[],[f6])).
fof(f326,plain,(
  spl0_17 | spl0_18 | spl0_1),
  inference(avatar_split_clause,[],[f199,f248,f323,f319])).
fof(f199,plain,(
  hskp27 | hskp31 | hskp29),
  inference(cnf_transformation,[],[f6])).
fof(f317,plain,(
  spl0_14 | spl0_16 | spl0_3),
  inference(avatar_split_clause,[],[f200,f256,f314,f305])).
fof(f200,plain,(
  hskp9 | hskp12 | hskp0),
  inference(cnf_transformation,[],[f6])).
fof(f312,plain,(
  spl0_14 | spl0_15 | spl0_9),
  inference(avatar_split_clause,[],[f201,f282,f309,f305])).
fof(f201,plain,(
  hskp25 | hskp21 | hskp0),
  inference(cnf_transformation,[],[f6])).
fof(f303,plain,(
  spl0_12 | spl0_6 | spl0_13),
  inference(avatar_split_clause,[],[f202,f300,f269,f296])).
fof(f202,plain,(
  hskp19 | hskp22 | hskp15),
  inference(cnf_transformation,[],[f6])).
fof(f294,plain,(
  spl0_10 | spl0_6 | spl0_11),
  inference(avatar_split_clause,[],[f203,f291,f269,f287])).
fof(f203,plain,(
  hskp17 | hskp22 | hskp28),
  inference(cnf_transformation,[],[f6])).
fof(f272,plain,(
  spl0_4 | spl0_5 | spl0_6),
  inference(avatar_split_clause,[],[f205,f269,f265,f261])).
fof(f205,plain,(
  hskp22 | hskp26 | hskp23),
  inference(cnf_transformation,[],[f6])).
fof(f259,plain,(
  spl0_1 | spl0_2 | spl0_3),
  inference(avatar_split_clause,[],[f206,f256,f252,f248])).
fof(f206,plain,(
  hskp9 | hskp7 | hskp27),
  inference(cnf_transformation,[],[f6])).
% SZS output end Proof for SYN474+1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 7803
% Time elapsed: 0.032 s
% ------------------------------
% ------------------------------
