% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for SYO650-1
% SZS output start Proof for SYO650-1
fof(f201,plain,(
  $false),
  inference(resolution,[],[f195,f28])).
fof(f28,axiom,(
  ~'LE'(f(z),'0')),
  file('./p/SYO650-1.tptp',unknown)).
fof(f195,plain,(
  ( ! [X1] : ('LE'(f(X1),'0')) )),
  inference(duplicate_literal_removal,[],[f190])).
fof(f190,plain,(
  ( ! [X1] : ('LE'(f(X1),'0') | 'LE'(f(X1),'0') | 'LE'(f(X1),'0')) )),
  inference(resolution,[],[f180,f68])).
fof(f68,plain,(
  ( ! [X0] : (iLEQ(suc(X0),suc(X0)) | 'LE'(f(X0),'0')) )),
  inference(subsumption_resolution,[],[f67,f47])).
fof(f47,plain,(
  ( ! [X3] : ('E'('0',f(suc(X3))) | 'LE'(f(X3),'0')) )),
  inference(subsumption_resolution,[],[f19,f10])).
fof(f10,axiom,(
  ( ! [X4] : ('LE'(f(X4),s('0'))) )),
  file('./p/SYO650-1.tptp',unknown)).
fof(f19,axiom,(
  ( ! [X3] : ('LE'(f(X3),'0') | 'E'('0',f(suc(X3))) | ~'LE'(f(suc(X3)),s('0'))) )),
  file('./p/SYO650-1.tptp',unknown)).
fof(f67,plain,(
  ( ! [X0] : (~'E'('0',f(suc(X0))) | iLEQ(suc(X0),suc(X0)) | 'LE'(f(X0),'0')) )),
  inference(subsumption_resolution,[],[f66,f41])).
fof(f41,plain,(
  ( ! [X3] : ('E'('0',f(suc(suc(X3)))) | 'LE'(f(X3),'0')) )),
  inference(subsumption_resolution,[],[f11,f10])).
fof(f11,axiom,(
  ( ! [X3] : ('LE'(f(X3),'0') | 'E'('0',f(suc(suc(X3)))) | ~'LE'(f(suc(suc(X3))),s('0'))) )),
  file('./p/SYO650-1.tptp',unknown)).
fof(f66,plain,(
  ( ! [X0] : (~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | iLEQ(suc(X0),suc(X0)) | 'LE'(f(X0),'0')) )),
  inference(subsumption_resolution,[],[f63,f58])).
fof(f58,plain,(
  ( ! [X0] : ('LE'(f(X0),'0') | 'E'('0',f(X0))) )),
  inference(resolution,[],[f7,f10])).
fof(f7,axiom,(
  ( ! [X3] : (~'LE'(f(X3),s('0')) | 'E'('0',f(X3)) | 'LE'(f(X3),'0')) )),
  file('./p/SYO650-1.tptp',unknown)).
fof(f63,plain,(
  ( ! [X0] : (~'E'('0',f(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | iLEQ(suc(X0),suc(X0)) | 'LE'(f(X0),'0')) )),
  inference(resolution,[],[f62,f49])).
fof(f49,plain,(
  ( ! [X3] : ('E'('0',f(suc(suc(suc(X3))))) | 'LE'(f(X3),'0')) )),
  inference(subsumption_resolution,[],[f22,f10])).
fof(f22,axiom,(
  ( ! [X3] : ('LE'(f(X3),'0') | 'E'('0',f(suc(suc(suc(X3))))) | ~'LE'(f(suc(suc(suc(X3)))),s('0'))) )),
  file('./p/SYO650-1.tptp',unknown)).
fof(f62,plain,(
  ( ! [X0] : (~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | iLEQ(suc(X0),suc(X0))) )),
  inference(subsumption_resolution,[],[f61,f57])).
fof(f57,plain,(
  ( ! [X1] : ('E'(f(X1),f(suc(X1))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(X1))) )),
  inference(subsumption_resolution,[],[f34,f21])).
fof(f21,axiom,(
  ( ! [X2,X0,X1] : (~iLEQ(suc(X1),suc(X0)) | 'E'(f(X0),f(suc(X0))) | 'E'(f(X1),f(suc(X1))) | ~'E'('0',f(suc(X0))) | ~'E'('0',f(X2)) | ~'E'('0',f(X0)) | ~iLEQ(suc(X0),suc(X2)) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(X2))) | 'E'(f(X2),f(suc(X2)))) )),
  file('./p/SYO650-1.tptp',unknown)).
fof(f34,axiom,(
  ( ! [X1] : (iLEQ(suc(X1),suc(X1)) | 'E'(f(X1),f(suc(X1))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(X1))) )),
  file('./p/SYO650-1.tptp',unknown)).
fof(f61,plain,(
  ( ! [X0] : (~'E'(f(X0),f(suc(X0))) | iLEQ(suc(X0),suc(X0)) | ~'E'('0',f(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | ~'E'('0',f(suc(suc(suc(X0)))))) )),
  inference(duplicate_literal_removal,[],[f60])).
fof(f60,plain,(
  ( ! [X0] : (~'E'(f(X0),f(suc(X0))) | iLEQ(suc(X0),suc(X0)) | ~'E'('0',f(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(X0)) | ~'E'(f(X0),f(suc(X0))) | ~'E'('0',f(suc(X0))) | ~'E'('0',f(suc(suc(X0))))) )),
  inference(resolution,[],[f14,f53])).
fof(f53,plain,(
  ( ! [X1] : ('E'(f(X1),f(suc(suc(X1)))) | ~'E'('0',f(X1)) | ~'E'(f(X1),f(suc(X1))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(suc(X1))))) )),
  inference(subsumption_resolution,[],[f26,f44])).
fof(f44,plain,(
  ( ! [X2,X0,X1] : (~'E'(f(X2),f(suc(X2))) | 'E'(f(X0),f(suc(suc(X0)))) | 'E'(f(X1),f(suc(suc(X1)))) | ~'E'('0',f(suc(X0))) | ~'E'(f(X0),f(suc(X0))) | ~'E'('0',f(X2)) | ~'E'('0',f(X0)) | ~iLEQ(suc(X0),suc(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(X1)))) | 'E'(f(X2),f(suc(suc(X2)))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(X2))) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(suc(suc(X0))))) )),
  inference(subsumption_resolution,[],[f15,f6])).
fof(f6,axiom,(
  ( ! [X2,X0,X1] : (~'E'(f(X2),f(suc(X2))) | 'E'(f(X0),f(suc(suc(X0)))) | 'E'(f(X1),f(suc(X1))) | ~'E'('0',f(suc(X0))) | ~'E'(f(X0),f(suc(X0))) | ~'E'('0',f(X2)) | ~'E'('0',f(X0)) | ~iLEQ(suc(X0),suc(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(X1)) | 'E'(f(X2),f(suc(suc(X2)))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(X2))) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(suc(suc(X0))))) )),
  file('./p/SYO650-1.tptp',unknown)).
fof(f15,axiom,(
  ( ! [X2,X0,X1] : ('E'(f(X2),f(suc(suc(X2)))) | 'E'(f(X0),f(suc(suc(X0)))) | 'E'(f(X1),f(suc(suc(X1)))) | ~'E'('0',f(suc(X0))) | ~'E'(f(X0),f(suc(X0))) | ~'E'('0',f(X2)) | ~'E'('0',f(X0)) | ~'E'(f(X1),f(suc(X1))) | ~iLEQ(suc(X0),suc(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'(f(X2),f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(X2))) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(suc(suc(X0))))) )),
  file('./p/SYO650-1.tptp',unknown)).
fof(f26,axiom,(
  ( ! [X1] : (iLEQ(suc(X1),suc(X1)) | 'E'(f(X1),f(suc(suc(X1)))) | ~'E'('0',f(X1)) | ~'E'(f(X1),f(suc(X1))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(suc(X1))))) )),
  file('./p/SYO650-1.tptp',unknown)).
fof(f14,axiom,(
  ( ! [X1] : (~'E'(f(X1),f(suc(suc(X1)))) | ~'E'(f(X1),f(suc(X1))) | iLEQ(suc(X1),suc(X1)) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(suc(suc(X1)))))) )),
  file('./p/SYO650-1.tptp',unknown)).
fof(f180,plain,(
  ( ! [X2,X3] : (~iLEQ(suc(X2),suc(X3)) | 'LE'(f(X2),'0') | 'LE'(f(X3),'0')) )),
  inference(duplicate_literal_removal,[],[f175])).
fof(f175,plain,(
  ( ! [X2,X3] : (~iLEQ(suc(X2),suc(X3)) | 'LE'(f(X2),'0') | 'LE'(f(X3),'0') | 'LE'(f(X2),'0') | 'LE'(f(X2),'0')) )),
  inference(resolution,[],[f156,f68])).
fof(f156,plain,(
  ( ! [X2,X0,X1] : (~iLEQ(suc(X2),suc(X0)) | ~iLEQ(suc(X0),suc(X1)) | 'LE'(f(X2),'0') | 'LE'(f(X1),'0') | 'LE'(f(X0),'0')) )),
  inference(subsumption_resolution,[],[f155,f47])).
fof(f155,plain,(
  ( ! [X2,X0,X1] : (~iLEQ(suc(X0),suc(X1)) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(X0))) | 'LE'(f(X2),'0') | 'LE'(f(X1),'0') | 'LE'(f(X0),'0')) )),
  inference(subsumption_resolution,[],[f154,f41])).
fof(f154,plain,(
  ( ! [X2,X0,X1] : (~iLEQ(suc(X0),suc(X1)) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | 'LE'(f(X2),'0') | 'LE'(f(X1),'0') | 'LE'(f(X0),'0')) )),
  inference(subsumption_resolution,[],[f153,f47])).
fof(f153,plain,(
  ( ! [X2,X0,X1] : (~iLEQ(suc(X0),suc(X1)) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | 'LE'(f(X2),'0') | 'LE'(f(X1),'0') | 'LE'(f(X0),'0')) )),
  inference(subsumption_resolution,[],[f152,f47])).
fof(f152,plain,(
  ( ! [X2,X0,X1] : (~iLEQ(suc(X0),suc(X1)) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | 'LE'(f(X2),'0') | 'LE'(f(X1),'0') | 'LE'(f(X0),'0')) )),
  inference(subsumption_resolution,[],[f151,f41])).
fof(f151,plain,(
  ( ! [X2,X0,X1] : (~iLEQ(suc(X0),suc(X1)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | 'LE'(f(X2),'0') | 'LE'(f(X1),'0') | 'LE'(f(X0),'0')) )),
  inference(subsumption_resolution,[],[f150,f41])).
fof(f150,plain,(
  ( ! [X2,X0,X1] : (~iLEQ(suc(X0),suc(X1)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | 'LE'(f(X2),'0') | 'LE'(f(X1),'0') | 'LE'(f(X0),'0')) )),
  inference(subsumption_resolution,[],[f147,f58])).
fof(f147,plain,(
  ( ! [X2,X0,X1] : (~iLEQ(suc(X0),suc(X1)) | ~'E'('0',f(X0)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | 'LE'(f(X2),'0') | 'LE'(f(X1),'0') | 'LE'(f(X0),'0')) )),
  inference(resolution,[],[f144,f49])).
fof(f144,plain,(
  ( ! [X2,X0,X1] : (~'E'('0',f(suc(suc(suc(X0))))) | ~iLEQ(suc(X0),suc(X1)) | ~'E'('0',f(X0)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | 'LE'(f(X2),'0') | 'LE'(f(X1),'0')) )),
  inference(subsumption_resolution,[],[f141,f58])).
fof(f141,plain,(
  ( ! [X2,X0,X1] : (~'E'('0',f(X0)) | ~iLEQ(suc(X0),suc(X1)) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | 'LE'(f(X2),'0') | 'LE'(f(X1),'0')) )),
  inference(resolution,[],[f138,f49])).
fof(f138,plain,(
  ( ! [X2,X0,X1] : (~'E'('0',f(suc(suc(suc(X1))))) | ~'E'('0',f(X0)) | ~iLEQ(suc(X0),suc(X1)) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | 'LE'(f(X2),'0')) )),
  inference(subsumption_resolution,[],[f135,f58])).
fof(f135,plain,(
  ( ! [X2,X0,X1] : (~'E'('0',f(suc(X0))) | ~'E'('0',f(X0)) | ~iLEQ(suc(X0),suc(X1)) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(suc(suc(X1))))) | 'LE'(f(X2),'0')) )),
  inference(resolution,[],[f134,f49])).
fof(f134,plain,(
  ( ! [X2,X0,X1] : (~'E'('0',f(suc(suc(suc(X2))))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(X1)) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(X0)) | ~'E'('0',f(suc(suc(suc(X1))))) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X0))) | ~iLEQ(suc(X2),suc(X1)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(suc(X0)))))) )),
  inference(subsumption_resolution,[],[f133,f57])).
fof(f133,plain,(
  ( ! [X2,X0,X1] : (~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(X1)) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(X0)) | ~'E'('0',f(suc(suc(suc(X1))))) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X0))) | ~iLEQ(suc(X2),suc(X1)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(suc(X2))))) | ~'E'(f(X1),f(suc(X1)))) )),
  inference(duplicate_literal_removal,[],[f132])).
fof(f132,plain,(
  ( ! [X2,X0,X1] : (~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(X1)) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(X0)) | ~'E'('0',f(suc(suc(suc(X1))))) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X0))) | ~iLEQ(suc(X2),suc(X1)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(suc(X2))))) | ~'E'('0',f(X1)) | ~'E'(f(X1),f(suc(X1))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(suc(X1))))) )),
  inference(resolution,[],[f131,f53])).
fof(f131,plain,(
  ( ! [X2,X0,X1] : (~'E'(f(X1),f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(X1)) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(X0)) | ~'E'('0',f(suc(suc(suc(X1))))) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X0))) | ~iLEQ(suc(X2),suc(X1)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(suc(X2)))))) )),
  inference(subsumption_resolution,[],[f130,f57])).
fof(f130,plain,(
  ( ! [X2,X0,X1] : (~'E'('0',f(suc(suc(suc(X0))))) | ~'E'(f(X1),f(suc(suc(X1)))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(X1)) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(X0)) | ~'E'('0',f(suc(suc(suc(X1))))) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X0))) | ~iLEQ(suc(X2),suc(X1)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(suc(X2))))) | ~'E'(f(X2),f(suc(X2)))) )),
  inference(duplicate_literal_removal,[],[f129])).
fof(f129,plain,(
  ( ! [X2,X0,X1] : (~'E'('0',f(suc(suc(suc(X0))))) | ~'E'(f(X1),f(suc(suc(X1)))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(X1)) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(X0)) | ~'E'('0',f(suc(suc(suc(X1))))) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X0))) | ~iLEQ(suc(X2),suc(X1)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(suc(X2))))) | ~'E'('0',f(X2)) | ~'E'(f(X2),f(suc(X2))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(suc(X2))))) )),
  inference(resolution,[],[f128,f53])).
fof(f128,plain,(
  ( ! [X2,X0,X1] : (~'E'(f(X2),f(suc(suc(X2)))) | ~'E'('0',f(suc(suc(suc(X1))))) | ~'E'(f(X0),f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | ~'E'('0',f(X0)) | ~iLEQ(suc(X0),suc(X1)) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(suc(suc(X2)))))) )),
  inference(subsumption_resolution,[],[f127,f57])).
fof(f127,plain,(
  ( ! [X2,X0,X1] : (~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(suc(X1))))) | ~'E'(f(X0),f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | ~'E'('0',f(X0)) | ~'E'(f(X2),f(suc(X2))) | ~iLEQ(suc(X0),suc(X1)) | ~'E'(f(X2),f(suc(suc(X2)))) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(suc(suc(X2)))))) )),
  inference(subsumption_resolution,[],[f126,f57])).
fof(f126,plain,(
  ( ! [X2,X0,X1] : (~'E'(f(X0),f(suc(X0))) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(suc(X1))))) | ~'E'(f(X0),f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | ~'E'('0',f(X0)) | ~'E'(f(X2),f(suc(X2))) | ~iLEQ(suc(X0),suc(X1)) | ~'E'(f(X2),f(suc(suc(X2)))) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(suc(suc(X2)))))) )),
  inference(subsumption_resolution,[],[f125,f42])).
fof(f42,plain,(
  ( ! [X2,X0,X1] : (~'E'(f(X1),f(suc(suc(X1)))) | ~'E'('0',f(suc(X0))) | ~'E'(f(X0),f(suc(X0))) | ~'E'('0',f(X2)) | ~'E'(f(X0),f(suc(suc(X0)))) | ~'E'('0',f(X0)) | ~iLEQ(suc(X0),suc(X2)) | 'E'(f(X2),f(suc(X2))) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(X2))) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(suc(suc(X1)))))) )),
  inference(subsumption_resolution,[],[f12,f1])).
fof(f1,axiom,(
  ( ! [X2,X0,X1] : (~'E'(f(X0),f(suc(suc(X0)))) | 'E'(f(X1),f(suc(X1))) | ~'E'('0',f(suc(X0))) | ~'E'(f(X0),f(suc(X0))) | ~'E'('0',f(X2)) | 'E'(f(X2),f(suc(X2))) | ~'E'('0',f(X0)) | ~iLEQ(suc(X0),suc(X2)) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(X2))) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(suc(suc(X0))))) )),
  file('./p/SYO650-1.tptp',unknown)).
fof(f12,axiom,(
  ( ! [X2,X0,X1] : ('E'(f(X2),f(suc(X2))) | ~'E'('0',f(suc(X0))) | ~'E'(f(X0),f(suc(X0))) | ~'E'('0',f(X2)) | ~'E'(f(X0),f(suc(suc(X0)))) | ~'E'('0',f(X0)) | ~'E'(f(X1),f(suc(X1))) | ~iLEQ(suc(X0),suc(X2)) | ~'E'(f(X1),f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(X2))) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(suc(suc(X1)))))) )),
  file('./p/SYO650-1.tptp',unknown)).
fof(f125,plain,(
  ( ! [X2,X0,X1] : (~'E'(f(X0),f(suc(X0))) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(suc(X1))))) | ~'E'(f(X0),f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | ~'E'('0',f(X0)) | ~'E'(f(X2),f(suc(X2))) | ~iLEQ(suc(X0),suc(X1)) | ~'E'(f(X2),f(suc(suc(X2)))) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(suc(suc(X2))))) | ~'E'(f(X1),f(suc(X1)))) )),
  inference(duplicate_literal_removal,[],[f124])).
fof(f124,plain,(
  ( ! [X2,X0,X1] : (~'E'(f(X0),f(suc(X0))) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(suc(X1))))) | ~'E'(f(X0),f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | ~'E'('0',f(X0)) | ~'E'(f(X2),f(suc(X2))) | ~iLEQ(suc(X0),suc(X1)) | ~'E'(f(X2),f(suc(suc(X2)))) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~iLEQ(suc(X2),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(suc(suc(X2))))) | ~'E'('0',f(X1)) | ~'E'(f(X1),f(suc(X1))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(suc(X1))))) )),
  inference(resolution,[],[f51,f53])).
fof(f51,plain,(
  ( ! [X2,X0,X1] : (~'E'(f(X2),f(suc(suc(X2)))) | ~'E'(f(X0),f(suc(X0))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(suc(X2))))) | ~'E'(f(X0),f(suc(suc(X0)))) | ~'E'('0',f(suc(X0))) | ~'E'('0',f(X0)) | ~'E'(f(X1),f(suc(X1))) | ~iLEQ(suc(X0),suc(X2)) | ~'E'(f(X1),f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(X2))) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(suc(suc(X1)))))) )),
  inference(subsumption_resolution,[],[f24,f42])).
fof(f24,axiom,(
  ( ! [X2,X0,X1] : (~'E'('0',f(suc(X0))) | ~'E'(f(X0),f(suc(X0))) | ~'E'('0',f(X2)) | ~'E'('0',f(suc(suc(suc(X2))))) | ~'E'(f(X0),f(suc(suc(X0)))) | ~'E'(f(X2),f(suc(suc(X2)))) | ~'E'('0',f(X0)) | ~'E'(f(X1),f(suc(X1))) | ~iLEQ(suc(X0),suc(X2)) | ~'E'(f(X1),f(suc(suc(X1)))) | ~'E'('0',f(suc(suc(suc(X0))))) | ~'E'('0',f(suc(suc(X2)))) | ~'E'('0',f(X1)) | ~'E'('0',f(suc(suc(X1)))) | ~'E'(f(X2),f(suc(X2))) | ~'E'('0',f(suc(X1))) | ~'E'('0',f(suc(X2))) | ~iLEQ(suc(X1),suc(X0)) | ~'E'('0',f(suc(suc(X0)))) | ~'E'('0',f(suc(suc(suc(X1)))))) )),
  file('./p/SYO650-1.tptp',unknown)).
% SZS output end Proof for SYO650-1
% ------------------------------
% Version: Vampire 4.4.0 (commit d98a3d53 on 2020-01-14 21:02:12 +0100)
% Termination reason: Refutation

% Memory used [KB]: 5117
% Time elapsed: 0.019 s
% ------------------------------
% ------------------------------
